Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.util_prob
% (56186)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 124
% Passive clauses: 2473
% Generated clauses: 19304
% Final active clauses: 120
% Final passive clauses: 2215
% Input formulas: 12
% Initial clauses: 17
% 
% Duplicate literals: 1109
% Trivial inequalities: 77
% Fw subsumption resolutions: 2350
% Fw demodulations: 3606
% Bw demodulations: 2
% 
% Simple tautologies: 280
% Equational tautologies: 26
% Forward subsumptions: 8385
% 
% Binary resolution: 165
% Forward superposition: 7323
% Backward superposition: 4812
% Self superposition: 162
% Equality factoring: 58
% Equality resolution: 23
% 
% Split clauses: 618
% Split components: 1255
% SAT solver clauses: 1132
% SAT solver unit clauses: 8
% SAT solver binary clauses: 82
% 
% Sat splits: 618
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1129
% TWLsolver calls for satisfiability: 148
% 
% Memory used [KB]: 11385
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (56210)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 96
% Passive clauses: 913
% Generated clauses: 12306
% Final active clauses: 89
% Final passive clauses: 756
% Input formulas: 12
% Initial clauses: 17
% 
% Duplicate literals: 549
% Trivial inequalities: 89
% Fw subsumption resolutions: 968
% Fw demodulations: 2088
% Bw demodulations: 61
% Condensations: 762
% 
% Simple tautologies: 161
% Equational tautologies: 25
% Forward subsumptions: 4886
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 122
% Forward superposition: 3979
% Backward superposition: 2567
% Self superposition: 119
% Equality factoring: 54
% Equality resolution: 13
% 
% Split clauses: 126
% Split components: 262
% SAT solver clauses: 1717
% SAT solver unit clauses: 21
% SAT solver binary clauses: 43
% 
% Sat splits: 126
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 1717
% TWLsolver calls for satisfiability: 137
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (56238)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 497
% Passive clauses: 18413
% Generated clauses: 140514
% Final active clauses: 92
% Final passive clauses: 1054
% Input formulas: 12
% Initial clauses: 17
% Discarded non-redundant clauses: 432
% 
% Duplicate literals: 3243
% Trivial inequalities: 816
% Fw subsumption resolutions: 8779
% Fw demodulations: 21714
% Bw demodulations: 14853
% Condensations: 7075
% Global subsumptions: 110
% 
% Simple tautologies: 957
% Equational tautologies: 754
% Forward subsumptions: 77596
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 1155
% 
% Unit resulting resolution: 4209
% Forward superposition: 32158
% Backward superposition: 21632
% Self superposition: 537
% Equality factoring: 150
% Equality resolution: 134
% 
% Split clauses: 440
% Split components: 891
% SAT solver clauses: 20025
% SAT solver unit clauses: 127
% SAT solver binary clauses: 1529
% 
% Sat splits: 440
% Sat splitting refutations: 620
% 
% Memory used [KB]: 64604
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% (56532)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1737
% Passive clauses: 30749
% Generated clauses: 113395
% Final active clauses: 123
% Final passive clauses: 18203
% Input formulas: 12
% Initial clauses: 15
% 
% Selected by SInE selection: 10
% SInE iterations: 1
% 
% Duplicate literals: 4225
% Trivial inequalities: 3517
% Fw subsumption resolutions: 11356
% Fw demodulations: 14206
% Bw demodulations: 957
% 
% Simple tautologies: 1096
% Equational tautologies: 145
% Forward subsumptions: 42387
% Bw demodulations to eq. taut.: 120
% 
% Binary resolution: 838
% Forward superposition: 39336
% Backward superposition: 32604
% Self superposition: 575
% Equality factoring: 357
% Equality resolution: 297
% 
% Split clauses: 4212
% Split components: 8883
% SAT solver clauses: 4428
% SAT solver unit clauses: 4
% SAT solver binary clauses: 211
% 
% Sat splits: 4212
% Sat splitting refutations: 161
% 
% Memory used [KB]: 44903
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.util_prob
% (56629)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 91
% Passive clauses: 1914
% Generated clauses: 11421
% Final active clauses: 66
% Final passive clauses: 1340
% Input formulas: 12
% Initial clauses: 17
% 
% Duplicate literals: 368
% Trivial inequalities: 241
% Fw subsumption resolutions: 1447
% Fw demodulations: 419
% Bw demodulations: 376
% Global subsumptions: 1
% 
% Simple tautologies: 429
% Equational tautologies: 62
% Forward subsumptions: 6199
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 157
% Forward superposition: 3953
% Backward superposition: 3628
% Self superposition: 137
% Equality factoring: 42
% Equality resolution: 19
% 
% Split clauses: 17
% Split components: 34
% SAT solver clauses: 1982
% SAT solver unit clauses: 37
% SAT solver binary clauses: 123
% 
% Sat splits: 17
% Sat splitting refutations: 48
% 
% Memory used [KB]: 10362
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.util_prob
% (56646)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 391
% Passive clauses: 3481
% Generated clauses: 5390
% Final active clauses: 364
% Final passive clauses: 2764
% Input formulas: 12
% Initial clauses: 17
% 
% Duplicate literals: 306
% Trivial inequalities: 121
% Fw subsumption resolutions: 47
% Bw subsumption resolutions: 152
% Fw demodulations: 18
% Bw demodulations: 1
% 
% Simple tautologies: 599
% Equational tautologies: 1
% Forward subsumptions: 874
% Backward subsumptions: 200
% 
% Binary resolution: 183
% Forward superposition: 4002
% Backward superposition: 262
% Equality resolution: 79
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 3487
% SAT solver unit clauses: 5
% SAT solver binary clauses: 44
% 
% Sat splits: 2
% Sat splitting refutations: 2
% 
% Memory used [KB]: 15095
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 120 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f353,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f41,f60,f64,f66,f67,f69,f62,f71,f72,f74,f75,f79,f77,f80,f84,f82,f85,f89,f90,f94,f95,f97,f99,f100,f102,f103,f105,f118,f98,f132,f134,f136,f137,f139,f149,f135,f163,f165,f167,f168,f170,f174,f190,f192,f194,f195,f197,f201,f210,f208,f226,f228,f229,f232,f248,f250,f258,f257,f256,f255,f259,f209,f260,f207,f261,f86,f263,f91,f268,f270,f87,f280,f282,f289,f291,f92,f301,f303,f310,f313,f227,f316,f319,f249,f336,f180,f337,f231,f182,f284,f181,f183,f185,f350,f349])).
fof(f349,plain,(
  $false | $spl44),
  inference(unit_resulting_resolution,[],[f44,f185])).
fof(f44,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f16])).
fof(f16,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.PREIMAGE_DIFF',aHL_FALSITY)).
fof(f350,plain,(
  $false | $spl44),
  inference(resolution,[],[f185,f44])).
fof(f185,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(cnf_transformation,[],[f185_D])).
fof(f185_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f183,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl42),
  inference(cnf_transformation,[],[f183_D])).
fof(f183_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f181,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl41),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f284,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | p(s(bool,'F_0')) | $spl38),
  inference(forward_demodulation,[],[f283,f57])).
fof(f57,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6,X7] : ! [X8,X9,X11] : s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,X11))),s(fun(X7,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))))),
  file('thm.util_prob.PREIMAGE_DIFF','thm.util_prob.IN_PREIMAGE')).
fof(f283,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl38),
  inference(forward_demodulation,[],[f274,f57])).
fof(f274,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl38),
  inference(superposition,[],[f55,f169])).
fof(f169,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl38),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f55,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X9,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9)))) & ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.util_prob.PREIMAGE_DIFF','thm.pred_set.IN_DIFF')).
fof(f182,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl43),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f231,plain,(
  ~p(s(bool,'T_0')) | $spl57),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f337,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl36),
  inference(superposition,[],[f54,f167])).
fof(f54,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f180,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl40),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f336,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl36),
  inference(superposition,[],[f53,f167])).
fof(f53,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f39])).
fof(f249,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl58),
  inference(cnf_transformation,[],[f249_D])).
fof(f249_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f319,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(trivial_inequality_removal,[],[f318])).
fof(f318,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(superposition,[],[f303,f46])).
fof(f46,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.PREIMAGE_DIFF',aHL_BOOL_CASES)).
fof(f316,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(unit_resulting_resolution,[],[f303,f46])).
fof(f227,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl54),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f313,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(trivial_inequality_removal,[],[f312])).
fof(f312,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(superposition,[],[f282,f46])).
fof(f310,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f282,f46])).
fof(f303,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(cnf_transformation,[],[f303_D])).
fof(f303_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f301,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl48),
  inference(global_subsumption,[],[f41,f196,f292])).
fof(f292,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl48),
  inference(superposition,[],[f52,f196])).
fof(f52,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X9,X10] : (s(fun(X6,bool),X9) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.util_prob.PREIMAGE_DIFF','thm.pred_set.EXTENSION')).
fof(f196,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl48),
  inference(cnf_transformation,[],[f196_D])).
fof(f196_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f92,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl22),
  inference(cnf_transformation,[],[f92_D])).
fof(f92_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f291,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(trivial_inequality_removal,[],[f290])).
fof(f290,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(superposition,[],[f197,f46])).
fof(f289,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(unit_resulting_resolution,[],[f197,f46])).
fof(f282,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(cnf_transformation,[],[f282_D])).
fof(f282_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f280,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl38),
  inference(global_subsumption,[],[f41,f169,f271])).
fof(f271,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl38),
  inference(superposition,[],[f52,f169])).
fof(f87,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl18),
  inference(cnf_transformation,[],[f87_D])).
fof(f87_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f270,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(trivial_inequality_removal,[],[f269])).
fof(f269,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(superposition,[],[f170,f46])).
fof(f268,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f170,f46])).
fof(f91,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl23),
  inference(cnf_transformation,[],[f91_D])).
fof(f91_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f263,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl9 | $spl46)),
  inference(backward_demodulation,[],[f194,f74])).
fof(f86,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl19),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f261,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl5 | $spl36)),
  inference(backward_demodulation,[],[f167,f69])).
fof(f207,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl50),
  inference(cnf_transformation,[],[f207_D])).
fof(f207_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f260,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | $spl22),
  inference(forward_demodulation,[],[f239,f57])).
fof(f239,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl22),
  inference(superposition,[],[f53,f92])).
fof(f209,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl53),
  inference(cnf_transformation,[],[f209_D])).
fof(f209_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f259,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'T_0')) | $spl22),
  inference(forward_demodulation,[],[f240,f57])).
fof(f240,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl22),
  inference(superposition,[],[f54,f92])).
fof(f255,plain,(
  $false | $spl57),
  inference(unit_resulting_resolution,[],[f58,f231])).
fof(f58,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f47])).
fof(f47,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f18])).
fof(f18,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.util_prob.PREIMAGE_DIFF','thm.bool.REFL_CLAUSE')).
fof(f256,plain,(
  $false | $spl57),
  inference(resolution,[],[f231,f58])).
fof(f257,plain,(
  $false | $spl57),
  inference(resolution,[],[f231,f43])).
fof(f43,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.PREIMAGE_DIFF','thm.bool.TRUTH')).
fof(f258,plain,(
  $false | $spl57),
  inference(resolution,[],[f231,f42])).
fof(f42,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.PREIMAGE_DIFF',aHL_TRUTH)).
fof(f250,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl59),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f248,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl22),
  inference(global_subsumption,[],[f41,f92,f238])).
fof(f238,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl22),
  inference(superposition,[],[f52,f92])).
fof(f232,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'T_0')) | $spl18),
  inference(forward_demodulation,[],[f218,f57])).
fof(f218,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl18),
  inference(superposition,[],[f54,f87])).
fof(f229,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | $spl18),
  inference(forward_demodulation,[],[f217,f57])).
fof(f217,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl18),
  inference(superposition,[],[f53,f87])).
fof(f228,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl55),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f226,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl18),
  inference(global_subsumption,[],[f41,f87,f216])).
fof(f216,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl18),
  inference(superposition,[],[f52,f87])).
fof(f208,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl51),
  inference(cnf_transformation,[],[f208_D])).
fof(f208_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f210,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl52),
  inference(cnf_transformation,[],[f210_D])).
fof(f210_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f201,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl24),
  inference(superposition,[],[f55,f93])).
fof(f93,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl24),
  inference(cnf_transformation,[],[f93_D])).
fof(f93_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f197,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(cnf_transformation,[],[f197_D])).
fof(f197_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f195,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f93,f74])).
fof(f194,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl46),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f192,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(trivial_inequality_removal,[],[f191])).
fof(f191,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(superposition,[],[f94,f46])).
fof(f190,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(unit_resulting_resolution,[],[f94,f46])).
fof(f174,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl20),
  inference(superposition,[],[f55,f88])).
fof(f88,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl20),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f170,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f168,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f88,f69])).
fof(f167,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl36),
  inference(cnf_transformation,[],[f167_D])).
fof(f167_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f165,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(trivial_inequality_removal,[],[f164])).
fof(f164,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(superposition,[],[f89,f46])).
fof(f163,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(unit_resulting_resolution,[],[f89,f46])).
fof(f135,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl33),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f149,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl14),
  inference(global_subsumption,[],[f41,f82,f144])).
fof(f144,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl14),
  inference(superposition,[],[f56,f82])).
fof(f56,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.util_prob.PREIMAGE_DIFF',aHL_EXT)).
fof(f139,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl35),
  inference(cnf_transformation,[],[f139_D])).
fof(f139_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f137,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f83,f71])).
fof(f83,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f136,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f136_D])).
fof(f136_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f134,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(trivial_inequality_removal,[],[f133])).
fof(f133,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(superposition,[],[f84,f46])).
fof(f132,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f84,f46])).
fof(f98,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl27),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f118,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl10),
  inference(global_subsumption,[],[f41,f77,f113])).
fof(f113,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl10),
  inference(superposition,[],[f56,f77])).
fof(f105,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl31),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f103,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl29)),
  inference(forward_demodulation,[],[f102,f77])).
fof(f102,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f100,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f78,f66])).
fof(f78,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f78_D])).
fof(f78_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f99,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f97,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(trivial_inequality_removal,[],[f96])).
fof(f96,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(superposition,[],[f79,f46])).
fof(f95,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f79,f46])).
fof(f94,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(cnf_transformation,[],[f94_D])).
fof(f94_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f90,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl9),
  inference(superposition,[],[f74,f46])).
fof(f89,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(cnf_transformation,[],[f89_D])).
fof(f89_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f85,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl5),
  inference(superposition,[],[f69,f46])).
fof(f82,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f84,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f80,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl7),
  inference(superposition,[],[f71,f46])).
fof(f77,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f77_D])).
fof(f77_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f79,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f79_D])).
fof(f79_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f75,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl3),
  inference(superposition,[],[f66,f46])).
fof(f74,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl9),
  inference(cnf_transformation,[],[f74_D])).
fof(f74_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f72,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(forward_demodulation,[],[f61,f57])).
fof(f61,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f60,f52])).
fof(f71,plain,(
  s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f62,plain,(
  s(bool,f(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f60,f56])).
fof(f69,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl5),
  inference(cnf_transformation,[],[f69_D])).
fof(f69_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f67,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(forward_demodulation,[],[f63,f57])).
fof(f63,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f60,f52])).
fof(f66,plain,(
  s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f66_D])).
fof(f66_D,plain,(
  s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f64,plain,(
  s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f60,f56])).
fof(f60,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl1),
  inference(cnf_transformation,[],[f60_D])).
fof(f60_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f41,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.DIFF_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.DIFF_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f29])).
fof(f29,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.DIFF_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) != s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.DIFF_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.DIFF_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'const.pred_set.DIFF_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'const.pred_set.DIFF_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) = s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))))),
  file('thm.util_prob.PREIMAGE_DIFF','thm.util_prob.PREIMAGE_DIFF')).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 46
% Passive clauses: 44
% Generated clauses: 283
% Final active clauses: 28
% Final passive clauses: 6
% Input formulas: 12
% Initial clauses: 17
% 
% Trivial inequalities: 9
% Fw demodulations: 62
% Bw demodulations: 8
% Global subsumptions: 6
% 
% Simple tautologies: 4
% Equational tautologies: 3
% Forward subsumptions: 105
% 
% Binary resolution: 4
% Unit resulting resolution: 22
% Forward superposition: 33
% Backward superposition: 87
% 
% Split clauses: 13
% Split components: 29
% SAT solver clauses: 153
% SAT solver unit clauses: 8
% SAT solver binary clauses: 96
% 
% Sat splits: 13
% Sat splitting refutations: 6
% 
% Memory used [KB]: 1407
% Time elapsed: 0.045 s
% ------------------------------
% ------------------------------
% Success in time 17.973 s
