Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.pred_set
% (102586)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 208
% Passive clauses: 3892
% Generated clauses: 26155
% Final active clauses: 133
% Final passive clauses: 3263
% Input formulas: 68
% Initial clauses: 140
% 
% Duplicate literals: 741
% Trivial inequalities: 299
% Fw subsumption resolutions: 2397
% Fw demodulations: 3765
% Bw demodulations: 63
% 
% Simple tautologies: 367
% Equational tautologies: 44
% Forward subsumptions: 13522
% 
% Binary resolution: 800
% Factoring: 1
% Forward superposition: 11801
% Backward superposition: 5739
% Self superposition: 148
% Equality factoring: 80
% Equality resolution: 30
% 
% Split clauses: 545
% Split components: 1109
% SAT solver clauses: 1374
% SAT solver unit clauses: 49
% SAT solver binary clauses: 370
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 545
% Sat splitting refutations: 24
% 
% TWLsolver clauses: 1357
% TWLsolver calls for satisfiability: 246
% 
% Memory used [KB]: 15095
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f11319,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f11254,f262,f411,f409,f487,f489,f276,f423,f263,f412,f410,f10217,f10234,f353,f456,f264,f414,f415,f417,f354,f455,f11056,f280,f425])).
fof(f425,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f425_D])).
fof(f425_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f280,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f11056,plain,(
  $false | ($spl3 | $spl14 | $spl2233)),
  inference(subsumption_resolution,[],[f11040,f10626])).
fof(f10626,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | $spl2233),
  inference(cnf_transformation,[],[f10626_D])).
fof(f10626_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) <=> ~$spl2233),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2233])])).
fof(f11040,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | ($spl3 | $spl14)),
  inference(resolution,[],[f412,f425])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) ) | $spl50),
  inference(cnf_transformation,[],[f455_D])).
fof(f455_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) & (s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X9,X10] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10)))) <=> s(fun(X6,bool),X9) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10))))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|0','thm.pred_set.SUBSET_INTER_ABSORPTION')).
fof(f417,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f415,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f267,f264])).
fof(f267,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f193])).
fof(f193,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f46])).
fof(f46,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f414,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f264,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|0',aHL_TRUTH)).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) ) | $spl52),
  inference(cnf_transformation,[],[f456_D])).
fof(f456_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f353,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f238])).
fof(f10234,plain,(
  $false | ($spl3 | $spl4 | $spl112)),
  inference(subsumption_resolution,[],[f10233,f414])).
fof(f10233,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl112)),
  inference(forward_demodulation,[],[f412,f489])).
fof(f10217,plain,(
  $false | ($spl1 | $spl4 | $spl52 | $spl112)),
  inference(subsumption_resolution,[],[f10216,f410])).
fof(f10216,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | ($spl4 | $spl52 | $spl112)),
  inference(subsumption_resolution,[],[f10167,f414])).
fof(f10167,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | ($spl52 | $spl112)),
  inference(superposition,[],[f456,f489])).
fof(f410,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | $spl1),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f412,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | $spl3),
  inference(cnf_transformation,[],[f412_D])).
fof(f412_D,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f263,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  (s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))))) & (s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f191])).
fof(f191,plain,(
  ? [X0,X1,X2] : ((s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))) & (s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) | ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f175])).
fof(f175,plain,(
  ? [X0,X1,X2] : (s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <~> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1,X2] : (s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ~! [X0] : ! [X1,X2] : (s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7))) <=> p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|0','thm.pred_set.INTER_SUBSET_EQN|split|0')).
fof(f423,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f276,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.INTER_SUBSET_EQN|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f489,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | $spl112),
  inference(cnf_transformation,[],[f489_D])).
fof(f489_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f487,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | ($spl2 | $spl10)),
  inference(resolution,[],[f423,f411])).
fof(f409,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | $spl0),
  inference(cnf_transformation,[],[f409_D])).
fof(f409_D,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f411,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | $spl2),
  inference(cnf_transformation,[],[f411_D])).
fof(f411_D,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f262,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f192])).
fof(f11254,plain,(
  $false | ($spl0 | $spl7 | $spl50 | $spl2232)),
  inference(subsumption_resolution,[],[f11253,f417])).
fof(f11253,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl50 | $spl2232)),
  inference(forward_demodulation,[],[f11252,f10625])).
fof(f10625,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) | $spl2232),
  inference(cnf_transformation,[],[f10625_D])).
fof(f10625_D,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1))) <=> ~$spl2232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2232])])).
fof(f11252,plain,(
  p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl50)),
  inference(trivial_inequality_removal,[],[f11193])).
fof(f11193,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),sK2) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),sK1)))) | ($spl0 | $spl50)),
  inference(superposition,[],[f455,f409])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 125
% Passive clauses: 1460
% Generated clauses: 12162
% Final active clauses: 90
% Final passive clauses: 976
% Input formulas: 68
% Initial clauses: 140
% 
% Duplicate literals: 184
% Trivial inequalities: 75
% Fw subsumption resolutions: 1067
% Fw demodulations: 1904
% Bw demodulations: 374
% Condensations: 599
% Global subsumptions: 17
% 
% Simple tautologies: 168
% Equational tautologies: 9
% Forward subsumptions: 4416
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 596
% Factoring: 1
% Forward superposition: 3295
% Backward superposition: 1401
% Self superposition: 73
% Equality factoring: 18
% Equality resolution: 16
% 
% Split clauses: 78
% Split components: 157
% SAT solver clauses: 1853
% SAT solver unit clauses: 159
% SAT solver binary clauses: 62
% SAT solver learnt clauses: 36
% SAT solver learnt literals: 1
% 
% Sat splits: 78
% Sat splitting refutations: 41
% 
% TWLsolver clauses: 1783
% TWLsolver calls for satisfiability: 496
% 
% Memory used [KB]: 9850
% Time elapsed: 0.489 s
% ------------------------------
% ------------------------------
% Success in time 1.11 s
