Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.util_prob
% (73253)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 178
% Passive clauses: 3121
% Generated clauses: 17019
% Final active clauses: 170
% Final passive clauses: 2756
% Input formulas: 18
% Initial clauses: 38
% 
% Duplicate literals: 744
% Trivial inequalities: 310
% Fw subsumption resolutions: 2425
% Fw demodulations: 2841
% Bw demodulations: 4
% 
% Simple tautologies: 357
% Equational tautologies: 21
% Forward subsumptions: 6018
% 
% Binary resolution: 821
% Forward superposition: 6644
% Backward superposition: 3056
% Self superposition: 96
% Equality factoring: 48
% Equality resolution: 42
% 
% Split clauses: 762
% Split components: 1535
% SAT solver clauses: 1368
% SAT solver unit clauses: 36
% SAT solver binary clauses: 320
% 
% Sat splits: 762
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1367
% TWLsolver calls for satisfiability: 261
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (73271)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 117
% Passive clauses: 1154
% Generated clauses: 10814
% Final active clauses: 111
% Final passive clauses: 1025
% Input formulas: 18
% Initial clauses: 38
% 
% Duplicate literals: 364
% Trivial inequalities: 195
% Fw subsumption resolutions: 866
% Fw demodulations: 1708
% Bw demodulations: 16
% Condensations: 559
% 
% Simple tautologies: 184
% Equational tautologies: 18
% Forward subsumptions: 3993
% 
% Binary resolution: 391
% Forward superposition: 3459
% Backward superposition: 1933
% Self superposition: 92
% Equality factoring: 46
% Equality resolution: 12
% 
% Split clauses: 125
% Split components: 255
% SAT solver clauses: 1767
% SAT solver unit clauses: 36
% SAT solver binary clauses: 24
% 
% Sat splits: 125
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1767
% TWLsolver calls for satisfiability: 276
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (73293)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 540
% Passive clauses: 7116
% Generated clauses: 176604
% Final active clauses: 138
% Final passive clauses: 794
% Input formulas: 18
% Initial clauses: 38
% 
% Duplicate literals: 7428
% Trivial inequalities: 1115
% Fw subsumption resolutions: 18796
% Fw demodulations: 25853
% Bw demodulations: 4202
% Condensations: 16089
% Global subsumptions: 457
% 
% Simple tautologies: 1481
% Equational tautologies: 498
% Forward subsumptions: 90060
% Fw demodulations to eq. taut.: 36
% Bw demodulations to eq. taut.: 242
% 
% Unit resulting resolution: 3226
% Forward superposition: 42213
% Backward superposition: 33920
% Self superposition: 710
% Equality factoring: 329
% Equality resolution: 131
% 
% Split clauses: 657
% Split components: 1328
% SAT solver clauses: 9261
% SAT solver unit clauses: 53
% SAT solver binary clauses: 1960
% 
% Sat splits: 657
% Sat splitting refutations: 371
% 
% Memory used [KB]: 87759
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% (73573)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 608
% Passive clauses: 19803
% Generated clauses: 98227
% Final active clauses: 349
% Final passive clauses: 16044
% Input formulas: 18
% Initial clauses: 34
% 
% Selected by SInE selection: 14
% SInE iterations: 1
% 
% Duplicate literals: 1905
% Trivial inequalities: 2231
% Fw subsumption resolutions: 17674
% Fw demodulations: 9621
% Bw demodulations: 185
% 
% Simple tautologies: 1838
% Equational tautologies: 71
% Forward subsumptions: 40854
% Bw demodulations to eq. taut.: 15
% 
% Binary resolution: 13338
% Factoring: 13
% Forward superposition: 27613
% Backward superposition: 22300
% Self superposition: 261
% Equality factoring: 181
% Equality resolution: 135
% 
% Split clauses: 6192
% Split components: 13269
% SAT solver clauses: 6328
% SAT solver unit clauses: 24
% SAT solver binary clauses: 1022
% 
% Sat splits: 6192
% Sat splitting refutations: 26
% 
% Memory used [KB]: 47845
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.util_prob
% (73675)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 1519
% Generated clauses: 13532
% Final active clauses: 82
% Final passive clauses: 1398
% Input formulas: 18
% Initial clauses: 38
% 
% Duplicate literals: 487
% Trivial inequalities: 409
% Fw subsumption resolutions: 1905
% Fw demodulations: 1121
% Bw demodulations: 14
% Global subsumptions: 1
% 
% Simple tautologies: 583
% Equational tautologies: 80
% Forward subsumptions: 7112
% 
% Unit resulting resolution: 591
% Forward superposition: 4619
% Backward superposition: 4152
% Self superposition: 109
% Equality factoring: 52
% Equality resolution: 12
% 
% Split clauses: 16
% Split components: 33
% SAT solver clauses: 1599
% SAT solver unit clauses: 77
% SAT solver binary clauses: 92
% 
% Sat splits: 16
% Sat splitting refutations: 61
% 
% Memory used [KB]: 9722
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f1430,plain,(
  $false),
  inference(subsumption_resolution,[],[f1429,f71])).
fof(f71,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.SCHROEDER_CLOSE_SUBSET',aHL_TRUTH)).
fof(f1429,plain,(
  ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f1411,f155])).
fof(f155,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))),
  inference(resolution,[],[f119,f75])).
fof(f75,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.util_prob.SCHROEDER_CLOSE_SUBSET','thm.bool.EQ_CLAUSES|split|1')).
fof(f119,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))),
  inference(resolution,[],[f69,f86])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : ! [X8,X12] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X12)))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8)))) => p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X12))))))),
  file('thm.util_prob.SCHROEDER_CLOSE_SUBSET','thm.pred_set.SUBSET_DEF')).
fof(f69,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2))))))),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(sK0,bool),sK2),s(fun(sK0,bool),'const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f45])).
fof(f45,plain,(
  ? [X0,X1,X2] : ~p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ~! [X0] : ! [X1,X2] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X2),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.util_prob.schroeder_close_2'(s(fun(X6,X6),X7),s(fun(X6,bool),X8))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X6,bool),X8),s(fun(X6,bool),'const.util_prob.schroeder_close_2'(s(fun(X6,X6),X7),s(fun(X6,bool),X8))))))),
  file('thm.util_prob.SCHROEDER_CLOSE_SUBSET','thm.util_prob.SCHROEDER_CLOSE_SUBSET')).
fof(f1411,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))))),
  inference(superposition,[],[f423,f82])).
fof(f82,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.arithmetic.FUNPOW_3'(s(fun(X0,X0),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X7,X9] : s(X6,X9) = s(X6,'const.arithmetic.FUNPOW_3'(s(fun(X6,X6),X7),s('type.num.num','const.num.|48|_0'),s(X6,X9)))),
  file('thm.util_prob.SCHROEDER_CLOSE_SUBSET','thm.arithmetic.FUNPOW|split|0')).
fof(f423,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(sK0,bool),fun(sK0,bool)),'const.pred_set.IMAGE_1'(s(fun(sK0,sK0),sK1))),s('type.num.num',X0),s(fun(sK0,bool),sK2))))))) )),
  inference(resolution,[],[f120,f90])).
fof(f90,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X4),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',sK4(X3,X2,X1,X0)),s(fun(X0,bool),X3))))))) & (! [X4] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X4),s(fun(X0,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3)))))) | ? [X5] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X5),s(fun(X0,bool),X3))))))) & (! [X4] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X4),s(fun(X0,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3))))))))),
  inference(rectify,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3)))))) | ? [X4] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X4),s(fun(X0,bool),X3))))))) & (! [X4] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X4),s(fun(X0,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3)))))) <=> ? [X4] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X4),s(fun(X0,bool),X3)))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.util_prob.schroeder_close_2'(s(fun(X0,X0),X2),s(fun(X0,bool),X3)))))) <=> ? [X4] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X0,bool),fun(X0,bool)),'const.pred_set.IMAGE_1'(s(fun(X0,X0),X2))),s('type.num.num',X4),s(fun(X0,bool),X3)))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X9,X7,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.util_prob.schroeder_close_2'(s(fun(X6,X6),X7),s(fun(X6,bool),X8)))))) <=> ? [X10] : p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.arithmetic.FUNPOW_3'(s(fun(fun(X6,bool),fun(X6,bool)),'const.pred_set.IMAGE_1'(s(fun(X6,X6),X7))),s('type.num.num',X10),s(fun(X6,bool),X8)))))))),
  file('thm.util_prob.SCHROEDER_CLOSE_SUBSET','thm.util_prob.SCHROEDER_CLOSE')).
fof(f120,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.util_prob.schroeder_close_2'(s(fun(sK0,sK0),sK1),s(fun(sK0,bool),sK2))))))),
  inference(resolution,[],[f69,f87])).
fof(f87,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.pred_set.SUBSET_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f60])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 69
% Passive clauses: 282
% Generated clauses: 1312
% Final active clauses: 65
% Final passive clauses: 239
% Input formulas: 18
% Initial clauses: 38
% 
% Duplicate literals: 15
% Trivial inequalities: 57
% Fw subsumption resolutions: 92
% Fw demodulations: 168
% Bw demodulations: 5
% Condensations: 11
% 
% Simple tautologies: 12
% Forward subsumptions: 690
% 
% Binary resolution: 193
% Forward superposition: 444
% Backward superposition: 337
% Equality factoring: 1
% 
% SAT solver clauses: 284
% SAT solver unit clauses: 15
% SAT solver binary clauses: 114
% 
% Memory used [KB]: 2430
% Time elapsed: 0.061 s
% ------------------------------
% ------------------------------
% Success in time 16.465 s
