Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.numeral
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.numeral
% SZS output start Proof for thm.numeral
fof(f297,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f288])).
fof(f288,plain,(
  s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) != s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))),s('type.num.num',sK1)))),
  inference(superposition,[],[f28,f269])).
fof(f269,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f268,f217])).
fof(f217,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f216,f35])).
fof(f35,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numeral.internal_mult_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numeral.internal_mult_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8,X9] : s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X8),s('type.num.num',X9))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numeral.internal_mult_0'),s('type.num.num',X8))),s('type.num.num',X9)))),
  file('thm.numeral.internal_mult_characterisation|split|2','ARITY_0_2_const.numeral.internal_mult')).
fof(f216,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numeral.internal_mult_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(forward_demodulation,[],[f36,f30])).
fof(f30,plain,(
  s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numeral.internal_mult_0') = s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')),
  inference(cnf_transformation,[],[f11])).
fof(f11,axiom,(
  s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.numeral.internal_mult_0') = s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0')),
  file('thm.numeral.internal_mult_characterisation|split|2','thm.numeral.internal_mult_def')).
fof(f36,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8,X9] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))) = s('type.num.num',happ(s(fun('type.num.num','type.num.num'),happ(s(fun('type.num.num',fun('type.num.num','type.num.num')),'const.arithmetic.*_0'),s('type.num.num',X8))),s('type.num.num',X9)))),
  file('thm.numeral.internal_mult_characterisation|split|2','ARITY_0_2_const.arithmetic.*')).
fof(f268,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f37,f217])).
fof(f37,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',X1)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num',X7))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num',X7))))),s('type.num.num',X7)))))),
  file('thm.numeral.internal_mult_characterisation|split|2','thm.numeral.numeral_mult|split|2')).
fof(f28,plain,(
  s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) != s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',sK0),s('type.num.num',sK1))))),s('type.num.num',sK1)))))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',sK0))),s('type.num.num',sK1))) != s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',sK0),s('type.num.num',sK1))))),s('type.num.num',sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f24])).
fof(f24,plain,(
  ? [X0,X1] : s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1))) != s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',X1)))))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1] : s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num',X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num',X7))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X6),s('type.num.num',X7))))),s('type.num.num',X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6))),s('type.num.num',X7))) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.numeral.iDUB_1'(s('type.num.num','const.numeral.internal_mult_2'(s('type.num.num',X6),s('type.num.num',X7))))),s('type.num.num',X7)))))),
  file('thm.numeral.internal_mult_characterisation|split|2','thm.numeral.internal_mult_characterisation|split|2')).
% SZS output end Proof for thm.numeral
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 19
% Passive clauses: 43
% Generated clauses: 241
% Final active clauses: 17
% Final passive clauses: 15
% Input formulas: 12
% Initial clauses: 12
% 
% Duplicate literals: 37
% Trivial inequalities: 3
% Fw demodulations: 6
% 
% Simple tautologies: 16
% Equational tautologies: 14
% Forward subsumptions: 117
% 
% Binary resolution: 4
% Forward superposition: 72
% Backward superposition: 61
% Self superposition: 34
% Equality factoring: 26
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 2
% SAT solver unit clauses: 1
% SAT solver binary clauses: 1
% 
% Sat splits: 1
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2
% TWLsolver calls for satisfiability: 2
% 
% Memory used [KB]: 767
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
% Success in time 0.019 s
