Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.list
% (86981)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 150
% Passive clauses: 2912
% Generated clauses: 22394
% Final active clauses: 128
% Final passive clauses: 2463
% Input formulas: 32
% Initial clauses: 78
% 
% Duplicate literals: 849
% Trivial inequalities: 463
% Fw subsumption resolutions: 2995
% Fw demodulations: 3589
% Bw demodulations: 17
% 
% Simple tautologies: 510
% Equational tautologies: 42
% Forward subsumptions: 10471
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 355
% Factoring: 2
% Forward superposition: 7802
% Backward superposition: 5860
% Self superposition: 165
% Equality factoring: 89
% Equality resolution: 42
% 
% Split clauses: 359
% Split components: 742
% SAT solver clauses: 745
% SAT solver unit clauses: 22
% SAT solver binary clauses: 128
% 
% Sat splits: 359
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 737
% TWLsolver calls for satisfiability: 131
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (87005)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 96
% Passive clauses: 588
% Generated clauses: 10870
% Final active clauses: 88
% Final passive clauses: 466
% Input formulas: 32
% Initial clauses: 78
% 
% Duplicate literals: 543
% Trivial inequalities: 160
% Fw subsumption resolutions: 971
% Fw demodulations: 1095
% Bw demodulations: 29
% Condensations: 1158
% Global subsumptions: 9
% 
% Simple tautologies: 233
% Equational tautologies: 31
% Forward subsumptions: 4984
% 
% Binary resolution: 198
% Factoring: 1
% Forward superposition: 2806
% Backward superposition: 3244
% Self superposition: 120
% Equality factoring: 72
% Equality resolution: 26
% 
% Split clauses: 81
% Split components: 170
% SAT solver clauses: 1171
% SAT solver unit clauses: 85
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 18
% 
% Sat splits: 81
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1135
% TWLsolver calls for satisfiability: 377
% 
% Memory used [KB]: 8955
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f37508,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f37489])).
fof(f37489,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),sK1))))) != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),sK1)))))),
  inference(superposition,[],[f126,f2920])).
fof(f2920,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),X0))))) = s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),X0)))))) )),
  inference(forward_demodulation,[],[f2445,f176])).
fof(f176,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X10,X11] : s(X6,X10) = s(X6,'const.bool.COND_3'(s(bool,'T_0'),s(X6,X10),s(X6,X11)))),
  file('thm.list.FILTER_COND_REWRITE|split|1','thm.bool.COND_CLAUSES|split|0')).
fof(f2445,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),X0))))) = s('type.list.list'(sK0),'const.bool.COND_3'(s(bool,'T_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),X0))))),s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),X0)))))) )),
  inference(superposition,[],[f179,f223])).
fof(f223,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))),
  inference(unit_resulting_resolution,[],[f125,f133])).
fof(f133,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.list.FILTER_COND_REWRITE|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f125,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3))))),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  p(s(bool,happ(s(fun(sK0,bool),sK2),s(sK0,sK3)))) & s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),sK1))))) != s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f75])).
fof(f75,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))))))),
  inference(ennf_transformation,[],[f35])).
fof(f35,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ~! [X0] : ! [X1,X2,X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),X1))))) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X2),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) => s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X7))))) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9)))) => s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X7))))) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),X7))))))),
  file('thm.list.FILTER_COND_REWRITE|split|1','thm.list.FILTER_COND_REWRITE|split|1')).
fof(f179,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2,X3] : s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))) = s('type.list.list'(X0),'const.bool.COND_3'(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))),s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ! [X8,X9,X12] : s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),X12))))) = s('type.list.list'(X6),'const.bool.COND_3'(s(bool,happ(s(fun(X6,bool),X8),s(X6,X9))),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X9),s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X12))))),s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X8),s('type.list.list'(X6),X12)))))),
  file('thm.list.FILTER_COND_REWRITE|split|1','thm.list.FILTER|split|1')).
fof(f126,plain,(
  s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),sK1))))) != s('type.list.list'(sK0),'const.list.FILTER_2'(s(fun(sK0,bool),sK2),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK3),s('type.list.list'(sK0),sK1)))))),
  inference(cnf_transformation,[],[f91])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 89
% Passive clauses: 384
% Generated clauses: 35785
% Final active clauses: 81
% Final passive clauses: 282
% Input formulas: 32
% Initial clauses: 78
% 
% Duplicate literals: 327
% Trivial inequalities: 102
% Fw subsumption resolutions: 428
% Fw demodulations: 22226
% Bw demodulations: 20
% Condensations: 552
% Global subsumptions: 6
% 
% Simple tautologies: 155
% Equational tautologies: 77
% Forward subsumptions: 7482
% Fw demodulations to eq. taut.: 4002
% 
% Unit resulting resolution: 10172
% Factoring: 2
% Forward superposition: 1716
% Backward superposition: 1498
% Self superposition: 70
% Equality factoring: 32
% Equality resolution: 22
% 
% Split clauses: 20
% Split components: 45
% SAT solver clauses: 437
% SAT solver unit clauses: 23
% SAT solver binary clauses: 71
% 
% Sat splits: 20
% Sat splitting refutations: 5
% 
% Memory used [KB]: 14328
% Time elapsed: 0.774 s
% ------------------------------
% ------------------------------
% Success in time 2.097 s
