Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.seq
% (49599)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 175
% Passive clauses: 2493
% Generated clauses: 21147
% Final active clauses: 161
% Final passive clauses: 2122
% Input formulas: 27
% Initial clauses: 50
% 
% Duplicate literals: 1244
% Trivial inequalities: 152
% Fw subsumption resolutions: 2480
% Fw demodulations: 4161
% Bw demodulations: 12
% 
% Simple tautologies: 303
% Equational tautologies: 21
% Forward subsumptions: 9071
% 
% Binary resolution: 235
% Forward superposition: 8530
% Backward superposition: 3954
% Self superposition: 109
% Equality factoring: 42
% Equality resolution: 24
% 
% Split clauses: 588
% Split components: 1183
% SAT solver clauses: 1119
% SAT solver unit clauses: 31
% SAT solver binary clauses: 136
% 
% Sat splits: 588
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1117
% TWLsolver calls for satisfiability: 183
% 
% Memory used [KB]: 14456
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.seq
% (49627)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 141
% Passive clauses: 956
% Generated clauses: 12143
% Final active clauses: 83
% Final passive clauses: 362
% Input formulas: 27
% Initial clauses: 50
% 
% Duplicate literals: 287
% Trivial inequalities: 148
% Fw subsumption resolutions: 785
% Fw demodulations: 1291
% Bw demodulations: 459
% Condensations: 320
% Global subsumptions: 4
% 
% Simple tautologies: 98
% Equational tautologies: 22
% Forward subsumptions: 6462
% Bw demodulations to eq. taut.: 46
% 
% Binary resolution: 124
% Forward superposition: 4667
% Backward superposition: 2737
% Self superposition: 159
% Equality factoring: 39
% Equality resolution: 1
% 
% Split clauses: 84
% Split components: 168
% SAT solver clauses: 1811
% SAT solver unit clauses: 62
% SAT solver binary clauses: 60
% SAT solver learnt clauses: 9
% SAT solver learnt literals: 1
% 
% Sat splits: 84
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1794
% TWLsolver calls for satisfiability: 290
% 
% Memory used [KB]: 15863
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.seq
% (49669)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 459
% Passive clauses: 9077
% Generated clauses: 148829
% Final active clauses: 124
% Final passive clauses: 1046
% Input formulas: 27
% Initial clauses: 50
% Discarded non-redundant clauses: 1276
% 
% Duplicate literals: 4802
% Trivial inequalities: 1371
% Fw subsumption resolutions: 14146
% Fw demodulations: 25039
% Bw demodulations: 5532
% Condensations: 9886
% Global subsumptions: 1863
% 
% Simple tautologies: 1408
% Equational tautologies: 344
% Forward subsumptions: 78553
% Fw demodulations to eq. taut.: 17
% Bw demodulations to eq. taut.: 182
% 
% Unit resulting resolution: 1450
% Forward superposition: 39655
% Backward superposition: 28475
% Self superposition: 642
% Equality factoring: 282
% Equality resolution: 95
% 
% Split clauses: 212
% Split components: 434
% SAT solver clauses: 13429
% SAT solver unit clauses: 241
% SAT solver binary clauses: 3443
% 
% Sat splits: 212
% Sat splitting refutations: 2062
% 
% Memory used [KB]: 102983
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.seq
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.seq
% SZS output start Proof for thm.seq
fof(f48070,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f149,f169,f173,f156,f175,f171,f213,f382,f388,f437,f384,f516,f11434,f11435,f11436,f11484,f12703,f12704,f12706,f12707,f12709,f12716,f12697,f12696,f12753,f12934,f12935,f12936,f12937,f13737,f15268,f15269,f16915,f16916,f16918,f16919,f18714,f18902,f18716,f18903,f18720,f18722,f18757,f18912,f18761,f18913,f18765,f18914,f18769,f18915,f18771,f18916,f18775,f18918,f18776,f18919,f18779,f18920,f18783,f18922,f18784,f18923,f18924,f18925,f18791,f11482,f18792,f18905,f18796,f18926,f18927,f18928,f18901,f20381,f20384,f20387,f20389,f24812,f25125,f24814,f24815,f25126,f25127,f26792,f20383,f26793,f20385,f26795,f20388,f26796,f20390,f27170,f27171,f27172,f27173,f11483,f20382,f27183,f27184,f27185,f27186,f27187,f27188,f27189,f27190,f27191,f27192,f27193,f27194,f27519,f27524,f29515,f151,f150,f33896,f33897,f33899,f33900,f33902,f33903,f33904,f33905,f33906,f33907,f33909,f33910,f33911,f33912,f33913,f33914,f33915,f33916,f33917,f33918,f33919,f33920,f33921,f33922,f33923,f33924,f34324,f34325,f34326,f34956,f34957,f34958,f36239,f36240,f36241,f34376,f36891,f36945,f36906,f36951,f36907,f36952,f36908,f36953,f36915,f36957,f36918,f36946,f36960,f36971,f36972,f37035,f37031,f37030,f12718,f37474,f37475,f37539,f37534,f37533,f37535,f37629,f37633,f37628,f37627,f37634,f12717,f38735,f38737,f38738,f38822,f38823,f38824,f38736,f38951,f38952,f38953,f38954,f38955,f38956,f38957,f38958,f38960,f38961,f38963,f12710,f38964,f38965,f38966,f38967,f38968,f38969,f39898,f39900,f39901,f39899,f40006,f40008,f40009,f40007,f40124,f40126,f40127,f40125,f40232,f40234,f40235,f40233,f40320,f40322,f40323,f40420,f34375,f40421,f34377,f40422,f34378,f40321,f40823,f40824,f41009,f41010,f41362,f41363,f42670,f42671,f42672,f42673,f42674,f42675,f42750,f42751,f42752,f42753,f42754,f42755,f43933,f43936,f43942,f44029,f43935,f44031,f43937,f44038,f43943,f46301,f46302,f46316,f46320,f46318,f47299,f12705,f47300,f47314,f43934,f47318,f47316,f44030,f47988,f16917,f47990,f48069])).
fof(f48069,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f47754,f41227])).
fof(f41227,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))),
  inference(resolution,[],[f41219,f117])).
fof(f117,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.seq.SEQ_BCONV','thm.bool.EQ_CLAUSES|split|0')).
fof(f41219,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))))),
  inference(subsumption_resolution,[],[f41218,f101])).
fof(f101,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X10] : ((p(s(bool,X10)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.seq.SEQ_BCONV','thm.bool.AND_CLAUSES|split|3')).
fof(f41218,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))))),
  inference(forward_demodulation,[],[f41217,f159])).
fof(f159,plain,(
  s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0)))),
  inference(resolution,[],[f86,f115])).
fof(f115,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.seq.SEQ_BCONV','thm.bool.EQ_CLAUSES|split|3')).
fof(f86,plain,(
  ~p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0)))) & p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0)))) & ~p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f57])).
fof(f57,plain,(
  ? [X0] : (p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) & p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) & ~p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ? [X0] : ((p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) & p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))) & ~p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ~! [X0] : ((p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) & p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))) => p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ((p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X6)))) & p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X6))))) => p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ((p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X6)))) & p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X6))))) => p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X6)))))),
  file('thm.seq.SEQ_BCONV','thm.seq.SEQ_BCONV')).
fof(f41217,plain,(
  p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))))),
  inference(subsumption_resolution,[],[f41209,f121])).
fof(f121,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.seq.SEQ_BCONV','thm.bool.TRUTH')).
fof(f41209,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))))),
  inference(superposition,[],[f29350,f238])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0)))),
  inference(resolution,[],[f84,f119])).
fof(f119,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.seq.SEQ_BCONV','thm.bool.EQ_CLAUSES|split|1')).
fof(f84,plain,(
  p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(cnf_transformation,[],[f61])).
fof(f29350,plain,(
  ( ! [X1] : (~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X1)))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X1)))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X1)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X1)))))))) )),
  inference(forward_demodulation,[],[f29298,f89])).
fof(f89,plain,(
  ( ! [X0] : (s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0))) = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0] : s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0))) = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X6))) = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X6)))))),
  file('thm.seq.SEQ_BCONV','thm.seq.SEQ_NEG_CONV')).
fof(f29298,plain,(
  ( ! [X1] : (~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X1)))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X1)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X1)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X1)))))))) )),
  inference(superposition,[],[f29078,f123])).
fof(f123,plain,(
  ( ! [X0] : (s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))) = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))) = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X6))) = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X6)))))),
  file('thm.seq.SEQ_BCONV','thm.seq.SEQ_NEG_BOUNDED')).
fof(f29078,plain,(
  ( ! [X0] : (~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1(X0)),s('type.num.num',sK2(X0)))))) )),
  inference(forward_demodulation,[],[f87,f124])).
fof(f124,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X13,X14] : s(bool,'const.arithmetic.>=_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X13))),s('type.num.num',X14)))),
  file('thm.seq.SEQ_BCONV','ARITY_0_2_const.arithmetic.>=')).
fof(f87,plain,(
  ( ! [X0] : (p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',sK1(X0)))),s('type.num.num',sK2(X0))))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) | (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',sK1(X0)))),s('type.num.num',sK2(X0))))) & ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK1(X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK2(X0)))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2])],[f59])).
fof(f59,plain,(
  ! [X0] : (~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) | ? [X1,X2] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X1))),s('type.num.num',X2)))) & ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) | ? [X1,X2] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X1))),s('type.num.num',X2)))) & ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : ((p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) & ! [X1,X2] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X1))),s('type.num.num',X2)))) => p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))) => p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ((p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X6)))) & ! [X8,X7] : (p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X8))),s('type.num.num',X7)))) => p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X7)))))))) => p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X6)))))),
  file('thm.seq.SEQ_BCONV','thm.seq.SEQ_ICONV')).
fof(f47754,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl0),
  inference(subsumption_resolution,[],[f47753,f121])).
fof(f47753,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl0),
  inference(forward_demodulation,[],[f47752,f238])).
fof(f47752,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0)))) | $spl0),
  inference(subsumption_resolution,[],[f47751,f120])).
fof(f120,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.seq.SEQ_BCONV',aHL_FALSITY)).
fof(f47751,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0)))) | $spl0),
  inference(forward_demodulation,[],[f47719,f159])).
fof(f47719,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0)))) | $spl0),
  inference(resolution,[],[f577,f29349])).
fof(f29349,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f29348,f27976])).
fof(f27976,plain,(
  ( ! [X2,X3] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',X2),s('type.realax.real',X3))) = s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) )),
  inference(superposition,[],[f27848,f132])).
fof(f132,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s(bool,'const.real.real_ge_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X11,X12] : s(bool,'const.real.real_ge_2'(s('type.realax.real',X11),s('type.realax.real',X12))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X11)))),
  file('thm.seq.SEQ_BCONV','thm.real.real_ge')).
fof(f27848,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(forward_demodulation,[],[f131,f132])).
fof(f131,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11,X12] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X11))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12)))))),
  file('thm.seq.SEQ_BCONV','thm.real.REAL_LE_NEG')).
fof(f29348,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))))))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f29347,f129])).
fof(f129,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X13] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',X13)))),
  file('thm.seq.SEQ_BCONV','ARITY_0_1_const.realax.real_neg')).
fof(f29347,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))))))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f29346,f128])).
fof(f128,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),X3),s(X1,happ(s(fun(X2,X1),X4),s(X2,X5))))) = s(X0,happ(s(fun(X2,X0),'_dst_x0x1_2'(s(fun(X1,X0),X3),s(fun(X2,X1),X4))),s(X2,X5)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X15,X16,X17] : ! [X18,X19,X20] : s(X15,happ(s(fun(X16,X15),X18),s(X16,happ(s(fun(X17,X16),X19),s(X17,X20))))) = s(X15,happ(s(fun(X17,X15),'_dst_x0x1_2'(s(fun(X16,X15),X18),s(fun(X17,X16),X19))),s(X17,X20)))),
  file('thm.seq.SEQ_BCONV','_dst_x0x1_ax')).
fof(f29346,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))))))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f29345,f129])).
fof(f29345,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))))))))))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f29344,f128])).
fof(f29344,plain,(
  ( ! [X0] : (p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))))))) )),
  inference(forward_demodulation,[],[f29297,f89])).
fof(f29297,plain,(
  ( ! [X0] : (~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0)))) | p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0))),s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X0)))))))))) )),
  inference(superposition,[],[f88,f123])).
fof(f88,plain,(
  ( ! [X0] : (p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK1(X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK2(X0))))))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f62])).
fof(f577,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) ) | $spl0),
  inference(resolution,[],[f150,f118])).
fof(f118,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f47990,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))) | $spl220),
  inference(cnf_transformation,[],[f47990_D])).
fof(f47990_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f16917,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,X60)) ) | $spl34),
  inference(cnf_transformation,[],[f16917_D])).
fof(f16917_D,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,X60)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f47988,plain,(
  ( ! [X17] : (s(bool,'F_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f47850,f127])).
fof(f127,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.seq.SEQ_BCONV',aHL_BOOL_CASES)).
fof(f47850,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK2('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0)))),s('type.num.num',sK1('_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) ) | $spl14),
  inference(superposition,[],[f35542,f41227])).
fof(f35542,plain,(
  ( ! [X70,X68,X69] : (s(bool,X70) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X69),s('type.num.num',X68))) | s(bool,'F_0') = s(bool,X70) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X68),s('type.num.num',X69)))) ) | $spl14),
  inference(superposition,[],[f34939,f130])).
fof(f130,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X8] : s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',X8))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))),
  file('thm.seq.SEQ_BCONV','thm.arithmetic.GREATER_EQ')).
fof(f34939,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl14),
  inference(subsumption_resolution,[],[f34873,f120])).
fof(f34873,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X18)) ) | $spl14),
  inference(superposition,[],[f11483,f127])).
fof(f44030,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X43),s('type.num.num',X42)))) ) | $spl210),
  inference(cnf_transformation,[],[f44030_D])).
fof(f44030_D,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X43),s('type.num.num',X42)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f47316,plain,(
  s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl219),
  inference(cnf_transformation,[],[f47316_D])).
fof(f47316_D,plain,(
  s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f47318,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X72),s('type.num.num',X71)))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f47206,f27549])).
fof(f27549,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X11),s('type.num.num',X10)))) )),
  inference(superposition,[],[f130,f127])).
fof(f47206,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X71),s('type.num.num',X72))) | s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X72),s('type.num.num',X71)))) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f37620,f46267])).
fof(f46267,plain,(
  s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl24),
  inference(trivial_inequality_removal,[],[f46147])).
fof(f46147,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | $spl24),
  inference(superposition,[],[f16970,f46135])).
fof(f46135,plain,(
  s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | $spl24),
  inference(trivial_inequality_removal,[],[f46121])).
fof(f46121,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | $spl24),
  inference(superposition,[],[f16970,f159])).
fof(f16970,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X13))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X13)))))) ) | $spl24),
  inference(forward_demodulation,[],[f16947,f368])).
fof(f368,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0))))))),
  inference(resolution,[],[f252,f115])).
fof(f252,plain,(
  ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))))),
  inference(subsumption_resolution,[],[f237,f86])).
fof(f237,plain,(
  p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))))),
  inference(resolution,[],[f84,f88])).
fof(f16947,plain,(
  ( ! [X13] : (s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X13))) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X13)))))) ) | $spl24),
  inference(superposition,[],[f12710,f89])).
fof(f37620,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X62),s('type.num.num',X63))) | s(bool,'T_0') != s(bool,X64) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X63),s('type.num.num',X62)))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f37593])).
fof(f37593,plain,(
  ( ! [X64,X62,X63] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X64) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X62),s('type.num.num',X63))) | s(bool,'T_0') != s(bool,X64) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X63),s('type.num.num',X62)))) ) | $spl14),
  inference(superposition,[],[f37473,f27549])).
fof(f37473,plain,(
  ( ! [X37,X36] : (s(bool,'F_0') != s(bool,X36) | s(bool,X36) != s(bool,X37) | s(bool,'T_0') != s(bool,X37)) ) | $spl14),
  inference(resolution,[],[f35813,f120])).
fof(f35813,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X13)) | s(bool,X12) != s(bool,X13) | s(bool,X12) != s(bool,X14) | s(bool,'T_0') != s(bool,X14)) ) | $spl14),
  inference(resolution,[],[f35088,f116])).
fof(f116,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f81])).
fof(f35088,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl14),
  inference(superposition,[],[f13097,f119])).
fof(f13097,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,X8) != s(bool,X9) | p(s(bool,X8))) ) | $spl14),
  inference(resolution,[],[f12944,f116])).
fof(f12944,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,X17)) ) | $spl14),
  inference(subsumption_resolution,[],[f11635,f216])).
fof(f216,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f215,f135])).
fof(f135,plain,(
  s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0)))),
  inference(resolution,[],[f85,f119])).
fof(f85,plain,(
  p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(cnf_transformation,[],[f61])).
fof(f215,plain,(
  s(bool,'F_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0)))),
  inference(subsumption_resolution,[],[f191,f121])).
fof(f191,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0)))),
  inference(superposition,[],[f112,f135])).
fof(f112,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.seq.SEQ_BCONV','thm.bool.EQ_CLAUSES|split|2')).
fof(f11635,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,X17) | s(bool,'T_0') = s(bool,'F_0') | p(s(bool,X17)) | ~p(s(bool,X16))) ) | $spl14),
  inference(superposition,[],[f11483,f117])).
fof(f43934,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X43),s('type.realax.real',X42)))) ) | $spl202),
  inference(cnf_transformation,[],[f43934_D])).
fof(f43934_D,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X43),s('type.realax.real',X42)))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f47314,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X70),s('type.realax.real',X69)))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f47204,f27632])).
fof(f27632,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X11))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X11),s('type.realax.real',X10)))) )),
  inference(superposition,[],[f132,f127])).
fof(f47204,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X69),s('type.realax.real',X70))) | s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X70),s('type.realax.real',X69)))) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f37617,f46267])).
fof(f37617,plain,(
  ( ! [X101,X99,X100] : (s(bool,X101) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X99),s('type.realax.real',X100))) | s(bool,'T_0') != s(bool,X101) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X100),s('type.realax.real',X99)))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f37604])).
fof(f37604,plain,(
  ( ! [X101,X99,X100] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X101) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X99),s('type.realax.real',X100))) | s(bool,'T_0') != s(bool,X101) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X100),s('type.realax.real',X99)))) ) | $spl14),
  inference(superposition,[],[f37473,f27632])).
fof(f47300,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'F_0') = s(bool,X36)) ) | $spl216),
  inference(cnf_transformation,[],[f47300_D])).
fof(f47300_D,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'F_0') = s(bool,X36)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f12705,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,X46)) ) | $spl20),
  inference(cnf_transformation,[],[f12705_D])).
fof(f12705_D,plain,(
  ( ! [X46] : (s(bool,'F_0') != s(bool,X46)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f47299,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'F_0') = s(bool,X36)) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f47180,f37614])).
fof(f37614,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') != s(bool,X0)) ) | $spl14),
  inference(equality_resolution,[],[f37473])).
fof(f47180,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X36)) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f35642,f46267])).
fof(f35642,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X2) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f34977,f127])).
fof(f34977,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X9) | s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | $spl14),
  inference(resolution,[],[f11962,f116])).
fof(f11962,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | $spl14),
  inference(subsumption_resolution,[],[f11925,f216])).
fof(f11925,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X14))) ) | $spl14),
  inference(superposition,[],[f11702,f119])).
fof(f11702,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X18)) ) | $spl14),
  inference(subsumption_resolution,[],[f11665,f120])).
fof(f11665,plain,(
  ( ! [X19,X18] : (s(bool,X18) != s(bool,X19) | s(bool,'F_0') = s(bool,X19) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X18)) ) | $spl14),
  inference(superposition,[],[f11483,f127])).
fof(f46318,plain,(
  s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | $spl215),
  inference(cnf_transformation,[],[f46318_D])).
fof(f46318_D,plain,(
  s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) <=> ~$spl215),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl215])])).
fof(f46320,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X72),s('type.num.num',X71)))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f46225,f27549])).
fof(f46225,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X71),s('type.num.num',X72))) | s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X72),s('type.num.num',X71)))) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f37620,f46135])).
fof(f46316,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X70),s('type.realax.real',X69)))) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f46223,f27632])).
fof(f46223,plain,(
  ( ! [X70,X69] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X69),s('type.realax.real',X70))) | s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X70),s('type.realax.real',X69)))) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f37617,f46135])).
fof(f46302,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'F_0') = s(bool,X36)) ) | $spl212),
  inference(cnf_transformation,[],[f46302_D])).
fof(f46302_D,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'F_0') = s(bool,X36)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f46301,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'F_0') = s(bool,X36)) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f46199,f37614])).
fof(f46199,plain,(
  ( ! [X35,X36] : (s(bool,'F_0') != s(bool,X35) | s(bool,X36) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),sK0))))) | s(bool,'T_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X36)) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f35642,f46135])).
fof(f43943,plain,(
  ( ! [X189,X188] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X188))),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X188),s('type.num.num',X189)))) ) | $spl208),
  inference(cnf_transformation,[],[f43943_D])).
fof(f43943_D,plain,(
  ( ! [X189,X188] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X188))),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X188),s('type.num.num',X189)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f44038,plain,(
  ( ! [X191,X189,X190,X188] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X188))),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X191),s('type.num.num',X190))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X188),s('type.num.num',X189)))) ) | $spl14),
  inference(subsumption_resolution,[],[f43994,f27549])).
fof(f43994,plain,(
  ( ! [X191,X189,X190,X188] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X190),s('type.num.num',X191))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X188))),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X191),s('type.num.num',X190))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X188),s('type.num.num',X189)))) ) | $spl14),
  inference(superposition,[],[f37620,f28279])).
fof(f28279,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X10))),s('type.num.num',X11))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num',X11)))) )),
  inference(superposition,[],[f124,f127])).
fof(f43937,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X54))))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X54)))) ) | $spl206),
  inference(cnf_transformation,[],[f43937_D])).
fof(f43937_D,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X54))))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X54)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f44031,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X54))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X56),s('type.num.num',X55))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X54)))) ) | $spl14),
  inference(subsumption_resolution,[],[f43961,f27549])).
fof(f43961,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X55),s('type.num.num',X56))) | s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X54))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X56),s('type.num.num',X55))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X54)))) ) | $spl14),
  inference(superposition,[],[f37620,f28097])).
fof(f28097,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X5))))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X5)))) )),
  inference(superposition,[],[f89,f127])).
fof(f43935,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X41))))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X41)))) ) | $spl204),
  inference(cnf_transformation,[],[f43935_D])).
fof(f43935_D,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X41))))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X41)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f44029,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X41))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X43),s('type.num.num',X42))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X41)))) ) | $spl14),
  inference(subsumption_resolution,[],[f43958,f27549])).
fof(f43958,plain,(
  ( ! [X43,X41,X42] : (s(bool,'F_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X42),s('type.num.num',X43))) | s(bool,'T_0') != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X41))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X43),s('type.num.num',X42))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X41)))) ) | $spl14),
  inference(superposition,[],[f37620,f29295])).
fof(f29295,plain,(
  ( ! [X5] : (s(bool,'F_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X5))))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X5)))) )),
  inference(superposition,[],[f123,f127])).
fof(f43942,plain,(
  ( ! [X191,X189,X190,X188] : (s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X188))),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X191),s('type.realax.real',X190))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X188),s('type.num.num',X189)))) ) | $spl14),
  inference(subsumption_resolution,[],[f43896,f27632])).
fof(f43896,plain,(
  ( ! [X191,X189,X190,X188] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X190),s('type.realax.real',X191))) | s(bool,'T_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X188))),s('type.num.num',X189))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X191),s('type.realax.real',X190))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X188),s('type.num.num',X189)))) ) | $spl14),
  inference(superposition,[],[f37617,f28279])).
fof(f43936,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X54))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X56),s('type.realax.real',X55))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X54)))) ) | $spl14),
  inference(subsumption_resolution,[],[f43863,f27632])).
fof(f43863,plain,(
  ( ! [X54,X56,X55] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X55),s('type.realax.real',X56))) | s(bool,'T_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X54))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X56),s('type.realax.real',X55))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X54)))) ) | $spl14),
  inference(superposition,[],[f37617,f28097])).
fof(f43933,plain,(
  ( ! [X43,X41,X42] : (s(bool,'T_0') != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X41))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X43),s('type.realax.real',X42))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X41)))) ) | $spl14),
  inference(subsumption_resolution,[],[f43860,f27632])).
fof(f43860,plain,(
  ( ! [X43,X41,X42] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X42),s('type.realax.real',X43))) | s(bool,'T_0') != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X41))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X43),s('type.realax.real',X42))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X41)))) ) | $spl14),
  inference(superposition,[],[f37617,f29295])).
fof(f42755,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))))) ) | $spl200),
  inference(cnf_transformation,[],[f42755_D])).
fof(f42755_D,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f42754,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') != s(bool,X20) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f42737,f37614])).
fof(f42737,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') != s(bool,X20) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))) | s(bool,'T_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35591,f928])).
fof(f928,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))),s('type.num.num',X5)))) ) | $spl0),
  inference(superposition,[],[f681,f128])).
fof(f681,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) ) | $spl0),
  inference(resolution,[],[f579,f119])).
fof(f579,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num',X4)))) ) | $spl0),
  inference(resolution,[],[f150,f115])).
fof(f35591,plain,(
  ( ! [X70,X68,X69] : (s(bool,X70) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X69),s('type.num.num',X68))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X68),s('type.num.num',X69))) | s(bool,'T_0') = s(bool,X70)) ) | $spl14),
  inference(superposition,[],[f34939,f130])).
fof(f42753,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X8,'type.num.num'),'_dst_x0x1_2'(s(fun(X9,'type.num.num'),X10),s(fun(X8,X9),X11))),s(X8,X12))),s('type.num.num',X7))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X9,'type.num.num'),X10),s(X9,happ(s(fun(X8,X9),X11),s(X8,X12))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X7)))))) ) | $spl198),
  inference(cnf_transformation,[],[f42753_D])).
fof(f42753_D,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X8,'type.num.num'),'_dst_x0x1_2'(s(fun(X9,'type.num.num'),X10),s(fun(X8,X9),X11))),s(X8,X12))),s('type.num.num',X7))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X9,'type.num.num'),X10),s(X9,happ(s(fun(X8,X9),X11),s(X8,X12))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X7)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f42752,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') != s(bool,X13) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X8,'type.num.num'),'_dst_x0x1_2'(s(fun(X9,'type.num.num'),X10),s(fun(X8,X9),X11))),s(X8,X12))),s('type.num.num',X7))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X9,'type.num.num'),X10),s(X9,happ(s(fun(X8,X9),X11),s(X8,X12))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X7)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f42736,f37614])).
fof(f42736,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') != s(bool,X13) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X8,'type.num.num'),'_dst_x0x1_2'(s(fun(X9,'type.num.num'),X10),s(fun(X8,X9),X11))),s(X8,X12))),s('type.num.num',X7))) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X9,'type.num.num'),X10),s(X9,happ(s(fun(X8,X9),X11),s(X8,X12))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X7)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35591,f927])).
fof(f927,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X7,'type.num.num'),X8),s(X7,happ(s(fun(X6,X7),X9),s(X6,X10))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X11))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X11),s('type.num.num',happ(s(fun(X6,'type.num.num'),'_dst_x0x1_2'(s(fun(X7,'type.num.num'),X8),s(fun(X6,X7),X9))),s(X6,X10)))))) ) | $spl0),
  inference(superposition,[],[f681,f128])).
fof(f42751,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))),s('type.num.num',X0))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X3,'type.num.num'),'_dst_x0x1_2'(s(fun(X1,'type.num.num'),X2),s(fun(X3,X1),X4))),s(X3,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) ) | $spl196),
  inference(cnf_transformation,[],[f42751_D])).
fof(f42751_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))),s('type.num.num',X0))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X3,'type.num.num'),'_dst_x0x1_2'(s(fun(X1,'type.num.num'),X2),s(fun(X3,X1),X4))),s(X3,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f42750,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') != s(bool,X6) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))),s('type.num.num',X0))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X3,'type.num.num'),'_dst_x0x1_2'(s(fun(X1,'type.num.num'),X2),s(fun(X3,X1),X4))),s(X3,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f42735,f37614])).
fof(f42735,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') != s(bool,X6) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(X3,X1),X4),s(X3,X5))))),s('type.num.num',X0))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X3,'type.num.num'),'_dst_x0x1_2'(s(fun(X1,'type.num.num'),X2),s(fun(X3,X1),X4))),s(X3,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35591,f926])).
fof(f926,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4)))))))) ) | $spl0),
  inference(superposition,[],[f681,f128])).
fof(f42675,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))) ) | $spl194),
  inference(cnf_transformation,[],[f42675_D])).
fof(f42675_D,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f42674,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'F_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f42641,f127])).
fof(f42641,plain,(
  ( ! [X14,X19,X17,X15,X20,X18,X16] : (s(bool,'T_0') != s(bool,X20) | s(bool,'F_0') = s(bool,X20) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X19))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X14,'type.num.num'),X15),s(X14,happ(s(fun(X16,X14),X17),s(X16,X18))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X19),s('type.num.num',happ(s(fun(X16,'type.num.num'),'_dst_x0x1_2'(s(fun(X14,'type.num.num'),X15),s(fun(X16,X14),X17))),s(X16,X18)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35553,f927])).
fof(f35553,plain,(
  ( ! [X107,X105,X106] : (s(bool,X107) != s(bool,'const.real.real_ge_2'(s('type.realax.real',X106),s('type.realax.real',X105))) | s(bool,'F_0') = s(bool,X107) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X105),s('type.realax.real',X106)))) ) | $spl14),
  inference(superposition,[],[f34939,f132])).
fof(f42673,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X12))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X7,'type.num.num'),'_dst_x0x1_2'(s(fun(X8,'type.num.num'),X9),s(fun(X7,X8),X10))),s(X7,X11))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',happ(s(fun(X8,'type.num.num'),X9),s(X8,happ(s(fun(X7,X8),X10),s(X7,X11)))))))) ) | $spl192),
  inference(cnf_transformation,[],[f42673_D])).
fof(f42673_D,plain,(
  ( ! [X12,X10,X8,X7,X11,X9] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X12))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X7,'type.num.num'),'_dst_x0x1_2'(s(fun(X8,'type.num.num'),X9),s(fun(X7,X8),X10))),s(X7,X11))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',happ(s(fun(X8,'type.num.num'),X9),s(X8,happ(s(fun(X7,X8),X10),s(X7,X11)))))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f42672,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'F_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X12))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X7,'type.num.num'),'_dst_x0x1_2'(s(fun(X8,'type.num.num'),X9),s(fun(X7,X8),X10))),s(X7,X11))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',happ(s(fun(X8,'type.num.num'),X9),s(X8,happ(s(fun(X7,X8),X10),s(X7,X11)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f42640,f127])).
fof(f42640,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s(bool,'T_0') != s(bool,X13) | s(bool,'F_0') = s(bool,X13) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X12))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X7,'type.num.num'),'_dst_x0x1_2'(s(fun(X8,'type.num.num'),X9),s(fun(X7,X8),X10))),s(X7,X11))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',happ(s(fun(X8,'type.num.num'),X9),s(X8,happ(s(fun(X7,X8),X10),s(X7,X11)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35553,f926])).
fof(f42671,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X1,'type.num.num'),'_dst_x0x1_2'(s(fun(X2,'type.num.num'),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),X3),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))),s('type.num.num',X0)))) ) | $spl190),
  inference(cnf_transformation,[],[f42671_D])).
fof(f42671_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X1,'type.num.num'),'_dst_x0x1_2'(s(fun(X2,'type.num.num'),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),X3),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))),s('type.num.num',X0)))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f42670,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X1,'type.num.num'),'_dst_x0x1_2'(s(fun(X2,'type.num.num'),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),X3),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))),s('type.num.num',X0)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f42639,f127])).
fof(f42639,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X1,'type.num.num'),'_dst_x0x1_2'(s(fun(X2,'type.num.num'),X3),s(fun(X1,X2),X4))),s(X1,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),X3),s(X2,happ(s(fun(X1,X2),X4),s(X1,X5))))),s('type.num.num',X0)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35553,f928])).
fof(f41363,plain,(
  ( ! [X70,X72] : (s(bool,X72) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X70))))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X70))) | s(bool,'F_0') = s(bool,X72)) ) | $spl188),
  inference(cnf_transformation,[],[f41363_D])).
fof(f41363_D,plain,(
  ( ! [X70,X72] : (s(bool,X72) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X70))))) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X70))) | s(bool,'F_0') = s(bool,X72)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f41362,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') != s(bool,X71) | s(bool,X72) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X70))))) | s(bool,'F_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X70)))) ) | $spl14),
  inference(subsumption_resolution,[],[f41296,f37614])).
fof(f41296,plain,(
  ( ! [X70,X72,X71] : (s(bool,'F_0') != s(bool,X71) | s(bool,X72) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X70))))) | s(bool,'T_0') = s(bool,X71) | s(bool,'F_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),X70)))) ) | $spl14),
  inference(superposition,[],[f35642,f29295])).
fof(f41010,plain,(
  ( ! [X116,X114,X113] : (s(bool,X116) != s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X113))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X114))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X114),s('type.realax.real',X113))) | s(bool,'F_0') = s(bool,X116)) ) | $spl186),
  inference(cnf_transformation,[],[f41010_D])).
fof(f41010_D,plain,(
  ( ! [X116,X114,X113] : (s(bool,X116) != s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X113))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X114))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X114),s('type.realax.real',X113))) | s(bool,'F_0') = s(bool,X116)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f41009,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'F_0') != s(bool,X115) | s(bool,X116) != s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X113))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X114))))) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X114),s('type.realax.real',X113)))) ) | $spl14),
  inference(subsumption_resolution,[],[f40939,f37614])).
fof(f40939,plain,(
  ( ! [X116,X114,X115,X113] : (s(bool,'F_0') != s(bool,X115) | s(bool,X116) != s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X113))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X114))))) | s(bool,'T_0') = s(bool,X115) | s(bool,'F_0') = s(bool,X116) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X114),s('type.realax.real',X113)))) ) | $spl14),
  inference(superposition,[],[f35642,f32150])).
fof(f32150,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X4),s('type.realax.real',X5)))) )),
  inference(subsumption_resolution,[],[f32077,f216])).
fof(f32077,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X4),s('type.realax.real',X5))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4)))))) )),
  inference(superposition,[],[f27981,f27633])).
fof(f27633,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X13))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X13),s('type.realax.real',X12)))) )),
  inference(superposition,[],[f132,f127])).
fof(f27981,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X12),s('type.realax.real',X13)))) )),
  inference(superposition,[],[f27848,f127])).
fof(f40824,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'F_0') = s(bool,X20)) ) | $spl184),
  inference(cnf_transformation,[],[f40824_D])).
fof(f40824_D,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'F_0') = s(bool,X20)) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f40823,plain,(
  ( ! [X19,X20] : (s(bool,'F_0') != s(bool,X19) | s(bool,X20) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'F_0') = s(bool,X20)) ) | ($spl14 | $spl24)),
  inference(subsumption_resolution,[],[f40752,f37614])).
fof(f40752,plain,(
  ( ! [X19,X20] : (s(bool,'F_0') != s(bool,X19) | s(bool,X20) != s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X20)) ) | ($spl14 | $spl24)),
  inference(superposition,[],[f35642,f17543])).
fof(f17543,plain,(
  s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | $spl24),
  inference(trivial_inequality_removal,[],[f17532])).
fof(f17532,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | $spl24),
  inference(superposition,[],[f16974,f368])).
fof(f16974,plain,(
  ( ! [X24,X23] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',X24),s('type.realax.real',X23))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) ) | $spl24),
  inference(forward_demodulation,[],[f16954,f368])).
fof(f16954,plain,(
  ( ! [X24,X23] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',X24),s('type.realax.real',X23))) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X24)))) ) | $spl24),
  inference(superposition,[],[f12710,f132])).
fof(f40321,plain,(
  ( ! [X83] : (s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83)))) ) | $spl182),
  inference(cnf_transformation,[],[f40321_D])).
fof(f40321_D,plain,(
  ( ! [X83] : (s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83)))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f34378,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X92,'type.num.num'),X93),s(X92,happ(s(fun(X94,X92),X95),s(X94,X96))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X97)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X97),s('type.num.num',happ(s(fun(X94,'type.num.num'),'_dst_x0x1_2'(s(fun(X92,'type.num.num'),X93),s(fun(X94,X92),X95))),s(X94,X96)))))) ) | $spl134),
  inference(cnf_transformation,[],[f34378_D])).
fof(f34378_D,plain,(
  ( ! [X94,X92,X97,X95,X93,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X92,'type.num.num'),X93),s(X92,happ(s(fun(X94,X92),X95),s(X94,X96))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X97)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X97),s('type.num.num',happ(s(fun(X94,'type.num.num'),'_dst_x0x1_2'(s(fun(X92,'type.num.num'),X93),s(fun(X94,X92),X95))),s(X94,X96)))))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f40422,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40378,f28174])).
fof(f28174,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X4)),s('type.num.num',sK5(X4)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4)))) )),
  inference(resolution,[],[f27726,f117])).
fof(f27726,plain,(
  ( ! [X4] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4)))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X4)),s('type.num.num',sK5(X4))))) )),
  inference(resolution,[],[f27706,f117])).
fof(f27706,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X0)),s('type.num.num',sK5(X0))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f93,f130])).
fof(f93,plain,(
  ( ! [X0] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X0)),s('type.num.num',sK6(X0)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ! [X7,X8] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X8))))))) | ! [X5,X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6)))))))) & (((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X0)),s('type.num.num',sK4(X0))))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK3(X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK4(X0)))))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X0)),s('type.num.num',sK6(X0))))) & ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK5(X0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',sK6(X0))))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5,sK6])],[f65])).
fof(f65,plain,(
  ! [X0] : ((~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ! [X7,X8] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X8))))))) | ! [X5,X6] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6)))))))) & ((? [X3,X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))))))) & ? [X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) & ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(rectify,[],[f64])).
fof(f64,plain,(
  ! [X0] : ((~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) | ! [X3,X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))))))) | ! [X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))) & ((? [X3,X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))))))) & ? [X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) & ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ((~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) | (! [X3,X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))))))) | ! [X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2))))))))) & ((? [X3,X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) & ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))))))) & ? [X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) & ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) <=> (! [X3,X4] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))))))) | ! [X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ! [X0] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) <=> (! [X3,X4] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X4))))))) | ! [X1,X2] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) => p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X2)))))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X6)))) <=> (! [X8,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X7))))))) | ! [X8,X7] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X8),s('type.num.num',X7)))) => p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X6),s('type.num.num',X7)))))))))),
  file('thm.seq.SEQ_BCONV','thm.seq.mono')).
fof(f40378,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X97)),s('type.num.num',sK5(X97)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34496,f927])).
fof(f34496,plain,(
  ( ! [X37,X36] : (s(bool,X36) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X37)),s('type.num.num',sK5(X37)))) | p(s(bool,X36)) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X37)))) ) | $spl14),
  inference(resolution,[],[f12944,f29474])).
fof(f29474,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X1)),s('type.num.num',sK5(X1))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X1)))) )),
  inference(trivial_inequality_removal,[],[f29468])).
fof(f29468,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X1))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X1)),s('type.num.num',sK5(X1)))))) )),
  inference(superposition,[],[f28077,f115])).
fof(f28077,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X4)),s('type.num.num',sK5(X4)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4)))) )),
  inference(resolution,[],[f27724,f117])).
fof(f27724,plain,(
  ( ! [X2] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X2)))) | s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X2)),s('type.num.num',sK5(X2))))) )),
  inference(resolution,[],[f27706,f112])).
fof(f34377,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X85,'type.num.num'),'_dst_x0x1_2'(s(fun(X86,'type.num.num'),X87),s(fun(X85,X86),X88))),s(X85,X89))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X90)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X90),s('type.num.num',happ(s(fun(X86,'type.num.num'),X87),s(X86,happ(s(fun(X85,X86),X88),s(X85,X89)))))))) ) | $spl132),
  inference(cnf_transformation,[],[f34377_D])).
fof(f34377_D,plain,(
  ( ! [X90,X88,X87,X85,X89,X86] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X85,'type.num.num'),'_dst_x0x1_2'(s(fun(X86,'type.num.num'),X87),s(fun(X85,X86),X88))),s(X85,X89))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X90)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X90),s('type.num.num',happ(s(fun(X86,'type.num.num'),X87),s(X86,happ(s(fun(X85,X86),X88),s(X85,X89)))))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f40421,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40377,f28174])).
fof(f40377,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X90)),s('type.num.num',sK5(X90)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34496,f926])).
fof(f34375,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X78))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X79,'type.num.num'),'_dst_x0x1_2'(s(fun(X80,'type.num.num'),X81),s(fun(X79,X80),X82))),s(X79,X83)))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X80,'type.num.num'),X81),s(X80,happ(s(fun(X79,X80),X82),s(X79,X83))))),s('type.num.num',X78)))) ) | $spl128),
  inference(cnf_transformation,[],[f34375_D])).
fof(f34375_D,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X78))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X79,'type.num.num'),'_dst_x0x1_2'(s(fun(X80,'type.num.num'),X81),s(fun(X79,X80),X82))),s(X79,X83)))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X80,'type.num.num'),X81),s(X80,happ(s(fun(X79,X80),X82),s(X79,X83))))),s('type.num.num',X78)))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f40420,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40376,f28174])).
fof(f40376,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34496,f928])).
fof(f40323,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40278,f28059])).
fof(f28059,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X4)),s('type.num.num',sK3(X4)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4)))) )),
  inference(resolution,[],[f27711,f117])).
fof(f27711,plain,(
  ( ! [X4] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4)))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X4)),s('type.num.num',sK3(X4))))) )),
  inference(resolution,[],[f27705,f117])).
fof(f27705,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X0)),s('type.num.num',sK3(X0))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f91,f130])).
fof(f91,plain,(
  ( ! [X0] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X0)),s('type.num.num',sK4(X0)))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f40278,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X97)),s('type.num.num',sK3(X97)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34493,f927])).
fof(f34493,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X31)),s('type.num.num',sK3(X31)))) | p(s(bool,X30)) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X31)))) ) | $spl14),
  inference(resolution,[],[f12944,f29188])).
fof(f29188,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X1)),s('type.num.num',sK3(X1))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X1)))) )),
  inference(trivial_inequality_removal,[],[f29182])).
fof(f29182,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X1))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X1)),s('type.num.num',sK3(X1)))))) )),
  inference(superposition,[],[f27956,f115])).
fof(f27956,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X4)),s('type.num.num',sK3(X4)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4)))) )),
  inference(resolution,[],[f27709,f117])).
fof(f27709,plain,(
  ( ! [X2] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X2)))) | s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X2)),s('type.num.num',sK3(X2))))) )),
  inference(resolution,[],[f27705,f112])).
fof(f40322,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40277,f28059])).
fof(f40277,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X90)),s('type.num.num',sK3(X90)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34493,f926])).
fof(f40320,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40276,f28059])).
fof(f40276,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34493,f928])).
fof(f40233,plain,(
  ( ! [X83] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83)))))) ) | $spl180),
  inference(cnf_transformation,[],[f40233_D])).
fof(f40233_D,plain,(
  ( ! [X83] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83)))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f40235,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X97)),s('type.num.num',sK3(X97))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40184,f29188])).
fof(f40184,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X97)),s('type.num.num',sK3(X97))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f40153,f927])).
fof(f40153,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4))) | p(s(bool,X5)) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X4)),s('type.num.num',sK3(X4)))))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f40147])).
fof(f40147,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X5) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4))) | p(s(bool,X5)) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X4)),s('type.num.num',sK3(X4)))))) ) | $spl14),
  inference(superposition,[],[f34491,f115])).
fof(f34491,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X27)),s('type.num.num',sK3(X27)))) | s(bool,X26) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X27))) | p(s(bool,X26))) ) | $spl14),
  inference(resolution,[],[f12944,f27709])).
fof(f40234,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X90)),s('type.num.num',sK3(X90))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40183,f29188])).
fof(f40183,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X90)),s('type.num.num',sK3(X90))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f40153,f926])).
fof(f40232,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40182,f29188])).
fof(f40182,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f40153,f928])).
fof(f40125,plain,(
  ( ! [X83] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83))))) ) | $spl178),
  inference(cnf_transformation,[],[f40125_D])).
fof(f40125_D,plain,(
  ( ! [X83] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f40127,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X97)),s('type.num.num',sK3(X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40076,f28059])).
fof(f40076,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X97)),s('type.num.num',sK3(X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34490,f927])).
fof(f34490,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X25))) | p(s(bool,X24)) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X25)),s('type.num.num',sK3(X25))))) ) | $spl14),
  inference(resolution,[],[f12944,f27711])).
fof(f40126,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X90)),s('type.num.num',sK3(X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40075,f28059])).
fof(f40075,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X90)),s('type.num.num',sK3(X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34490,f926])).
fof(f40124,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f40074,f28059])).
fof(f40074,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34490,f928])).
fof(f40007,plain,(
  ( ! [X83] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83)))))) ) | $spl176),
  inference(cnf_transformation,[],[f40007_D])).
fof(f40007_D,plain,(
  ( ! [X83] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83)))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f40009,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X97)),s('type.num.num',sK5(X97))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f39958,f29474])).
fof(f39958,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X97)),s('type.num.num',sK5(X97))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f39927,f927])).
fof(f39927,plain,(
  ( ! [X4,X5] : (s(bool,X5) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4))) | p(s(bool,X5)) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X4)),s('type.num.num',sK5(X4)))))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f39921])).
fof(f39921,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X5) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X4))) | p(s(bool,X5)) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X4)),s('type.num.num',sK5(X4)))))) ) | $spl14),
  inference(superposition,[],[f34489,f115])).
fof(f34489,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X23)),s('type.num.num',sK5(X23)))) | s(bool,X22) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X23))) | p(s(bool,X22))) ) | $spl14),
  inference(resolution,[],[f12944,f27724])).
fof(f40008,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X90)),s('type.num.num',sK5(X90))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f39957,f29474])).
fof(f39957,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X90)),s('type.num.num',sK5(X90))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f39927,f926])).
fof(f40006,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f39956,f29474])).
fof(f39956,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f39927,f928])).
fof(f39899,plain,(
  ( ! [X83] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83))))) ) | $spl174),
  inference(cnf_transformation,[],[f39899_D])).
fof(f39899_D,plain,(
  ( ! [X83] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f39901,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X97)),s('type.num.num',sK5(X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f39850,f28174])).
fof(f39850,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X97)),s('type.num.num',sK5(X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34488,f927])).
fof(f34488,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X21))) | p(s(bool,X20)) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X21)),s('type.num.num',sK5(X21))))) ) | $spl14),
  inference(resolution,[],[f12944,f27726])).
fof(f39900,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X90)),s('type.num.num',sK5(X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f39849,f28174])).
fof(f39849,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X90)),s('type.num.num',sK5(X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34488,f926])).
fof(f39898,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f39848,f28174])).
fof(f39848,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34488,f928])).
fof(f38969,plain,(
  ( ! [X173,X171,X170] : (s(bool,X173) != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X170))),s('type.num.num',X171))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X170),s('type.num.num',X171))) | s(bool,'F_0') = s(bool,X173)) ) | $spl172),
  inference(cnf_transformation,[],[f38969_D])).
fof(f38969_D,plain,(
  ( ! [X173,X171,X170] : (s(bool,X173) != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X170))),s('type.num.num',X171))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X170),s('type.num.num',X171))) | s(bool,'F_0') = s(bool,X173)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f38968,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') != s(bool,X172) | s(bool,X173) != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X170))),s('type.num.num',X171))) | s(bool,'F_0') = s(bool,X173) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X170),s('type.num.num',X171)))) ) | $spl14),
  inference(subsumption_resolution,[],[f38888,f37614])).
fof(f38888,plain,(
  ( ! [X173,X171,X172,X170] : (s(bool,'F_0') != s(bool,X172) | s(bool,X173) != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X170))),s('type.num.num',X171))) | s(bool,'T_0') = s(bool,X172) | s(bool,'F_0') = s(bool,X173) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X170),s('type.num.num',X171)))) ) | $spl14),
  inference(superposition,[],[f35642,f28279])).
fof(f38967,plain,(
  ( ! [X158,X161,X159] : (s(bool,X161) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X158))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X159))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X158),s('type.realax.real',X159))) | s(bool,'F_0') = s(bool,X161)) ) | $spl170),
  inference(cnf_transformation,[],[f38967_D])).
fof(f38967_D,plain,(
  ( ! [X158,X161,X159] : (s(bool,X161) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X158))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X159))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X158),s('type.realax.real',X159))) | s(bool,'F_0') = s(bool,X161)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f38966,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,'F_0') != s(bool,X160) | s(bool,X161) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X158))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X159))))) | s(bool,'F_0') = s(bool,X161) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X158),s('type.realax.real',X159)))) ) | $spl14),
  inference(subsumption_resolution,[],[f38885,f37614])).
fof(f38885,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,'F_0') != s(bool,X160) | s(bool,X161) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X158))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X159))))) | s(bool,'T_0') = s(bool,X160) | s(bool,'F_0') = s(bool,X161) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X158),s('type.realax.real',X159)))) ) | $spl14),
  inference(superposition,[],[f35642,f27981])).
fof(f38965,plain,(
  ( ! [X146,X149,X147] : (s(bool,X149) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X146),s('type.realax.real',X147))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X147),s('type.realax.real',X146))) | s(bool,'F_0') = s(bool,X149)) ) | $spl168),
  inference(cnf_transformation,[],[f38965_D])).
fof(f38965_D,plain,(
  ( ! [X146,X149,X147] : (s(bool,X149) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X146),s('type.realax.real',X147))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X147),s('type.realax.real',X146))) | s(bool,'F_0') = s(bool,X149)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f38964,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,'F_0') != s(bool,X148) | s(bool,X149) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X146),s('type.realax.real',X147))) | s(bool,'F_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X147),s('type.realax.real',X146)))) ) | $spl14),
  inference(subsumption_resolution,[],[f38882,f37614])).
fof(f38882,plain,(
  ( ! [X146,X149,X147,X148] : (s(bool,'F_0') != s(bool,X148) | s(bool,X149) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X146),s('type.realax.real',X147))) | s(bool,'T_0') = s(bool,X148) | s(bool,'F_0') = s(bool,X149) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X147),s('type.realax.real',X146)))) ) | $spl14),
  inference(superposition,[],[f35642,f27632])).
fof(f12710,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,X66)) ) | $spl24),
  inference(cnf_transformation,[],[f12710_D])).
fof(f12710_D,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,X66)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f38963,plain,(
  ( ! [X134,X135] : (s(bool,'F_0') != s(bool,X134) | s(bool,X135) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,X135)) ) | $spl14),
  inference(subsumption_resolution,[],[f38878,f37614])).
fof(f38878,plain,(
  ( ! [X134,X135] : (s(bool,'F_0') != s(bool,X134) | s(bool,X135) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'T_0') = s(bool,X134) | s(bool,'F_0') = s(bool,X135)) ) | $spl14),
  inference(superposition,[],[f35642,f368])).
fof(f38961,plain,(
  ( ! [X92,X95,X93] : (s(bool,X95) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X92),s('type.num.num',X93))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X93),s('type.num.num',X92))) | s(bool,'F_0') = s(bool,X95)) ) | $spl166),
  inference(cnf_transformation,[],[f38961_D])).
fof(f38961_D,plain,(
  ( ! [X92,X95,X93] : (s(bool,X95) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X92),s('type.num.num',X93))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X93),s('type.num.num',X92))) | s(bool,'F_0') = s(bool,X95)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f38960,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'F_0') != s(bool,X94) | s(bool,X95) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X92),s('type.num.num',X93))) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X93),s('type.num.num',X92)))) ) | $spl14),
  inference(subsumption_resolution,[],[f38869,f37614])).
fof(f38869,plain,(
  ( ! [X94,X92,X95,X93] : (s(bool,'F_0') != s(bool,X94) | s(bool,X95) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X92),s('type.num.num',X93))) | s(bool,'T_0') = s(bool,X94) | s(bool,'F_0') = s(bool,X95) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X93),s('type.num.num',X92)))) ) | $spl14),
  inference(superposition,[],[f35642,f27549])).
fof(f38958,plain,(
  ( ! [X70,X74,X72,X71,X77,X75,X73] : (s(bool,X77) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X70,'type.num.num'),X71),s(X70,happ(s(fun(X72,X70),X73),s(X72,X74))))),s('type.num.num',X75))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X75))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X72,'type.num.num'),'_dst_x0x1_2'(s(fun(X70,'type.num.num'),X71),s(fun(X72,X70),X73))),s(X72,X74))))))) | s(bool,'F_0') = s(bool,X77)) ) | $spl164),
  inference(cnf_transformation,[],[f38958_D])).
fof(f38958_D,plain,(
  ( ! [X70,X74,X72,X71,X77,X75,X73] : (s(bool,X77) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X70,'type.num.num'),X71),s(X70,happ(s(fun(X72,X70),X73),s(X72,X74))))),s('type.num.num',X75))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X75))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X72,'type.num.num'),'_dst_x0x1_2'(s(fun(X70,'type.num.num'),X71),s(fun(X72,X70),X73))),s(X72,X74))))))) | s(bool,'F_0') = s(bool,X77)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f38957,plain,(
  ( ! [X70,X76,X74,X72,X71,X77,X75,X73] : (s(bool,'F_0') != s(bool,X76) | s(bool,X77) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X70,'type.num.num'),X71),s(X70,happ(s(fun(X72,X70),X73),s(X72,X74))))),s('type.num.num',X75))) | s(bool,'F_0') = s(bool,X77) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X75))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X72,'type.num.num'),'_dst_x0x1_2'(s(fun(X70,'type.num.num'),X71),s(fun(X72,X70),X73))),s(X72,X74)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38863,f37614])).
fof(f38863,plain,(
  ( ! [X70,X76,X74,X72,X71,X77,X75,X73] : (s(bool,'F_0') != s(bool,X76) | s(bool,X77) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X70,'type.num.num'),X71),s(X70,happ(s(fun(X72,X70),X73),s(X72,X74))))),s('type.num.num',X75))) | s(bool,'T_0') = s(bool,X76) | s(bool,'F_0') = s(bool,X77) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X75))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X72,'type.num.num'),'_dst_x0x1_2'(s(fun(X70,'type.num.num'),X71),s(fun(X72,X70),X73))),s(X72,X74)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35642,f928])).
fof(f38956,plain,(
  ( ! [X66,X64,X62,X69,X67,X65,X63] : (s(bool,X69) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X62),s('type.num.num',happ(s(fun(X63,'type.num.num'),'_dst_x0x1_2'(s(fun(X64,'type.num.num'),X65),s(fun(X63,X64),X66))),s(X63,X67))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X64,'type.num.num'),X65),s(X64,happ(s(fun(X63,X64),X66),s(X63,X67))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X62))))) | s(bool,'F_0') = s(bool,X69)) ) | $spl162),
  inference(cnf_transformation,[],[f38956_D])).
fof(f38956_D,plain,(
  ( ! [X66,X64,X62,X69,X67,X65,X63] : (s(bool,X69) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X62),s('type.num.num',happ(s(fun(X63,'type.num.num'),'_dst_x0x1_2'(s(fun(X64,'type.num.num'),X65),s(fun(X63,X64),X66))),s(X63,X67))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X64,'type.num.num'),X65),s(X64,happ(s(fun(X63,X64),X66),s(X63,X67))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X62))))) | s(bool,'F_0') = s(bool,X69)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f38955,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (s(bool,'F_0') != s(bool,X68) | s(bool,X69) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X62),s('type.num.num',happ(s(fun(X63,'type.num.num'),'_dst_x0x1_2'(s(fun(X64,'type.num.num'),X65),s(fun(X63,X64),X66))),s(X63,X67))))) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X64,'type.num.num'),X65),s(X64,happ(s(fun(X63,X64),X66),s(X63,X67))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X62)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38862,f37614])).
fof(f38862,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (s(bool,'F_0') != s(bool,X68) | s(bool,X69) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X62),s('type.num.num',happ(s(fun(X63,'type.num.num'),'_dst_x0x1_2'(s(fun(X64,'type.num.num'),X65),s(fun(X63,X64),X66))),s(X63,X67))))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X64,'type.num.num'),X65),s(X64,happ(s(fun(X63,X64),X66),s(X63,X67))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X62)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35642,f927])).
fof(f38954,plain,(
  ( ! [X61,X59,X57,X54,X58,X56,X55] : (s(bool,X61) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X54),s('type.num.num',happ(s(fun(X55,'type.num.num'),X56),s(X55,happ(s(fun(X57,X55),X58),s(X57,X59))))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X57,'type.num.num'),'_dst_x0x1_2'(s(fun(X55,'type.num.num'),X56),s(fun(X57,X55),X58))),s(X57,X59))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X54))))) | s(bool,'F_0') = s(bool,X61)) ) | $spl160),
  inference(cnf_transformation,[],[f38954_D])).
fof(f38954_D,plain,(
  ( ! [X61,X59,X57,X54,X58,X56,X55] : (s(bool,X61) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X54),s('type.num.num',happ(s(fun(X55,'type.num.num'),X56),s(X55,happ(s(fun(X57,X55),X58),s(X57,X59))))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X57,'type.num.num'),'_dst_x0x1_2'(s(fun(X55,'type.num.num'),X56),s(fun(X57,X55),X58))),s(X57,X59))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X54))))) | s(bool,'F_0') = s(bool,X61)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f38953,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55] : (s(bool,'F_0') != s(bool,X60) | s(bool,X61) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X54),s('type.num.num',happ(s(fun(X55,'type.num.num'),X56),s(X55,happ(s(fun(X57,X55),X58),s(X57,X59))))))) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X57,'type.num.num'),'_dst_x0x1_2'(s(fun(X55,'type.num.num'),X56),s(fun(X57,X55),X58))),s(X57,X59))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X54)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38861,f37614])).
fof(f38861,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55] : (s(bool,'F_0') != s(bool,X60) | s(bool,X61) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X54),s('type.num.num',happ(s(fun(X55,'type.num.num'),X56),s(X55,happ(s(fun(X57,X55),X58),s(X57,X59))))))) | s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,X61) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X57,'type.num.num'),'_dst_x0x1_2'(s(fun(X55,'type.num.num'),X56),s(fun(X57,X55),X58))),s(X57,X59))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X54)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f35642,f926])).
fof(f38952,plain,(
  ( ! [X50,X48] : (s(bool,X50) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X48))))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X48))) | s(bool,'F_0') = s(bool,X50)) ) | $spl158),
  inference(cnf_transformation,[],[f38952_D])).
fof(f38952_D,plain,(
  ( ! [X50,X48] : (s(bool,X50) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X48))))) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X48))) | s(bool,'F_0') = s(bool,X50)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f38951,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,X50) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X48))))) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X48)))) ) | $spl14),
  inference(subsumption_resolution,[],[f38859,f37614])).
fof(f38859,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') != s(bool,X49) | s(bool,X50) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X48))))) | s(bool,'T_0') = s(bool,X49) | s(bool,'F_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X48)))) ) | $spl14),
  inference(superposition,[],[f35642,f28097])).
fof(f38736,plain,(
  ( ! [X83] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))))) ) | $spl156),
  inference(cnf_transformation,[],[f38736_D])).
fof(f38736_D,plain,(
  ( ! [X83] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f38824,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38781,f27726])).
fof(f38781,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X97)),s('type.num.num',sK5(X97)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34498,f927])).
fof(f34498,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X41)),s('type.num.num',sK5(X41)))) | p(s(bool,X40)) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X41))))) ) | $spl14),
  inference(resolution,[],[f12944,f27706])).
fof(f38823,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38780,f27726])).
fof(f38780,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X90)),s('type.num.num',sK5(X90)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34498,f926])).
fof(f38822,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38779,f27726])).
fof(f38779,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X83)),s('type.num.num',sK5(X83)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34498,f928])).
fof(f38738,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38694,f27711])).
fof(f38694,plain,(
  ( ! [X94,X92,X97,X95,X93,X91,X96] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X97)),s('type.num.num',sK3(X97)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X91,'type.num.num'),X92),s(X91,happ(s(fun(X93,X91),X94),s(X93,X95))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X96)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X97)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X96),s('type.num.num',happ(s(fun(X93,'type.num.num'),'_dst_x0x1_2'(s(fun(X91,'type.num.num'),X92),s(fun(X93,X91),X94))),s(X93,X95)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34495,f927])).
fof(f34495,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X35)),s('type.num.num',sK3(X35)))) | p(s(bool,X34)) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X35))))) ) | $spl14),
  inference(resolution,[],[f12944,f27705])).
fof(f38737,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38693,f27711])).
fof(f38693,plain,(
  ( ! [X90,X88,X87,X85,X89,X86,X84] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X90)),s('type.num.num',sK3(X90)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X84,'type.num.num'),'_dst_x0x1_2'(s(fun(X85,'type.num.num'),X86),s(fun(X84,X85),X87))),s(X84,X88))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X89)))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X90)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X89),s('type.num.num',happ(s(fun(X85,'type.num.num'),X86),s(X85,happ(s(fun(X84,X85),X87),s(X84,X88)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34495,f926])).
fof(f38735,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f38692,f27711])).
fof(f38692,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X83)),s('type.num.num',sK3(X83)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X77))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X78,'type.num.num'),'_dst_x0x1_2'(s(fun(X79,'type.num.num'),X80),s(fun(X78,X79),X81))),s(X78,X82)))))))) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X83)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X79,'type.num.num'),X80),s(X79,happ(s(fun(X78,X79),X81),s(X78,X82))))),s('type.num.num',X77)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f34495,f928])).
fof(f12717,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl26),
  inference(cnf_transformation,[],[f12717_D])).
fof(f12717_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f37634,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,'T_0') != s(bool,X16)) ) | $spl154),
  inference(cnf_transformation,[],[f37634_D])).
fof(f37634_D,plain,(
  ( ! [X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,'T_0') != s(bool,X16)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f37627,plain,(
  ( ! [X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') != s(bool,X20) | p(s(bool,X19))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f37577])).
fof(f37577,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X19) | s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') != s(bool,X20) | p(s(bool,X19))) ) | $spl14),
  inference(superposition,[],[f37473,f113])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f79])).
fof(f37628,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X18) | p(s(bool,X17))) ) | $spl14),
  inference(trivial_inequality_removal,[],[f37576])).
fof(f37576,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,X17) | s(bool,'F_0') != s(bool,X18) | s(bool,'T_0') != s(bool,X18) | p(s(bool,X17))) ) | $spl14),
  inference(superposition,[],[f37473,f115])).
fof(f37633,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) | $spl152),
  inference(cnf_transformation,[],[f37633_D])).
fof(f37633_D,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,X15)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f37629,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,'T_0') != s(bool,X16) | s(bool,'T_0') = s(bool,X15)) ) | $spl14),
  inference(trivial_inequality_removal,[],[f37575])).
fof(f37575,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X15) | s(bool,'F_0') != s(bool,X16) | s(bool,'T_0') != s(bool,X16) | s(bool,'T_0') = s(bool,X15)) ) | $spl14),
  inference(superposition,[],[f37473,f127])).
fof(f37535,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X38) | s(bool,'T_0') != s(bool,X39) | s(bool,X38) != s(bool,X39)) ) | $spl150),
  inference(cnf_transformation,[],[f37535_D])).
fof(f37535_D,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X38) | s(bool,'T_0') != s(bool,X39) | s(bool,X38) != s(bool,X39)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f37533,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,'F_0') != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'T_0') != s(bool,X35)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f37496])).
fof(f37496,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,'F_0') != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'T_0') != s(bool,X35) | p(s(bool,X33))) ) | $spl14),
  inference(superposition,[],[f35813,f113])).
fof(f37534,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'T_0') != s(bool,X32)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f37495])).
fof(f37495,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'T_0') != s(bool,X32) | p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f35813,f115])).
fof(f37539,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') != s(bool,X29)) ) | $spl14),
  inference(subsumption_resolution,[],[f37494,f116])).
fof(f37494,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') != s(bool,X29) | s(bool,'T_0') = s(bool,X27)) ) | $spl14),
  inference(superposition,[],[f35813,f127])).
fof(f37475,plain,(
  ( ! [X43,X41,X42] : (s(bool,'F_0') != s(bool,X41) | s(bool,X41) != s(bool,X42) | s(bool,'T_0') != s(bool,X42) | p(s(bool,X43))) ) | $spl14),
  inference(resolution,[],[f35813,f99])).
fof(f99,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) & p(s(bool,X10))) <=> p(s(bool,'F_0')))),
  file('thm.seq.SEQ_BCONV','thm.bool.AND_CLAUSES|split|2')).
fof(f37474,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,X38) | s(bool,X38) != s(bool,X39) | s(bool,'T_0') != s(bool,X39) | p(s(bool,X40))) ) | $spl14),
  inference(resolution,[],[f35813,f101])).
fof(f12718,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5)) ) | $spl28),
  inference(cnf_transformation,[],[f12718_D])).
fof(f12718_D,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f37030,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,'F_0') != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f36993])).
fof(f36993,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | s(bool,'F_0') != s(bool,X34) | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | p(s(bool,X33))) ) | $spl14),
  inference(superposition,[],[f35087,f113])).
fof(f35087,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f13097,f127])).
fof(f37031,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f36992])).
fof(f36992,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | s(bool,'F_0') != s(bool,X31) | s(bool,X31) != s(bool,X32) | s(bool,'F_0') = s(bool,X32) | p(s(bool,X30))) ) | $spl14),
  inference(superposition,[],[f35087,f115])).
fof(f37035,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,X29)) ) | $spl14),
  inference(subsumption_resolution,[],[f36991,f116])).
fof(f36991,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X27)) ) | $spl14),
  inference(superposition,[],[f35087,f127])).
fof(f36972,plain,(
  ( ! [X37,X38,X36] : (s(bool,'F_0') != s(bool,X36) | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X37) | p(s(bool,X38))) ) | $spl14),
  inference(resolution,[],[f35087,f99])).
fof(f36971,plain,(
  ( ! [X35,X33,X34] : (s(bool,'F_0') != s(bool,X33) | s(bool,X33) != s(bool,X34) | s(bool,'F_0') = s(bool,X34) | p(s(bool,X35))) ) | $spl14),
  inference(resolution,[],[f35087,f101])).
fof(f36960,plain,(
  ( ! [X118,X117] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X117))),s('type.num.num',X118))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X117),s('type.num.num',X118)))) ) | $spl148),
  inference(cnf_transformation,[],[f36960_D])).
fof(f36960_D,plain,(
  ( ! [X118,X117] : (s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X117))),s('type.num.num',X118))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X117),s('type.num.num',X118)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f36946,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) ) | $spl138),
  inference(cnf_transformation,[],[f36946_D])).
fof(f36946_D,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0')))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f36918,plain,(
  ( ! [X118,X119,X117] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X119),s(bool,'T_0'))) | s(bool,'F_0') != s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X117))),s('type.num.num',X118))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X117),s('type.num.num',X118)))) ) | $spl82),
  inference(superposition,[],[f35390,f28280])).
fof(f28280,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X12))),s('type.num.num',X13))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',X13)))) )),
  inference(superposition,[],[f124,f127])).
fof(f35390,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X4),s(bool,X5))) | s(bool,'F_0') != s(bool,X5)) ) | $spl82),
  inference(resolution,[],[f34794,f112])).
fof(f34794,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) ) | $spl82),
  inference(superposition,[],[f34066,f115])).
fof(f34066,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) ) | $spl82),
  inference(subsumption_resolution,[],[f34026,f120])).
fof(f34026,plain,(
  ( ! [X13] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) ) | $spl82),
  inference(superposition,[],[f116,f25125])).
fof(f36957,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X108))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X109))))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X108),s('type.realax.real',X109)))) ) | $spl146),
  inference(cnf_transformation,[],[f36957_D])).
fof(f36957_D,plain,(
  ( ! [X109,X108] : (s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X108))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X109))))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X108),s('type.realax.real',X109)))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f36915,plain,(
  ( ! [X109,X110,X108] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X110),s(bool,'T_0'))) | s(bool,'F_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X108))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X109))))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X108),s('type.realax.real',X109)))) ) | $spl82),
  inference(superposition,[],[f35390,f27982])).
fof(f27982,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15))))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X14),s('type.realax.real',X15)))) )),
  inference(superposition,[],[f27848,f127])).
fof(f36953,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X87,'type.num.num'),X88),s(X87,happ(s(fun(X89,X87),X90),s(X89,X91))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X92))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X92),s('type.num.num',happ(s(fun(X89,'type.num.num'),'_dst_x0x1_2'(s(fun(X87,'type.num.num'),X88),s(fun(X89,X87),X90))),s(X89,X91)))))) ) | $spl144),
  inference(cnf_transformation,[],[f36953_D])).
fof(f36953_D,plain,(
  ( ! [X92,X90,X88,X87,X91,X89] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X87,'type.num.num'),X88),s(X87,happ(s(fun(X89,X87),X90),s(X89,X91))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X92))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X92),s('type.num.num',happ(s(fun(X89,'type.num.num'),'_dst_x0x1_2'(s(fun(X87,'type.num.num'),X88),s(fun(X89,X87),X90))),s(X89,X91)))))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f36908,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X93),s(bool,'T_0'))) | s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X87,'type.num.num'),X88),s(X87,happ(s(fun(X89,X87),X90),s(X89,X91))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X92))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X92),s('type.num.num',happ(s(fun(X89,'type.num.num'),'_dst_x0x1_2'(s(fun(X87,'type.num.num'),X88),s(fun(X89,X87),X90))),s(X89,X91)))))) ) | ($spl0 | $spl82)),
  inference(superposition,[],[f35390,f927])).
fof(f36952,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X80,'type.num.num'),'_dst_x0x1_2'(s(fun(X81,'type.num.num'),X82),s(fun(X80,X81),X83))),s(X80,X84))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X85))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X85),s('type.num.num',happ(s(fun(X81,'type.num.num'),X82),s(X81,happ(s(fun(X80,X81),X83),s(X80,X84)))))))) ) | $spl142),
  inference(cnf_transformation,[],[f36952_D])).
fof(f36952_D,plain,(
  ( ! [X80,X85,X83,X81,X84,X82] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X80,'type.num.num'),'_dst_x0x1_2'(s(fun(X81,'type.num.num'),X82),s(fun(X80,X81),X83))),s(X80,X84))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X85))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X85),s('type.num.num',happ(s(fun(X81,'type.num.num'),X82),s(X81,happ(s(fun(X80,X81),X83),s(X80,X84)))))))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f36907,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X86),s(bool,'T_0'))) | s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X80,'type.num.num'),'_dst_x0x1_2'(s(fun(X81,'type.num.num'),X82),s(fun(X80,X81),X83))),s(X80,X84))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X85))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X85),s('type.num.num',happ(s(fun(X81,'type.num.num'),X82),s(X81,happ(s(fun(X80,X81),X83),s(X80,X84)))))))) ) | ($spl0 | $spl82)),
  inference(superposition,[],[f35390,f926])).
fof(f36951,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X73))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X74,'type.num.num'),'_dst_x0x1_2'(s(fun(X75,'type.num.num'),X76),s(fun(X74,X75),X77))),s(X74,X78))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X75,'type.num.num'),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))),s('type.num.num',X73)))) ) | $spl140),
  inference(cnf_transformation,[],[f36951_D])).
fof(f36951_D,plain,(
  ( ! [X78,X76,X74,X77,X75,X73] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X73))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X74,'type.num.num'),'_dst_x0x1_2'(s(fun(X75,'type.num.num'),X76),s(fun(X74,X75),X77))),s(X74,X78))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X75,'type.num.num'),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))),s('type.num.num',X73)))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f36906,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X79),s(bool,'T_0'))) | s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X73))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X74,'type.num.num'),'_dst_x0x1_2'(s(fun(X75,'type.num.num'),X76),s(fun(X74,X75),X77))),s(X74,X78))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X75,'type.num.num'),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))),s('type.num.num',X73)))) ) | ($spl0 | $spl82)),
  inference(superposition,[],[f35390,f928])).
fof(f36945,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X29))))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X29)))) ) | $spl136),
  inference(cnf_transformation,[],[f36945_D])).
fof(f36945_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X29))))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X29)))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f36891,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') != s(bool,happ(s(fun(bool,bool),X30),s(bool,'T_0'))) | s(bool,'F_0') != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X29))))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X29)))) ) | $spl82),
  inference(superposition,[],[f35390,f28098])).
fof(f28098,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X6))))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X6)))) )),
  inference(superposition,[],[f89,f127])).
fof(f34376,plain,(
  ( ! [X84] : (~p(s(bool,happ(s(fun(bool,bool),X84),s(bool,'T_0'))))) ) | $spl130),
  inference(cnf_transformation,[],[f34376_D])).
fof(f34376_D,plain,(
  ( ! [X84] : (~p(s(bool,happ(s(fun(bool,bool),X84),s(bool,'T_0'))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f36241,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (~p(s(bool,happ(s(fun(bool,bool),X93),s(bool,'T_0')))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X87,'type.num.num'),X88),s(X87,happ(s(fun(X89,X87),X90),s(X89,X91))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X92)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X92),s('type.num.num',happ(s(fun(X89,'type.num.num'),'_dst_x0x1_2'(s(fun(X87,'type.num.num'),X88),s(fun(X89,X87),X90))),s(X89,X91)))))) ) | ($spl0 | $spl82)),
  inference(superposition,[],[f36164,f927])).
fof(f36164,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X21),s(bool,X20)))) | p(s(bool,X20))) ) | $spl82),
  inference(subsumption_resolution,[],[f36108,f216])).
fof(f36108,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(bool,bool),X21),s(bool,X20)))) | s(bool,'T_0') = s(bool,'F_0') | p(s(bool,X20))) ) | $spl82),
  inference(superposition,[],[f36060,f115])).
fof(f36060,plain,(
  ( ! [X39,X38] : (~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X39)))) | s(bool,'T_0') = s(bool,X39)) ) | $spl82),
  inference(duplicate_literal_removal,[],[f36021])).
fof(f36021,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,X39) | ~p(s(bool,happ(s(fun(bool,bool),X38),s(bool,X39)))) | s(bool,'T_0') = s(bool,X39)) ) | $spl82),
  inference(superposition,[],[f119,f34002])).
fof(f34002,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl82),
  inference(superposition,[],[f25125,f127])).
fof(f36240,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X86),s(bool,'T_0')))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X80,'type.num.num'),'_dst_x0x1_2'(s(fun(X81,'type.num.num'),X82),s(fun(X80,X81),X83))),s(X80,X84))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X85)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X85),s('type.num.num',happ(s(fun(X81,'type.num.num'),X82),s(X81,happ(s(fun(X80,X81),X83),s(X80,X84)))))))) ) | ($spl0 | $spl82)),
  inference(superposition,[],[f36164,f926])).
fof(f36239,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,happ(s(fun(bool,bool),X79),s(bool,'T_0')))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X73))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X74,'type.num.num'),'_dst_x0x1_2'(s(fun(X75,'type.num.num'),X76),s(fun(X74,X75),X77))),s(X74,X78)))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X75,'type.num.num'),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))),s('type.num.num',X73)))) ) | ($spl0 | $spl82)),
  inference(superposition,[],[f36164,f928])).
fof(f34958,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'F_0') = s(bool,X93) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X87,'type.num.num'),X88),s(X87,happ(s(fun(X89,X87),X90),s(X89,X91))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X92)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X92),s('type.num.num',happ(s(fun(X89,'type.num.num'),'_dst_x0x1_2'(s(fun(X87,'type.num.num'),X88),s(fun(X89,X87),X90))),s(X89,X91)))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f34897,f127])).
fof(f34897,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89] : (s(bool,'T_0') != s(bool,X93) | s(bool,'F_0') = s(bool,X93) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X87,'type.num.num'),X88),s(X87,happ(s(fun(X89,X87),X90),s(X89,X91))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X92)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X92),s('type.num.num',happ(s(fun(X89,'type.num.num'),'_dst_x0x1_2'(s(fun(X87,'type.num.num'),X88),s(fun(X89,X87),X90))),s(X89,X91)))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f11483,f927])).
fof(f34957,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'F_0') = s(bool,X86) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X80,'type.num.num'),'_dst_x0x1_2'(s(fun(X81,'type.num.num'),X82),s(fun(X80,X81),X83))),s(X80,X84))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X85)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X85),s('type.num.num',happ(s(fun(X81,'type.num.num'),X82),s(X81,happ(s(fun(X80,X81),X83),s(X80,X84)))))))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f34896,f127])).
fof(f34896,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'T_0') != s(bool,X86) | s(bool,'F_0') = s(bool,X86) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X80,'type.num.num'),'_dst_x0x1_2'(s(fun(X81,'type.num.num'),X82),s(fun(X80,X81),X83))),s(X80,X84))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X85)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X85),s('type.num.num',happ(s(fun(X81,'type.num.num'),X82),s(X81,happ(s(fun(X80,X81),X83),s(X80,X84)))))))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f11483,f926])).
fof(f34956,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'F_0') = s(bool,X79) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X73))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X74,'type.num.num'),'_dst_x0x1_2'(s(fun(X75,'type.num.num'),X76),s(fun(X74,X75),X77))),s(X74,X78)))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X75,'type.num.num'),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))),s('type.num.num',X73)))) ) | ($spl0 | $spl14)),
  inference(subsumption_resolution,[],[f34895,f127])).
fof(f34895,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'T_0') != s(bool,X79) | s(bool,'F_0') = s(bool,X79) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X73))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X74,'type.num.num'),'_dst_x0x1_2'(s(fun(X75,'type.num.num'),X76),s(fun(X74,X75),X77))),s(X74,X78)))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X75,'type.num.num'),X76),s(X75,happ(s(fun(X74,X75),X77),s(X74,X78))))),s('type.num.num',X73)))) ) | ($spl0 | $spl14)),
  inference(superposition,[],[f11483,f928])).
fof(f34326,plain,(
  ( ! [X94,X92,X97,X95,X93,X98,X96] : (~p(s(bool,happ(s(fun(bool,bool),X98),s(bool,'T_0')))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X92,'type.num.num'),X93),s(X92,happ(s(fun(X94,X92),X95),s(X94,X96))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X97)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X97),s('type.num.num',happ(s(fun(X94,'type.num.num'),'_dst_x0x1_2'(s(fun(X92,'type.num.num'),X93),s(fun(X94,X92),X95))),s(X94,X96)))))) ) | ($spl0 | $spl86)),
  inference(superposition,[],[f33983,f927])).
fof(f33983,plain,(
  ( ! [X2,X3] : (~p(s(bool,happ(s(fun(bool,bool),X3),s(bool,X2)))) | p(s(bool,X2))) ) | $spl86),
  inference(superposition,[],[f25127,f115])).
fof(f34325,plain,(
  ( ! [X90,X88,X87,X85,X91,X89,X86] : (~p(s(bool,happ(s(fun(bool,bool),X91),s(bool,'T_0')))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X85,'type.num.num'),'_dst_x0x1_2'(s(fun(X86,'type.num.num'),X87),s(fun(X85,X86),X88))),s(X85,X89))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X90)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X90),s('type.num.num',happ(s(fun(X86,'type.num.num'),X87),s(X86,happ(s(fun(X85,X86),X88),s(X85,X89)))))))) ) | ($spl0 | $spl86)),
  inference(superposition,[],[f33983,f926])).
fof(f34324,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,happ(s(fun(bool,bool),X84),s(bool,'T_0')))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X78))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X79,'type.num.num'),'_dst_x0x1_2'(s(fun(X80,'type.num.num'),X81),s(fun(X79,X80),X82))),s(X79,X83)))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X80,'type.num.num'),X81),s(X80,happ(s(fun(X79,X80),X82),s(X79,X83))))),s('type.num.num',X78)))) ) | ($spl0 | $spl86)),
  inference(superposition,[],[f33983,f928])).
fof(f33924,plain,(
  ( ! [X177,X180,X178] : (s(bool,X180) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X177))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X178))))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X177),s('type.realax.real',X178))) | s(bool,'T_0') = s(bool,X180)) ) | $spl126),
  inference(cnf_transformation,[],[f33924_D])).
fof(f33924_D,plain,(
  ( ! [X177,X180,X178] : (s(bool,X180) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X177))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X178))))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X177),s('type.realax.real',X178))) | s(bool,'T_0') = s(bool,X180)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f33923,plain,(
  ( ! [X177,X180,X178,X179] : (s(bool,'F_0') = s(bool,X179) | s(bool,X180) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X177))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X178))))) | s(bool,'T_0') = s(bool,X180) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X177),s('type.realax.real',X178)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33831,f127])).
fof(f33831,plain,(
  ( ! [X177,X180,X178,X179] : (s(bool,'T_0') != s(bool,X179) | s(bool,'F_0') = s(bool,X179) | s(bool,X180) != s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X177))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X178))))) | s(bool,'T_0') = s(bool,X180) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X177),s('type.realax.real',X178)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f27982])).
fof(f30166,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30022,f127])).
fof(f30022,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X4) | s(bool,'F_0') = s(bool,X5) | s(bool,X4) != s(bool,X5)) ) | ($spl72 | $spl84)),
  inference(resolution,[],[f29994,f112])).
fof(f29994,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | ($spl72 | $spl84)),
  inference(forward_demodulation,[],[f29927,f29715])).
fof(f29715,plain,(
  ( ! [X83,X84,X82] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X82),s('type.num.num',happ(s(fun(bool,'type.num.num'),X83),s(bool,X84)))))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f29714,f27854])).
fof(f27854,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',happ(s(fun(bool,'type.num.num'),X5),s(bool,X3)))))) ) | $spl84),
  inference(superposition,[],[f25126,f119])).
fof(f29714,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,X84)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X82),s('type.num.num',happ(s(fun(bool,'type.num.num'),X83),s(bool,X84)))))) ) | $spl72),
  inference(subsumption_resolution,[],[f29664,f116])).
fof(f29664,plain,(
  ( ! [X83,X84,X82] : (p(s(bool,X84)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X82),s('type.num.num',happ(s(fun(bool,'type.num.num'),X83),s(bool,X84))))) | s(bool,'T_0') = s(bool,X84)) ) | $spl72),
  inference(superposition,[],[f113,f27741])).
fof(f27741,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl72),
  inference(superposition,[],[f20382,f127])).
fof(f29927,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) != s(bool,X5) | s(bool,X5) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num',happ(s(fun(bool,'type.num.num'),X7),s(bool,X5))))) | p(s(bool,X4))) ) | $spl72),
  inference(superposition,[],[f28369,f115])).
fof(f28369,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',happ(s(fun(bool,'type.num.num'),X8),s(bool,X6)))))) ) | $spl72),
  inference(resolution,[],[f27742,f112])).
fof(f27742,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',happ(s(fun(bool,'type.num.num'),X5),s(bool,X3)))))) ) | $spl72),
  inference(superposition,[],[f20382,f115])).
fof(f33922,plain,(
  ( ! [X165,X166,X168] : (s(bool,X168) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X165),s('type.realax.real',X166))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X166),s('type.realax.real',X165))) | s(bool,'T_0') = s(bool,X168)) ) | $spl124),
  inference(cnf_transformation,[],[f33922_D])).
fof(f33922_D,plain,(
  ( ! [X165,X166,X168] : (s(bool,X168) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X165),s('type.realax.real',X166))) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X166),s('type.realax.real',X165))) | s(bool,'T_0') = s(bool,X168)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f33921,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,'F_0') = s(bool,X167) | s(bool,X168) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X165),s('type.realax.real',X166))) | s(bool,'T_0') = s(bool,X168) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X166),s('type.realax.real',X165)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33828,f127])).
fof(f33828,plain,(
  ( ! [X167,X165,X166,X168] : (s(bool,'T_0') != s(bool,X167) | s(bool,'F_0') = s(bool,X167) | s(bool,X168) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X165),s('type.realax.real',X166))) | s(bool,'T_0') = s(bool,X168) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',X166),s('type.realax.real',X165)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f27633])).
fof(f33920,plain,(
  ( ! [X160] : (s(bool,X160) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'T_0') = s(bool,X160)) ) | $spl122),
  inference(cnf_transformation,[],[f33920_D])).
fof(f33920_D,plain,(
  ( ! [X160] : (s(bool,X160) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'T_0') = s(bool,X160)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f33919,plain,(
  ( ! [X159,X160] : (s(bool,'F_0') = s(bool,X159) | s(bool,X160) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'T_0') = s(bool,X160)) ) | ($spl0 | $spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33826,f127])).
fof(f33826,plain,(
  ( ! [X159,X160] : (s(bool,'T_0') != s(bool,X159) | s(bool,'F_0') = s(bool,X159) | s(bool,X160) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | s(bool,'T_0') = s(bool,X160)) ) | ($spl0 | $spl72 | $spl84)),
  inference(superposition,[],[f30166,f597])).
fof(f597,plain,(
  s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))))) | $spl0),
  inference(resolution,[],[f595,f119])).
fof(f595,plain,(
  p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0))))))) | $spl0),
  inference(subsumption_resolution,[],[f587,f121])).
fof(f587,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0))))))) | $spl0),
  inference(superposition,[],[f150,f310])).
fof(f310,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0))))),
  inference(resolution,[],[f251,f119])).
fof(f251,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0)))))),
  inference(forward_demodulation,[],[f250,f124])).
fof(f250,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',sK1(sK0)))),s('type.num.num',sK2(sK0)))))),
  inference(subsumption_resolution,[],[f236,f86])).
fof(f236,plain,(
  p(s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0)))) | p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',sK1(sK0)))),s('type.num.num',sK2(sK0)))))),
  inference(resolution,[],[f84,f87])).
fof(f33918,plain,(
  ( ! [X156,X154,X152,X151,X149,X153,X150] : (s(bool,X156) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X149,'type.num.num'),X150),s(X149,happ(s(fun(X151,X149),X152),s(X151,X153))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X154))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X154),s('type.num.num',happ(s(fun(X151,'type.num.num'),'_dst_x0x1_2'(s(fun(X149,'type.num.num'),X150),s(fun(X151,X149),X152))),s(X151,X153))))) | s(bool,'T_0') = s(bool,X156)) ) | $spl120),
  inference(cnf_transformation,[],[f33918_D])).
fof(f33918_D,plain,(
  ( ! [X156,X154,X152,X151,X149,X153,X150] : (s(bool,X156) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X149,'type.num.num'),X150),s(X149,happ(s(fun(X151,X149),X152),s(X151,X153))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X154))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X154),s('type.num.num',happ(s(fun(X151,'type.num.num'),'_dst_x0x1_2'(s(fun(X149,'type.num.num'),X150),s(fun(X151,X149),X152))),s(X151,X153))))) | s(bool,'T_0') = s(bool,X156)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f33917,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (s(bool,'F_0') = s(bool,X155) | s(bool,X156) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X149,'type.num.num'),X150),s(X149,happ(s(fun(X151,X149),X152),s(X151,X153))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X154))))) | s(bool,'T_0') = s(bool,X156) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X154),s('type.num.num',happ(s(fun(X151,'type.num.num'),'_dst_x0x1_2'(s(fun(X149,'type.num.num'),X150),s(fun(X151,X149),X152))),s(X151,X153)))))) ) | ($spl0 | $spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33824,f127])).
fof(f33824,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (s(bool,'T_0') != s(bool,X155) | s(bool,'F_0') = s(bool,X155) | s(bool,X156) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X149,'type.num.num'),X150),s(X149,happ(s(fun(X151,X149),X152),s(X151,X153))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X154))))) | s(bool,'T_0') = s(bool,X156) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X154),s('type.num.num',happ(s(fun(X151,'type.num.num'),'_dst_x0x1_2'(s(fun(X149,'type.num.num'),X150),s(fun(X151,X149),X152))),s(X151,X153)))))) ) | ($spl0 | $spl72 | $spl84)),
  inference(superposition,[],[f30166,f927])).
fof(f33916,plain,(
  ( ! [X146,X144,X142,X145,X143,X141,X148] : (s(bool,X148) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X141,'type.num.num'),'_dst_x0x1_2'(s(fun(X142,'type.num.num'),X143),s(fun(X141,X142),X144))),s(X141,X145))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X146))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X146),s('type.num.num',happ(s(fun(X142,'type.num.num'),X143),s(X142,happ(s(fun(X141,X142),X144),s(X141,X145))))))) | s(bool,'T_0') = s(bool,X148)) ) | $spl118),
  inference(cnf_transformation,[],[f33916_D])).
fof(f33916_D,plain,(
  ( ! [X146,X144,X142,X145,X143,X141,X148] : (s(bool,X148) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X141,'type.num.num'),'_dst_x0x1_2'(s(fun(X142,'type.num.num'),X143),s(fun(X141,X142),X144))),s(X141,X145))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X146))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X146),s('type.num.num',happ(s(fun(X142,'type.num.num'),X143),s(X142,happ(s(fun(X141,X142),X144),s(X141,X145))))))) | s(bool,'T_0') = s(bool,X148)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f33915,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (s(bool,'F_0') = s(bool,X147) | s(bool,X148) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X141,'type.num.num'),'_dst_x0x1_2'(s(fun(X142,'type.num.num'),X143),s(fun(X141,X142),X144))),s(X141,X145))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X146))))) | s(bool,'T_0') = s(bool,X148) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X146),s('type.num.num',happ(s(fun(X142,'type.num.num'),X143),s(X142,happ(s(fun(X141,X142),X144),s(X141,X145)))))))) ) | ($spl0 | $spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33823,f127])).
fof(f33823,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (s(bool,'T_0') != s(bool,X147) | s(bool,'F_0') = s(bool,X147) | s(bool,X148) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X141,'type.num.num'),'_dst_x0x1_2'(s(fun(X142,'type.num.num'),X143),s(fun(X141,X142),X144))),s(X141,X145))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X146))))) | s(bool,'T_0') = s(bool,X148) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X146),s('type.num.num',happ(s(fun(X142,'type.num.num'),X143),s(X142,happ(s(fun(X141,X142),X144),s(X141,X145)))))))) ) | ($spl0 | $spl72 | $spl84)),
  inference(superposition,[],[f30166,f926])).
fof(f33914,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X137] : (s(bool,X140) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X133))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X134,'type.num.num'),'_dst_x0x1_2'(s(fun(X135,'type.num.num'),X136),s(fun(X134,X135),X137))),s(X134,X138))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X135,'type.num.num'),X136),s(X135,happ(s(fun(X134,X135),X137),s(X134,X138))))),s('type.num.num',X133))) | s(bool,'T_0') = s(bool,X140)) ) | $spl116),
  inference(cnf_transformation,[],[f33914_D])).
fof(f33914_D,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X137] : (s(bool,X140) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X133))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X134,'type.num.num'),'_dst_x0x1_2'(s(fun(X135,'type.num.num'),X136),s(fun(X134,X135),X137))),s(X134,X138))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X135,'type.num.num'),X136),s(X135,happ(s(fun(X134,X135),X137),s(X134,X138))))),s('type.num.num',X133))) | s(bool,'T_0') = s(bool,X140)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f33913,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (s(bool,'F_0') = s(bool,X139) | s(bool,X140) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X133))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X134,'type.num.num'),'_dst_x0x1_2'(s(fun(X135,'type.num.num'),X136),s(fun(X134,X135),X137))),s(X134,X138))))))) | s(bool,'T_0') = s(bool,X140) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X135,'type.num.num'),X136),s(X135,happ(s(fun(X134,X135),X137),s(X134,X138))))),s('type.num.num',X133)))) ) | ($spl0 | $spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33822,f127])).
fof(f33822,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (s(bool,'T_0') != s(bool,X139) | s(bool,'F_0') = s(bool,X139) | s(bool,X140) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X133))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X134,'type.num.num'),'_dst_x0x1_2'(s(fun(X135,'type.num.num'),X136),s(fun(X134,X135),X137))),s(X134,X138))))))) | s(bool,'T_0') = s(bool,X140) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X135,'type.num.num'),X136),s(X135,happ(s(fun(X134,X135),X137),s(X134,X138))))),s('type.num.num',X133)))) ) | ($spl0 | $spl72 | $spl84)),
  inference(superposition,[],[f30166,f928])).
fof(f33912,plain,(
  ( ! [X132,X130] : (s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X130)),s('type.num.num',sK6(X130)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X130))) | s(bool,'T_0') = s(bool,X132)) ) | $spl114),
  inference(cnf_transformation,[],[f33912_D])).
fof(f33912_D,plain,(
  ( ! [X132,X130] : (s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X130)),s('type.num.num',sK6(X130)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X130))) | s(bool,'T_0') = s(bool,X132)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f33911,plain,(
  ( ! [X132,X130,X131] : (s(bool,'F_0') = s(bool,X131) | s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X130)),s('type.num.num',sK6(X130)))) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X130)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33821,f127])).
fof(f33821,plain,(
  ( ! [X132,X130,X131] : (s(bool,'T_0') != s(bool,X131) | s(bool,'F_0') = s(bool,X131) | s(bool,X132) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X130)),s('type.num.num',sK6(X130)))) | s(bool,'T_0') = s(bool,X132) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X130)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f31505])).
fof(f31505,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X29)),s('type.num.num',sK6(X29)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X29)))) ) | ($spl72 | $spl84)),
  inference(trivial_inequality_removal,[],[f31495])).
fof(f31495,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X29)),s('type.num.num',sK6(X29)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X29)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30508,f28174])).
fof(f30508,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X49),s('type.num.num',X48))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X48),s('type.num.num',X49)))) ) | ($spl72 | $spl84)),
  inference(forward_demodulation,[],[f30459,f135])).
fof(f30459,plain,(
  ( ! [X48,X49] : (s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X49),s('type.num.num',X48))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X48),s('type.num.num',X49)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30417,f130])).
fof(f30417,plain,(
  ( ! [X27] : (s(bool,X27) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f30383,f121])).
fof(f30383,plain,(
  ( ! [X27] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X27) | s(bool,X27) != s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30339,f135])).
fof(f30339,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f30283,f216])).
fof(f30283,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,X15) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X15) | ~p(s(bool,X14))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30024,f119])).
fof(f30024,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X8)) ) | ($spl72 | $spl84)),
  inference(resolution,[],[f29994,f117])).
fof(f33910,plain,(
  ( ! [X127,X129] : (s(bool,X129) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X127)),s('type.num.num',sK4(X127)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X127))) | s(bool,'T_0') = s(bool,X129)) ) | $spl112),
  inference(cnf_transformation,[],[f33910_D])).
fof(f33910_D,plain,(
  ( ! [X127,X129] : (s(bool,X129) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X127)),s('type.num.num',sK4(X127)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X127))) | s(bool,'T_0') = s(bool,X129)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f33909,plain,(
  ( ! [X127,X128,X129] : (s(bool,'F_0') = s(bool,X128) | s(bool,X129) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X127)),s('type.num.num',sK4(X127)))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X127)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33820,f127])).
fof(f33820,plain,(
  ( ! [X127,X128,X129] : (s(bool,'T_0') != s(bool,X128) | s(bool,'F_0') = s(bool,X128) | s(bool,X129) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X127)),s('type.num.num',sK4(X127)))) | s(bool,'T_0') = s(bool,X129) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X127)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f31506])).
fof(f31506,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X28)),s('type.num.num',sK4(X28)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X28)))) ) | ($spl72 | $spl84)),
  inference(trivial_inequality_removal,[],[f31494])).
fof(f31494,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X28)),s('type.num.num',sK4(X28)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X28)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30508,f28059])).
fof(f33907,plain,(
  ( ! [X109,X107,X106] : (s(bool,X109) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X106),s('type.num.num',X107))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',X106))) | s(bool,'T_0') = s(bool,X109)) ) | $spl110),
  inference(cnf_transformation,[],[f33907_D])).
fof(f33907_D,plain,(
  ( ! [X109,X107,X106] : (s(bool,X109) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X106),s('type.num.num',X107))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',X106))) | s(bool,'T_0') = s(bool,X109)) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f33906,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'F_0') = s(bool,X108) | s(bool,X109) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X106),s('type.num.num',X107))) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',X106)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33814,f127])).
fof(f33814,plain,(
  ( ! [X109,X107,X108,X106] : (s(bool,'T_0') != s(bool,X108) | s(bool,'F_0') = s(bool,X108) | s(bool,X109) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',X106),s('type.num.num',X107))) | s(bool,'T_0') = s(bool,X109) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',X106)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f27550])).
fof(f27550,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X12),s('type.num.num',X13))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X13),s('type.num.num',X12)))) )),
  inference(superposition,[],[f130,f127])).
fof(f33905,plain,(
  ( ! [X103,X105] : (s(bool,X105) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X103)),s('type.num.num',sK5(X103)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X103))) | s(bool,'T_0') = s(bool,X105)) ) | $spl108),
  inference(cnf_transformation,[],[f33905_D])).
fof(f33905_D,plain,(
  ( ! [X103,X105] : (s(bool,X105) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X103)),s('type.num.num',sK5(X103)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X103))) | s(bool,'T_0') = s(bool,X105)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f33904,plain,(
  ( ! [X103,X105,X104] : (s(bool,'F_0') = s(bool,X104) | s(bool,X105) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X103)),s('type.num.num',sK5(X103)))) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X103)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33813,f127])).
fof(f33813,plain,(
  ( ! [X103,X105,X104] : (s(bool,'T_0') != s(bool,X104) | s(bool,'F_0') = s(bool,X104) | s(bool,X105) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK6(X103)),s('type.num.num',sK5(X103)))) | s(bool,'T_0') = s(bool,X105) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X103)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f28174])).
fof(f33903,plain,(
  ( ! [X102,X100] : (s(bool,X102) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X100)),s('type.num.num',sK3(X100)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X100))) | s(bool,'T_0') = s(bool,X102)) ) | $spl106),
  inference(cnf_transformation,[],[f33903_D])).
fof(f33903_D,plain,(
  ( ! [X102,X100] : (s(bool,X102) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X100)),s('type.num.num',sK3(X100)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X100))) | s(bool,'T_0') = s(bool,X102)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f33902,plain,(
  ( ! [X101,X102,X100] : (s(bool,'F_0') = s(bool,X101) | s(bool,X102) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X100)),s('type.num.num',sK3(X100)))) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X100)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33812,f127])).
fof(f33812,plain,(
  ( ! [X101,X102,X100] : (s(bool,'T_0') != s(bool,X101) | s(bool,'F_0') = s(bool,X101) | s(bool,X102) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X100)),s('type.num.num',sK3(X100)))) | s(bool,'T_0') = s(bool,X102) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X100)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f28059])).
fof(f33900,plain,(
  ( ! [X47,X45] : (s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X45))))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X45))) | s(bool,'T_0') = s(bool,X47)) ) | $spl104),
  inference(cnf_transformation,[],[f33900_D])).
fof(f33900_D,plain,(
  ( ! [X47,X45] : (s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X45))))) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X45))) | s(bool,'T_0') = s(bool,X47)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f33899,plain,(
  ( ! [X47,X45,X46] : (s(bool,'F_0') = s(bool,X46) | s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X45))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X45)))) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33800,f127])).
fof(f33800,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,X46) | s(bool,'F_0') = s(bool,X46) | s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),'_dst_x0x1_2'(s(fun('type.realax.real','type.realax.real'),'const.realax.real_neg_0'),s(fun('type.num.num','type.realax.real'),X45))))) | s(bool,'T_0') = s(bool,X47) | s(bool,'F_0') = s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),X45)))) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f28098])).
fof(f33897,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'T_0') = s(bool,X37)) ) | $spl102),
  inference(cnf_transformation,[],[f33897_D])).
fof(f33897_D,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f33896,plain,(
  ( ! [X37,X36] : (s(bool,'F_0') = s(bool,X36) | s(bool,X37) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl72 | $spl84)),
  inference(subsumption_resolution,[],[f33796,f127])).
fof(f33796,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X36) | s(bool,'F_0') = s(bool,X36) | s(bool,X37) != s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl72 | $spl84)),
  inference(superposition,[],[f30166,f238])).
fof(f150,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))))))) ) | $spl0),
  inference(cnf_transformation,[],[f150_D])).
fof(f150_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f151,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',X2)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))))))) ) | $spl2),
  inference(cnf_transformation,[],[f151_D])).
fof(f151_D,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',X2)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))))))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f29515,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(subsumption_resolution,[],[f29508,f108])).
fof(f108,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'T_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) & p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'T_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : ((p(s(bool,X10)) & p(s(bool,'T_0'))) <=> p(s(bool,X10)))),
  file('thm.seq.SEQ_BCONV','thm.bool.AND_CLAUSES|split|1')).
fof(f29508,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X2),s('type.num.num',X3)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(superposition,[],[f29466,f135])).
fof(f29466,plain,(
  ( ! [X6,X0,X8,X7,X5] : (~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num',X5)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X7))))) )),
  inference(forward_demodulation,[],[f29465,f130])).
fof(f29465,plain,(
  ( ! [X6,X0,X8,X7,X5] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X8))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num',X5)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f29464,f132])).
fof(f29464,plain,(
  ( ! [X6,X0,X8,X7,X5] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num',X5)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X8)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(forward_demodulation,[],[f90,f130])).
fof(f90,plain,(
  ( ! [X6,X0,X8,X7,X5] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X6)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X7))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),X0),s('type.num.num',X8)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num',X8)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X0))))) )),
  inference(cnf_transformation,[],[f66])).
fof(f27524,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f27523,f120])).
fof(f27523,plain,(
  p(s(bool,'F_0')) | $spl34),
  inference(forward_demodulation,[],[f27522,f16917])).
fof(f27522,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X0)),s('type.num.num',sK6(X0)))))) ) | $spl34),
  inference(subsumption_resolution,[],[f27214,f101])).
fof(f27214,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK5(X0)),s('type.num.num',sK6(X0)))))) ) | $spl34),
  inference(backward_demodulation,[],[f16917,f93])).
fof(f27519,plain,(
  $false | $spl34),
  inference(subsumption_resolution,[],[f27518,f120])).
fof(f27518,plain,(
  p(s(bool,'F_0')) | $spl34),
  inference(forward_demodulation,[],[f27517,f16917])).
fof(f27517,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X0)),s('type.num.num',sK4(X0)))))) ) | $spl34),
  inference(subsumption_resolution,[],[f27212,f101])).
fof(f27212,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK3(X0)),s('type.num.num',sK4(X0)))))) ) | $spl34),
  inference(backward_demodulation,[],[f16917,f91])).
fof(f27194,plain,(
  ( ! [X146,X151,X147,X150,X148] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X151))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X148),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X151),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'F_0')))))) ) | $spl100),
  inference(cnf_transformation,[],[f27194_D])).
fof(f27194_D,plain,(
  ( ! [X146,X151,X147,X150,X148] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X151))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X148),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X151),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'F_0')))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f27193,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X151))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X151),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X149) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X148),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,'F_0')))))) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f27153,f127])).
fof(f27153,plain,(
  ( ! [X146,X151,X149,X147,X150,X148] : (s(bool,'T_0') != s(bool,X149) | s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,X147))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X148))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X151))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X151),s('type.num.num',happ(s(fun(bool,'type.num.num'),X150),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X149) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X148),s('type.num.num',happ(s(fun(bool,'type.num.num'),X146),s(bool,'F_0')))))) ) | ($spl0 | $spl11 | $spl12)),
  inference(superposition,[],[f26828,f18969])).
fof(f18969,plain,(
  ( ! [X12,X10,X11] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X11),s(bool,X12))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X10))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num',happ(s(fun(bool,'type.num.num'),X11),s(bool,'F_0')))))) ) | ($spl0 | $spl11 | $spl12)),
  inference(resolution,[],[f18956,f117])).
fof(f18956,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X1),s(bool,X0))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun(bool,'type.num.num'),X1),s(bool,'F_0')))))) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f18939,f10754])).
fof(f10754,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X17),s('type.num.num',happ(s(fun(bool,'type.num.num'),X15),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X15),s(bool,X16))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X17))))))) ) | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f10734,f112])).
fof(f10734,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X17),s('type.num.num',happ(s(fun(bool,'type.num.num'),X15),s(bool,'F_0'))))) | p(s(bool,X16)) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X15),s(bool,X16))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X17))))))) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f10447,f115])).
fof(f10447,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X31),s(bool,X30))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X32))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X32),s('type.num.num',happ(s(fun(bool,'type.num.num'),X31),s(bool,'F_0'))))) | p(s(bool,X30))) ) | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f10411,f99])).
fof(f10411,plain,(
  ( ! [X30,X31,X32] : (s(bool,X30) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X31),s(bool,X30))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X32))))) | p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X32),s('type.num.num',happ(s(fun(bool,'type.num.num'),X31),s(bool,'F_0'))))) | p(s(bool,X30))) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f10333,f115])).
fof(f10333,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X7),s(bool,X8))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X6))))) | p(s(bool,X8)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num',happ(s(fun(bool,'type.num.num'),X7),s(bool,'F_0')))))) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f10163,f112])).
fof(f10163,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,X4))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,'F_0'))))) | p(s(bool,X4))) ) | ($spl0 | $spl11)),
  inference(forward_demodulation,[],[f10162,f520])).
fof(f520,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))) | $spl11),
  inference(resolution,[],[f384,f115])).
fof(f10162,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,X4))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))))))) | p(s(bool,X4))) ) | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f10130,f99])).
fof(f10130,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,X4))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))))))) | p(s(bool,X4))) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f5489,f520])).
fof(f5489,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X6),s(bool,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X7)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X9)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',happ(s(fun(bool,'type.num.num'),X6),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X9))))))) | p(s(bool,X5))) ) | $spl0),
  inference(superposition,[],[f5398,f115])).
fof(f5398,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f5368,f124])).
fof(f5368,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X0))),s('type.num.num',X1)))))))) ) | $spl0),
  inference(superposition,[],[f1660,f124])).
fof(f1660,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X28)))))) | p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X28),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1622,f128])).
fof(f1622,plain,(
  ( ! [X28,X26,X24,X27,X25] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X28)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X28),s('type.num.num',happ(s(fun(X24,'type.num.num'),'_dst_x0x1_2'(s(fun(bool,'type.num.num'),X27),s(fun(X24,bool),X25))),s(X24,X26))))) | p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))) ) | $spl0),
  inference(superposition,[],[f688,f115])).
fof(f688,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X7,'type.num.num'),X8),s(X7,happ(s(fun(X6,X7),X9),s(X6,X10))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X11)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X11),s('type.num.num',happ(s(fun(X6,'type.num.num'),'_dst_x0x1_2'(s(fun(X7,'type.num.num'),X8),s(fun(X6,X7),X9))),s(X6,X10)))))) ) | $spl0),
  inference(superposition,[],[f579,f128])).
fof(f18939,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X1),s(bool,X0))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun(bool,'type.num.num'),X1),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f11482,f127])).
fof(f26828,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | s(bool,X1) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,X1))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f19886,f127])).
fof(f19886,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') != s(bool,X14) | s(bool,X14) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X13),s(bool,X14))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X12))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',happ(s(fun(bool,'type.num.num'),X13),s(bool,'F_0')))))) ) | $spl12),
  inference(resolution,[],[f19210,f116])).
fof(f19210,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun(bool,'type.num.num'),X4),s(bool,'F_0'))))) | s(bool,X3) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X4),s(bool,X3))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5)))))) ) | $spl12),
  inference(superposition,[],[f18933,f119])).
fof(f18933,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X9),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X8))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',happ(s(fun(bool,'type.num.num'),X9),s(bool,'F_0')))))) ) | $spl12),
  inference(resolution,[],[f11482,f117])).
fof(f27192,plain,(
  ( ! [X138,X136,X135,X139] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X139))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X136),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X139),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'F_0')))))) ) | $spl98),
  inference(cnf_transformation,[],[f27192_D])).
fof(f27192_D,plain,(
  ( ! [X138,X136,X135,X139] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X139))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X136),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X139),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'F_0')))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f27191,plain,(
  ( ! [X138,X136,X135,X139,X137] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X139))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X139),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X137) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X136),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'F_0')))))) ) | $spl12),
  inference(subsumption_resolution,[],[f27151,f127])).
fof(f27151,plain,(
  ( ! [X138,X136,X135,X139,X137] : (s(bool,'T_0') != s(bool,X137) | s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X136))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X139))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X139),s('type.num.num',happ(s(fun(bool,'type.num.num'),X138),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X137) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X136),s('type.num.num',happ(s(fun(bool,'type.num.num'),X135),s(bool,'F_0')))))) ) | $spl12),
  inference(superposition,[],[f26828,f18933])).
fof(f27190,plain,(
  ( ! [X127,X134,X130,X128,X126,X133,X131,X129] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X134))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X131),s('type.num.num',happ(s(fun(X128,'type.num.num'),'_dst_x0x1_2'(s(fun(X126,'type.num.num'),X127),s(fun(X128,X126),X129))),s(X128,X130))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X134),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'F_0')))))) ) | $spl96),
  inference(cnf_transformation,[],[f27190_D])).
fof(f27190_D,plain,(
  ( ! [X127,X134,X130,X128,X126,X133,X131,X129] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X134))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X131),s('type.num.num',happ(s(fun(X128,'type.num.num'),'_dst_x0x1_2'(s(fun(X126,'type.num.num'),X127),s(fun(X128,X126),X129))),s(X128,X130))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X134),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'F_0')))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f27189,plain,(
  ( ! [X127,X134,X132,X130,X128,X126,X133,X131,X129] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X134))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X134),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X132) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X131),s('type.num.num',happ(s(fun(X128,'type.num.num'),'_dst_x0x1_2'(s(fun(X126,'type.num.num'),X127),s(fun(X128,X126),X129))),s(X128,X130)))))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f27150,f127])).
fof(f27150,plain,(
  ( ! [X127,X134,X132,X130,X128,X126,X133,X131,X129] : (s(bool,'T_0') != s(bool,X132) | s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X126,'type.num.num'),X127),s(X126,happ(s(fun(X128,X126),X129),s(X128,X130))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X131))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X134))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X134),s('type.num.num',happ(s(fun(bool,'type.num.num'),X133),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X132) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X131),s('type.num.num',happ(s(fun(X128,'type.num.num'),'_dst_x0x1_2'(s(fun(X126,'type.num.num'),X127),s(fun(X128,X126),X129))),s(X128,X130)))))) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f26828,f927])).
fof(f27188,plain,(
  ( ! [X125,X121,X118,X124,X122,X120,X119,X117] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X125))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X122),s('type.num.num',happ(s(fun(X118,'type.num.num'),X119),s(X118,happ(s(fun(X117,X118),X120),s(X117,X121))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X125),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'F_0')))))) ) | $spl94),
  inference(cnf_transformation,[],[f27188_D])).
fof(f27188_D,plain,(
  ( ! [X125,X121,X118,X124,X122,X120,X119,X117] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X125))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X122),s('type.num.num',happ(s(fun(X118,'type.num.num'),X119),s(X118,happ(s(fun(X117,X118),X120),s(X117,X121))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X125),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'F_0')))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f27187,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119,X117] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X125))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X125),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X123) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X122),s('type.num.num',happ(s(fun(X118,'type.num.num'),X119),s(X118,happ(s(fun(X117,X118),X120),s(X117,X121)))))))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f27149,f127])).
fof(f27149,plain,(
  ( ! [X125,X123,X121,X118,X124,X122,X120,X119,X117] : (s(bool,'T_0') != s(bool,X123) | s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X117,'type.num.num'),'_dst_x0x1_2'(s(fun(X118,'type.num.num'),X119),s(fun(X117,X118),X120))),s(X117,X121))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X122))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X125))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X125),s('type.num.num',happ(s(fun(bool,'type.num.num'),X124),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X123) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X122),s('type.num.num',happ(s(fun(X118,'type.num.num'),X119),s(X118,happ(s(fun(X117,X118),X120),s(X117,X121)))))))) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f26828,f926])).
fof(f27186,plain,(
  ( ! [X111,X109,X116,X112,X110,X108,X115,X113] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X116))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X110,'type.num.num'),X111),s(X110,happ(s(fun(X109,X110),X112),s(X109,X113))))),s('type.num.num',X108))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X116),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'F_0')))))) ) | $spl92),
  inference(cnf_transformation,[],[f27186_D])).
fof(f27186_D,plain,(
  ( ! [X111,X109,X116,X112,X110,X108,X115,X113] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X116))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X110,'type.num.num'),X111),s(X110,happ(s(fun(X109,X110),X112),s(X109,X113))))),s('type.num.num',X108))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X116),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'F_0')))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f27185,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X108,X115,X113] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X116))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X116),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X114) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X110,'type.num.num'),X111),s(X110,happ(s(fun(X109,X110),X112),s(X109,X113))))),s('type.num.num',X108)))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f27148,f127])).
fof(f27148,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X108,X115,X113] : (s(bool,'T_0') != s(bool,X114) | s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X108))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X109,'type.num.num'),'_dst_x0x1_2'(s(fun(X110,'type.num.num'),X111),s(fun(X109,X110),X112))),s(X109,X113))))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X116))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X116),s('type.num.num',happ(s(fun(bool,'type.num.num'),X115),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X114) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X110,'type.num.num'),X111),s(X110,happ(s(fun(X109,X110),X112),s(X109,X113))))),s('type.num.num',X108)))) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f26828,f928])).
fof(f27184,plain,(
  ( ! [X103,X107,X106,X104] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X107))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X104),s('type.num.num',X103))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'F_0')))))) ) | $spl90),
  inference(cnf_transformation,[],[f27184_D])).
fof(f27184_D,plain,(
  ( ! [X103,X107,X106,X104] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X107))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X104),s('type.num.num',X103))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'F_0')))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f27183,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X107))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X105) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X104),s('type.num.num',X103)))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f27147,f127])).
fof(f27147,plain,(
  ( ! [X103,X107,X105,X106,X104] : (s(bool,'T_0') != s(bool,X105) | s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X103))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X104))))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X107))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X107),s('type.num.num',happ(s(fun(bool,'type.num.num'),X106),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X105) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X104),s('type.num.num',X103)))) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f26828,f681])).
fof(f20382,plain,(
  ( ! [X399,X401] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X401),s('type.num.num',happ(s(fun(bool,'type.num.num'),X399),s(bool,'F_0')))))) ) | $spl72),
  inference(cnf_transformation,[],[f20382_D])).
fof(f20382_D,plain,(
  ( ! [X399,X401] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X401),s('type.num.num',happ(s(fun(bool,'type.num.num'),X399),s(bool,'F_0')))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f11483,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25))) ) | $spl14),
  inference(cnf_transformation,[],[f11483_D])).
fof(f11483_D,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,X24) | p(s(bool,X25))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f27173,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X47),s('type.num.num',happ(s(fun(bool,'type.num.num'),X46),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X45) | p(s(bool,X44))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f27135,f683])).
fof(f683,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X4))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',X5)))) ) | $spl0),
  inference(resolution,[],[f579,f114])).
fof(f114,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f27135,plain,(
  ( ! [X47,X45,X46,X44] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X46),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X47))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X47),s('type.num.num',happ(s(fun(bool,'type.num.num'),X46),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X45) | p(s(bool,X44))) ) | $spl12),
  inference(superposition,[],[f26828,f113])).
fof(f27172,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X43),s('type.num.num',happ(s(fun(bool,'type.num.num'),X42),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X41) | p(s(bool,X40))) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f27134,f683])).
fof(f27134,plain,(
  ( ! [X43,X41,X42,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X42),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X43))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X43),s('type.num.num',happ(s(fun(bool,'type.num.num'),X42),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X41) | p(s(bool,X40))) ) | $spl12),
  inference(superposition,[],[f26828,f115])).
fof(f27171,plain,(
  ( ! [X37,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X37)) ) | $spl88),
  inference(cnf_transformation,[],[f27171_D])).
fof(f27171_D,plain,(
  ( ! [X37,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X37)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f27170,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X39),s('type.num.num',happ(s(fun(bool,'type.num.num'),X38),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | ($spl0 | $spl12)),
  inference(subsumption_resolution,[],[f27133,f683])).
fof(f27133,plain,(
  ( ! [X39,X37,X38,X36] : (s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X38),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X39))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X39),s('type.num.num',happ(s(fun(bool,'type.num.num'),X38),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X36)) ) | $spl12),
  inference(superposition,[],[f26828,f127])).
fof(f20390,plain,(
  ( ! [X480,X482,X477,X479,X481] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X479),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,X477))))))) | s(bool,'F_0') = s(bool,X477) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X479))))))) ) | $spl80),
  inference(cnf_transformation,[],[f20390_D])).
fof(f20390_D,plain,(
  ( ! [X480,X482,X477,X479,X481] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X479),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,X477))))))) | s(bool,'F_0') = s(bool,X477) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X479))))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f26796,plain,(
  ( ! [X466,X472,X469,X471,X467,X468,X470] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X469),s('type.num.num',happ(s(fun(X470,'type.num.num'),X471),s(X470,happ(s(fun(bool,X470),X472),s(bool,X467))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X470,'type.num.num'),X471),s(X470,happ(s(fun(bool,X470),X472),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X469)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X468),s('type.num.num',happ(s(fun(bool,'type.num.num'),X466),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X467)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f26646,f18956])).
fof(f26646,plain,(
  ( ! [X466,X472,X469,X471,X467,X468,X470] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X469),s('type.num.num',happ(s(fun(X470,'type.num.num'),X471),s(X470,happ(s(fun(bool,X470),X472),s(bool,X467))))))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X466),s(bool,X467))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X468)))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X470,'type.num.num'),X471),s(X470,happ(s(fun(bool,X470),X472),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X469)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X468),s('type.num.num',happ(s(fun(bool,'type.num.num'),X466),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X467)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f3222,f19209])).
fof(f19209,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X1),s(bool,X0))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X2),s('type.num.num',happ(s(fun(bool,'type.num.num'),X1),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f18933,f127])).
fof(f3222,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X4))))))) | ~p(s(bool,X4)) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))))))) ) | $spl0),
  inference(resolution,[],[f3199,f118])).
fof(f3199,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X3)))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X4)))))) | ~p(s(bool,X3))) ) | $spl0),
  inference(superposition,[],[f1008,f128])).
fof(f1008,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X36),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X37,'type.num.num'),X38),s(fun(bool,X37),X39))),s(bool,X35)))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X37,'type.num.num'),X38),s(X37,happ(s(fun(bool,X37),X39),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X36)))))) | ~p(s(bool,X35))) ) | $spl0),
  inference(superposition,[],[f585,f119])).
fof(f585,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X5),s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4)))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5))))))) ) | $spl0),
  inference(superposition,[],[f150,f128])).
fof(f20388,plain,(
  ( ! [X461,X459,X460,X456,X458] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,X456))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X461)))))) | s(bool,'F_0') = s(bool,X456) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X461),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,'T_0')))))))) ) | $spl78),
  inference(cnf_transformation,[],[f20388_D])).
fof(f20388_D,plain,(
  ( ! [X461,X459,X460,X456,X458] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,X456))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X461)))))) | s(bool,'F_0') = s(bool,X456) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X461),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,'T_0')))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f26795,plain,(
  ( ! [X453,X455,X457,X452,X454,X456,X458] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X455,'type.num.num'),X456),s(X455,happ(s(fun(bool,X455),X457),s(bool,X453))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X458)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X458),s('type.num.num',happ(s(fun(X455,'type.num.num'),X456),s(X455,happ(s(fun(bool,X455),X457),s(bool,'T_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X454),s('type.num.num',happ(s(fun(bool,'type.num.num'),X452),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X453)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f26644,f18956])).
fof(f26644,plain,(
  ( ! [X453,X455,X457,X452,X454,X456,X458] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X455,'type.num.num'),X456),s(X455,happ(s(fun(bool,X455),X457),s(bool,X453))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X458)))))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X452),s(bool,X453))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X454)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X458),s('type.num.num',happ(s(fun(X455,'type.num.num'),X456),s(X455,happ(s(fun(bool,X455),X457),s(bool,'T_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X454),s('type.num.num',happ(s(fun(bool,'type.num.num'),X452),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X453)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f3057,f19209])).
fof(f3057,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X10,'type.num.num'),X11),s(X10,happ(s(fun(bool,X10),X12),s(bool,X13))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X14)))))) | ~p(s(bool,X13)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X14),s('type.num.num',happ(s(fun(X10,'type.num.num'),X11),s(X10,happ(s(fun(bool,X10),X12),s(bool,'T_0')))))))) ) | $spl0),
  inference(resolution,[],[f3047,f115])).
fof(f3047,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'T_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X4))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | ~p(s(bool,X4))) ) | $spl0),
  inference(superposition,[],[f1003,f128])).
fof(f1003,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X11),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X12,'type.num.num'),X13),s(fun(bool,X12),X14))),s(bool,'T_0')))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X12,'type.num.num'),X13),s(X12,happ(s(fun(bool,X12),X14),s(bool,X10))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X11)))))) | ~p(s(bool,X10))) ) | $spl0),
  inference(superposition,[],[f585,f119])).
fof(f20385,plain,(
  ( ! [X424,X426,X421,X423,X425] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,X421))))),s('type.num.num',X426))) | s(bool,'F_0') = s(bool,X421) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X426))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,'T_0'))))))))))) ) | $spl76),
  inference(cnf_transformation,[],[f20385_D])).
fof(f20385_D,plain,(
  ( ! [X424,X426,X421,X423,X425] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,X421))))),s('type.num.num',X426))) | s(bool,'F_0') = s(bool,X421) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X426))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,'T_0'))))))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f26793,plain,(
  ( ! [X420,X422,X418,X421,X423,X417,X419] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X420,'type.num.num'),X421),s(X420,happ(s(fun(bool,X420),X422),s(bool,X418))))),s('type.num.num',X423))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X423))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X420,'type.num.num'),X421),s(X420,happ(s(fun(bool,X420),X422),s(bool,'T_0')))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X419),s('type.num.num',happ(s(fun(bool,'type.num.num'),X417),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X418)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f26639,f18956])).
fof(f26639,plain,(
  ( ! [X420,X422,X418,X421,X423,X417,X419] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X420,'type.num.num'),X421),s(X420,happ(s(fun(bool,X420),X422),s(bool,X418))))),s('type.num.num',X423))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X417),s(bool,X418))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X419)))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X423))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X420,'type.num.num'),X421),s(X420,happ(s(fun(bool,X420),X422),s(bool,'T_0')))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X419),s('type.num.num',happ(s(fun(bool,'type.num.num'),X417),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X418)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f2580,f19209])).
fof(f2580,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(bool,X1),X3),s(bool,X4))))),s('type.num.num',X0))) | ~p(s(bool,X4)) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(bool,X1),X3),s(bool,'T_0'))))))))))) ) | $spl0),
  inference(resolution,[],[f2555,f118])).
fof(f2555,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X3))))),s('type.num.num',X4)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'T_0')))))))))) | ~p(s(bool,X3))) ) | $spl0),
  inference(superposition,[],[f755,f128])).
fof(f755,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X36,'type.num.num'),X37),s(fun(bool,X36),X38))),s(bool,X35))),s('type.num.num',X39)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X39))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X36,'type.num.num'),X37),s(X36,happ(s(fun(bool,X36),X38),s(bool,'T_0')))))))))) | ~p(s(bool,X35))) ) | $spl0),
  inference(superposition,[],[f583,f119])).
fof(f583,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X2,'type.num.num'),'_dst_x0x1_2'(s(fun(X0,'type.num.num'),X1),s(fun(X2,X0),X3))),s(X2,X4))),s('type.num.num',X5)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X5))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(X2,X0),X3),s(X2,X4))))))))))) ) | $spl0),
  inference(superposition,[],[f150,f128])).
fof(f20383,plain,(
  ( ! [X405,X403,X404,X400,X402] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X402))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,X400)))))))))) | s(bool,'F_0') = s(bool,X400) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,'T_0'))))),s('type.num.num',X402)))) ) | $spl74),
  inference(cnf_transformation,[],[f20383_D])).
fof(f20383_D,plain,(
  ( ! [X405,X403,X404,X400,X402] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X402))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,X400)))))))))) | s(bool,'F_0') = s(bool,X400) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,'T_0'))))),s('type.num.num',X402)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f26792,plain,(
  ( ! [X397,X399,X401,X396,X398,X400,X402] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X399))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X400,'type.num.num'),X401),s(X400,happ(s(fun(bool,X400),X402),s(bool,X397)))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X400,'type.num.num'),X401),s(X400,happ(s(fun(bool,X400),X402),s(bool,'T_0'))))),s('type.num.num',X399))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X398),s('type.num.num',happ(s(fun(bool,'type.num.num'),X396),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X397)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f26636,f18956])).
fof(f26636,plain,(
  ( ! [X397,X399,X401,X396,X398,X400,X402] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X399))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X400,'type.num.num'),X401),s(X400,happ(s(fun(bool,X400),X402),s(bool,X397)))))))))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X396),s(bool,X397))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X398)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X400,'type.num.num'),X401),s(X400,happ(s(fun(bool,X400),X402),s(bool,'T_0'))))),s('type.num.num',X399))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X398),s('type.num.num',happ(s(fun(bool,'type.num.num'),X396),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X397)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f1853,f19209])).
fof(f1853,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X10))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X11,'type.num.num'),X12),s(X11,happ(s(fun(bool,X11),X13),s(bool,X14)))))))))) | ~p(s(bool,X14)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X11,'type.num.num'),X12),s(X11,happ(s(fun(bool,X11),X13),s(bool,'T_0'))))),s('type.num.num',X10)))) ) | $spl0),
  inference(resolution,[],[f1841,f115])).
fof(f1841,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'T_0'))))),s('type.num.num',X3)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X4)))))))))) | ~p(s(bool,X4))) ) | $spl0),
  inference(superposition,[],[f750,f128])).
fof(f750,plain,(
  ( ! [X14,X12,X10,X13,X11] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X11,'type.num.num'),X12),s(fun(bool,X11),X13))),s(bool,'T_0'))),s('type.num.num',X14)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X14))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X11,'type.num.num'),X12),s(X11,happ(s(fun(bool,X11),X13),s(bool,X10)))))))))) | ~p(s(bool,X10))) ) | $spl0),
  inference(superposition,[],[f583,f119])).
fof(f25127,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))))) ) | $spl86),
  inference(cnf_transformation,[],[f25127_D])).
fof(f25127_D,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f25126,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'T_0')))))) ) | $spl84),
  inference(cnf_transformation,[],[f25126_D])).
fof(f25126_D,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'T_0')))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f24815,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num',happ(s(fun(bool,'type.num.num'),X11),s(bool,'T_0'))))) | ~p(s(bool,happ(s(fun(bool,bool),X9),s(bool,'F_0'))))) ) | $spl72),
  inference(superposition,[],[f20502,f117])).
fof(f20502,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'F_0')))))))) ) | $spl72),
  inference(superposition,[],[f20382,f128])).
fof(f24814,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X7),s('type.num.num',happ(s(fun(bool,'type.num.num'),X8),s(bool,'T_0'))))) | ~p(s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))))) ) | $spl72),
  inference(superposition,[],[f20502,f119])).
fof(f25125,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | $spl82),
  inference(cnf_transformation,[],[f25125_D])).
fof(f25125_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f24812,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'T_0'))))) | s(bool,'F_0') = s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) ) | $spl72),
  inference(superposition,[],[f20502,f127])).
fof(f20389,plain,(
  ( ! [X476,X478,X480,X482,X477,X479,X481] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X479),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,X477))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X479)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X478),s('type.num.num',happ(s(fun(bool,'type.num.num'),X476),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X477)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f20225,f18956])).
fof(f20225,plain,(
  ( ! [X476,X478,X480,X482,X477,X479,X481] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X479),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,X477))))))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X476),s(bool,X477))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X478)))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X480,'type.num.num'),X481),s(X480,happ(s(fun(bool,X480),X482),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X479)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X478),s('type.num.num',happ(s(fun(bool,'type.num.num'),X476),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X477)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f3222,f19209])).
fof(f20387,plain,(
  ( ! [X455,X461,X457,X459,X460,X456,X458] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,X456))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X461)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X461),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,'T_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X457),s('type.num.num',happ(s(fun(bool,'type.num.num'),X455),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X456)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f20222,f18956])).
fof(f20222,plain,(
  ( ! [X455,X461,X457,X459,X460,X456,X458] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,X456))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X461)))))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X455),s(bool,X456))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X457)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X461),s('type.num.num',happ(s(fun(X458,'type.num.num'),X459),s(X458,happ(s(fun(bool,X458),X460),s(bool,'T_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X457),s('type.num.num',happ(s(fun(bool,'type.num.num'),X455),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X456)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f3057,f19209])).
fof(f20384,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X425] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,X421))))),s('type.num.num',X426))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X426))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,'T_0')))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X422),s('type.num.num',happ(s(fun(bool,'type.num.num'),X420),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X421)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f20217,f18956])).
fof(f20217,plain,(
  ( ! [X420,X422,X424,X426,X421,X423,X425] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,X421))))),s('type.num.num',X426))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X420),s(bool,X421))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X422)))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X426))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X423,'type.num.num'),X424),s(X423,happ(s(fun(bool,X423),X425),s(bool,'T_0')))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X422),s('type.num.num',happ(s(fun(bool,'type.num.num'),X420),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X421)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f2580,f19209])).
fof(f20381,plain,(
  ( ! [X399,X405,X401,X403,X404,X400,X402] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X402))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,X400)))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,'T_0'))))),s('type.num.num',X402))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X401),s('type.num.num',happ(s(fun(bool,'type.num.num'),X399),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X400)) ) | ($spl0 | $spl11 | $spl12)),
  inference(subsumption_resolution,[],[f20214,f18956])).
fof(f20214,plain,(
  ( ! [X399,X405,X401,X403,X404,X400,X402] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X402))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,X400)))))))))) | ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X399),s(bool,X400))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X401)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X403,'type.num.num'),X404),s(X403,happ(s(fun(bool,X403),X405),s(bool,'T_0'))))),s('type.num.num',X402))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X401),s('type.num.num',happ(s(fun(bool,'type.num.num'),X399),s(bool,'F_0'))))) | s(bool,'F_0') = s(bool,X400)) ) | ($spl0 | $spl12)),
  inference(superposition,[],[f1853,f19209])).
fof(f18901,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X14)),s('type.num.num',sK3(X14))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X14)))) ) | $spl38),
  inference(cnf_transformation,[],[f18901_D])).
fof(f18901_D,plain,(
  ( ! [X14] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X14)),s('type.num.num',sK3(X14))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X14)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f18928,plain,(
  ( ! [X318,X320,X319] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X319),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,X318))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X319))))))) ) | $spl70),
  inference(cnf_transformation,[],[f18928_D])).
fof(f18928_D,plain,(
  ( ! [X318,X320,X319] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X319),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,X318))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X319))))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f18927,plain,(
  ( ! [X318,X320,X317,X319] : (s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X319),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,X318))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X319)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X317)),s('type.num.num',sK3(X317))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X317)))) ) | ($spl0 | $spl11 | $spl24)),
  inference(subsumption_resolution,[],[f18802,f10383])).
fof(f10383,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X28),s(bool,X27))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X29)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X29),s('type.num.num',happ(s(fun(bool,'type.num.num'),X28),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl0 | $spl11)),
  inference(subsumption_resolution,[],[f10350,f99])).
fof(f10350,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X28),s(bool,X27))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X29)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X29),s('type.num.num',happ(s(fun(bool,'type.num.num'),X28),s(bool,'F_0'))))) | p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,X27)) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f10163,f127])).
fof(f18802,plain,(
  ( ! [X318,X320,X317,X319] : (s(bool,'T_0') != s(bool,X318) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X319),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X320),s(bool,X318))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X319)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X317)),s('type.num.num',sK3(X317))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X317)))) ) | ($spl0 | $spl11 | $spl24)),
  inference(superposition,[],[f10757,f17983])).
fof(f17983,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X9)),s('type.num.num',sK3(X9)))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X9)))) ) | $spl24),
  inference(resolution,[],[f17870,f117])).
fof(f17870,plain,(
  ( ! [X12] : (p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X12)))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X12)),s('type.num.num',sK3(X12))))) ) | $spl24),
  inference(subsumption_resolution,[],[f17712,f99])).
fof(f17712,plain,(
  ( ! [X12] : (p(s(bool,'F_0')) | p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X12)))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X12)),s('type.num.num',sK3(X12))))) ) | $spl24),
  inference(superposition,[],[f91,f17143])).
fof(f17143,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X2),s('type.num.num',X3)))) ) | $spl24),
  inference(trivial_inequality_removal,[],[f17136])).
fof(f17136,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X2),s('type.num.num',X3)))) ) | $spl24),
  inference(superposition,[],[f16971,f127])).
fof(f16971,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num',X14))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X15)))) ) | $spl24),
  inference(forward_demodulation,[],[f16948,f368])).
fof(f16948,plain,(
  ( ! [X14,X15] : (s(bool,'const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num',X14))) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X15)))) ) | $spl24),
  inference(superposition,[],[f12710,f130])).
fof(f10757,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num',happ(s(fun(bool,'type.num.num'),X7),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X7),s(bool,X5))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X6)))))) | p(s(bool,X4))) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f10754,f115])).
fof(f18926,plain,(
  ( ! [X300,X299] : (s(bool,'T_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X299),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X300))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X300),s('type.num.num',happ(s(fun(bool,'type.num.num'),X299),s(bool,'F_0')))))) ) | $spl68),
  inference(cnf_transformation,[],[f18926_D])).
fof(f18926_D,plain,(
  ( ! [X300,X299] : (s(bool,'T_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X299),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X300))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X300),s('type.num.num',happ(s(fun(bool,'type.num.num'),X299),s(bool,'F_0')))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f18796,plain,(
  ( ! [X300,X298,X299] : (s(bool,'T_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X299),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X300))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X300),s('type.num.num',happ(s(fun(bool,'type.num.num'),X299),s(bool,'F_0'))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X298)),s('type.num.num',sK3(X298))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X298)))) ) | ($spl0 | $spl11 | $spl24)),
  inference(superposition,[],[f10447,f17983])).
fof(f18905,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X39),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X40))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X40),s('type.num.num',happ(s(fun(bool,'type.num.num'),X39),s(bool,'F_0')))))) ) | $spl44),
  inference(cnf_transformation,[],[f18905_D])).
fof(f18905_D,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X39),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X40))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X40),s('type.num.num',happ(s(fun(bool,'type.num.num'),X39),s(bool,'F_0')))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f18792,plain,(
  ( ! [X288,X287,X286] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X287),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X288))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X286)),s('type.num.num',sK3(X286))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X288),s('type.num.num',happ(s(fun(bool,'type.num.num'),X287),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X286)))) ) | ($spl0 | $spl11 | $spl24)),
  inference(superposition,[],[f10333,f17983])).
fof(f11482,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X26)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X26),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'F_0')))))) ) | $spl12),
  inference(cnf_transformation,[],[f11482_D])).
fof(f11482_D,plain,(
  ( ! [X26,X27] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X26)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X26),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'F_0')))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f18791,plain,(
  ( ! [X285,X283,X284] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X284),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X285)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X285),s('type.num.num',happ(s(fun(bool,'type.num.num'),X284),s(bool,'F_0'))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X283)),s('type.num.num',sK3(X283))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X283)))) ) | ($spl0 | $spl11 | $spl24)),
  inference(superposition,[],[f10163,f17983])).
fof(f18925,plain,(
  ( ! [X280,X282,X279,X281] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X280)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X280),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X281),s('type.num.num',X282)))))))) ) | $spl66),
  inference(cnf_transformation,[],[f18925_D])).
fof(f18925_D,plain,(
  ( ! [X280,X282,X279,X281] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X280)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X280),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X281),s('type.num.num',X282)))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f18924,plain,(
  ( ! [X280,X282,X279,X281,X278] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X280)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X280),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X281),s('type.num.num',X282))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X278)),s('type.num.num',sK3(X278))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X278)))) ) | ($spl0 | $spl24)),
  inference(subsumption_resolution,[],[f18790,f5114])).
fof(f5114,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))))))) ) | $spl0),
  inference(forward_demodulation,[],[f5088,f124])).
fof(f5088,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X0))),s('type.num.num',X1)))))))) ) | $spl0),
  inference(superposition,[],[f1658,f124])).
fof(f1658,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,happ(s(fun(X14,bool),X15),s(X14,X16)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X17),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X18)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num',happ(s(fun(bool,'type.num.num'),X17),s(bool,happ(s(fun(X14,bool),X15),s(X14,X16)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1620,f128])).
fof(f1620,plain,(
  ( ! [X14,X17,X15,X18,X16] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X17),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X18)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num',happ(s(fun(X14,'type.num.num'),'_dst_x0x1_2'(s(fun(bool,'type.num.num'),X17),s(fun(X14,bool),X15))),s(X14,X16))))) | ~p(s(bool,happ(s(fun(X14,bool),X15),s(X14,X16))))) ) | $spl0),
  inference(superposition,[],[f688,f119])).
fof(f18790,plain,(
  ( ! [X280,X282,X279,X281,X278] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X280)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X281),s('type.num.num',X282)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X280),s('type.num.num',happ(s(fun(bool,'type.num.num'),X279),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X281),s('type.num.num',X282))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X278)),s('type.num.num',sK3(X278))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X278)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f5489,f17983])).
fof(f18923,plain,(
  ( ! [X251,X249,X252,X250] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X249),s('type.num.num',happ(s(fun(X250,'type.num.num'),X251),s(X250,happ(s(fun(bool,X250),X252),s(bool,'T_0'))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X250,'type.num.num'),X251),s(X250,happ(s(fun(bool,X250),X252),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X249))))))) ) | $spl64),
  inference(cnf_transformation,[],[f18923_D])).
fof(f18923_D,plain,(
  ( ! [X251,X249,X252,X250] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X249),s('type.num.num',happ(s(fun(X250,'type.num.num'),X251),s(X250,happ(s(fun(bool,X250),X252),s(bool,'T_0'))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X250,'type.num.num'),X251),s(X250,happ(s(fun(bool,X250),X252),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X249))))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f18784,plain,(
  ( ! [X251,X249,X252,X250,X248] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X249),s('type.num.num',happ(s(fun(X250,'type.num.num'),X251),s(X250,happ(s(fun(bool,X250),X252),s(bool,'T_0'))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X248)),s('type.num.num',sK3(X248))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X250,'type.num.num'),X251),s(X250,happ(s(fun(bool,X250),X252),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X249)))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X248)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f3357,f17983])).
fof(f3357,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X4))))))) | p(s(bool,X4)) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))))))) ) | $spl0),
  inference(resolution,[],[f3332,f118])).
fof(f3332,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X3)))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X4)))))) | p(s(bool,X3))) ) | $spl0),
  inference(superposition,[],[f1011,f128])).
fof(f1011,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X51),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X52,'type.num.num'),X53),s(fun(bool,X52),X54))),s(bool,X50)))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X52,'type.num.num'),X53),s(X52,happ(s(fun(bool,X52),X54),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X51)))))) | p(s(bool,X50))) ) | $spl0),
  inference(superposition,[],[f585,f115])).
fof(f18922,plain,(
  ( ! [X246,X244,X247,X245] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X244),s('type.num.num',happ(s(fun(X245,'type.num.num'),X246),s(X245,happ(s(fun(bool,X245),X247),s(bool,'T_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X245,'type.num.num'),X246),s(X245,happ(s(fun(bool,X245),X247),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X244))))))) ) | $spl62),
  inference(cnf_transformation,[],[f18922_D])).
fof(f18922_D,plain,(
  ( ! [X246,X244,X247,X245] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X244),s('type.num.num',happ(s(fun(X245,'type.num.num'),X246),s(X245,happ(s(fun(bool,X245),X247),s(bool,'T_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X245,'type.num.num'),X246),s(X245,happ(s(fun(bool,X245),X247),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X244))))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f18783,plain,(
  ( ! [X243,X246,X244,X247,X245] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X244),s('type.num.num',happ(s(fun(X245,'type.num.num'),X246),s(X245,happ(s(fun(bool,X245),X247),s(bool,'T_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X245,'type.num.num'),X246),s(X245,happ(s(fun(bool,X245),X247),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X244)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X243)),s('type.num.num',sK3(X243))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X243)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f3332,f17983])).
fof(f18920,plain,(
  ( ! [X227,X225,X226,X224] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X224,'type.num.num'),X225),s(X224,happ(s(fun(bool,X224),X226),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X227)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X227),s('type.num.num',happ(s(fun(X224,'type.num.num'),X225),s(X224,happ(s(fun(bool,X224),X226),s(bool,'F_0')))))))) ) | $spl60),
  inference(cnf_transformation,[],[f18920_D])).
fof(f18920_D,plain,(
  ( ! [X227,X225,X226,X224] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X224,'type.num.num'),X225),s(X224,happ(s(fun(bool,X224),X226),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X227)))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X227),s('type.num.num',happ(s(fun(X224,'type.num.num'),X225),s(X224,happ(s(fun(bool,X224),X226),s(bool,'F_0')))))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f18779,plain,(
  ( ! [X227,X225,X223,X226,X224] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X224,'type.num.num'),X225),s(X224,happ(s(fun(bool,X224),X226),s(bool,'T_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X227)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X223)),s('type.num.num',sK3(X223))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X227),s('type.num.num',happ(s(fun(X224,'type.num.num'),X225),s(X224,happ(s(fun(bool,X224),X226),s(bool,'F_0'))))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X223)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f3122,f17983])).
fof(f3122,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X10,'type.num.num'),X11),s(X10,happ(s(fun(bool,X10),X12),s(bool,X13))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X14)))))) | p(s(bool,X13)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X14),s('type.num.num',happ(s(fun(X10,'type.num.num'),X11),s(X10,happ(s(fun(bool,X10),X12),s(bool,'F_0')))))))) ) | $spl0),
  inference(resolution,[],[f3110,f115])).
fof(f3110,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'F_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X4))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | p(s(bool,X4))) ) | $spl0),
  inference(superposition,[],[f1005,f128])).
fof(f1005,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X21),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X22,'type.num.num'),X23),s(fun(bool,X22),X24))),s(bool,'F_0')))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X22,'type.num.num'),X23),s(X22,happ(s(fun(bool,X22),X24),s(bool,X20))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X21)))))) | p(s(bool,X20))) ) | $spl0),
  inference(superposition,[],[f585,f115])).
fof(f18919,plain,(
  ( ! [X212,X210,X211,X209] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X209,'type.num.num'),X210),s(X209,happ(s(fun(bool,X209),X211),s(bool,'T_0'))))),s('type.num.num',X212))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X212))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X209,'type.num.num'),X210),s(X209,happ(s(fun(bool,X209),X211),s(bool,'F_0'))))))))))) ) | $spl58),
  inference(cnf_transformation,[],[f18919_D])).
fof(f18919_D,plain,(
  ( ! [X212,X210,X211,X209] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X209,'type.num.num'),X210),s(X209,happ(s(fun(bool,X209),X211),s(bool,'T_0'))))),s('type.num.num',X212))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X212))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X209,'type.num.num'),X210),s(X209,happ(s(fun(bool,X209),X211),s(bool,'F_0'))))))))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f18776,plain,(
  ( ! [X212,X210,X208,X211,X209] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X209,'type.num.num'),X210),s(X209,happ(s(fun(bool,X209),X211),s(bool,'T_0'))))),s('type.num.num',X212))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X208)),s('type.num.num',sK3(X208))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X212))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X209,'type.num.num'),X210),s(X209,happ(s(fun(bool,X209),X211),s(bool,'F_0')))))))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X208)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f2715,f17983])).
fof(f2715,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(bool,X1),X3),s(bool,X4))))),s('type.num.num',X0))) | p(s(bool,X4)) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X1,'type.num.num'),X2),s(X1,happ(s(fun(bool,X1),X3),s(bool,'F_0'))))))))))) ) | $spl0),
  inference(resolution,[],[f2688,f118])).
fof(f2688,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X3))))),s('type.num.num',X4)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X4))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'F_0')))))))))) | p(s(bool,X3))) ) | $spl0),
  inference(superposition,[],[f758,f128])).
fof(f758,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X51,'type.num.num'),X52),s(fun(bool,X51),X53))),s(bool,X50))),s('type.num.num',X54)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X54))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X51,'type.num.num'),X52),s(X51,happ(s(fun(bool,X51),X53),s(bool,'F_0')))))))))) | p(s(bool,X50))) ) | $spl0),
  inference(superposition,[],[f583,f115])).
fof(f18918,plain,(
  ( ! [X206,X204,X207,X205] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X204,'type.num.num'),X205),s(X204,happ(s(fun(bool,X204),X206),s(bool,'T_0'))))),s('type.num.num',X207)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X207))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X204,'type.num.num'),X205),s(X204,happ(s(fun(bool,X204),X206),s(bool,'F_0'))))))))))) ) | $spl56),
  inference(cnf_transformation,[],[f18918_D])).
fof(f18918_D,plain,(
  ( ! [X206,X204,X207,X205] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X204,'type.num.num'),X205),s(X204,happ(s(fun(bool,X204),X206),s(bool,'T_0'))))),s('type.num.num',X207)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X207))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X204,'type.num.num'),X205),s(X204,happ(s(fun(bool,X204),X206),s(bool,'F_0'))))))))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f18775,plain,(
  ( ! [X206,X204,X207,X205,X203] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X204,'type.num.num'),X205),s(X204,happ(s(fun(bool,X204),X206),s(bool,'T_0'))))),s('type.num.num',X207)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X207))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X204,'type.num.num'),X205),s(X204,happ(s(fun(bool,X204),X206),s(bool,'F_0')))))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X203)),s('type.num.num',sK3(X203))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X203)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f2688,f17983])).
fof(f18916,plain,(
  ( ! [X187,X185,X186,X184] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X184))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X185,'type.num.num'),X186),s(X185,happ(s(fun(bool,X185),X187),s(bool,'T_0')))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X185,'type.num.num'),X186),s(X185,happ(s(fun(bool,X185),X187),s(bool,'F_0'))))),s('type.num.num',X184)))) ) | $spl54),
  inference(cnf_transformation,[],[f18916_D])).
fof(f18916_D,plain,(
  ( ! [X187,X185,X186,X184] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X184))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X185,'type.num.num'),X186),s(X185,happ(s(fun(bool,X185),X187),s(bool,'T_0')))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X185,'type.num.num'),X186),s(X185,happ(s(fun(bool,X185),X187),s(bool,'F_0'))))),s('type.num.num',X184)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f18771,plain,(
  ( ! [X187,X185,X186,X184,X183] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X184))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X185,'type.num.num'),X186),s(X185,happ(s(fun(bool,X185),X187),s(bool,'T_0')))))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X183)),s('type.num.num',sK3(X183))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X185,'type.num.num'),X186),s(X185,happ(s(fun(bool,X185),X187),s(bool,'F_0'))))),s('type.num.num',X184))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X183)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f2417,f17983])).
fof(f2417,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X10))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X11,'type.num.num'),X12),s(X11,happ(s(fun(bool,X11),X13),s(bool,X14)))))))))) | p(s(bool,X14)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X11,'type.num.num'),X12),s(X11,happ(s(fun(bool,X11),X13),s(bool,'F_0'))))),s('type.num.num',X10)))) ) | $spl0),
  inference(resolution,[],[f2403,f115])).
fof(f2403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,'F_0'))))),s('type.num.num',X3)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,happ(s(fun(bool,X0),X2),s(bool,X4)))))))))) | p(s(bool,X4))) ) | $spl0),
  inference(superposition,[],[f752,f128])).
fof(f752,plain,(
  ( ! [X24,X23,X21,X22,X20] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X21,'type.num.num'),X22),s(fun(bool,X21),X23))),s(bool,'F_0'))),s('type.num.num',X24)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X24))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X21,'type.num.num'),X22),s(X21,happ(s(fun(bool,X21),X23),s(bool,X20)))))))))) | p(s(bool,X20))) ) | $spl0),
  inference(superposition,[],[f583,f115])).
fof(f18915,plain,(
  ( ! [X177,X175,X176,X174] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X174),s('type.num.num',happ(s(fun(X175,'type.num.num'),X176),s(X175,happ(s(fun(bool,X175),X177),s(bool,'T_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X175,'type.num.num'),X176),s(fun(bool,X175),X177))),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X174))))))) ) | $spl52),
  inference(cnf_transformation,[],[f18915_D])).
fof(f18915_D,plain,(
  ( ! [X177,X175,X176,X174] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X174),s('type.num.num',happ(s(fun(X175,'type.num.num'),X176),s(X175,happ(s(fun(bool,X175),X177),s(bool,'T_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X175,'type.num.num'),X176),s(fun(bool,X175),X177))),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X174))))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f18769,plain,(
  ( ! [X177,X175,X173,X176,X174] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X174),s('type.num.num',happ(s(fun(X175,'type.num.num'),X176),s(X175,happ(s(fun(bool,X175),X177),s(bool,'T_0')))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X175,'type.num.num'),X176),s(fun(bool,X175),X177))),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X174)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X173)),s('type.num.num',sK3(X173))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X173)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f1077,f17983])).
fof(f1077,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X51),s('type.num.num',happ(s(fun(X52,'type.num.num'),X53),s(X52,happ(s(fun(bool,X52),X54),s(bool,X50)))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X52,'type.num.num'),X53),s(fun(bool,X52),X54))),s(bool,'F_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X51)))))) | p(s(bool,X50))) ) | $spl0),
  inference(superposition,[],[f586,f115])).
fof(f586,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X11),s('type.num.num',happ(s(fun(X7,'type.num.num'),X8),s(X7,happ(s(fun(X6,X7),X9),s(X6,X10)))))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X6,'type.num.num'),'_dst_x0x1_2'(s(fun(X7,'type.num.num'),X8),s(fun(X6,X7),X9))),s(X6,X10))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X11))))))) ) | $spl0),
  inference(superposition,[],[f150,f128])).
fof(f18914,plain,(
  ( ! [X156,X154,X157,X155] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X154),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X155,'type.num.num'),X156),s(fun(bool,X155),X157))),s(bool,'T_0')))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X155,'type.num.num'),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X154))))))) ) | $spl50),
  inference(cnf_transformation,[],[f18914_D])).
fof(f18914_D,plain,(
  ( ! [X156,X154,X157,X155] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X154),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X155,'type.num.num'),X156),s(fun(bool,X155),X157))),s(bool,'T_0')))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X155,'type.num.num'),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X154))))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f18765,plain,(
  ( ! [X156,X154,X157,X155,X153] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X154),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X155,'type.num.num'),X156),s(fun(bool,X155),X157))),s(bool,'T_0')))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X155,'type.num.num'),X156),s(X155,happ(s(fun(bool,X155),X157),s(bool,'F_0'))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X154)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X153)),s('type.num.num',sK3(X153))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X153)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f1011,f17983])).
fof(f18913,plain,(
  ( ! [X134,X136,X135,X137] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X134,'type.num.num'),X135),s(X134,happ(s(fun(bool,X134),X136),s(bool,'T_0'))))),s('type.num.num',X137)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X137))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X134,'type.num.num'),X135),s(fun(bool,X134),X136))),s(bool,'F_0'))))))))) ) | $spl48),
  inference(cnf_transformation,[],[f18913_D])).
fof(f18913_D,plain,(
  ( ! [X134,X136,X135,X137] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X134,'type.num.num'),X135),s(X134,happ(s(fun(bool,X134),X136),s(bool,'T_0'))))),s('type.num.num',X137)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X137))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X134,'type.num.num'),X135),s(fun(bool,X134),X136))),s(bool,'F_0'))))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f18761,plain,(
  ( ! [X134,X136,X135,X133,X137] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X134,'type.num.num'),X135),s(X134,happ(s(fun(bool,X134),X136),s(bool,'T_0'))))),s('type.num.num',X137)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X137))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X134,'type.num.num'),X135),s(fun(bool,X134),X136))),s(bool,'F_0')))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X133)),s('type.num.num',sK3(X133))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X133)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f861,f17983])).
fof(f861,plain,(
  ( ! [X54,X52,X50,X53,X51] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X51,'type.num.num'),X52),s(X51,happ(s(fun(bool,X51),X53),s(bool,X50))))),s('type.num.num',X54)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X54))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X51,'type.num.num'),X52),s(fun(bool,X51),X53))),s(bool,'F_0')))))))) | p(s(bool,X50))) ) | $spl0),
  inference(superposition,[],[f584,f115])).
fof(f584,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(X7,'type.num.num'),X8),s(X7,happ(s(fun(X6,X7),X9),s(X6,X10))))),s('type.num.num',X11)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X11))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X6,'type.num.num'),'_dst_x0x1_2'(s(fun(X7,'type.num.num'),X8),s(fun(X6,X7),X9))),s(X6,X10))))))))) ) | $spl0),
  inference(superposition,[],[f150,f128])).
fof(f18912,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X114,'type.num.num'),X115),s(fun(bool,X114),X116))),s(bool,'T_0'))),s('type.num.num',X117)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X117))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X114,'type.num.num'),X115),s(X114,happ(s(fun(bool,X114),X116),s(bool,'F_0'))))))))))) ) | $spl46),
  inference(cnf_transformation,[],[f18912_D])).
fof(f18912_D,plain,(
  ( ! [X116,X114,X117,X115] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X114,'type.num.num'),X115),s(fun(bool,X114),X116))),s(bool,'T_0'))),s('type.num.num',X117)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X117))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X114,'type.num.num'),X115),s(X114,happ(s(fun(bool,X114),X116),s(bool,'F_0'))))))))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f18757,plain,(
  ( ! [X116,X114,X117,X115,X113] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),'_dst_x0x1_2'(s(fun(X114,'type.num.num'),X115),s(fun(bool,X114),X116))),s(bool,'T_0'))),s('type.num.num',X117)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X117))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X114,'type.num.num'),X115),s(X114,happ(s(fun(bool,X114),X116),s(bool,'F_0')))))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X113)),s('type.num.num',sK3(X113))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X113)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f758,f17983])).
fof(f18722,plain,(
  ( ! [X39,X38,X40] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X39),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X40))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X40),s('type.num.num',happ(s(fun(bool,'type.num.num'),X39),s(bool,'F_0'))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X38)),s('type.num.num',sK3(X38))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X38)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f4847,f17983])).
fof(f4847,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X11),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X9))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X10))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num',happ(s(fun(bool,'type.num.num'),X11),s(bool,'F_0'))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X9))))) ) | $spl0),
  inference(resolution,[],[f4841,f114])).
fof(f4841,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'F_0')))))) ) | $spl0),
  inference(forward_demodulation,[],[f4804,f124])).
fof(f4804,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X0))),s('type.num.num',X1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',happ(s(fun(bool,'type.num.num'),X2),s(bool,'F_0')))))) ) | $spl0),
  inference(superposition,[],[f1220,f124])).
fof(f1220,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X10),s(bool,happ(s(fun(X11,bool),X12),s(X11,X13))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X14)))))) | p(s(bool,happ(s(fun(X11,bool),X12),s(X11,X13)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X14),s('type.num.num',happ(s(fun(bool,'type.num.num'),X10),s(bool,'F_0')))))) ) | $spl0),
  inference(resolution,[],[f1136,f115])).
fof(f1136,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X27),s('type.num.num',happ(s(fun(bool,'type.num.num'),X28),s(bool,'F_0')))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X28),s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X27)))))) | p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))) ) | $spl0),
  inference(forward_demodulation,[],[f1101,f128])).
fof(f1101,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X27),s('type.num.num',happ(s(fun(bool,'type.num.num'),X28),s(bool,'F_0')))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X24,'type.num.num'),'_dst_x0x1_2'(s(fun(bool,'type.num.num'),X28),s(fun(X24,bool),X25))),s(X24,X26))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X27)))))) | p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))) ) | $spl0),
  inference(superposition,[],[f586,f115])).
fof(f18720,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X33),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X34)))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X32)),s('type.num.num',sK3(X32))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X34),s('type.num.num',happ(s(fun(bool,'type.num.num'),X33),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X32)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f4841,f17983])).
fof(f18903,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X21))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X22),s(bool,'T_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X22),s(bool,'F_0'))),s('type.num.num',X21)))) ) | $spl42),
  inference(cnf_transformation,[],[f18903_D])).
fof(f18903_D,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X21))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X22),s(bool,'T_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X22),s(bool,'F_0'))),s('type.num.num',X21)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f18716,plain,(
  ( ! [X21,X22,X20] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X21))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X22),s(bool,'T_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X22),s(bool,'F_0'))),s('type.num.num',X21))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X20)),s('type.num.num',sK3(X20))))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X20)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f4213,f17983])).
fof(f4213,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X11))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X10),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X9))))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X10),s(bool,'F_0'))),s('type.num.num',X11))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num',X9))))) ) | $spl0),
  inference(resolution,[],[f4150,f114])).
fof(f4150,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,'F_0'))),s('type.num.num',X2)))) ) | $spl0),
  inference(forward_demodulation,[],[f4115,f124])).
fof(f4115,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num',X1)))))))))) | p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',X0))),s('type.num.num',X1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X3),s(bool,'F_0'))),s('type.num.num',X2)))) ) | $spl0),
  inference(superposition,[],[f1167,f124])).
fof(f1167,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X10))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X11),s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))))))))) | p(s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X11),s(bool,'F_0'))),s('type.num.num',X10)))) ) | $spl0),
  inference(resolution,[],[f921,f115])).
fof(f921,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'F_0'))),s('type.num.num',X28)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X28))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,happ(s(fun(X24,bool),X25),s(X24,X26)))))))))) | p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))) ) | $spl0),
  inference(forward_demodulation,[],[f884,f128])).
fof(f884,plain,(
  ( ! [X28,X26,X24,X27,X25] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'F_0'))),s('type.num.num',X28)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X28))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(X24,'type.num.num'),'_dst_x0x1_2'(s(fun(bool,'type.num.num'),X27),s(fun(X24,bool),X25))),s(X24,X26)))))))) | p(s(bool,happ(s(fun(X24,bool),X25),s(X24,X26))))) ) | $spl0),
  inference(superposition,[],[f584,f115])).
fof(f18902,plain,(
  ( ! [X15,X16] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X15))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X16),s(bool,'T_0')))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X16),s(bool,'F_0'))),s('type.num.num',X15)))) ) | $spl40),
  inference(cnf_transformation,[],[f18902_D])).
fof(f18902_D,plain,(
  ( ! [X15,X16] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X15))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X16),s(bool,'T_0')))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X16),s(bool,'F_0'))),s('type.num.num',X15)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f18714,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X15))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X16),s(bool,'T_0')))))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK4(X14)),s('type.num.num',sK3(X14))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',happ(s(fun(bool,'type.num.num'),X16),s(bool,'F_0'))),s('type.num.num',X15))) | s(bool,'T_0') = s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),X14)))) ) | ($spl0 | $spl24)),
  inference(superposition,[],[f4150,f17983])).
fof(f16919,plain,(
  ( ! [X66,X64,X62,X63] : (s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,X63))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X64),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,'F_0'))))) | p(s(bool,X66))) ) | $spl36),
  inference(cnf_transformation,[],[f16919_D])).
fof(f16919_D,plain,(
  ( ! [X66,X64,X62,X63] : (s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,X63))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X64))))) | s(bool,'T_0') = s(bool,X63) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X64),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,'F_0'))))) | p(s(bool,X66))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f16918,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'F_0') = s(bool,X65) | s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,X63))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X64))))) | p(s(bool,X66)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X64),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl0 | $spl11 | $spl28)),
  inference(subsumption_resolution,[],[f16864,f127])).
fof(f16864,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'T_0') != s(bool,X65) | s(bool,'F_0') = s(bool,X65) | s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,X63))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X64))))) | p(s(bool,X66)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X64),s('type.num.num',happ(s(fun(bool,'type.num.num'),X62),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X63)) ) | ($spl0 | $spl11 | $spl28)),
  inference(superposition,[],[f13283,f10465])).
fof(f10465,plain,(
  ( ! [X14,X12,X13] : (s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X13),s(bool,X14))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X12))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',happ(s(fun(bool,'type.num.num'),X13),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f10335,f117])).
fof(f10335,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num',happ(s(fun(bool,'type.num.num'),X13),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X13),s(bool,X14))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X12)))))) ) | ($spl0 | $spl11)),
  inference(resolution,[],[f10163,f117])).
fof(f13283,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,X3) != s(bool,X4) | p(s(bool,X3))) ) | $spl28),
  inference(superposition,[],[f12718,f115])).
fof(f16916,plain,(
  ( ! [X61,X59,X58] : (s(bool,X61) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X58))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X59))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X59),s('type.num.num',X58))) | p(s(bool,X61))) ) | $spl32),
  inference(cnf_transformation,[],[f16916_D])).
fof(f16916_D,plain,(
  ( ! [X61,X59,X58] : (s(bool,X61) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X58))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X59))))) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X59),s('type.num.num',X58))) | p(s(bool,X61))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f16915,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'F_0') = s(bool,X60) | s(bool,X61) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X58))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X59))))) | p(s(bool,X61)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X59),s('type.num.num',X58)))) ) | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f16863,f127])).
fof(f16863,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') != s(bool,X60) | s(bool,'F_0') = s(bool,X60) | s(bool,X61) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X58))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X59))))) | p(s(bool,X61)) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X59),s('type.num.num',X58)))) ) | ($spl0 | $spl28)),
  inference(superposition,[],[f13283,f681])).
fof(f15269,plain,(
  ( ! [X192] : (s(bool,X192) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0)))) | s(bool,'F_0') = s(bool,X192)) ) | $spl30),
  inference(cnf_transformation,[],[f15269_D])).
fof(f15269_D,plain,(
  ( ! [X192] : (s(bool,X192) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0)))) | s(bool,'F_0') = s(bool,X192)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f15268,plain,(
  ( ! [X191,X192] : (s(bool,'F_0') != s(bool,X191) | s(bool,X192) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0)))) | s(bool,'F_0') = s(bool,X192)) ) | ($spl11 | $spl14 | $spl28)),
  inference(subsumption_resolution,[],[f15134,f112])).
fof(f15134,plain,(
  ( ! [X191,X192] : (s(bool,'F_0') != s(bool,X191) | s(bool,X192) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0)))) | p(s(bool,X191)) | s(bool,'F_0') = s(bool,X192)) ) | ($spl11 | $spl14 | $spl28)),
  inference(superposition,[],[f13220,f15044])).
fof(f15044,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0)))) | ($spl11 | $spl28)),
  inference(trivial_inequality_removal,[],[f15034])).
fof(f15034,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0)))) | ($spl11 | $spl28)),
  inference(superposition,[],[f13486,f520])).
fof(f13486,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num',X14))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X15)))) ) | $spl28),
  inference(forward_demodulation,[],[f13464,f159])).
fof(f13464,plain,(
  ( ! [X14,X15] : (s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num',X14))) | s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num',X15)))) ) | $spl28),
  inference(superposition,[],[f13320,f130])).
fof(f13320,plain,(
  ( ! [X27] : (s(bool,X27) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'F_0') = s(bool,X27)) ) | $spl28),
  inference(trivial_inequality_removal,[],[f13304])).
fof(f13304,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X27) | s(bool,X27) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0)))) ) | $spl28),
  inference(superposition,[],[f12718,f159])).
fof(f13220,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f13097,f127])).
fof(f13737,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X65) | s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,X66)) ) | $spl14),
  inference(subsumption_resolution,[],[f13690,f112])).
fof(f13690,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X65) | s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | p(s(bool,X65)) | s(bool,'F_0') = s(bool,X66)) ) | $spl14),
  inference(superposition,[],[f13220,f368])).
fof(f12937,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f12921])).
fof(f12921,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl20),
  inference(superposition,[],[f12705,f159])).
fof(f12936,plain,(
  $false | ($spl11 | $spl20)),
  inference(trivial_inequality_removal,[],[f12923])).
fof(f12923,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl11 | $spl20)),
  inference(superposition,[],[f12705,f520])).
fof(f12935,plain,(
  $false | $spl20),
  inference(trivial_inequality_removal,[],[f12927])).
fof(f12927,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl20),
  inference(superposition,[],[f12705,f368])).
fof(f12934,plain,(
  $false | $spl20),
  inference(equality_resolution,[],[f12705])).
fof(f12753,plain,(
  $false | $spl26),
  inference(resolution,[],[f12717,f120])).
fof(f12696,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | p(s(bool,X15)) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X17)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f12671])).
fof(f12671,plain,(
  ( ! [X17,X15,X16] : (s(bool,'F_0') != s(bool,X16) | p(s(bool,X15)) | s(bool,X16) != s(bool,X17) | s(bool,'F_0') = s(bool,X17) | p(s(bool,X15))) ) | $spl16),
  inference(superposition,[],[f11720,f113])).
fof(f11720,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(superposition,[],[f11503,f127])).
fof(f11503,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,X8) | p(s(bool,X9)) | s(bool,X8) != s(bool,X9)) ) | $spl16),
  inference(resolution,[],[f11484,f116])).
fof(f12697,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X14)) ) | $spl16),
  inference(duplicate_literal_removal,[],[f12670])).
fof(f12670,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | s(bool,X13) != s(bool,X14) | s(bool,'F_0') = s(bool,X14) | p(s(bool,X12))) ) | $spl16),
  inference(superposition,[],[f11720,f115])).
fof(f12716,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,X3)) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5)) ) | $spl16),
  inference(subsumption_resolution,[],[f12667,f116])).
fof(f12667,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X4) | p(s(bool,X3)) | s(bool,X4) != s(bool,X5) | s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X3)) ) | $spl16),
  inference(superposition,[],[f11720,f127])).
fof(f12709,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X65) | s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,X66)) ) | $spl16),
  inference(subsumption_resolution,[],[f12659,f112])).
fof(f12659,plain,(
  ( ! [X66,X65] : (s(bool,'F_0') != s(bool,X65) | p(s(bool,X65)) | s(bool,X66) != s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0)))))) | s(bool,'F_0') = s(bool,X66)) ) | $spl16),
  inference(superposition,[],[f11720,f368])).
fof(f12707,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))) | s(bool,'F_0') = s(bool,X51)) ) | $spl22),
  inference(cnf_transformation,[],[f12707_D])).
fof(f12707_D,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))) | s(bool,'F_0') = s(bool,X51)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f12706,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') != s(bool,X50) | s(bool,X51) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))) | s(bool,'F_0') = s(bool,X51)) ) | ($spl11 | $spl16)),
  inference(subsumption_resolution,[],[f12655,f112])).
fof(f12655,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') != s(bool,X50) | p(s(bool,X50)) | s(bool,X51) != s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))) | s(bool,'F_0') = s(bool,X51)) ) | ($spl11 | $spl16)),
  inference(superposition,[],[f11720,f520])).
fof(f12704,plain,(
  ( ! [X47] : (s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'F_0') = s(bool,X47)) ) | $spl18),
  inference(cnf_transformation,[],[f12704_D])).
fof(f12704_D,plain,(
  ( ! [X47] : (s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'F_0') = s(bool,X47)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f12703,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X46) | s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'F_0') = s(bool,X47)) ) | $spl16),
  inference(subsumption_resolution,[],[f12653,f112])).
fof(f12653,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,X46) | p(s(bool,X46)) | s(bool,X47) != s(bool,'const.seq.convergent_1'(s(fun('type.num.num','type.realax.real'),sK0))) | s(bool,'F_0') = s(bool,X47)) ) | $spl16),
  inference(superposition,[],[f11720,f159])).
fof(f11484,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | s(bool,X28) != s(bool,X29) | p(s(bool,X29))) ) | $spl16),
  inference(cnf_transformation,[],[f11484_D])).
fof(f11484_D,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | s(bool,X28) != s(bool,X29) | p(s(bool,X29))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f11436,plain,(
  ( ! [X35,X33,X34,X32] : (s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X34),s('type.num.num',happ(s(fun(bool,'type.num.num'),X35),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X35),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X34)))))) | p(s(bool,X33)) | ~p(s(bool,X32))) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f10757,f117])).
fof(f11435,plain,(
  ( ! [X30,X28,X31,X29] : (s(bool,X28) != s(bool,X29) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X30),s('type.num.num',happ(s(fun(bool,'type.num.num'),X31),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X31),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X30)))))) | p(s(bool,X29)) | ~p(s(bool,X28))) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f10757,f119])).
fof(f11434,plain,(
  ( ! [X26,X24,X27,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'F_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X26),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'F_0'))))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',happ(s(fun(bool,'type.num.num'),X27),s(bool,'T_0'))))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X26)))))) | p(s(bool,X25)) | s(bool,'F_0') = s(bool,X24)) ) | ($spl0 | $spl11)),
  inference(superposition,[],[f10757,f127])).
fof(f516,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f515,f120])).
fof(f515,plain,(
  p(s(bool,'F_0')) | $spl2),
  inference(forward_demodulation,[],[f514,f368])).
fof(f514,plain,(
  p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0))))))) | $spl2),
  inference(subsumption_resolution,[],[f507,f121])).
fof(f507,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0))))))) | $spl2),
  inference(superposition,[],[f151,f310])).
fof(f384,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) | $spl11),
  inference(cnf_transformation,[],[f384_D])).
fof(f384_D,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f437,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))))) )),
  inference(subsumption_resolution,[],[f436,f108])).
fof(f436,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))))) )),
  inference(forward_demodulation,[],[f435,f135])).
fof(f435,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(forward_demodulation,[],[f434,f130])).
fof(f434,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(forward_demodulation,[],[f433,f132])).
fof(f433,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(forward_demodulation,[],[f432,f130])).
fof(f432,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(subsumption_resolution,[],[f405,f99])).
fof(f405,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(superposition,[],[f90,f368])).
fof(f388,plain,(
  $false | $spl5),
  inference(subsumption_resolution,[],[f171,f84])).
fof(f382,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))))) )),
  inference(subsumption_resolution,[],[f381,f108])).
fof(f381,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0)))))) )),
  inference(forward_demodulation,[],[f380,f135])).
fof(f380,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(forward_demodulation,[],[f379,f130])).
fof(f379,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(forward_demodulation,[],[f378,f132])).
fof(f378,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK2(sK0)),s('type.num.num',sK1(sK0))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(forward_demodulation,[],[f365,f130])).
fof(f365,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.seq.mono_1'(s(fun('type.num.num','type.realax.real'),sK0))))) )),
  inference(resolution,[],[f252,f90])).
fof(f213,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',X2)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f212,f130])).
fof(f212,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f211,f132])).
fof(f211,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f210,f130])).
fof(f210,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(subsumption_resolution,[],[f181,f108])).
fof(f181,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(superposition,[],[f90,f135])).
fof(f171,plain,(
  ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0)))) | $spl5),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f175,plain,(
  ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0))))))) | $spl9),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f156,plain,(
  ~p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK1(sK0)))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',sK2(sK0))))))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(resolution,[],[f86,f88])).
fof(f173,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0))))) | $spl6),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f169,plain,(
  p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',sK1(sK0)),s('type.num.num',sK2(sK0))))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(forward_demodulation,[],[f155,f124])).
fof(f155,plain,(
  p(s(bool,happ(s(fun('type.num.num',bool),happ(s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'),s('type.num.num',sK1(sK0)))),s('type.num.num',sK2(sK0))))) | ~p(s(bool,'const.nets.bounded_2'(s('type.pair.prod'('type.topology.metric'('type.realax.real'),fun('type.num.num',fun('type.num.num',bool))),'const.pair.,_2'(s('type.topology.metric'('type.realax.real'),'const.topology.mr|49|_0'),s(fun('type.num.num',fun('type.num.num',bool)),'const.arithmetic.>=_0'))),s(fun('type.num.num','type.realax.real'),sK0))))),
  inference(resolution,[],[f86,f87])).
fof(f149,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X3),s('type.num.num',X2)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1))))))) )),
  inference(forward_demodulation,[],[f148,f130])).
fof(f148,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f147,f132])).
fof(f147,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(forward_demodulation,[],[f134,f130])).
fof(f134,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.real.real_ge_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X0))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X2))),s('type.realax.real',happ(s(fun('type.num.num','type.realax.real'),sK0),s('type.num.num',X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))))) )),
  inference(resolution,[],[f85,f90])).
% SZS output end Proof for thm.seq
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 640
% Passive clauses: 14463
% Generated clauses: 48234
% Final active clauses: 257
% Final passive clauses: 10578
% Input formulas: 27
% Initial clauses: 49
% 
% Selected by SInE selection: 26
% SInE iterations: 1
% 
% Duplicate literals: 501
% Trivial inequalities: 1249
% Fw subsumption resolutions: 2989
% Fw demodulations: 9129
% Bw demodulations: 191
% 
% Simple tautologies: 1032
% Forward subsumptions: 19488
% Bw demodulations to eq. taut.: 59
% 
% Binary resolution: 4855
% Factoring: 3
% Forward superposition: 20757
% Backward superposition: 7774
% Self superposition: 26
% Equality resolution: 112
% 
% Split clauses: 155
% Split components: 310
% SAT solver clauses: 166
% SAT solver unit clauses: 9
% SAT solver binary clauses: 8
% 
% Sat splits: 155
% Sat splitting refutations: 10
% 
% Memory used [KB]: 45542
% Time elapsed: 2.488 s
% ------------------------------
% ------------------------------
% Success in time 14.213 s
