Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f12954,plain,(
  $false),
  inference(subsumption_resolution,[],[f12951,f191])).
fof(f191,plain,(
  s('type.integer.int',sK11) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  s('type.integer.int',sK11) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))),s('type.integer.int',sK11))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK11),s('type.integer.int',sK10))),s('type.integer.int',sK11)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11])],[f92])).
fof(f92,plain,(
  ? [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X1))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X1))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1] : (~s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X1),s('type.integer.int',X0))),s('type.integer.int',X1))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (~s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X7))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (~s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X7),s('type.integer.int',X6))),s('type.integer.int',X7))))),
  file('thm.integer.INT_MOD_NEG_NUMERATOR','thm.integer.INT_MOD_NEG_NUMERATOR')).
fof(f12951,plain,(
  s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(trivial_inequality_removal,[],[f12932])).
fof(f12932,plain,(
  s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))),s('type.integer.int',sK11))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))),s('type.integer.int',sK11))) | s('type.integer.int',sK11) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f192,f4926])).
fof(f4926,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1))),s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3548,f236])).
fof(f236,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6,X10] : s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',X6),s('type.integer.int',X10))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X10)))))),
  file('thm.integer.INT_MOD_NEG_NUMERATOR','thm.integer.int_sub')).
fof(f3548,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X1),s('type.integer.int',X0))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int',X0))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f267,f200])).
fof(f200,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.integer.int',X6)))),
  file('thm.integer.INT_MOD_NEG_NUMERATOR','thm.integer.INT_MUL_LID')).
fof(f267,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))) | s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1,X2] : (s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2] : (~s('type.integer.int',X2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X1),s('type.integer.int',X2))),s('type.integer.int',X0))),s('type.integer.int',X2))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11,X12,X7] : (~s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int',X11),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int',X12),s('type.integer.int',X7))),s('type.integer.int',X11))),s('type.integer.int',X7))))),
  file('thm.integer.INT_MOD_NEG_NUMERATOR','thm.integer.INT_MOD_ADD_MULTIPLES')).
fof(f192,plain,(
  s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK10))),s('type.integer.int',sK11))) != s('type.integer.int','const.integer.int_mod_2'(s('type.integer.int','const.integer.int_sub_2'(s('type.integer.int',sK11),s('type.integer.int',sK10))),s('type.integer.int',sK11)))),
  inference(cnf_transformation,[],[f133])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 478
% Passive clauses: 2865
% Generated clauses: 12429
% Final active clauses: 472
% Final passive clauses: 2175
% Input formulas: 36
% Initial clauses: 134
% 
% Duplicate literals: 441
% Trivial inequalities: 54
% Fw subsumption resolutions: 1922
% Fw demodulations: 188
% Bw demodulations: 6
% 
% Simple tautologies: 193
% Equational tautologies: 60
% Forward subsumptions: 5262
% 
% Binary resolution: 3540
% Factoring: 1018
% Forward superposition: 1437
% Backward superposition: 3110
% Self superposition: 139
% Equality factoring: 108
% Equality resolution: 12
% 
% Split clauses: 1374
% Split components: 2904
% SAT solver clauses: 1650
% SAT solver unit clauses: 32
% SAT solver binary clauses: 136
% 
% Sat splits: 1374
% 
% TWLsolver clauses: 1644
% TWLsolver calls for satisfiability: 245
% 
% Memory used [KB]: 6524
% Time elapsed: 0.395 s
% ------------------------------
% ------------------------------
% Success in time 0.398 s
