Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.util_prob
% (87449)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 77
% Passive clauses: 3483
% Generated clauses: 13885
% Final active clauses: 74
% Final passive clauses: 3309
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 821
% Trivial inequalities: 149
% Fw subsumption resolutions: 1667
% Fw demodulations: 745
% Bw demodulations: 1
% 
% Simple tautologies: 272
% Equational tautologies: 93
% Forward subsumptions: 4945
% 
% Binary resolution: 81
% Forward superposition: 6286
% Backward superposition: 3392
% Self superposition: 316
% Equality factoring: 241
% Equality resolution: 14
% 
% Split clauses: 801
% Split components: 1614
% SAT solver clauses: 1738
% SAT solver unit clauses: 296
% SAT solver binary clauses: 508
% 
% Sat splits: 801
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1730
% TWLsolver calls for satisfiability: 83
% 
% Memory used [KB]: 8699
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (87464)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 51
% Passive clauses: 504
% Generated clauses: 8633
% Final active clauses: 47
% Final passive clauses: 444
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 1214
% Trivial inequalities: 85
% Fw subsumption resolutions: 409
% Fw demodulations: 126
% Bw demodulations: 7
% Condensations: 1795
% 
% Simple tautologies: 152
% Equational tautologies: 68
% Forward subsumptions: 2792
% 
% Binary resolution: 33
% Forward superposition: 2295
% Backward superposition: 1787
% Self superposition: 228
% Equality factoring: 198
% Equality resolution: 8
% 
% Split clauses: 86
% Split components: 173
% SAT solver clauses: 1106
% SAT solver unit clauses: 19
% SAT solver binary clauses: 87
% 
% Sat splits: 86
% Sat splitting refutations: 14
% 
% TWLsolver clauses: 1106
% TWLsolver calls for satisfiability: 98
% 
% Memory used [KB]: 7803
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (87485)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 292
% Passive clauses: 5019
% Generated clauses: 95207
% Final active clauses: 102
% Final passive clauses: 595
% Input formulas: 11
% Initial clauses: 16
% Discarded non-redundant clauses: 1720
% 
% Duplicate literals: 6938
% Trivial inequalities: 632
% Fw subsumption resolutions: 7811
% Fw demodulations: 6062
% Bw demodulations: 3029
% Condensations: 14624
% Global subsumptions: 37
% 
% Simple tautologies: 953
% Equational tautologies: 2074
% Forward subsumptions: 51018
% Bw demodulations to eq. taut.: 98
% 
% Unit resulting resolution: 4942
% Forward superposition: 27942
% Backward superposition: 18277
% Self superposition: 626
% Equality factoring: 406
% Equality resolution: 55
% 
% Split clauses: 141
% Split components: 283
% SAT solver clauses: 5427
% SAT solver unit clauses: 100
% SAT solver binary clauses: 464
% 
% Sat splits: 141
% Sat splitting refutations: 87
% 
% Memory used [KB]: 38378
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% (87772)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2430
% Passive clauses: 28336
% Generated clauses: 95822
% Final active clauses: 1334
% Final passive clauses: 12527
% Input formulas: 11
% Initial clauses: 16
% 
% Selected by SInE selection: 11
% SInE iterations: 1
% 
% Duplicate literals: 4214
% Trivial inequalities: 11449
% Fw subsumption resolutions: 3032
% Fw demodulations: 3700
% Bw demodulations: 80
% 
% Simple tautologies: 3796
% Equational tautologies: 1
% Forward subsumptions: 40355
% 
% Binary resolution: 1183
% Factoring: 1
% Forward superposition: 30828
% Backward superposition: 38485
% Equality factoring: 32
% Equality resolution: 544
% 
% Split clauses: 2242
% Split components: 5616
% SAT solver clauses: 2260
% SAT solver binary clauses: 202
% 
% Sat splits: 2242
% Sat splitting refutations: 17
% 
% Memory used [KB]: 39530
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.util_prob
% (87871)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 88
% Passive clauses: 1949
% Generated clauses: 10541
% Final active clauses: 64
% Final passive clauses: 1366
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 355
% Trivial inequalities: 251
% Fw subsumption resolutions: 1318
% Fw demodulations: 124
% Bw demodulations: 405
% Global subsumptions: 1
% 
% Simple tautologies: 428
% Equational tautologies: 61
% Forward subsumptions: 5750
% Bw demodulations to eq. taut.: 9
% 
% Unit resulting resolution: 123
% Forward superposition: 3789
% Backward superposition: 3482
% Self superposition: 135
% Equality factoring: 43
% Equality resolution: 20
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 2008
% SAT solver unit clauses: 44
% SAT solver binary clauses: 127
% 
% Sat splits: 19
% Sat splitting refutations: 37
% 
% Memory used [KB]: 9338
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.util_prob
% (87888)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 427
% Passive clauses: 3749
% Generated clauses: 5612
% Final active clauses: 427
% Final passive clauses: 3333
% Input formulas: 11
% Initial clauses: 16
% 
% Duplicate literals: 279
% Trivial inequalities: 126
% Fw subsumption resolutions: 41
% Fw demodulations: 42
% Condensations: 5
% 
% Simple tautologies: 613
% Equational tautologies: 1
% Forward subsumptions: 809
% 
% Binary resolution: 216
% Forward superposition: 4626
% Backward superposition: 246
% Equality resolution: 94
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 3753
% SAT solver unit clauses: 3
% SAT solver binary clauses: 44
% 
% Sat splits: 1
% 
% Memory used [KB]: 17526
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f337,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f39,f57,f61,f63,f64,f66,f59,f68,f69,f71,f72,f76,f74,f77,f81,f79,f82,f86,f87,f91,f92,f94,f96,f97,f99,f100,f102,f114,f95,f127,f129,f131,f132,f134,f143,f130,f156,f158,f160,f161,f163,f166,f167,f182,f184,f186,f187,f189,f192,f199,f193,f201,f216,f218,f220,f176,f172,f237,f239,f247,f246,f245,f244,f249,f200,f198,f250,f83,f252,f88,f257,f259,f84,f268,f270,f277,f279,f89,f288,f290,f297,f300,f217,f303,f306,f238,f319,f318,f322,f326,f325,f331,f324,f329,f332,f222,f271,f173,f334,f333,f272,f177,f175])).
fof(f175,plain,(
  p(s(bool,'F_0')) | $spl42),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f177,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl45),
  inference(cnf_transformation,[],[f177_D])).
fof(f177_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f272,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | p(s(bool,'F_0')) | $spl38),
  inference(forward_demodulation,[],[f263,f54])).
fof(f54,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X4))),s(fun(X1,bool),X3))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3)))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6,X7] : ! [X8,X9,X11] : s(bool,'const.bool.IN_2'(s(X7,happ(s(fun(X6,X7),X8),s(X6,X11))),s(fun(X7,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9)))))),
  file('thm.util_prob.PREIMAGE_UNION','thm.util_prob.IN_PREIMAGE')).
fof(f263,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl38),
  inference(superposition,[],[f52,f162])).
fof(f162,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl38),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f52,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X9,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X9),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9)))) | p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.util_prob.PREIMAGE_UNION','thm.pred_set.IN_UNION')).
fof(f333,plain,(
  $false | $spl42),
  inference(unit_resulting_resolution,[],[f42,f175])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f15])).
fof(f15,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.PREIMAGE_UNION',aHL_FALSITY)).
fof(f334,plain,(
  $false | $spl42),
  inference(resolution,[],[f175,f42])).
fof(f173,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl41),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f271,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | p(s(bool,'F_0')) | $spl38),
  inference(forward_demodulation,[],[f262,f54])).
fof(f262,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | $spl38),
  inference(superposition,[],[f51,f162])).
fof(f51,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f222,plain,(
  ~p(s(bool,'T_0')) | $spl57),
  inference(cnf_transformation,[],[f222_D])).
fof(f222_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f332,plain,(
  ~p(s(bool,'T_0')) | ($spl36 | $spl41 | $spl45)),
  inference(forward_demodulation,[],[f323,f160])).
fof(f323,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))))) | ($spl41 | $spl45)),
  inference(unit_resulting_resolution,[],[f173,f177,f50])).
fof(f50,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f329,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)))))) | $spl67),
  inference(cnf_transformation,[],[f329_D])).
fof(f329_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f324,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f177,f177,f50])).
fof(f331,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK3)))))) | $spl69),
  inference(cnf_transformation,[],[f331_D])).
fof(f331_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK3)))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f325,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK3)))))) | ($spl41 | $spl45)),
  inference(unit_resulting_resolution,[],[f173,f177,f50])).
fof(f326,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK4),s(fun(sK1,bool),sK4)))))) | $spl45),
  inference(unit_resulting_resolution,[],[f177,f177,f50])).
fof(f322,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)))))) | $spl65),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)))))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f318,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)))))) | $spl41),
  inference(unit_resulting_resolution,[],[f173,f173,f50])).
fof(f319,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK3)))))) | $spl41),
  inference(unit_resulting_resolution,[],[f173,f173,f50])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl58),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f306,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(trivial_inequality_removal,[],[f305])).
fof(f305,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(superposition,[],[f290,f43])).
fof(f43,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.PREIMAGE_UNION',aHL_BOOL_CASES)).
fof(f303,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(unit_resulting_resolution,[],[f290,f43])).
fof(f217,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl54),
  inference(cnf_transformation,[],[f217_D])).
fof(f217_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f300,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(trivial_inequality_removal,[],[f299])).
fof(f299,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(superposition,[],[f270,f43])).
fof(f297,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f270,f43])).
fof(f290,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl63),
  inference(cnf_transformation,[],[f290_D])).
fof(f290_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f288,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl48),
  inference(global_subsumption,[],[f39,f188,f280])).
fof(f280,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl48),
  inference(superposition,[],[f49,f188])).
fof(f49,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK5(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f34])).
fof(f34,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f20])).
fof(f20,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f19])).
fof(f19,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X9,X10] : (s(fun(X6,bool),X9) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X9))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.util_prob.PREIMAGE_UNION','thm.pred_set.EXTENSION')).
fof(f188,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl48),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f89,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl22),
  inference(cnf_transformation,[],[f89_D])).
fof(f89_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f279,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(trivial_inequality_removal,[],[f278])).
fof(f278,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(superposition,[],[f189,f43])).
fof(f277,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(unit_resulting_resolution,[],[f189,f43])).
fof(f270,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl61),
  inference(cnf_transformation,[],[f270_D])).
fof(f270_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f268,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl38),
  inference(global_subsumption,[],[f39,f162,f260])).
fof(f260,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl38),
  inference(superposition,[],[f49,f162])).
fof(f84,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl18),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f259,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(trivial_inequality_removal,[],[f258])).
fof(f258,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(superposition,[],[f163,f43])).
fof(f257,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(unit_resulting_resolution,[],[f163,f43])).
fof(f88,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl23),
  inference(cnf_transformation,[],[f88_D])).
fof(f88_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f252,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl9 | $spl46)),
  inference(backward_demodulation,[],[f186,f71])).
fof(f83,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl19),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f250,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl5 | $spl36)),
  inference(backward_demodulation,[],[f160,f66])).
fof(f198,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl50),
  inference(cnf_transformation,[],[f198_D])).
fof(f198_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f200,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl52),
  inference(cnf_transformation,[],[f200_D])).
fof(f200_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f249,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | $spl22),
  inference(forward_demodulation,[],[f248,f54])).
fof(f248,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl22),
  inference(forward_demodulation,[],[f229,f54])).
fof(f229,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl22),
  inference(superposition,[],[f50,f89])).
fof(f244,plain,(
  $false | $spl57),
  inference(unit_resulting_resolution,[],[f55,f222])).
fof(f55,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f44])).
fof(f44,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.util_prob.PREIMAGE_UNION','thm.bool.REFL_CLAUSE')).
fof(f245,plain,(
  $false | $spl57),
  inference(resolution,[],[f222,f55])).
fof(f246,plain,(
  $false | $spl57),
  inference(resolution,[],[f222,f41])).
fof(f41,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.PREIMAGE_UNION','thm.bool.TRUTH')).
fof(f247,plain,(
  $false | $spl57),
  inference(resolution,[],[f222,f40])).
fof(f40,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.PREIMAGE_UNION',aHL_TRUTH)).
fof(f239,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl59),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f237,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl22),
  inference(global_subsumption,[],[f39,f89,f228])).
fof(f228,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl22),
  inference(superposition,[],[f49,f89])).
fof(f172,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl40),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f176,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl44),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f220,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | $spl18),
  inference(forward_demodulation,[],[f219,f54])).
fof(f219,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl18),
  inference(forward_demodulation,[],[f208,f54])).
fof(f208,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3)))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))) | $spl18),
  inference(superposition,[],[f50,f84])).
fof(f218,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl55),
  inference(cnf_transformation,[],[f218_D])).
fof(f218_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f216,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | $spl18),
  inference(global_subsumption,[],[f39,f84,f207])).
fof(f207,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl18),
  inference(superposition,[],[f49,f84])).
fof(f201,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl53),
  inference(cnf_transformation,[],[f201_D])).
fof(f201_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f193,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl24),
  inference(superposition,[],[f52,f90])).
fof(f90,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl24),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f199,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl51),
  inference(cnf_transformation,[],[f199_D])).
fof(f199_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f192,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl24),
  inference(superposition,[],[f51,f90])).
fof(f189,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl49),
  inference(cnf_transformation,[],[f189_D])).
fof(f189_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f187,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f90,f71])).
fof(f186,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl46),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f184,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(trivial_inequality_removal,[],[f183])).
fof(f183,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(superposition,[],[f91,f43])).
fof(f182,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(unit_resulting_resolution,[],[f91,f43])).
fof(f167,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK4)))) | $spl20),
  inference(superposition,[],[f52,f85])).
fof(f85,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl20),
  inference(cnf_transformation,[],[f85_D])).
fof(f85_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f166,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),sK3)))) | $spl20),
  inference(superposition,[],[f51,f85])).
fof(f163,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl39),
  inference(cnf_transformation,[],[f163_D])).
fof(f163_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f161,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f85,f66])).
fof(f160,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl36),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f158,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(trivial_inequality_removal,[],[f157])).
fof(f157,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(superposition,[],[f86,f43])).
fof(f156,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(unit_resulting_resolution,[],[f86,f43])).
fof(f130,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl33),
  inference(cnf_transformation,[],[f130_D])).
fof(f130_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f143,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl14),
  inference(global_subsumption,[],[f39,f79,f138])).
fof(f138,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl14),
  inference(superposition,[],[f53,f79])).
fof(f53,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK6(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK6(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.util_prob.PREIMAGE_UNION',aHL_EXT)).
fof(f134,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl35),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f132,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f80,f68])).
fof(f80,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f131,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl32),
  inference(cnf_transformation,[],[f131_D])).
fof(f131_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f129,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(trivial_inequality_removal,[],[f128])).
fof(f128,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(superposition,[],[f81,f43])).
fof(f127,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f81,f43])).
fof(f95,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl27),
  inference(cnf_transformation,[],[f95_D])).
fof(f95_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f114,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl10),
  inference(global_subsumption,[],[f39,f74,f109])).
fof(f109,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) = s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl10),
  inference(superposition,[],[f53,f74])).
fof(f102,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl31),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f100,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl29)),
  inference(forward_demodulation,[],[f99,f74])).
fof(f99,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f99_D])).
fof(f99_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f97,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f75,f63])).
fof(f75,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f75_D])).
fof(f75_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f96,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl26),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f94,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(trivial_inequality_removal,[],[f93])).
fof(f93,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(superposition,[],[f76,f43])).
fof(f92,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f76,f43])).
fof(f91,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl25),
  inference(cnf_transformation,[],[f91_D])).
fof(f91_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f87,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl9),
  inference(superposition,[],[f71,f43])).
fof(f86,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | $spl21),
  inference(cnf_transformation,[],[f86_D])).
fof(f86_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f82,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl5),
  inference(superposition,[],[f66,f43])).
fof(f79,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f79_D])).
fof(f79_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f81,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f81_D])).
fof(f81_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f77,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl7),
  inference(superposition,[],[f68,f43])).
fof(f74,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f74_D])).
fof(f74_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f76,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f76_D])).
fof(f76_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f72,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl3),
  inference(superposition,[],[f63,f43])).
fof(f71,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl9),
  inference(cnf_transformation,[],[f71_D])).
fof(f71_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f69,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(forward_demodulation,[],[f58,f54])).
fof(f58,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f57,f49])).
fof(f68,plain,(
  s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f68_D])).
fof(f68_D,plain,(
  s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f59,plain,(
  s(bool,f(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f57,f53])).
fof(f66,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl5),
  inference(cnf_transformation,[],[f66_D])).
fof(f66_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f64,plain,(
  s(bool,'const.bool.IN_2'(s(sK1,happ(s(fun(sK0,sK1),sK2),s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)))),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(forward_demodulation,[],[f60,f54])).
fof(f60,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK5('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),sK0)),s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f57,f49])).
fof(f63,plain,(
  s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f63_D])).
fof(f63_D,plain,(
  s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f61,plain,(
  s(bool,f(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) != s(bool,g(s(sK0,sK6('const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4)))),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f57,f53])).
fof(f57,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) | $spl1),
  inference(cnf_transformation,[],[f57_D])).
fof(f57_D,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f39,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),'const.pred_set.UNION_2'(s(fun(sK1,bool),sK3),s(fun(sK1,bool),sK4))))) != s(fun(sK0,bool),'const.pred_set.UNION_2'(s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK3))),s(fun(sK0,bool),'const.util_prob.PREIMAGE_2'(s(fun(sK0,sK1),sK2),s(fun(sK1,bool),sK4)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f27])).
fof(f27,plain,(
  ? [X0,X1,X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.UNION_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) != s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ~! [X0,X1,X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.UNION_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(flattening,[],[f13])).
fof(f13,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),'const.pred_set.UNION_2'(s(fun(X1,bool),X3),s(fun(X1,bool),X4))))) = s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X3))),s(fun(X0,bool),'const.util_prob.PREIMAGE_2'(s(fun(X0,X1),X2),s(fun(X1,bool),X4)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'const.pred_set.UNION_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),'const.pred_set.UNION_2'(s(fun(X7,bool),X9),s(fun(X7,bool),X10))))) = s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X9))),s(fun(X6,bool),'const.util_prob.PREIMAGE_2'(s(fun(X6,X7),X8),s(fun(X7,bool),X10)))))),
  file('thm.util_prob.PREIMAGE_UNION','thm.util_prob.PREIMAGE_UNION')).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 44
% Passive clauses: 47
% Generated clauses: 270
% Final active clauses: 24
% Final passive clauses: 10
% Input formulas: 11
% Initial clauses: 16
% 
% Trivial inequalities: 9
% Fw demodulations: 59
% Bw demodulations: 8
% Global subsumptions: 6
% 
% Simple tautologies: 4
% Equational tautologies: 3
% Forward subsumptions: 87
% 
% Binary resolution: 4
% Unit resulting resolution: 26
% Forward superposition: 32
% Backward superposition: 68
% 
% Split clauses: 12
% Split components: 26
% SAT solver clauses: 166
% SAT solver unit clauses: 8
% SAT solver binary clauses: 107
% 
% Sat splits: 12
% Sat splitting refutations: 6
% 
% Memory used [KB]: 1407
% Time elapsed: 0.021 s
% ------------------------------
% ------------------------------
% Success in time 18.073 s
