Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.quotient
% (82980)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 1998
% Generated clauses: 21093
% Final active clauses: 91
% Final passive clauses: 1713
% Input formulas: 18
% Initial clauses: 35
% 
% Duplicate literals: 1136
% Trivial inequalities: 307
% Fw subsumption resolutions: 3212
% Fw demodulations: 2476
% Bw demodulations: 42
% 
% Simple tautologies: 456
% Equational tautologies: 57
% Forward subsumptions: 9749
% 
% Binary resolution: 282
% Forward superposition: 6380
% Backward superposition: 6854
% Self superposition: 218
% Equality factoring: 162
% Equality resolution: 21
% 
% Split clauses: 1189
% Split components: 2473
% SAT solver clauses: 1854
% SAT solver unit clauses: 28
% SAT solver binary clauses: 183
% 
% Sat splits: 1189
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1851
% TWLsolver calls for satisfiability: 154
% 
% Memory used [KB]: 9850
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f996,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f986,f66,f106,f68,f108,f95,f120,f123,f126,f62,f104,f71,f110,f124,f128,f64,f100,f947,f949,f146,f148,f63,f102,f212,f215,f65,f98,f75,f112])).
fof(f112,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f75,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.quotient.quotient_28_prove_ltb','thm.bool.EQ_CLAUSES|split|3')).
fof(f98,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),X9),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3)))) | ~p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),X9),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5))))) ) | $spl0),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),X9),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3)))) | ~p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),X9),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f65,plain,(
  ( ! [X9] : (~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),X9),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3)))) | ~p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),X9),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5))))) )),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4)))) & p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5)))) & p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3)))) & ! [X9] : (~p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),X9),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5)))) | ~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),X9),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f43])).
fof(f43,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X8),s(fun(X2,X1),X6),s(fun(X1,X2),X4)))) & p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X8),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) & p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X8),s(X2,X3))),s(X2,X3)))) & ! [X9] : (~p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X9),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X9),s(X2,happ(s(fun(X1,X2),X4),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3))))))),s(X2,X3))))))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ? [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X8),s(fun(X2,X1),X6),s(fun(X1,X2),X4)))) & (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X8),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) & (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X8),s(X2,X3))),s(X2,X3)))) & ! [X9] : (~p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X9),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X9),s(X2,happ(s(fun(X1,X2),X4),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3))))))),s(X2,X3))))))))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X8),s(fun(X2,X1),X6),s(fun(X1,X2),X4)))) => (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X8),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X8),s(X2,X3))),s(X2,X3)))) => ? [X9] : (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X9),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) & p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X9),s(X2,happ(s(fun(X1,X2),X4),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3))))))),s(X2,X3))))))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X8),s(fun(X2,X1),X6),s(fun(X1,X2),X4)))) => (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X8),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X8),s(X2,X3))),s(X2,X3)))) => ? [X9] : (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X2,X0),fun(fun(X2,X0),bool)),'const.quotient.===>_2'(s(fun(X2,fun(X2,bool)),X9),s(fun(X0,fun(X0,bool)),X7))),s(fun(X2,X0),X5)))) & p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X9),s(X2,happ(s(fun(X1,X2),X4),s(X1,happ(s(fun(X2,X1),X6),s(X2,X3))))))),s(X2,X3))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X14),s(fun(X8,X7),X12),s(fun(X7,X8),X10)))) => (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X8,X6),fun(fun(X8,X6),bool)),'const.quotient.===>_2'(s(fun(X8,fun(X8,bool)),X14),s(fun(X6,fun(X6,bool)),X13))),s(fun(X8,X6),X11)))) => (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X14),s(X8,X9))),s(X8,X9)))) => ? [X15] : (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X8,X6),fun(fun(X8,X6),bool)),'const.quotient.===>_2'(s(fun(X8,fun(X8,bool)),X15),s(fun(X6,fun(X6,bool)),X13))),s(fun(X8,X6),X11)))) & p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X15),s(X8,happ(s(fun(X7,X8),X10),s(X7,happ(s(fun(X8,X7),X12),s(X8,X9))))))),s(X8,X9))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11,X12,X13,X14] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X14),s(fun(X8,X7),X12),s(fun(X7,X8),X10)))) => (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X8,X6),fun(fun(X8,X6),bool)),'const.quotient.===>_2'(s(fun(X8,fun(X8,bool)),X14),s(fun(X6,fun(X6,bool)),X13))),s(fun(X8,X6),X11)))) => (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X14),s(X8,X9))),s(X8,X9)))) => ? [X15] : (p(s(bool,'const.quotient.respects_2'(s(fun(fun(X8,X6),fun(fun(X8,X6),bool)),'const.quotient.===>_2'(s(fun(X8,fun(X8,bool)),X15),s(fun(X6,fun(X6,bool)),X13))),s(fun(X8,X6),X11)))) & p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X15),s(X8,happ(s(fun(X7,X8),X10),s(X7,happ(s(fun(X8,X7),X12),s(X8,X9))))))),s(X8,X9))))))))),
  file('thm.quotient.quotient_28_prove_ltb','thm.quotient.quotient_28_prove_ltb')).
fof(f215,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),X0),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),X0),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3)))) ) | $spl52),
  inference(cnf_transformation,[],[f215_D])).
fof(f215_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),X0),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),X0),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3)))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f212,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),X0),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),X0),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3)))) ) | ($spl0 | $spl18)),
  inference(resolution,[],[f98,f112])).
fof(f102,plain,(
  p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5)))) | $spl4),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f63,plain,(
  p(s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5))))),
  inference(cnf_transformation,[],[f46])).
fof(f148,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5))) | $spl36),
  inference(cnf_transformation,[],[f148_D])).
fof(f148_D,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f146,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.respects_2'(s(fun(fun(sK2,sK0),fun(fun(sK2,sK0),bool)),'const.quotient.===>_2'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK0,fun(sK0,bool)),sK7))),s(fun(sK2,sK0),sK5))) | ($spl4 | $spl14)),
  inference(resolution,[],[f102,f110])).
fof(f949,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3))) | $spl178),
  inference(cnf_transformation,[],[f949_D])).
fof(f949_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3))) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f947,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3))) | ($spl8 | $spl36 | $spl52)),
  inference(subsumption_resolution,[],[f938,f106])).
fof(f938,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,happ(s(fun(sK1,sK2),sK4),s(sK1,happ(s(fun(sK2,sK1),sK6),s(sK2,sK3))))))),s(sK2,sK3))) | ($spl36 | $spl52)),
  inference(superposition,[],[f215,f148])).
fof(f100,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3)))) | $spl2),
  inference(cnf_transformation,[],[f100_D])).
fof(f100_D,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f64,plain,(
  p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3))))),
  inference(cnf_transformation,[],[f46])).
fof(f128,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3))) | $spl32),
  inference(cnf_transformation,[],[f128_D])).
fof(f128_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f124,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3))) | ($spl2 | $spl14)),
  inference(resolution,[],[f110,f100])).
fof(f110,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f71,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.quotient.quotient_28_prove_ltb','thm.bool.EQ_CLAUSES|split|1')).
fof(f104,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4)))) | $spl6),
  inference(cnf_transformation,[],[f104_D])).
fof(f104_D,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4)))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f62,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4))))),
  inference(cnf_transformation,[],[f46])).
fof(f126,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4))) | $spl30),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f123,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4))) | ($spl6 | $spl14)),
  inference(resolution,[],[f110,f104])).
fof(f120,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))))))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) ) | $spl26),
  inference(cnf_transformation,[],[f120_D])).
fof(f120_D,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))))))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f95,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))))))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))))))),s(X1,X5))))))),
  inference(ennf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))))))),s(X1,X5))))))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,happ(s(fun(X0,X1),X4),s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))))))),s(X1,X5))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X8] : ! [X17,X18,X19] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X8,fun(X8,bool)),X17),s(fun(X8,X6),X18),s(fun(X6,X8),X19)))) => ! [X20] : (p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X17),s(X8,X20))),s(X8,X20)))) => p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X17),s(X8,happ(s(fun(X6,X8),X19),s(X6,happ(s(fun(X8,X6),X18),s(X8,X20))))))),s(X8,X20))))))),
  file('thm.quotient.quotient_28_prove_ltb','thm.quotient.QUOTIENT_REP_ABS')).
fof(f108,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f68,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.quotient_28_prove_ltb',aHL_FALSITY)).
fof(f106,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f66,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.quotient_28_prove_ltb',aHL_TRUTH)).
fof(f986,plain,(
  $false | ($spl8 | $spl11 | $spl26 | $spl30 | $spl32 | $spl178)),
  inference(subsumption_resolution,[],[f985,f106])).
fof(f985,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl11 | $spl26 | $spl30 | $spl32 | $spl178)),
  inference(forward_demodulation,[],[f984,f126])).
fof(f984,plain,(
  ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4)))) | ($spl8 | $spl11 | $spl26 | $spl32 | $spl178)),
  inference(subsumption_resolution,[],[f983,f106])).
fof(f983,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4)))) | ($spl11 | $spl26 | $spl32 | $spl178)),
  inference(forward_demodulation,[],[f982,f128])).
fof(f982,plain,(
  ~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4)))) | ($spl11 | $spl26 | $spl178)),
  inference(subsumption_resolution,[],[f958,f108])).
fof(f958,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK8),s(sK2,sK3))),s(sK2,sK3)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK2,fun(sK2,bool)),sK8),s(fun(sK2,sK1),sK6),s(fun(sK1,sK2),sK4)))) | ($spl26 | $spl178)),
  inference(superposition,[],[f120,f949])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 33
% Passive clauses: 91
% Generated clauses: 905
% Final active clauses: 30
% Final passive clauses: 58
% Input formulas: 18
% Initial clauses: 35
% 
% Duplicate literals: 76
% Trivial inequalities: 20
% Fw subsumption resolutions: 45
% Fw demodulations: 59
% Bw demodulations: 3
% Condensations: 94
% Global subsumptions: 4
% 
% Simple tautologies: 54
% Equational tautologies: 15
% Forward subsumptions: 292
% 
% Binary resolution: 18
% Forward superposition: 161
% Backward superposition: 246
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 5
% 
% SAT solver clauses: 197
% SAT solver unit clauses: 44
% SAT solver binary clauses: 20
% SAT solver learnt clauses: 9
% 
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 180
% TWLsolver calls for satisfiability: 119
% 
% Memory used [KB]: 2046
% Time elapsed: 0.043 s
% ------------------------------
% ------------------------------
% Success in time 0.657 s
