Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (86268)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 721
% Passive clauses: 13288
% Generated clauses: 23690
% Final active clauses: 634
% Final passive clauses: 11994
% Input formulas: 69
% Initial clauses: 247
% 
% Duplicate literals: 282
% Trivial inequalities: 136
% Fw subsumption resolutions: 1904
% Fw demodulations: 316
% Bw demodulations: 86
% 
% Simple tautologies: 258
% Equational tautologies: 69
% Forward subsumptions: 7224
% 
% Binary resolution: 13239
% Factoring: 1006
% Forward superposition: 2945
% Backward superposition: 3223
% Self superposition: 42
% Equality factoring: 26
% Equality resolution: 33
% 
% Split clauses: 344
% Split components: 688
% SAT solver clauses: 698
% SAT solver unit clauses: 20
% SAT solver binary clauses: 180
% 
% Sat splits: 344
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 696
% TWLsolver calls for satisfiability: 385
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (86303)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 449
% Passive clauses: 3334
% Generated clauses: 14044
% Final active clauses: 352
% Final passive clauses: 1681
% Input formulas: 69
% Initial clauses: 247
% 
% Duplicate literals: 207
% Trivial inequalities: 27
% Fw subsumption resolutions: 1364
% Fw demodulations: 774
% Bw demodulations: 1231
% Condensations: 267
% Global subsumptions: 15
% 
% Simple tautologies: 301
% Equational tautologies: 22
% Forward subsumptions: 4875
% Bw demodulations to eq. taut.: 44
% 
% Binary resolution: 2892
% Factoring: 70
% Forward superposition: 1351
% Backward superposition: 2184
% Self superposition: 43
% Equality factoring: 30
% Equality resolution: 7
% 
% Split clauses: 155
% Split components: 311
% SAT solver clauses: 2763
% SAT solver unit clauses: 215
% SAT solver binary clauses: 83
% SAT solver learnt clauses: 30
% SAT solver learnt literals: 10
% 
% Sat splits: 155
% Sat splitting refutations: 12
% 
% TWLsolver clauses: 2703
% TWLsolver calls for satisfiability: 1874
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (86330)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1671
% Passive clauses: 17605
% Generated clauses: 88407
% Final active clauses: 1274
% Final passive clauses: 3893
% Input formulas: 69
% Initial clauses: 247
% Discarded non-redundant clauses: 15854
% 
% Duplicate literals: 318
% Trivial inequalities: 101
% Fw subsumption resolutions: 3096
% Fw demodulations: 11067
% Bw demodulations: 1842
% Condensations: 980
% Global subsumptions: 2
% 
% Simple tautologies: 200
% Equational tautologies: 453
% Forward subsumptions: 45492
% Bw demodulations to eq. taut.: 28
% 
% Unit resulting resolution: 41853
% Factoring: 5
% Forward superposition: 14692
% Backward superposition: 12635
% Self superposition: 129
% Equality factoring: 34
% Equality resolution: 17
% 
% Split clauses: 75
% Split components: 155
% SAT solver clauses: 19832
% SAT solver unit clauses: 1456
% SAT solver binary clauses: 9722
% 
% Sat splits: 75
% Sat splitting refutations: 2072
% 
% Memory used [KB]: 93900
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f120873,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f688,f689,f690,f691,f692,f693,f694,f695,f696,f697,f698,f699,f1075,f1079,f1063,f1064,f1134,f1136,f1137,f1138,f2669,f2664,f2750,f2751,f2752,f2753,f2928,f2944,f2946,f3061,f3056,f3141,f3142,f3143,f3144,f3239,f3238,f3237,f3345,f3346,f3347,f3348,f3349,f3378,f3406,f3407,f3408,f3580,f3640,f3686,f3731,f4372,f4386,f4406,f4419,f4422,f4429,f4430,f4421,f4437,f4440,f4456,f2929,f4470,f4471,f4439,f6046,f6165,f6168,f6158,f6157,f6156,f6155,f6298,f6300,f6291,f6290,f6289,f6288,f8473,f8476,f8474,f10720,f10723,f10732,f10722,f10733,f10724,f11201,f11217,f11203,f12602,f13858,f13860,f13921,f13922,f14137,f14636,f14638,f14637,f14696,f14697,f14698,f14699,f14700,f14701,f14702,f14703,f14704,f14839,f14841,f14840,f15211,f15213,f15212,f15202,f15201,f15200,f15199,f15322,f15324,f15323,f15355,f15357,f15356,f15605,f15621,f15607,f16255,f16257,f16256,f16246,f16245,f16244,f16243,f16468,f16470,f16469,f16544,f16545,f16546,f16547,f16549,f16550,f16551,f16552,f16553,f16733,f16735,f16734,f16738,f16739,f16740,f16741,f16744,f16745,f16746,f16747,f17013,f17014,f17015,f17016,f17020,f17021,f17022,f17288,f17289,f17290,f17291,f17296,f17297,f17616,f17653,f17655,f17656,f17657,f17658,f17659,f18669,f18668,f18684,f18685,f20078,f20079,f20080,f20081,f20083,f20084,f20085,f20086,f20087,f20272,f20273,f20274,f20275,f20278,f20279,f20280,f20281,f20472,f20473,f20474,f20475,f20479,f20480,f20481,f20701,f20702,f20703,f20704,f20709,f20710,f25429,f25872,f25875,f25877,f25879,f25881,f25883,f25885,f25887,f25889,f25593,f25893,f25894,f25895,f25896,f25897,f25898,f25899,f25900,f25902,f25904,f25907,f25981,f25982,f25983,f25984,f25986,f16671,f25987,f25988,f25989,f25990,f26113,f26114,f26115,f26116,f26119,f26120,f26121,f26122,f26245,f26246,f26247,f26248,f26252,f26253,f26254,f26377,f26378,f26379,f26380,f26385,f26386,f27731,f27909,f27910,f27936,f27941,f28061,f28062,f28066,f28067,f28935,f28966,f28936,f28969,f29756,f29828,f29869,f29829,f29870,f29838,f29839,f29840,f29841,f29842,f29843,f29844,f29845,f29848,f29888,f30075,f28965,f30076,f28968,f30079,f30080,f30081,f30082,f30087,f30088,f30089,f30090,f30094,f33372,f33373,f33374,f33375,f33376,f33384,f33386,f33483,f33481,f33488,f33489,f33470,f33469,f33478,f34839,f34840,f34841,f34842,f34843,f34953,f34954,f34938,f34937,f34946,f36451,f36452,f36453,f36454,f36455,f36566,f36567,f36551,f36550,f36559,f38266,f38267,f38268,f38269,f38270,f38382,f38383,f38367,f38366,f38375,f39524,f39525,f39526,f39527,f39532,f39533,f39534,f39535,f39539,f39565,f39583,f39584,f39585,f39586,f39590,f39859,f39889,f39890,f39892,f39906,f39957,f39974,f44473,f44474,f44475,f44476,f44477,f44592,f44593,f44577,f44576,f44585,f46076,f46611,f46612,f46613,f46614,f46615,f46731,f46732,f46716,f46715,f46724,f47672,f47678,f48016,f48020,f48321,f48699,f48700,f48701,f48702,f48703,f48820,f48821,f48805,f48804,f48813,f50415,f50416,f50417,f50418,f50419,f50537,f50538,f50522,f50521,f50530,f52012,f52013,f52014,f52035,f52036,f52037,f52047,f52048,f52049,f52057,f52058,f52059,f52075,f52076,f52077,f52088,f52090,f52103,f52104,f52105,f52138,f52139,f52140,f52151,f52152,f52153,f52158,f52171,f52168,f52169,f52020,f52170,f52022,f52181,f52182,f52183,f52204,f52205,f52206,f52216,f52217,f52218,f52226,f52227,f52228,f52244,f52245,f52246,f52257,f52259,f52272,f52273,f52274,f52307,f52308,f52309,f52320,f52321,f52322,f52327,f52340,f52337,f52338,f52189,f52339,f52191,f52350,f52351,f52352,f52500,f52501,f52502,f52512,f52513,f52514,f52522,f52523,f52524,f52540,f52541,f52542,f52553,f52555,f52568,f52569,f52570,f52603,f52604,f52605,f52616,f52617,f52618,f52623,f52636,f52633,f52634,f52358,f52635,f52360,f52670,f52671,f52672,f52693,f52694,f52695,f52705,f52706,f52707,f52716,f52717,f52718,f52735,f52736,f52737,f52748,f52750,f52763,f52764,f52765,f52798,f52799,f52800,f52811,f52812,f52813,f52818,f52831,f52828,f52829,f52678,f52830,f52680,f53038,f53040,f53042,f53044,f53182,f53184,f53186,f53188,f53558,f53559,f53560,f53561,f53836,f53838,f53840,f53842,f53888,f53890,f53892,f53894,f54473,f54475,f54477,f54479,f54525,f54527,f54529,f54531,f54615,f54617,f54619,f54621,f54687,f53082,f54689,f53084,f54691,f53086,f54693,f53088,f55203,f55205,f55207,f55209,f55279,f55281,f55283,f55285,f55659,f55660,f55661,f55662,f55994,f55996,f55998,f56000,f56046,f56048,f56050,f56052,f56772,f56774,f56776,f56778,f57010,f57069,f57070,f57071,f57072,f57073,f57074,f57081,f57083,f57085,f57087,f57171,f57173,f57175,f57177,f57243,f55247,f57245,f55249,f57247,f55251,f57249,f55253,f58264,f58263,f58280,f58281,f59363,f59364,f59366,f59392,f59592,f59614,f60170,f60171,f60173,f60199,f60315,f25430,f60383,f60405,f61088,f61089,f61090,f61091,f61092,f61122,f61123,f61124,f61125,f61126,f61138,f61139,f61140,f61141,f61142,f61151,f61152,f61153,f61154,f61155,f61176,f61177,f61178,f61179,f61180,f61191,f61193,f61207,f61208,f61209,f61210,f61211,f61249,f61250,f61251,f61252,f61253,f61265,f61266,f61267,f61268,f61269,f61276,f61294,f61289,f61290,f61100,f61291,f61102,f61292,f61104,f61293,f61106,f61480,f61481,f61482,f61483,f61484,f61514,f61515,f61516,f61517,f61518,f61530,f61531,f61532,f61533,f61534,f61543,f61544,f61545,f61546,f61547,f61702,f61703,f61704,f61705,f61706,f61717,f61719,f61733,f61734,f61735,f61736,f61737,f61775,f61776,f61777,f61778,f61779,f61791,f61792,f61793,f61794,f61795,f61802,f61820,f61815,f61816,f61492,f61817,f61494,f61818,f61496,f61819,f61498,f61831,f61832,f61833,f61834,f61835,f61865,f61866,f61867,f61868,f61869,f61881,f61882,f61883,f61884,f61885,f61894,f61895,f61896,f61897,f61898,f61919,f61920,f61921,f61922,f61923,f61934,f61936,f61950,f61951,f61952,f61953,f61954,f61992,f61993,f61994,f61995,f61996,f62008,f62009,f62010,f62011,f62012,f62130,f62148,f62166,f62161,f62162,f61843,f62163,f61845,f62164,f61847,f62165,f61849,f62177,f62178,f62179,f62180,f62181,f62211,f62212,f62213,f62214,f62215,f62317,f62320,f62318,f62332,f62131,f62358,f62360,f62362,f62364,f62366,f62368,f62370,f62518,f62519,f62520,f62521,f62522,f62531,f62532,f62533,f62534,f62535,f62556,f62557,f62558,f62559,f62560,f62644,f62694,f62696,f62710,f62711,f62712,f62713,f62714,f62752,f62753,f62754,f62755,f62756,f62768,f62769,f62770,f62771,f62772,f62780,f62799,f62793,f62794,f62795,f62189,f62796,f62191,f62797,f62193,f62798,f62195,f62884,f62645,f63581,f64307,f63582,f64955,f52110,f64957,f52279,f64959,f52575,f64961,f52770,f64963,f61218,f64965,f61744,f64967,f61961,f64969,f62721,f65076,f65077,f65079,f65105,f65139,f52024,f65141,f52193,f65143,f52362,f65145,f52682,f65147,f61108,f65149,f61500,f65151,f61851,f65153,f62197,f65159,f39905,f65175,f30121,f65176,f65177,f59391,f65178,f60198,f65254,f65276,f66085,f66086,f66088,f66115,f66095,f66233,f66255,f66867,f66870,f62455,f66873,f66685,f66537,f66534,f66531,f65789,f65641,f65638,f65635,f57285,f57213,f57212,f54729,f54657,f54656,f66884,f66885,f39605,f39604,f66889,f66890,f32694,f32695,f66891,f66892,f66893,f66894,f31792,f31539,f66895,f66896,f66897,f66898,f30111,f30110,f29779,f29849,f29146,f29026,f28973,f66902,f28972,f66903,f28538,f66904,f66905,f66907,f66908,f66909,f66906,f28460,f66910,f28459,f66911,f66912,f66914,f66913,f66915,f66917,f66918,f66919,f28177,f28125,f27960,f25379,f66927,f25378,f66928,f66929,f66931,f66932,f66933,f66930,f25064,f25065,f706,f25066,f25067,f25088,f25086,f1076,f25050,f25049,f66886,f25036,f66934,f25075,f25073,f66937,f66938,f66916,f66939,f66940,f66961,f66962,f66963,f66964,f7282,f66965,f66966,f66967,f4525,f4526,f4527,f4528,f66969,f66970,f66971,f66972,f1081,f66975,f62479,f62477,f62462,f62464,f62466,f62472,f66977,f62356,f62399,f62353,f31701,f28213,f25277,f25278,f25279,f25280,f25325,f25323,f66986,f25300,f25296,f25294,f15214,f25292,f25290,f710,f25288,f25286,f66989,f66992,f66993,f66994,f66995,f67040,f67041,f704,f67042,f67043,f705,f67044,f67045,f703,f67046,f67047,f702,f67048,f67049,f67050,f709,f711,f67051,f707,f713,f67214,f67215,f1077,f67277,f67293,f67282,f712,f67815,f67828,f68672,f68671,f68635,f68634,f68633,f68747,f68748,f68749,f68750,f68751,f68790,f69031,f68791,f69061,f69072,f69125,f69126,f69127,f69238,f69239,f69938,f69941,f69942,f69943,f69944,f70443,f70446,f70447,f70448,f70449,f70661,f70666,f70773,f70774,f70766,f70765,f70764,f70763,f6170,f70798,f70834,f70844,f70845,f70837,f70880,f70882,f70881,f71092,f71094,f71093,f71113,f71118,f71171,f70811,f71265,f71270,f71326,f71364,f71372,f71373,f71367,f71374,f71382,f71383,f71377,f70843,f71384,f71394,f71395,f71387,f71581,f71582,f6299,f71574,f71573,f6169,f71572,f71571,f6301,f71689,f71691,f71690,f71696,f71701,f71754,f71339,f71799,f71801,f71800,f72172,f72175,f72176,f72177,f72178,f72466,f72487,f72495,f72496,f72490,f72497,f72505,f72506,f72500,f71393,f73125,f73145,f73128,f73129,f73130,f73146,f73131,f73147,f73186,f73197,f73281,f72192,f73284,f73285,f73286,f72193,f73287,f72194,f73358,f73369,f73913,f73933,f73916,f73917,f73918,f73934,f73919,f73935,f73974,f73985,f74048,f74068,f74051,f74052,f74053,f74069,f74054,f74070,f74138,f74149,f74157,f74158,f74189,f73144,f74190,f72191,f74191,f73932,f74192,f74067,f74193,f66114,f74194,f65104,f75148,f75149,f75150,f75151,f75270,f75281,f75282,f75268,f75283,f75284,f75402,f75271,f77824,f78101,f78104,f78105,f78106,f78107,f78158,f78164,f78350,f70679,f78355,f78356,f70678,f78374,f78401,f25906,f78666,f69958,f78669,f78670,f78671,f69959,f78672,f69960,f78778,f78170,f78784,f69957,f79709,f80040,f80043,f80044,f80045,f80046,f80101,f80108,f80270,f71283,f80275,f80276,f71282,f80295,f80323,f25909,f80570,f70463,f80573,f80574,f80575,f70464,f80576,f70465,f80660,f80115,f80667,f70462,f81554,f81555,f83866,f83883,f83868,f83898,f83915,f83900,f87840,f88754,f88755,f88756,f88757,f88758,f88759,f88760,f88761,f88762,f88472,f88765,f25874,f88766,f25876,f88767,f25880,f88768,f25878,f88769,f25882,f88770,f25884,f88771,f25886,f88772,f25888,f88774,f25873,f25890,f89464,f89468,f89579,f89581,f89600,f89603,f89605,f89466,f89660,f89662,f89465,f89960,f89965,f89966,f90076,f89661,f89659,f90093,f90627,f90628,f90629,f90630,f90640,f90648,f90649,f90666,f90674,f90675,f90092,f90902,f90924,f90905,f90906,f90907,f90925,f90908,f90926,f90985,f90997,f90991,f90923,f91080,f91102,f91083,f91084,f91085,f91103,f91086,f91104,f91130,f91131,f6058,f91171,f89601,f91177,f91101,f91409,f89981,f91414,f91415,f89980,f91493,f90089,f91508,f102516,f102518,f102525,f102517,f102526,f87841,f102728,f102731,f102732,f102733,f102734,f102735,f102736,f102737,f102738,f102739,f102740,f102860,f102867,f102868,f102869,f102870,f102871,f102872,f102873,f102874,f102875,f102876,f103043,f103044,f103045,f103046,f103047,f103048,f103049,f103050,f103051,f103052,f103250,f103251,f103252,f103253,f103254,f102855,f103255,f103256,f103257,f103258,f103259,f103448,f103449,f103450,f103451,f103453,f20205,f103454,f103455,f103456,f103457,f103655,f103656,f103657,f103658,f103661,f20206,f103662,f103663,f103664,f103873,f103874,f103875,f103876,f103880,f20207,f103881,f103882,f104091,f20201,f104092,f20202,f104093,f20203,f104094,f20204,f104099,f20208,f104100,f20209,f105759,f106057,f106082,f106060,f106061,f106062,f106083,f106063,f106084,f106245,f106260,f106252,f106081,f106457,f106482,f106460,f106461,f106462,f106483,f106463,f106484,f106540,f106556,f106547,f106481,f106704,f106705,f106706,f106707,f106708,f106710,f16672,f106711,f106712,f106713,f106878,f106881,f106922,f106923,f106924,f106925,f106926,f106929,f16673,f106930,f106931,f107094,f107095,f107134,f107135,f107136,f107137,f107138,f106828,f107142,f107143,f107307,f12603,f107417,f107418,f107420,f107435,f107481,f107495,f107488,f107434,f107570,f107571,f107573,f107588,f107634,f107648,f107641,f107587,f107723,f107724,f107726,f107741,f107787,f107801,f107794,f107740,f107876,f107877,f107879,f39559,f107938,f29883,f107945,f30114,f108026,f108027,f108029,f39560,f108088,f29884,f108095,f30115,f108303,f108304,f108306,f39561,f108365,f29885,f108372,f30116,f108453,f108454,f108456,f39562,f108535,f29886,f108542,f30117,f110578,f110641,f110642,f17654,f110643,f110644,f110645,f110646,f6166,f110651,f16667,f110652,f16668,f110653,f16669,f110654,f16670,f110659,f16674,f110660,f16675,f112886,f112887,f112889,f39552,f113140,f113139,f113154,f113155,f18683,f113184,f29880,f113191,f30107,f113274,f113275,f113277,f39553,f113336,f29881,f113343,f30108,f113552,f113553,f113555,f39554,f113614,f29882,f113621,f30109,f114409,f114744,f114747,f114745,f114756,f114410,f115026,f115090,f115027,f117208,f117547,f117209,f10721,f118302,f118303,f6167,f118305,f39551,f118491,f29879,f118498,f30106,f119363,f8475,f119364,f119369,f119406,f1135,f120450,f120493,f1078,f120864,f120866])).
fof(f120866,plain,(
  $false | ($spl1 | $spl8 | $spl568)),
  inference(subsumption_resolution,[],[f120865,f582])).
fof(f582,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.EXISTS_CONG',aHL_TRUTH)).
fof(f120865,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl8 | $spl568)),
  inference(forward_demodulation,[],[f120747,f83900])).
fof(f120747,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12))))) | ($spl1 | $spl8)),
  inference(resolution,[],[f67054,f67586])).
fof(f67586,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl8),
  inference(subsumption_resolution,[],[f67317,f595])).
fof(f595,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f523])).
fof(f523,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : ! [X11] : (s(X6,X11) = s(X6,X11) <=> p(s(bool,'T_0')))),
  file('thm.list.EXISTS_CONG','thm.bool.REFL_CLAUSE')).
fof(f67317,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl8),
  inference(superposition,[],[f354,f708])).
fof(f708,plain,(
  s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl8),
  inference(cnf_transformation,[],[f708_D])).
fof(f708_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f354,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f237])).
fof(f237,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f77])).
fof(f77,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X9,X14] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X14)))) <=> ? [X15] : (p(s(bool,'const.bool.IN_2'(s(X6,X15),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X14)))))) & p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X15))))))),
  file('thm.list.EXISTS_CONG','thm.list.EXISTS_MEM')).
fof(f67054,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))))) ) | $spl1),
  inference(resolution,[],[f702,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f120864,plain,(
  $false | ($spl1 | $spl8 | $spl566)),
  inference(subsumption_resolution,[],[f120863,f582])).
fof(f120863,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl8 | $spl566)),
  inference(forward_demodulation,[],[f120746,f83868])).
fof(f120746,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12))))) | ($spl1 | $spl8)),
  inference(resolution,[],[f67054,f67584])).
fof(f67584,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl8),
  inference(subsumption_resolution,[],[f67315,f595])).
fof(f67315,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl8),
  inference(superposition,[],[f655,f708])).
fof(f655,plain,(
  ( ! [X13] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X13),s('type.list.list'(sK12),sK13)))) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,X13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(forward_demodulation,[],[f607,f347])).
fof(f347,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14)),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  s('type.list.list'(sK12),sK13) = s('type.list.list'(sK12),sK14) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X5))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X5)))) & s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK14)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14,sK15,sK16])],[f176])).
fof(f176,plain,(
  ? [X0,X1,X2,X3,X4] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5)))) & s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) != s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f175])).
fof(f175,plain,(
  ? [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) & s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) != s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3,X4] : ((s('type.list.list'(X0),X1) = s('type.list.list'(X0),X2) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => s(bool,happ(s(fun(X0,bool),X3),s(X0,X5))) = s(bool,happ(s(fun(X0,bool),X4),s(X0,X5))))) => s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X3),s('type.list.list'(X0),X1))) = s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X4),s('type.list.list'(X0),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9,X10] : ((s('type.list.list'(X6),X7) = s('type.list.list'(X6),X8) & ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) => s(bool,happ(s(fun(X6,bool),X9),s(X6,X11))) = s(bool,happ(s(fun(X6,bool),X10),s(X6,X11))))) => s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7))) = s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X10),s('type.list.list'(X6),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9,X10] : ((s('type.list.list'(X6),X7) = s('type.list.list'(X6),X8) & ! [X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) => s(bool,happ(s(fun(X6,bool),X9),s(X6,X11))) = s(bool,happ(s(fun(X6,bool),X10),s(X6,X11))))) => s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X7))) = s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X10),s('type.list.list'(X6),X8))))),
  file('thm.list.EXISTS_CONG','thm.list.EXISTS_CONG')).
fof(f607,plain,(
  ( ! [X13] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,X13,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X13),s('type.list.list'(sK12),sK14))))) )),
  inference(superposition,[],[f354,f347])).
fof(f1078,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | $spl16),
  inference(cnf_transformation,[],[f1078_D])).
fof(f1078_D,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f120493,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(sK12,bool),X3),s(sK12,sK17(sK14,X3,sK12))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X3),s('type.list.list'(sK12),sK13)))) ) | $spl656),
  inference(cnf_transformation,[],[f120493_D])).
fof(f120493_D,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun(sK12,bool),X3),s(sK12,sK17(sK14,X3,sK12))))) | ~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X3),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl656),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl656])])).
fof(f120450,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun(sK12,bool),X3),s(sK12,sK17(sK14,X3,sK12))))) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X3),s('type.list.list'(sK12),sK13)))) ) | $spl22),
  inference(resolution,[],[f4847,f29192])).
fof(f29192,plain,(
  ( ! [X4,X5,X3] : (~sP7(X4,X3,X5) | ~p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f29149])).
fof(f29149,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~sP7(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(resolution,[],[f29056,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f106,f228,f227,f226,f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f226,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f227,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f228,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f106,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X26,X24,X25] : ((p(s(bool,X25)) <=> s(bool,X24) = s(bool,X26)) <=> ((p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,X26))) & (p(s(bool,X25)) | ~p(s(bool,X26)) | ~p(s(bool,X24))) & (p(s(bool,X24)) | ~p(s(bool,X26)) | ~p(s(bool,X25))) & (p(s(bool,X26)) | ~p(s(bool,X24)) | ~p(s(bool,X25)))))),
  file('thm.list.EXISTS_CONG','thm.sat.sat_2|49|_prove_ltb')).
fof(f29056,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X8,X6) | ~p(s(bool,X6)) | ~sP7(X8,X6,X7) | p(s(bool,X7))) ) | $spl22),
  inference(resolution,[],[f29052,f455])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f286])).
fof(f286,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f285])).
fof(f285,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f227])).
fof(f29052,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X3,X5) | ~p(s(bool,X3)) | sP9(X4,X5,X3) | ~sP7(X5,X3,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f29044,f562])).
fof(f562,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.EQ_CLAUSES|split|1')).
fof(f29044,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | sP9(X4,X5,X3) | ~p(s(bool,X3)) | ~sP8(X4,X3,X5) | ~sP7(X5,X3,X4)) ) | $spl22),
  inference(resolution,[],[f29028,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f283])).
fof(f283,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f228])).
fof(f29028,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'T_0') != s(bool,X7)) ) | $spl22),
  inference(subsumption_resolution,[],[f29027,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f225])).
fof(f29027,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl22),
  inference(subsumption_resolution,[],[f28982,f462])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f292])).
fof(f28982,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl22),
  inference(resolution,[],[f28891,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f284])).
fof(f28891,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X8) | p(s(bool,X6))) ) | $spl22),
  inference(resolution,[],[f28875,f455])).
fof(f28875,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,X6,X7) | sP9(X8,X7,X6) | p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl22),
  inference(subsumption_resolution,[],[f28874,f459])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f289])).
fof(f289,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f226])).
fof(f28874,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | sP9(X8,X7,X6) | p(s(bool,X7)) | ~sP8(X8,X6,X7) | ~sP7(X7,X6,X8)) ) | $spl22),
  inference(subsumption_resolution,[],[f28864,f558])).
fof(f558,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.EQ_CLAUSES|split|2')).
fof(f28864,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,'F_0') != s(bool,X7) | sP9(X8,X7,X6) | p(s(bool,X7)) | ~sP8(X8,X6,X7) | ~sP7(X7,X6,X8)) ) | $spl22),
  inference(resolution,[],[f28844,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f28844,plain,(
  ( ! [X6,X8,X7] : (sP6(X7,X6,X8) | s(bool,'T_0') != s(bool,X7) | s(bool,'F_0') != s(bool,X6)) ) | $spl22),
  inference(subsumption_resolution,[],[f28802,f461])).
fof(f28802,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') != s(bool,X7) | p(s(bool,X8)) | sP6(X7,X6,X8)) ) | $spl22),
  inference(resolution,[],[f4491,f445])).
fof(f4491,plain,(
  ( ! [X4,X5,X3] : (sP9(X5,X3,X4) | s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl22),
  inference(superposition,[],[f2929,f464])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f4847,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,happ(s(fun(sK12,bool),X163),s(sK12,sK17(sK14,X163,sK12))))) | sP7('const.list.EXISTS_2'(s(fun(sK12,bool),X163),s('type.list.list'(sK12),sK13)),X164,X165)) )),
  inference(resolution,[],[f609,f459])).
fof(f609,plain,(
  ( ! [X15] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X15),s('type.list.list'(sK12),sK13)))) | p(s(bool,happ(s(fun(sK12,bool),X15),s(sK12,sK17(sK14,X15,sK12)))))) )),
  inference(superposition,[],[f355,f347])).
fof(f355,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK17(X2,X1,X0))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f239])).
fof(f1135,plain,(
  ( ! [X10] : (~p(s(bool,X10))) ) | $spl20),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  ( ! [X10] : (~p(s(bool,X10))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f119406,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK12,bool),X29),s(sK12,sK17(sK14,X29,sK12))))) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X29),s('type.list.list'(sK12),sK13)))) ) | $spl654),
  inference(cnf_transformation,[],[f119406_D])).
fof(f119406_D,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK12,bool),X29),s(sK12,sK17(sK14,X29,sK12))))) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X29),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl654),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl654])])).
fof(f119369,plain,(
  ( ! [X45,X44] : (p(s(bool,happ(s(fun(sK12,bool),X44),s(sK12,sK17(sK14,X44,sK12))))) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X44),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X45))) ) | ($spl8 | $spl10)),
  inference(resolution,[],[f4843,f69231])).
fof(f69231,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f69230,f528])).
fof(f528,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f319])).
fof(f319,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f318])).
fof(f318,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.IMP_CLAUSES|split|4')).
fof(f69230,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP8(X4,X5,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f69229,f711])).
fof(f69229,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP8(X4,X5,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl8),
  inference(subsumption_resolution,[],[f69228,f461])).
fof(f69228,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP8(X4,X5,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP6(X5,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f69188,f459])).
fof(f69188,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP8(X4,X5,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X5,X4) | ~sP6(X5,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X4)) ) | $spl8),
  inference(resolution,[],[f67779,f450])).
fof(f67779,plain,(
  ( ! [X6,X7] : (~sP9(X7,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X6) | ~p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f67295,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f293])).
fof(f67295,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | $spl8),
  inference(backward_demodulation,[],[f708,f679])).
fof(f679,plain,(
  s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(forward_demodulation,[],[f349,f347])).
fof(f349,plain,(
  s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK14)))),
  inference(cnf_transformation,[],[f234])).
fof(f4843,plain,(
  ( ! [X152,X151,X153] : (p(s(bool,happ(s(fun(sK12,bool),X151),s(sK12,sK17(sK14,X151,sK12))))) | sP8(X152,'const.list.EXISTS_2'(s(fun(sK12,bool),X151),s('type.list.list'(sK12),sK13)),X153)) )),
  inference(resolution,[],[f609,f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f287])).
fof(f119364,plain,(
  ( ! [X33,X32] : (p(s(bool,happ(s(fun(sK12,bool),X32),s(sK12,sK17(sK14,X32,sK12))))) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X32),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X33))) ) | $spl19),
  inference(resolution,[],[f4843,f73790])).
fof(f73790,plain,(
  ( ! [X10,X9] : (~sP8(X9,X10,'F_0') | s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X9))) ) | $spl19),
  inference(subsumption_resolution,[],[f73789,f461])).
fof(f73789,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X10) | ~sP8(X9,X10,'F_0') | ~sP6(X10,'F_0',X9)) ) | $spl19),
  inference(subsumption_resolution,[],[f73788,f528])).
fof(f73788,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X10) | p(s(bool,'F_0')) | ~sP8(X9,X10,'F_0') | ~sP6(X10,'F_0',X9)) ) | $spl19),
  inference(subsumption_resolution,[],[f73744,f459])).
fof(f73744,plain,(
  ( ! [X10,X9] : (~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X10) | p(s(bool,'F_0')) | ~sP8(X9,X10,'F_0') | ~sP7('F_0',X10,X9) | ~sP6(X10,'F_0',X9)) ) | $spl19),
  inference(resolution,[],[f67015,f450])).
fof(f67015,plain,(
  ( ! [X10,X11] : (~sP9(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl19),
  inference(superposition,[],[f1081,f466])).
fof(f8475,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17)) ) | $spl44),
  inference(cnf_transformation,[],[f8475_D])).
fof(f8475_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f119363,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,happ(s(fun(sK12,bool),X29),s(sK12,sK17(sK14,X29,sK12))))) | ~p(s(bool,X30)) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X29),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X31)) ) | $spl22),
  inference(resolution,[],[f4843,f30061])).
fof(f30061,plain,(
  ( ! [X4,X5,X3] : (~sP8(X4,X5,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X3)) ) | $spl22),
  inference(subsumption_resolution,[],[f30060,f461])).
fof(f30060,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP8(X4,X5,X3) | ~sP6(X5,X3,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f30059,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f290])).
fof(f30059,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | ~sP8(X4,X5,X3) | ~sP7(X3,X5,X4) | ~sP6(X5,X3,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f29972,f557])).
fof(f557,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f338])).
fof(f29972,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X3)) | ~sP8(X4,X5,X3) | ~sP7(X3,X5,X4) | ~sP6(X5,X3,X4)) ) | $spl22),
  inference(resolution,[],[f4493,f450])).
fof(f4493,plain,(
  ( ! [X10,X11,X9] : (~sP9(X11,X9,X10) | s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | $spl22),
  inference(superposition,[],[f2929,f466])).
fof(f30106,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X22,sK16,sK12)))) | ~p(s(bool,X22))) ) | $spl180),
  inference(cnf_transformation,[],[f30106_D])).
fof(f30106_D,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X22,sK16,sK12)))) | ~p(s(bool,X22))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f118498,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X20,sK16,sK12))))) ) | ($spl19 | $spl22 | $spl566)),
  inference(resolution,[],[f110723,f73680])).
fof(f73680,plain,(
  ( ! [X12,X11] : (~sP6(X12,'F_0',X11) | p(s(bool,X11)) | s(bool,'F_0') != s(bool,X12)) ) | ($spl19 | $spl22)),
  inference(duplicate_literal_removal,[],[f73658])).
fof(f73658,plain,(
  ( ! [X12,X11] : (p(s(bool,X11)) | ~sP6(X12,'F_0',X11) | p(s(bool,X11)) | s(bool,'F_0') != s(bool,X12)) ) | ($spl19 | $spl22)),
  inference(resolution,[],[f73649,f467])).
fof(f73649,plain,(
  ( ! [X8,X9] : (sP9(X8,'F_0',X9) | p(s(bool,X8)) | ~sP6(X9,'F_0',X8)) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f73644,f455])).
fof(f73644,plain,(
  ( ! [X8,X9] : (sP9(X8,'F_0',X9) | p(s(bool,X8)) | ~sP8(X8,X9,'F_0') | ~sP6(X9,'F_0',X8)) ) | ($spl19 | $spl22)),
  inference(resolution,[],[f73631,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X0)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f284])).
fof(f73631,plain,(
  ( ! [X4,X3] : (sP7('F_0',X3,X4)) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f73630,f458])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f290])).
fof(f73630,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | sP7('F_0',X3,X4)) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f73600,f457])).
fof(f73600,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | sP7('F_0',X3,X4)) ) | ($spl19 | $spl22)),
  inference(resolution,[],[f73586,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f284])).
fof(f73586,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | ~p(s(bool,X6)) | p(s(bool,X7))) ) | ($spl19 | $spl22)),
  inference(resolution,[],[f73579,f455])).
fof(f73579,plain,(
  ( ! [X4,X5] : (~sP8(X5,X4,'F_0') | ~p(s(bool,X4)) | sP9(X5,'F_0',X4)) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f73578,f562])).
fof(f73578,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,'F_0')) ) | ($spl19 | $spl22)),
  inference(subsumption_resolution,[],[f73571,f29028])).
fof(f73571,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP9(X5,'F_0',X4) | ~p(s(bool,X4)) | ~sP8(X5,X4,'F_0') | ~sP6(X4,'F_0',X5)) ) | $spl19),
  inference(resolution,[],[f73556,f449])).
fof(f73556,plain,(
  ( ! [X4,X3] : (sP7('F_0',X3,X4) | s(bool,'T_0') != s(bool,X3)) ) | $spl19),
  inference(subsumption_resolution,[],[f73528,f457])).
fof(f73528,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,X3) | p(s(bool,X4)) | sP7('F_0',X3,X4)) ) | $spl19),
  inference(resolution,[],[f67013,f446])).
fof(f67013,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | $spl19),
  inference(superposition,[],[f1081,f464])).
fof(f110723,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X210,sK16,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl566),
  inference(resolution,[],[f87254,f463])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f292])).
fof(f87254,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X7,sK16,sK12))))) | ~p(s(bool,X7))) ) | $spl566),
  inference(subsumption_resolution,[],[f86974,f595])).
fof(f86974,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X7,sK16,sK12))))) | ~p(s(bool,X7))) ) | $spl566),
  inference(superposition,[],[f364,f83868])).
fof(f364,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22])],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6] : ! [X9,X29] : (? [X11] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))) & p(s(bool,X29))) <=> (? [X11] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))) & p(s(bool,X29))))),
  file('thm.list.EXISTS_CONG','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f29879,plain,(
  ( ! [X159,X157] : (p(s(bool,X157)) | ~p(s(bool,X159)) | s(bool,X157) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X159,sK16,sK12))))) ) | $spl162),
  inference(cnf_transformation,[],[f29879_D])).
fof(f29879_D,plain,(
  ( ! [X159,X157] : (p(s(bool,X157)) | ~p(s(bool,X159)) | s(bool,X157) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X159,sK16,sK12))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f118491,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X0,sK16,sK12))))) ) | ($spl22 | $spl566)),
  inference(resolution,[],[f110723,f29451])).
fof(f29451,plain,(
  ( ! [X4,X5,X3] : (~sP6(X4,X3,X5) | p(s(bool,X3)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f29408])).
fof(f29408,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | ~sP6(X4,X3,X5) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(resolution,[],[f29288,f467])).
fof(f29288,plain,(
  ( ! [X6,X8,X7] : (sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP6(X8,X6,X7) | p(s(bool,X7))) ) | $spl22),
  inference(resolution,[],[f29283,f455])).
fof(f29283,plain,(
  ( ! [X6,X8,X7] : (~sP8(X7,X8,X6) | p(s(bool,X6)) | sP9(X7,X6,X8) | ~sP6(X8,X6,X7)) ) | $spl22),
  inference(subsumption_resolution,[],[f29278,f558])).
fof(f29278,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | sP9(X7,X6,X8) | p(s(bool,X6)) | ~sP8(X7,X8,X6) | ~sP6(X8,X6,X7)) ) | $spl22),
  inference(resolution,[],[f29259,f450])).
fof(f29259,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'F_0') != s(bool,X10)) ) | $spl22),
  inference(subsumption_resolution,[],[f29258,f458])).
fof(f29258,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | sP7(X10,X9,X11)) ) | $spl22),
  inference(subsumption_resolution,[],[f29215,f457])).
fof(f29215,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,'F_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl22),
  inference(resolution,[],[f28933,f446])).
fof(f28933,plain,(
  ( ! [X6,X8,X7] : (sP9(X6,X7,X8) | ~p(s(bool,X8)) | s(bool,'F_0') != s(bool,X7) | p(s(bool,X6))) ) | $spl22),
  inference(resolution,[],[f28884,f455])).
fof(f28884,plain,(
  ( ! [X4,X5,X3] : (~sP8(X5,X3,X4) | sP9(X5,X4,X3) | ~p(s(bool,X3)) | s(bool,'F_0') != s(bool,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f28883,f562])).
fof(f28883,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') != s(bool,X4) | sP9(X5,X4,X3) | ~p(s(bool,X3)) | ~sP8(X5,X3,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f28877,f28844])).
fof(f28877,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,'F_0') != s(bool,X4) | sP9(X5,X4,X3) | ~p(s(bool,X3)) | ~sP8(X5,X3,X4) | ~sP6(X3,X4,X5)) ) | $spl22),
  inference(resolution,[],[f28845,f449])).
fof(f28845,plain,(
  ( ! [X10,X11,X9] : (sP7(X9,X10,X11) | s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') != s(bool,X9)) ) | $spl22),
  inference(subsumption_resolution,[],[f28803,f457])).
fof(f28803,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X9,X10,X11)) ) | $spl22),
  inference(resolution,[],[f4491,f446])).
fof(f39551,plain,(
  ( ! [X21,X22] : (s(bool,X21) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X22,sK16,sK12)))) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X21)) ) | $spl210),
  inference(cnf_transformation,[],[f39551_D])).
fof(f39551_D,plain,(
  ( ! [X21,X22] : (s(bool,X21) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X22,sK16,sK12)))) | ~p(s(bool,X22)) | s(bool,'F_0') != s(bool,X21)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f118305,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X33,sK16,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl22 | $spl566)),
  inference(resolution,[],[f110717,f28678])).
fof(f28678,plain,(
  ( ! [X24,X23,X22] : (~sP8(X23,'T_0',X24) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | s(bool,'F_0') != s(bool,X22)) ) | $spl22),
  inference(subsumption_resolution,[],[f28677,f595])).
fof(f28677,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | ~p(s(bool,'T_0')) | ~sP8(X23,'T_0',X24)) ) | $spl22),
  inference(subsumption_resolution,[],[f28676,f463])).
fof(f28676,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | ~p(s(bool,'T_0')) | ~sP8(X23,'T_0',X24) | ~sP6('T_0',X24,X23)) ) | $spl22),
  inference(subsumption_resolution,[],[f28628,f457])).
fof(f28628,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X24) | ~p(s(bool,'T_0')) | ~sP8(X23,'T_0',X24) | ~sP7(X24,'T_0',X23) | ~sP6('T_0',X24,X23)) ) | $spl22),
  inference(resolution,[],[f4478,f449])).
fof(f4478,plain,(
  ( ! [X14,X15,X16] : (~sP9(X16,X14,'T_0') | s(bool,'F_0') != s(bool,X15) | ~p(s(bool,X16)) | s(bool,X14) != s(bool,X15)) ) | $spl22),
  inference(superposition,[],[f2929,f466])).
fof(f110717,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X192,sK16,sK12)))) | ~p(s(bool,X192))) ) | $spl566),
  inference(resolution,[],[f87254,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f287])).
fof(f6167,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | $spl32),
  inference(cnf_transformation,[],[f6167_D])).
fof(f6167_D,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f118303,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X27,sK16,sK12))))) ) | ($spl22 | $spl566)),
  inference(resolution,[],[f110717,f30061])).
fof(f118302,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X24,sK16,sK12))))) ) | ($spl22 | $spl566)),
  inference(resolution,[],[f110717,f30058])).
fof(f30058,plain,(
  ( ! [X2,X0,X1] : (~sP8(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,'F_0') != s(bool,X0)) ) | $spl22),
  inference(subsumption_resolution,[],[f30057,f463])).
fof(f30057,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0) | ~sP6(X2,X0,X1)) ) | $spl22),
  inference(subsumption_resolution,[],[f30056,f457])).
fof(f30056,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0) | ~sP7(X0,X2,X1) | ~sP6(X2,X0,X1)) ) | $spl22),
  inference(subsumption_resolution,[],[f29971,f562])).
fof(f29971,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X1)) | s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2)) | ~sP8(X1,X2,X0) | ~sP7(X0,X2,X1) | ~sP6(X2,X0,X1)) ) | $spl22),
  inference(resolution,[],[f4493,f449])).
fof(f10721,plain,(
  ( ! [X21] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21)))) ) | $spl46),
  inference(cnf_transformation,[],[f10721_D])).
fof(f10721_D,plain,(
  ( ! [X21] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f117209,plain,(
  ( ! [X55] : (s(bool,'F_0') != s(bool,X55) | s(bool,X55) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl652),
  inference(cnf_transformation,[],[f117209_D])).
fof(f117209_D,plain,(
  ( ! [X55] : (s(bool,'F_0') != s(bool,X55) | s(bool,X55) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl652),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl652])])).
fof(f117547,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') != s(bool,X59) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X58))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X58))) | s(bool,X59) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl10),
  inference(subsumption_resolution,[],[f117489,f864])).
fof(f864,plain,(
  ( ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X5))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X5)))) )),
  inference(forward_demodulation,[],[f348,f347])).
fof(f348,plain,(
  ( ! [X5] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X5))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X5))) | ~p(s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK14))))))) )),
  inference(cnf_transformation,[],[f234])).
fof(f117489,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'F_0') != s(bool,X59) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X58))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X58))) | s(bool,X59) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl10),
  inference(resolution,[],[f67992,f976])).
fof(f976,plain,(
  ( ! [X175,X176,X174] : (sP9('const.bool.IN_2'(s(sK12,X174),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X175,X176) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X174))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X174))) | s(bool,X175) = s(bool,X176)) )),
  inference(resolution,[],[f864,f464])).
fof(f67992,plain,(
  ( ! [X265,X266] : (~sP9(X266,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X265) | p(s(bool,X266)) | s(bool,'F_0') != s(bool,X265)) ) | $spl10),
  inference(superposition,[],[f467,f711])).
fof(f117208,plain,(
  ( ! [X54,X55] : (s(bool,'F_0') != s(bool,X55) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | s(bool,X55) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl10),
  inference(subsumption_resolution,[],[f117143,f864])).
fof(f117143,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.bool.IN_2'(s(sK12,X54),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'F_0') != s(bool,X55) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | s(bool,X55) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl10),
  inference(resolution,[],[f67991,f976])).
fof(f67991,plain,(
  ( ! [X263,X264] : (~sP9(X264,X263,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X264)) | s(bool,'F_0') != s(bool,X263)) ) | $spl10),
  inference(superposition,[],[f467,f711])).
fof(f115027,plain,(
  ( ! [X38] : (s(bool,'F_0') = s(bool,X38) | s(bool,X38) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl650),
  inference(cnf_transformation,[],[f115027_D])).
fof(f115027_D,plain,(
  ( ! [X38] : (s(bool,'F_0') = s(bool,X38) | s(bool,X38) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl650),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl650])])).
fof(f115090,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') = s(bool,X39) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X40))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X40)))) ) | $spl10),
  inference(subsumption_resolution,[],[f115062,f864])).
fof(f115062,plain,(
  ( ! [X39,X40] : (s(bool,'F_0') = s(bool,X39) | p(s(bool,'const.bool.IN_2'(s(sK12,X40),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X40))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X40)))) ) | $spl10),
  inference(resolution,[],[f67840,f977])).
fof(f977,plain,(
  ( ! [X177,X178,X179] : (~sP9('const.bool.IN_2'(s(sK12,X177),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X179,X178) | s(bool,X178) != s(bool,X179) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X177))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X177)))) )),
  inference(resolution,[],[f864,f467])).
fof(f67840,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3))) ) | $spl10),
  inference(superposition,[],[f711,f464])).
fof(f115026,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,X38) | s(bool,X38) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X39))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X39)))) ) | $spl10),
  inference(subsumption_resolution,[],[f114997,f864])).
fof(f114997,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') = s(bool,X38) | p(s(bool,'const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X38) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X39))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X39)))) ) | $spl10),
  inference(resolution,[],[f67839,f977])).
fof(f67839,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X1))) ) | $spl10),
  inference(superposition,[],[f711,f464])).
fof(f114410,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X51) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl646),
  inference(cnf_transformation,[],[f114410_D])).
fof(f114410_D,plain,(
  ( ! [X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,X51) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl646),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl646])])).
fof(f114756,plain,(
  ( ! [X54,X55] : (s(bool,'T_0') != s(bool,X55) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | s(bool,X55) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl8),
  inference(subsumption_resolution,[],[f114688,f864])).
fof(f114688,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.bool.IN_2'(s(sK12,X54),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'T_0') != s(bool,X55) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | s(bool,X55) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl8),
  inference(resolution,[],[f67450,f976])).
fof(f67450,plain,(
  ( ! [X263,X264] : (~sP9(X264,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X263) | p(s(bool,X264)) | s(bool,'T_0') != s(bool,X263)) ) | $spl8),
  inference(superposition,[],[f467,f708])).
fof(f114745,plain,(
  ( ! [X36] : (~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X36) | p(s(bool,X36))) ) | $spl648),
  inference(cnf_transformation,[],[f114745_D])).
fof(f114745_D,plain,(
  ( ! [X36] : (~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X36) | p(s(bool,X36))) ) <=> ~$spl648),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl648])])).
fof(f114747,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | p(s(bool,X38)) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl8 | $spl22)),
  inference(forward_demodulation,[],[f114746,f708])).
fof(f114746,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f114709,f455])).
fof(f114709,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP8(X38,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl8 | $spl22)),
  inference(trivial_inequality_removal,[],[f114676])).
fof(f114676,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP8(X38,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f67450,f28175])).
fof(f28175,plain,(
  ( ! [X4,X5,X3] : (sP9(X5,X4,'T_0') | s(bool,X3) != s(bool,X4) | ~sP8(X5,'T_0',X4) | ~sP7(X4,'T_0',X5)) ) | $spl22),
  inference(subsumption_resolution,[],[f28168,f595])).
fof(f28168,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | sP9(X5,X4,'T_0') | ~p(s(bool,'T_0')) | ~sP8(X5,'T_0',X4) | ~sP7(X4,'T_0',X5)) ) | $spl22),
  inference(resolution,[],[f28166,f449])).
fof(f28166,plain,(
  ( ! [X6,X8,X7] : (sP6('T_0',X7,X8) | s(bool,X6) != s(bool,X7)) ) | $spl22),
  inference(subsumption_resolution,[],[f28137,f461])).
fof(f28137,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | sP6('T_0',X7,X8) | p(s(bool,X8))) ) | $spl22),
  inference(resolution,[],[f28060,f455])).
fof(f28060,plain,(
  ( ! [X6,X8,X7] : (~sP8(X8,'T_0',X7) | s(bool,X6) != s(bool,X7) | sP6('T_0',X7,X8)) ) | $spl22),
  inference(subsumption_resolution,[],[f28033,f462])).
fof(f28033,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~sP8(X8,'T_0',X7) | sP6('T_0',X7,X8)) ) | $spl22),
  inference(resolution,[],[f27940,f445])).
fof(f27940,plain,(
  ( ! [X6,X8,X7] : (sP9(X8,X7,'T_0') | s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~sP8(X8,'T_0',X7)) ) | $spl22),
  inference(subsumption_resolution,[],[f27939,f8572])).
fof(f8572,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | p(s(bool,X53)) | s(bool,'F_0') = s(bool,X52)) ) | $spl22),
  inference(subsumption_resolution,[],[f8531,f595])).
fof(f8531,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X53) | ~p(s(bool,'T_0')) | p(s(bool,X53)) | s(bool,'F_0') = s(bool,X52)) ) | $spl22),
  inference(superposition,[],[f8469,f588])).
fof(f588,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.list.EXISTS_CONG',aHL_BOOL_CASES)).
fof(f8469,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | p(s(bool,X30))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f8438])).
fof(f8438,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | ~p(s(bool,X31)) | p(s(bool,X30))) ) | $spl22),
  inference(superposition,[],[f4474,f560])).
fof(f560,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X16] : (s(bool,'F_0') = s(bool,X16) <=> ~p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.EQ_CLAUSES|split|3')).
fof(f4474,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X4) != s(bool,X5) | ~p(s(bool,X4))) ) | $spl22),
  inference(superposition,[],[f2929,f564])).
fof(f564,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) <=> p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.EQ_CLAUSES|split|0')).
fof(f27939,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'F_0') != s(bool,X6) | sP9(X8,X7,'T_0') | p(s(bool,X7)) | ~sP8(X8,'T_0',X7)) ) | $spl22),
  inference(subsumption_resolution,[],[f27930,f459])).
fof(f27930,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'F_0') != s(bool,X6) | sP9(X8,X7,'T_0') | p(s(bool,X7)) | ~sP8(X8,'T_0',X7) | ~sP7(X7,'T_0',X8)) ) | $spl22),
  inference(resolution,[],[f27911,f450])).
fof(f27911,plain,(
  ( ! [X6,X8,X7] : (sP6('T_0',X7,X8) | s(bool,X6) != s(bool,X7) | s(bool,'F_0') != s(bool,X6)) ) | $spl22),
  inference(subsumption_resolution,[],[f27883,f461])).
fof(f27883,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) != s(bool,X7) | p(s(bool,X8)) | sP6('T_0',X7,X8)) ) | $spl22),
  inference(resolution,[],[f4476,f445])).
fof(f4476,plain,(
  ( ! [X10,X8,X9] : (sP9(X10,X8,'T_0') | s(bool,'F_0') != s(bool,X9) | s(bool,X8) != s(bool,X9) | p(s(bool,X10))) ) | $spl22),
  inference(superposition,[],[f2929,f464])).
fof(f114744,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl8 | $spl22)),
  inference(forward_demodulation,[],[f114711,f708])).
fof(f114711,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,X37) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl8 | $spl22)),
  inference(trivial_inequality_removal,[],[f114710])).
fof(f114710,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X37) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl8 | $spl22)),
  inference(duplicate_literal_removal,[],[f114675])).
fof(f114675,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X37) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X36)) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f67450,f28176])).
fof(f28176,plain,(
  ( ! [X10,X11,X9] : (sP9(X11,X10,'T_0') | s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | ~sP7(X10,'T_0',X11)) ) | $spl22),
  inference(subsumption_resolution,[],[f28170,f455])).
fof(f28170,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,X10) | sP9(X11,X10,'T_0') | p(s(bool,X11)) | ~sP8(X11,'T_0',X10) | ~sP7(X10,'T_0',X11)) ) | $spl22),
  inference(resolution,[],[f28166,f451])).
fof(f114409,plain,(
  ( ! [X50,X51] : (s(bool,'T_0') != s(bool,X51) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X50))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X50))) | s(bool,X51) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl8),
  inference(subsumption_resolution,[],[f114357,f864])).
fof(f114357,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.bool.IN_2'(s(sK12,X50),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'T_0') != s(bool,X51) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X50))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X50))) | s(bool,X51) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl8),
  inference(resolution,[],[f67449,f976])).
fof(f67449,plain,(
  ( ! [X261,X262] : (~sP9(X262,X261,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X262)) | s(bool,'T_0') != s(bool,X261)) ) | $spl8),
  inference(superposition,[],[f467,f708])).
fof(f30109,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X31,sK12)))) | ~p(s(bool,X31))) ) | $spl186),
  inference(cnf_transformation,[],[f30109_D])).
fof(f30109_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X31,sK12)))) | ~p(s(bool,X31))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f113621,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X20,sK12))))) ) | ($spl19 | $spl22 | $spl86)),
  inference(resolution,[],[f107206,f73680])).
fof(f107206,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl86),
  inference(resolution,[],[f16673,f463])).
fof(f29882,plain,(
  ( ! [X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X168)) | s(bool,X166) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X168,sK12))))) ) | $spl168),
  inference(cnf_transformation,[],[f29882_D])).
fof(f29882_D,plain,(
  ( ! [X166,X168] : (p(s(bool,X166)) | ~p(s(bool,X168)) | s(bool,X166) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X168,sK12))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f113614,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X0,sK12))))) ) | ($spl22 | $spl86)),
  inference(resolution,[],[f107206,f29451])).
fof(f39554,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X31,sK12)))) | ~p(s(bool,X31)) | s(bool,'F_0') != s(bool,X30)) ) | $spl216),
  inference(cnf_transformation,[],[f39554_D])).
fof(f39554_D,plain,(
  ( ! [X30,X31] : (s(bool,X30) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X31,sK12)))) | ~p(s(bool,X31)) | s(bool,'F_0') != s(bool,X30)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f113555,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X33,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl22 | $spl86)),
  inference(resolution,[],[f107200,f28678])).
fof(f107200,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl86),
  inference(resolution,[],[f16673,f453])).
fof(f113553,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X27,sK12))))) ) | ($spl22 | $spl86)),
  inference(resolution,[],[f107200,f30061])).
fof(f113552,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X24,sK12))))) ) | ($spl22 | $spl86)),
  inference(resolution,[],[f107200,f30058])).
fof(f30108,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X28,sK12)))) | ~p(s(bool,X28))) ) | $spl184),
  inference(cnf_transformation,[],[f30108_D])).
fof(f30108_D,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X28,sK12)))) | ~p(s(bool,X28))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f113343,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X20,sK12))))) ) | ($spl19 | $spl22 | $spl84)),
  inference(resolution,[],[f106994,f73680])).
fof(f106994,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl84),
  inference(resolution,[],[f16672,f463])).
fof(f29881,plain,(
  ( ! [X165,X163] : (p(s(bool,X163)) | ~p(s(bool,X165)) | s(bool,X163) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X165,sK12))))) ) | $spl166),
  inference(cnf_transformation,[],[f29881_D])).
fof(f29881_D,plain,(
  ( ! [X165,X163] : (p(s(bool,X163)) | ~p(s(bool,X165)) | s(bool,X163) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X165,sK12))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f113336,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X0,sK12))))) ) | ($spl22 | $spl84)),
  inference(resolution,[],[f106994,f29451])).
fof(f39553,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X28,sK12)))) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X27)) ) | $spl214),
  inference(cnf_transformation,[],[f39553_D])).
fof(f39553_D,plain,(
  ( ! [X28,X27] : (s(bool,X27) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X28,sK12)))) | ~p(s(bool,X28)) | s(bool,'F_0') != s(bool,X27)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f113277,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X33,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl22 | $spl84)),
  inference(resolution,[],[f106988,f28678])).
fof(f106988,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl84),
  inference(resolution,[],[f16672,f453])).
fof(f113275,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X27,sK12))))) ) | ($spl22 | $spl84)),
  inference(resolution,[],[f106988,f30061])).
fof(f113274,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X24,sK12))))) ) | ($spl22 | $spl84)),
  inference(resolution,[],[f106988,f30058])).
fof(f30107,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X25,sK16,sK12)))) | ~p(s(bool,X25))) ) | $spl182),
  inference(cnf_transformation,[],[f30107_D])).
fof(f30107_D,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X25,sK16,sK12)))) | ~p(s(bool,X25))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f113191,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X20,sK16,sK12))))) ) | ($spl19 | $spl22 | $spl82)),
  inference(resolution,[],[f106776,f73680])).
fof(f106776,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X210,sK16,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl82),
  inference(resolution,[],[f16671,f463])).
fof(f29880,plain,(
  ( ! [X162,X160] : (p(s(bool,X160)) | ~p(s(bool,X162)) | s(bool,X160) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X162,sK16,sK12))))) ) | $spl164),
  inference(cnf_transformation,[],[f29880_D])).
fof(f29880_D,plain,(
  ( ! [X162,X160] : (p(s(bool,X160)) | ~p(s(bool,X162)) | s(bool,X160) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X162,sK16,sK12))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f113184,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X0,sK16,sK12))))) ) | ($spl22 | $spl82)),
  inference(resolution,[],[f106776,f29451])).
fof(f18683,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43)))) ) | $spl96),
  inference(cnf_transformation,[],[f18683_D])).
fof(f18683_D,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f113155,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X56)))) | ~p(s(bool,X54))) )),
  inference(subsumption_resolution,[],[f113076,f557])).
fof(f113076,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X56)))) | ~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f1027,f588])).
fof(f1027,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X31),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X32),s(fun(bool,bool),X33)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X31))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X31))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X31),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X33)))) | ~p(s(bool,X32))) )),
  inference(superposition,[],[f864,f583])).
fof(f583,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | s(X0,X1) = s(X0,X2) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X6] : ! [X11,X31,X33] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X31),s(fun(X6,bool),X33)))))) <=> (s(X6,X11) = s(X6,X31) | p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X33))))))),
  file('thm.list.EXISTS_CONG','thm.pred_set.IN_INSERT')).
fof(f113154,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X50)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X49))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X49))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X50)))) | ~p(s(bool,X48))) )),
  inference(subsumption_resolution,[],[f113074,f557])).
fof(f113074,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X50)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X49))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X49))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X50)))) | ~p(s(bool,X48)) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f1027,f587])).
fof(f587,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,'F_0') = s(bool,X16))),
  file('thm.list.EXISTS_CONG','thm.bool.BOOL_CASES_AX')).
fof(f113139,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X47)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X47)))) | ~p(s(bool,X45))) )),
  inference(duplicate_literal_removal,[],[f113073])).
fof(f113073,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X47)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X47)))) | ~p(s(bool,X45)) | ~p(s(bool,X45))) )),
  inference(superposition,[],[f1027,f564])).
fof(f113140,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | ~p(s(bool,X42))) )),
  inference(duplicate_literal_removal,[],[f113072])).
fof(f113072,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | ~p(s(bool,X42)) | ~p(s(bool,X42))) )),
  inference(superposition,[],[f1027,f562])).
fof(f39552,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X25,sK16,sK12)))) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X24)) ) | $spl212),
  inference(cnf_transformation,[],[f39552_D])).
fof(f39552_D,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X25,sK16,sK12)))) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,X24)) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f112889,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X33,sK16,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl22 | $spl82)),
  inference(resolution,[],[f106770,f28678])).
fof(f106770,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X192,sK16,sK12)))) | ~p(s(bool,X192))) ) | $spl82),
  inference(resolution,[],[f16671,f453])).
fof(f112887,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X27,sK16,sK12))))) ) | ($spl22 | $spl82)),
  inference(resolution,[],[f106770,f30061])).
fof(f112886,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X24,sK16,sK12))))) ) | ($spl22 | $spl82)),
  inference(resolution,[],[f106770,f30058])).
fof(f16675,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl90),
  inference(cnf_transformation,[],[f16675_D])).
fof(f16675_D,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f110660,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl566),
  inference(resolution,[],[f87254,f373])).
fof(f373,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK25,sK26])],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X29,X9] : (! [X11] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))) => p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X11))))) => (? [X11] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))) => ? [X11] : p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X11))))))),
  file('thm.list.EXISTS_CONG','thm.bool.MONO_EXISTS')).
fof(f16674,plain,(
  ( ! [X21] : (p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl88),
  inference(cnf_transformation,[],[f16674_D])).
fof(f16674_D,plain,(
  ( ! [X21] : (p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f110659,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl566),
  inference(resolution,[],[f87254,f372])).
fof(f372,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK26(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK25(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f253])).
fof(f16670,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl80),
  inference(cnf_transformation,[],[f16670_D])).
fof(f16670_D,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f110654,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl566),
  inference(resolution,[],[f87254,f363])).
fof(f363,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19,sK20])],[f242])).
fof(f242,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f241])).
fof(f241,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f240])).
fof(f240,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ! [X9,X29] : (? [X11] : (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))) | p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X11))))) <=> (? [X11] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))) | ? [X11] : p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X11))))))),
  file('thm.list.EXISTS_CONG','thm.bool.EXISTS_OR_THM')).
fof(f16669,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl78),
  inference(cnf_transformation,[],[f16669_D])).
fof(f16669_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f110653,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl566),
  inference(resolution,[],[f87254,f362])).
fof(f362,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK20(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK20(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f16668,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl76),
  inference(cnf_transformation,[],[f16668_D])).
fof(f16668_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f110652,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl566),
  inference(resolution,[],[f87254,f361])).
fof(f361,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f16667,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl74),
  inference(cnf_transformation,[],[f16667_D])).
fof(f16667_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f110651,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl566),
  inference(resolution,[],[f87254,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK19(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK18(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f243])).
fof(f6166,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X24)))) | ~p(s(bool,X22)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X24))))))) ) | $spl30),
  inference(cnf_transformation,[],[f6166_D])).
fof(f6166_D,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X24)))) | ~p(s(bool,X22)) | ~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X24))))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f110646,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X62)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X60))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X60))) | p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),X62)))) | ~p(s(bool,X61))) )),
  inference(subsumption_resolution,[],[f110551,f1000])).
fof(f1000,plain,(
  ( ! [X216] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X216),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X216))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X216)))) )),
  inference(resolution,[],[f864,f561])).
fof(f561,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f340])).
fof(f110551,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X62)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X60))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X60))) | p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),X62)))) | ~p(s(bool,X61)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X60),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1026,f588])).
fof(f1026,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK12,X28),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X30)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X28))) | p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),X30)))) | ~p(s(bool,X29))) )),
  inference(superposition,[],[f864,f583])).
fof(f110645,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),X56)))) | ~p(s(bool,X55))) )),
  inference(subsumption_resolution,[],[f110549,f1000])).
fof(f110549,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),X56)))) | ~p(s(bool,X55)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X54),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1026,f587])).
fof(f110644,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42))) | p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),X44)))) | ~p(s(bool,X43))) )),
  inference(subsumption_resolution,[],[f110545,f864])).
fof(f110545,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42))) | p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),X44)))) | ~p(s(bool,X43)) | p(s(bool,'const.bool.IN_2'(s(sK12,X42),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f560])).
fof(f110643,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X41)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X39))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X39))) | p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),X41)))) | ~p(s(bool,X40))) )),
  inference(subsumption_resolution,[],[f110544,f864])).
fof(f110544,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X41)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X39))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X39))) | p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),X41)))) | ~p(s(bool,X40)) | p(s(bool,'const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f558])).
fof(f17654,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X24),s(fun(bool,bool),X26)))))) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X26))))) ) | $spl94),
  inference(cnf_transformation,[],[f17654_D])).
fof(f17654_D,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X24),s(fun(bool,bool),X26)))))) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X26))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f110642,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X28),s(fun(bool,bool),X30)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),X30)))) | ~p(s(bool,X29)) | p(s(bool,X28))) )),
  inference(subsumption_resolution,[],[f110541,f864])).
fof(f110541,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X28),s(fun(bool,bool),X30)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),X30)))) | ~p(s(bool,X29)) | p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | p(s(bool,X28))) )),
  inference(superposition,[],[f1026,f515])).
fof(f515,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f202])).
fof(f202,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X27,X28] : ((p(s(bool,X27)) => p(s(bool,X28))) => ((p(s(bool,X28)) => p(s(bool,X27))) => s(bool,X27) = s(bool,X28)))),
  file('thm.list.EXISTS_CONG','thm.bool.IMP_ANTISYM_AX')).
fof(f110641,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X24),s(fun(bool,bool),X26)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23))) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X26)))) | ~p(s(bool,X25)) | p(s(bool,X24))) )),
  inference(subsumption_resolution,[],[f110540,f864])).
fof(f110540,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X24),s(fun(bool,bool),X26)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23))) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X26)))) | ~p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,'const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f515])).
fof(f110578,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f110535])).
fof(f110535,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | ~p(s(bool,X1)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0)))) )),
  inference(superposition,[],[f1026,f998])).
fof(f998,plain,(
  ( ! [X214] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK12,X214),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X214))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X214)))) )),
  inference(resolution,[],[f864,f558])).
fof(f30117,plain,(
  ( ! [X51] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X51,sK12)))) | ~p(s(bool,X51))) ) | $spl194),
  inference(cnf_transformation,[],[f30117_D])).
fof(f30117_D,plain,(
  ( ! [X51] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X51,sK12)))) | ~p(s(bool,X51))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f108542,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X20,sK12))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f104163,f73680])).
fof(f104163,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl8),
  inference(resolution,[],[f72645,f463])).
fof(f72645,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X9,sK12))))) | ~p(s(bool,X9))) ) | $spl8),
  inference(resolution,[],[f67583,f371])).
fof(f371,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK24(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK23,sK24])],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f248])).
fof(f248,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X6] : ! [X9,X29] : (? [X11] : (p(s(bool,X9)) & p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X11))))) <=> (p(s(bool,X9)) & ? [X11] : p(s(bool,happ(s(fun(X6,bool),X29),s(X6,X11))))))),
  file('thm.list.EXISTS_CONG','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f67583,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) | $spl8),
  inference(subsumption_resolution,[],[f67314,f595])).
fof(f67314,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) | $spl8),
  inference(superposition,[],[f609,f708])).
fof(f29886,plain,(
  ( ! [X180,X178] : (p(s(bool,X178)) | ~p(s(bool,X180)) | s(bool,X178) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X180,sK12))))) ) | $spl176),
  inference(cnf_transformation,[],[f29886_D])).
fof(f29886_D,plain,(
  ( ! [X180,X178] : (p(s(bool,X178)) | ~p(s(bool,X180)) | s(bool,X178) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X180,sK12))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f108535,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X0,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f104163,f29451])).
fof(f39562,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X51,sK12)))) | ~p(s(bool,X51)) | s(bool,'F_0') != s(bool,X50)) ) | $spl224),
  inference(cnf_transformation,[],[f39562_D])).
fof(f39562_D,plain,(
  ( ! [X50,X51] : (s(bool,X50) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X51,sK12)))) | ~p(s(bool,X51)) | s(bool,'F_0') != s(bool,X50)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f108456,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X33,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f104157,f28678])).
fof(f104157,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl8),
  inference(resolution,[],[f72645,f453])).
fof(f108454,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X27,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f104157,f30061])).
fof(f108453,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X24,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f104157,f30058])).
fof(f30116,plain,(
  ( ! [X48] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X48,sK12)))) | ~p(s(bool,X48))) ) | $spl192),
  inference(cnf_transformation,[],[f30116_D])).
fof(f30116_D,plain,(
  ( ! [X48] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X48,sK12)))) | ~p(s(bool,X48))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f108372,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X20,sK12))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f103945,f73680])).
fof(f103945,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl8),
  inference(resolution,[],[f72644,f463])).
fof(f72644,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X8,sK12))))) | ~p(s(bool,X8))) ) | $spl8),
  inference(resolution,[],[f67583,f369])).
fof(f369,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f251])).
fof(f29885,plain,(
  ( ! [X177,X175] : (p(s(bool,X175)) | ~p(s(bool,X177)) | s(bool,X175) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X177,sK12))))) ) | $spl174),
  inference(cnf_transformation,[],[f29885_D])).
fof(f29885_D,plain,(
  ( ! [X177,X175] : (p(s(bool,X175)) | ~p(s(bool,X177)) | s(bool,X175) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X177,sK12))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f108365,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X0,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103945,f29451])).
fof(f39561,plain,(
  ( ! [X47,X48] : (s(bool,X47) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X48,sK12)))) | ~p(s(bool,X48)) | s(bool,'F_0') != s(bool,X47)) ) | $spl222),
  inference(cnf_transformation,[],[f39561_D])).
fof(f39561_D,plain,(
  ( ! [X47,X48] : (s(bool,X47) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X48,sK12)))) | ~p(s(bool,X48)) | s(bool,'F_0') != s(bool,X47)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f108306,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X33,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103939,f28678])).
fof(f103939,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl8),
  inference(resolution,[],[f72644,f453])).
fof(f108304,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X27,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103939,f30061])).
fof(f108303,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X24,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103939,f30058])).
fof(f30115,plain,(
  ( ! [X45] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X45,sK15,sK12)))) | ~p(s(bool,X45))) ) | $spl190),
  inference(cnf_transformation,[],[f30115_D])).
fof(f30115_D,plain,(
  ( ! [X45] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X45,sK15,sK12)))) | ~p(s(bool,X45))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f108095,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X20,sK15,sK12))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f103727,f73680])).
fof(f103727,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X210,sK15,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl8),
  inference(resolution,[],[f72643,f463])).
fof(f72643,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X7,sK15,sK12))))) | ~p(s(bool,X7))) ) | $spl8),
  inference(resolution,[],[f67583,f366])).
fof(f366,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK22(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f247])).
fof(f29884,plain,(
  ( ! [X174,X172] : (p(s(bool,X172)) | ~p(s(bool,X174)) | s(bool,X172) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X174,sK15,sK12))))) ) | $spl172),
  inference(cnf_transformation,[],[f29884_D])).
fof(f29884_D,plain,(
  ( ! [X174,X172] : (p(s(bool,X172)) | ~p(s(bool,X174)) | s(bool,X172) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X174,sK15,sK12))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f108088,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X0,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103727,f29451])).
fof(f39560,plain,(
  ( ! [X45,X44] : (s(bool,X44) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X45,sK15,sK12)))) | ~p(s(bool,X45)) | s(bool,'F_0') != s(bool,X44)) ) | $spl220),
  inference(cnf_transformation,[],[f39560_D])).
fof(f39560_D,plain,(
  ( ! [X45,X44] : (s(bool,X44) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X45,sK15,sK12)))) | ~p(s(bool,X45)) | s(bool,'F_0') != s(bool,X44)) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f108029,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X33,sK15,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103721,f28678])).
fof(f103721,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X192,sK15,sK12)))) | ~p(s(bool,X192))) ) | $spl8),
  inference(resolution,[],[f72643,f453])).
fof(f108027,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X27,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103721,f30061])).
fof(f108026,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X24,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103721,f30058])).
fof(f30114,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X42,sK15,sK12)))) | ~p(s(bool,X42))) ) | $spl188),
  inference(cnf_transformation,[],[f30114_D])).
fof(f30114_D,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X42,sK15,sK12)))) | ~p(s(bool,X42))) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f107945,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X20,sK15,sK12))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f103520,f73680])).
fof(f103520,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X210,sK15,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl8),
  inference(resolution,[],[f72642,f463])).
fof(f72642,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X6,sK15,sK12))))) | ~p(s(bool,X6))) ) | $spl8),
  inference(resolution,[],[f67583,f364])).
fof(f29883,plain,(
  ( ! [X171,X169] : (p(s(bool,X169)) | ~p(s(bool,X171)) | s(bool,X169) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X171,sK15,sK12))))) ) | $spl170),
  inference(cnf_transformation,[],[f29883_D])).
fof(f29883_D,plain,(
  ( ! [X171,X169] : (p(s(bool,X169)) | ~p(s(bool,X171)) | s(bool,X169) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X171,sK15,sK12))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f107938,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X0,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103520,f29451])).
fof(f39559,plain,(
  ( ! [X41,X42] : (s(bool,X41) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X42,sK15,sK12)))) | ~p(s(bool,X42)) | s(bool,'F_0') != s(bool,X41)) ) | $spl218),
  inference(cnf_transformation,[],[f39559_D])).
fof(f39559_D,plain,(
  ( ! [X41,X42] : (s(bool,X41) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X42,sK15,sK12)))) | ~p(s(bool,X42)) | s(bool,'F_0') != s(bool,X41)) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f107879,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X33,sK15,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103514,f28678])).
fof(f103514,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X192,sK15,sK12)))) | ~p(s(bool,X192))) ) | $spl8),
  inference(resolution,[],[f72642,f453])).
fof(f107877,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X27,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103514,f30061])).
fof(f107876,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X24,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f103514,f30058])).
fof(f107740,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X24,sK12))))) ) | $spl640),
  inference(cnf_transformation,[],[f107740_D])).
fof(f107740_D,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X24,sK12))))) ) <=> ~$spl640),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl640])])).
fof(f107794,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X20,sK12))))) ) | ($spl19 | $spl22 | $spl134)),
  inference(resolution,[],[f103322,f73680])).
fof(f103322,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl134),
  inference(resolution,[],[f25888,f463])).
fof(f107801,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X0,sK12))))) ) | $spl644),
  inference(cnf_transformation,[],[f107801_D])).
fof(f107801_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X0,sK12))))) ) <=> ~$spl644),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl644])])).
fof(f107787,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X0,sK12))))) ) | ($spl22 | $spl134)),
  inference(resolution,[],[f103322,f29451])).
fof(f107741,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s(bool,'F_0') != s(bool,X35) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X33,sK12))))) ) | $spl642),
  inference(cnf_transformation,[],[f107741_D])).
fof(f107741_D,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s(bool,'F_0') != s(bool,X35) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X33,sK12))))) ) <=> ~$spl642),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl642])])).
fof(f107726,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X33,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl22 | $spl134)),
  inference(resolution,[],[f103316,f28678])).
fof(f103316,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl134),
  inference(resolution,[],[f25888,f453])).
fof(f107724,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X27,sK12))))) ) | ($spl22 | $spl134)),
  inference(resolution,[],[f103316,f30061])).
fof(f107723,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X24,sK12))))) ) | ($spl22 | $spl134)),
  inference(resolution,[],[f103316,f30058])).
fof(f107587,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X24,sK12))))) ) | $spl634),
  inference(cnf_transformation,[],[f107587_D])).
fof(f107587_D,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X24,sK12))))) ) <=> ~$spl634),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl634])])).
fof(f107641,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X20,sK12))))) ) | ($spl19 | $spl22 | $spl132)),
  inference(resolution,[],[f103115,f73680])).
fof(f103115,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl132),
  inference(resolution,[],[f25886,f463])).
fof(f107648,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X0,sK12))))) ) | $spl638),
  inference(cnf_transformation,[],[f107648_D])).
fof(f107648_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X0,sK12))))) ) <=> ~$spl638),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl638])])).
fof(f107634,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X0,sK12))))) ) | ($spl22 | $spl132)),
  inference(resolution,[],[f103115,f29451])).
fof(f107588,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s(bool,'F_0') != s(bool,X35) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X33,sK12))))) ) | $spl636),
  inference(cnf_transformation,[],[f107588_D])).
fof(f107588_D,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s(bool,'F_0') != s(bool,X35) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X33,sK12))))) ) <=> ~$spl636),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl636])])).
fof(f107573,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X33,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl22 | $spl132)),
  inference(resolution,[],[f103109,f28678])).
fof(f103109,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl132),
  inference(resolution,[],[f25886,f453])).
fof(f107571,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X27,sK12))))) ) | ($spl22 | $spl132)),
  inference(resolution,[],[f103109,f30061])).
fof(f107570,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X24,sK12))))) ) | ($spl22 | $spl132)),
  inference(resolution,[],[f103109,f30058])).
fof(f107434,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X24,sK16,sK12))))) ) | $spl628),
  inference(cnf_transformation,[],[f107434_D])).
fof(f107434_D,plain,(
  ( ! [X24] : (~p(s(bool,X24)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X24,sK16,sK12))))) ) <=> ~$spl628),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl628])])).
fof(f107488,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X20,sK16,sK12))))) ) | ($spl19 | $spl22 | $spl130)),
  inference(resolution,[],[f102939,f73680])).
fof(f102939,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X210,sK16,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl130),
  inference(resolution,[],[f25884,f463])).
fof(f107495,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X0,sK16,sK12))))) ) | $spl632),
  inference(cnf_transformation,[],[f107495_D])).
fof(f107495_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X0,sK16,sK12))))) ) <=> ~$spl632),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl632])])).
fof(f107481,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X0,sK16,sK12))))) ) | ($spl22 | $spl130)),
  inference(resolution,[],[f102939,f29451])).
fof(f107435,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s(bool,'F_0') != s(bool,X35) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X33,sK16,sK12))))) ) | $spl630),
  inference(cnf_transformation,[],[f107435_D])).
fof(f107435_D,plain,(
  ( ! [X35,X33] : (~p(s(bool,X33)) | s(bool,'F_0') != s(bool,X35) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X33,sK16,sK12))))) ) <=> ~$spl630),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl630])])).
fof(f107420,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X33,sK16,sK12)))) | s(bool,'F_0') != s(bool,X35)) ) | ($spl22 | $spl130)),
  inference(resolution,[],[f102933,f28678])).
fof(f102933,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X192,sK16,sK12)))) | ~p(s(bool,X192))) ) | $spl130),
  inference(resolution,[],[f25884,f453])).
fof(f107418,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X29) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X27,sK16,sK12))))) ) | ($spl22 | $spl130)),
  inference(resolution,[],[f102933,f30061])).
fof(f107417,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X24,sK16,sK12))))) ) | ($spl22 | $spl130)),
  inference(resolution,[],[f102933,f30058])).
fof(f12603,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X5))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X5)))) ) | $spl54),
  inference(cnf_transformation,[],[f12603_D])).
fof(f12603_D,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X5))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X5)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f107307,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f107270])).
fof(f107270,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f1013,f465])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f1013,plain,(
  ( ! [X10,X11,X9] : (~sP9(X11,'const.bool.IN_2'(s(sK12,X9),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X10) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X9))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X9))) | ~p(s(bool,X11)) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f864,f466])).
fof(f107143,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl86),
  inference(resolution,[],[f16673,f373])).
fof(f107142,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl86),
  inference(resolution,[],[f16673,f372])).
fof(f106828,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X13,sK16,sK12))))) | ~p(s(bool,X13))) ) | $spl626),
  inference(cnf_transformation,[],[f106828_D])).
fof(f106828_D,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X13,sK16,sK12))))) | ~p(s(bool,X13))) ) <=> ~$spl626),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl626])])).
fof(f107138,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X13,sK16,sK12))))) | ~p(s(bool,X13))) ) | $spl86),
  inference(resolution,[],[f16673,f364])).
fof(f107137,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl86),
  inference(resolution,[],[f16673,f363])).
fof(f107136,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl86),
  inference(resolution,[],[f16673,f362])).
fof(f107135,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl86),
  inference(resolution,[],[f16673,f361])).
fof(f107134,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl86),
  inference(resolution,[],[f16673,f360])).
fof(f107095,plain,(
  ( ! [X54,X52,X53] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X52))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X52))) | ~p(s(bool,X53)) | s(bool,X53) != s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54)))) )),
  inference(subsumption_resolution,[],[f107073,f864])).
fof(f107073,plain,(
  ( ! [X54,X52,X53] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X52))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X52))) | ~p(s(bool,X53)) | p(s(bool,'const.bool.IN_2'(s(sK12,X54),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X53) != s(bool,'const.bool.IN_2'(s(sK12,X52),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54)))) )),
  inference(resolution,[],[f1011,f977])).
fof(f1011,plain,(
  ( ! [X4,X5,X3] : (sP9(X5,'const.bool.IN_2'(s(sK12,X3),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X4) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3))) | ~p(s(bool,X4)) | p(s(bool,X5))) )),
  inference(superposition,[],[f864,f464])).
fof(f107094,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f107058])).
fof(f107058,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1011,f467])).
fof(f106931,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl84),
  inference(resolution,[],[f16672,f373])).
fof(f106930,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl84),
  inference(resolution,[],[f16672,f372])).
fof(f16673,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl86),
  inference(cnf_transformation,[],[f16673_D])).
fof(f16673_D,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f106929,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl84),
  inference(resolution,[],[f16672,f371])).
fof(f106926,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X13,sK16,sK12))))) | ~p(s(bool,X13))) ) | $spl84),
  inference(resolution,[],[f16672,f364])).
fof(f106925,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl84),
  inference(resolution,[],[f16672,f363])).
fof(f106924,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl84),
  inference(resolution,[],[f16672,f362])).
fof(f106923,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl84),
  inference(resolution,[],[f16672,f361])).
fof(f106922,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl84),
  inference(resolution,[],[f16672,f360])).
fof(f106881,plain,(
  ( ! [X54,X55,X53] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X53))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X53))) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK12,X53),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55)))) )),
  inference(subsumption_resolution,[],[f106857,f864])).
fof(f106857,plain,(
  ( ! [X54,X55,X53] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X53))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X53))) | ~p(s(bool,X54)) | p(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK12,X53),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55)))) )),
  inference(resolution,[],[f1010,f977])).
fof(f1010,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(superposition,[],[f864,f464])).
fof(f106878,plain,(
  ( ! [X19,X20,X18] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X18))) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK12,X18),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f106842])).
fof(f106842,plain,(
  ( ! [X19,X20,X18] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X18))) | ~p(s(bool,X19)) | p(s(bool,X20)) | p(s(bool,X20)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK12,X18),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1010,f467])).
fof(f106713,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl82),
  inference(resolution,[],[f16671,f373])).
fof(f106712,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl82),
  inference(resolution,[],[f16671,f372])).
fof(f106711,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl82),
  inference(resolution,[],[f16671,f371])).
fof(f16672,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl84),
  inference(cnf_transformation,[],[f16672_D])).
fof(f16672_D,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X17,sK12))))) | ~p(s(bool,X17))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f106710,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl82),
  inference(resolution,[],[f16671,f369])).
fof(f106708,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X13,sK16,sK12))))) | ~p(s(bool,X13))) ) | $spl82),
  inference(resolution,[],[f16671,f364])).
fof(f106707,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl82),
  inference(resolution,[],[f16671,f363])).
fof(f106706,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl82),
  inference(resolution,[],[f16671,f362])).
fof(f106705,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl82),
  inference(resolution,[],[f16671,f361])).
fof(f106704,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl82),
  inference(resolution,[],[f16671,f360])).
fof(f106481,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | $spl617),
  inference(cnf_transformation,[],[f106481_D])).
fof(f106481_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) <=> ~$spl617),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl617])])).
fof(f106547,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f75347,f73680])).
fof(f75347,plain,(
  ( ! [X134,X135] : (sP6('const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X134,X135)) ) | $spl8),
  inference(resolution,[],[f67586,f463])).
fof(f106556,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl624),
  inference(cnf_transformation,[],[f106556_D])).
fof(f106556_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl624),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl624])])).
fof(f106540,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75347,f29451])).
fof(f106484,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl622),
  inference(cnf_transformation,[],[f106484_D])).
fof(f106484_D,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl622),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl622])])).
fof(f106463,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,'F_0') != s(bool,X25)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75341,f28678])).
fof(f75341,plain,(
  ( ! [X123,X122] : (sP8(X122,X123,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl8),
  inference(resolution,[],[f67586,f453])).
fof(f106483,plain,(
  ( ! [X23] : (sP9(X23,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),'T_0')) ) | $spl620),
  inference(cnf_transformation,[],[f106483_D])).
fof(f106483_D,plain,(
  ( ! [X23] : (sP9(X23,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),'T_0')) ) <=> ~$spl620),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl620])])).
fof(f106462,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X22) | sP9(X23,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),'T_0') | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75341,f28027])).
fof(f28027,plain,(
  ( ! [X4,X5,X3] : (~sP8(X5,'T_0',X4) | s(bool,'F_0') != s(bool,X3) | sP9(X5,X4,'T_0') | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f28026,f595])).
fof(f28026,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'F_0') != s(bool,X3) | sP9(X5,X4,'T_0') | ~p(s(bool,'T_0')) | ~sP8(X5,'T_0',X4)) ) | $spl22),
  inference(subsumption_resolution,[],[f28021,f463])).
fof(f28021,plain,(
  ( ! [X4,X5,X3] : (s(bool,X3) != s(bool,X4) | s(bool,'F_0') != s(bool,X3) | sP9(X5,X4,'T_0') | ~p(s(bool,'T_0')) | ~sP8(X5,'T_0',X4) | ~sP6('T_0',X4,X5)) ) | $spl22),
  inference(resolution,[],[f27912,f449])).
fof(f27912,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,'T_0',X11) | s(bool,X9) != s(bool,X10) | s(bool,'F_0') != s(bool,X9)) ) | $spl22),
  inference(subsumption_resolution,[],[f27884,f457])).
fof(f27884,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,X9) != s(bool,X10) | p(s(bool,X11)) | sP7(X10,'T_0',X11)) ) | $spl22),
  inference(resolution,[],[f4476,f446])).
fof(f106461,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') != s(bool,X21) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75341,f30061])).
fof(f106460,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75341,f30058])).
fof(f106482,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13))) ) | $spl618),
  inference(cnf_transformation,[],[f106482_D])).
fof(f106482_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13))) ) <=> ~$spl618),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl618])])).
fof(f106457,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75341,f28884])).
fof(f106081,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | $spl607),
  inference(cnf_transformation,[],[f106081_D])).
fof(f106081_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) <=> ~$spl607),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl607])])).
fof(f106252,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f75214,f73680])).
fof(f75214,plain,(
  ( ! [X134,X135] : (sP6('const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X134,X135)) ) | $spl8),
  inference(resolution,[],[f67584,f463])).
fof(f106260,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl614),
  inference(cnf_transformation,[],[f106260_D])).
fof(f106260_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl614),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl614])])).
fof(f106245,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75214,f29451])).
fof(f106084,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl612),
  inference(cnf_transformation,[],[f106084_D])).
fof(f106084_D,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl612),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl612])])).
fof(f106063,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,'F_0') != s(bool,X25)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75208,f28678])).
fof(f75208,plain,(
  ( ! [X123,X122] : (sP8(X122,X123,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl8),
  inference(resolution,[],[f67584,f453])).
fof(f106083,plain,(
  ( ! [X23] : (sP9(X23,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),'T_0')) ) | $spl610),
  inference(cnf_transformation,[],[f106083_D])).
fof(f106083_D,plain,(
  ( ! [X23] : (sP9(X23,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),'T_0')) ) <=> ~$spl610),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl610])])).
fof(f106062,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X22) | sP9(X23,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),'T_0') | s(bool,X22) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75208,f28027])).
fof(f106061,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | s(bool,'T_0') != s(bool,X21) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75208,f30061])).
fof(f106060,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75208,f30058])).
fof(f106082,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13))) ) | $spl608),
  inference(cnf_transformation,[],[f106082_D])).
fof(f106082_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13))) ) <=> ~$spl608),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl608])])).
fof(f106057,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f75208,f28884])).
fof(f105759,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f105725])).
fof(f105725,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X12))) )),
  inference(resolution,[],[f964,f465])).
fof(f964,plain,(
  ( ! [X140,X138,X139] : (~sP9(X139,X140,'const.bool.IN_2'(s(sK12,X138),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X139)) | ~p(s(bool,X140)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X138))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X138)))) )),
  inference(resolution,[],[f864,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f284])).
fof(f20209,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl114),
  inference(cnf_transformation,[],[f20209_D])).
fof(f20209_D,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f104100,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72645,f373])).
fof(f20208,plain,(
  ( ! [X21] : (p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl112),
  inference(cnf_transformation,[],[f20208_D])).
fof(f20208_D,plain,(
  ( ! [X21] : (p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f104099,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72645,f372])).
fof(f20204,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl104),
  inference(cnf_transformation,[],[f20204_D])).
fof(f20204_D,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f104094,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72645,f363])).
fof(f20203,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl102),
  inference(cnf_transformation,[],[f20203_D])).
fof(f20203_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f104093,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72645,f362])).
fof(f20202,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl100),
  inference(cnf_transformation,[],[f20202_D])).
fof(f20202_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f104092,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72645,f361])).
fof(f20201,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl98),
  inference(cnf_transformation,[],[f20201_D])).
fof(f20201_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f104091,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72645,f360])).
fof(f103882,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72644,f373])).
fof(f103881,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72644,f372])).
fof(f20207,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl110),
  inference(cnf_transformation,[],[f20207_D])).
fof(f20207_D,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f103880,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl8),
  inference(resolution,[],[f72644,f371])).
fof(f103876,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72644,f363])).
fof(f103875,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72644,f362])).
fof(f103874,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72644,f361])).
fof(f103873,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72644,f360])).
fof(f103664,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72643,f373])).
fof(f103663,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72643,f372])).
fof(f103662,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl8),
  inference(resolution,[],[f72643,f371])).
fof(f20206,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl108),
  inference(cnf_transformation,[],[f20206_D])).
fof(f20206_D,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f103661,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl8),
  inference(resolution,[],[f72643,f369])).
fof(f103658,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72643,f363])).
fof(f103657,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72643,f362])).
fof(f103656,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72643,f361])).
fof(f103655,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72643,f360])).
fof(f103457,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72642,f373])).
fof(f103456,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72642,f372])).
fof(f103455,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl8),
  inference(resolution,[],[f72642,f371])).
fof(f103454,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl8),
  inference(resolution,[],[f72642,f369])).
fof(f20205,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) | $spl106),
  inference(cnf_transformation,[],[f20205_D])).
fof(f20205_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f103453,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) | $spl8),
  inference(resolution,[],[f72642,f366])).
fof(f103451,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72642,f363])).
fof(f103450,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72642,f362])).
fof(f103449,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72642,f361])).
fof(f103448,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl8),
  inference(resolution,[],[f72642,f360])).
fof(f103259,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl134),
  inference(resolution,[],[f25888,f373])).
fof(f103258,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl134),
  inference(resolution,[],[f25888,f372])).
fof(f103257,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl134),
  inference(resolution,[],[f25888,f371])).
fof(f103256,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl134),
  inference(resolution,[],[f25888,f369])).
fof(f103255,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) | $spl134),
  inference(resolution,[],[f25888,f366])).
fof(f102855,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X13,sK15,sK12))))) | ~p(s(bool,X13))) ) | $spl604),
  inference(cnf_transformation,[],[f102855_D])).
fof(f102855_D,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X13,sK15,sK12))))) | ~p(s(bool,X13))) ) <=> ~$spl604),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl604])])).
fof(f103254,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X13,sK15,sK12))))) | ~p(s(bool,X13))) ) | $spl134),
  inference(resolution,[],[f25888,f364])).
fof(f103253,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl134),
  inference(resolution,[],[f25888,f363])).
fof(f103252,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl134),
  inference(resolution,[],[f25888,f362])).
fof(f103251,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl134),
  inference(resolution,[],[f25888,f361])).
fof(f103250,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl134),
  inference(resolution,[],[f25888,f360])).
fof(f103052,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl132),
  inference(resolution,[],[f25886,f373])).
fof(f103051,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl132),
  inference(resolution,[],[f25886,f372])).
fof(f103050,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl132),
  inference(resolution,[],[f25886,f371])).
fof(f103049,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl132),
  inference(resolution,[],[f25886,f369])).
fof(f103048,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) | $spl132),
  inference(resolution,[],[f25886,f366])).
fof(f103047,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X13,sK15,sK12))))) | ~p(s(bool,X13))) ) | $spl132),
  inference(resolution,[],[f25886,f364])).
fof(f103046,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl132),
  inference(resolution,[],[f25886,f363])).
fof(f103045,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl132),
  inference(resolution,[],[f25886,f362])).
fof(f103044,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl132),
  inference(resolution,[],[f25886,f361])).
fof(f103043,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl132),
  inference(resolution,[],[f25886,f360])).
fof(f102876,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl130),
  inference(resolution,[],[f25884,f373])).
fof(f102875,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl130),
  inference(resolution,[],[f25884,f372])).
fof(f102874,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl130),
  inference(resolution,[],[f25884,f371])).
fof(f102873,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl130),
  inference(resolution,[],[f25884,f369])).
fof(f102872,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) | $spl130),
  inference(resolution,[],[f25884,f366])).
fof(f102871,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X13,sK15,sK12))))) | ~p(s(bool,X13))) ) | $spl130),
  inference(resolution,[],[f25884,f364])).
fof(f102870,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl130),
  inference(resolution,[],[f25884,f363])).
fof(f102869,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl130),
  inference(resolution,[],[f25884,f362])).
fof(f102868,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl130),
  inference(resolution,[],[f25884,f361])).
fof(f102867,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl130),
  inference(resolution,[],[f25884,f360])).
fof(f102860,plain,(
  $false | ($spl128 | $spl573)),
  inference(subsumption_resolution,[],[f102859,f582])).
fof(f102859,plain,(
  ~p(s(bool,'T_0')) | ($spl128 | $spl573)),
  inference(subsumption_resolution,[],[f102834,f572])).
fof(f572,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,'F_0'))) => ~p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.IMP_F')).
fof(f102834,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl128 | $spl573)),
  inference(superposition,[],[f25882,f91265])).
fof(f91265,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | $spl573),
  inference(resolution,[],[f89465,f558])).
fof(f102740,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl128),
  inference(resolution,[],[f25882,f373])).
fof(f102739,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl128),
  inference(resolution,[],[f25882,f372])).
fof(f102738,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl128),
  inference(resolution,[],[f25882,f371])).
fof(f102737,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl128),
  inference(resolution,[],[f25882,f369])).
fof(f102736,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) | $spl128),
  inference(resolution,[],[f25882,f366])).
fof(f102735,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X13,sK15,sK12))))) | ~p(s(bool,X13))) ) | $spl128),
  inference(resolution,[],[f25882,f364])).
fof(f102734,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl128),
  inference(resolution,[],[f25882,f363])).
fof(f102733,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl128),
  inference(resolution,[],[f25882,f362])).
fof(f102732,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl128),
  inference(resolution,[],[f25882,f361])).
fof(f102731,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl128),
  inference(resolution,[],[f25882,f360])).
fof(f102728,plain,(
  $false | ($spl573 | $spl595)),
  inference(subsumption_resolution,[],[f91101,f91265])).
fof(f87841,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,X31) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl570),
  inference(cnf_transformation,[],[f87841_D])).
fof(f87841_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,X31) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl570),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl570])])).
fof(f102526,plain,(
  ( ! [X94,X95] : (s(bool,X94) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X95))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X95))) | s(bool,'F_0') != s(bool,X94)) ) | $spl7),
  inference(subsumption_resolution,[],[f102488,f864])).
fof(f102488,plain,(
  ( ! [X94,X95] : (s(bool,X94) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X95))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X95))) | s(bool,'F_0') != s(bool,X94) | p(s(bool,'const.bool.IN_2'(s(sK12,X95),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl7),
  inference(resolution,[],[f977,f67193])).
fof(f67193,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl7),
  inference(superposition,[],[f707,f464])).
fof(f102517,plain,(
  ( ! [X70] : (s(bool,X70) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X70)) ) | $spl602),
  inference(cnf_transformation,[],[f102517_D])).
fof(f102517_D,plain,(
  ( ! [X70] : (s(bool,X70) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X70)) ) <=> ~$spl602),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl602])])).
fof(f102525,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X93))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X93))) | s(bool,'T_0') = s(bool,X92)) ) | $spl8),
  inference(subsumption_resolution,[],[f102487,f864])).
fof(f102487,plain,(
  ( ! [X92,X93] : (s(bool,X92) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X93))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X93))) | s(bool,'T_0') = s(bool,X92) | p(s(bool,'const.bool.IN_2'(s(sK12,X93),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl8),
  inference(resolution,[],[f977,f67301])).
fof(f67301,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f708,f464])).
fof(f102518,plain,(
  ( ! [X72,X73] : (s(bool,X72) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X73))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X73))) | s(bool,'F_0') != s(bool,X72)) ) | $spl7),
  inference(subsumption_resolution,[],[f102477,f864])).
fof(f102477,plain,(
  ( ! [X72,X73] : (s(bool,X72) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X73))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X73))) | s(bool,'F_0') != s(bool,X72) | p(s(bool,'const.bool.IN_2'(s(sK12,X73),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl7),
  inference(resolution,[],[f977,f67192])).
fof(f67192,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl7),
  inference(superposition,[],[f707,f464])).
fof(f102516,plain,(
  ( ! [X70,X71] : (s(bool,X70) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X71))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X71))) | s(bool,'T_0') = s(bool,X70)) ) | $spl8),
  inference(subsumption_resolution,[],[f102476,f864])).
fof(f102476,plain,(
  ( ! [X70,X71] : (s(bool,X70) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X71))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X71))) | s(bool,'T_0') = s(bool,X70) | p(s(bool,'const.bool.IN_2'(s(sK12,X71),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl8),
  inference(resolution,[],[f977,f67300])).
fof(f67300,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) ) | $spl8),
  inference(superposition,[],[f708,f464])).
fof(f91508,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl22),
  inference(resolution,[],[f972,f29192])).
fof(f972,plain,(
  ( ! [X163,X164,X162] : (sP7('const.bool.IN_2'(s(sK12,X162),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X163,X164) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X162))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X162)))) )),
  inference(resolution,[],[f864,f459])).
fof(f90089,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl580),
  inference(cnf_transformation,[],[f90089_D])).
fof(f90089_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) <=> ~$spl580),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl580])])).
fof(f91493,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl573)),
  inference(resolution,[],[f89774,f29192])).
fof(f89774,plain,(
  ( ! [X107,X108] : (sP7(happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),X107,X108)) ) | $spl573),
  inference(resolution,[],[f89465,f459])).
fof(f89980,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | $spl577),
  inference(cnf_transformation,[],[f89980_D])).
fof(f89980_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) <=> ~$spl577),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl577])])).
fof(f91415,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X20))) ) | ($spl19 | $spl573)),
  inference(resolution,[],[f89770,f73790])).
fof(f89770,plain,(
  ( ! [X99,X100] : (sP8(X99,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),X100)) ) | $spl573),
  inference(resolution,[],[f89465,f454])).
fof(f91414,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl573)),
  inference(resolution,[],[f89770,f30061])).
fof(f89981,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X9))) ) | $spl578),
  inference(cnf_transformation,[],[f89981_D])).
fof(f89981_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X9))) ) <=> ~$spl578),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl578])])).
fof(f91409,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl573)),
  inference(resolution,[],[f89770,f28875])).
fof(f91101,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | $spl595),
  inference(cnf_transformation,[],[f91101_D])).
fof(f91101_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) <=> ~$spl595),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl595])])).
fof(f91177,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl19 | $spl22 | $spl572)),
  inference(resolution,[],[f90168,f73680])).
fof(f90168,plain,(
  ( ! [X136,X137] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),X136,X137)) ) | $spl572),
  inference(resolution,[],[f89466,f463])).
fof(f89601,plain,(
  ( ! [X319] : (s(bool,X319) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X319))) ) | $spl574),
  inference(cnf_transformation,[],[f89601_D])).
fof(f89601_D,plain,(
  ( ! [X319] : (s(bool,X319) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X319))) ) <=> ~$spl574),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl574])])).
fof(f91171,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl572)),
  inference(resolution,[],[f90168,f29451])).
fof(f6058,plain,(
  ( ! [X14] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl28),
  inference(cnf_transformation,[],[f6058_D])).
fof(f6058_D,plain,(
  ( ! [X14] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f91131,plain,(
  ( ! [X30,X31] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X30))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X30))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X30),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X31))) ) | $spl19),
  inference(resolution,[],[f968,f73790])).
fof(f968,plain,(
  ( ! [X152,X151,X150] : (sP8(X151,'const.bool.IN_2'(s(sK12,X150),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X152) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X150))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X150)))) )),
  inference(resolution,[],[f864,f454])).
fof(f91130,plain,(
  ( ! [X28,X29,X27] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,'F_0') != s(bool,X29)) ) | $spl22),
  inference(resolution,[],[f968,f30061])).
fof(f91104,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl600),
  inference(cnf_transformation,[],[f91104_D])).
fof(f91104_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) <=> ~$spl600),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl600])])).
fof(f91086,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl572)),
  inference(resolution,[],[f90162,f28678])).
fof(f90162,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl572),
  inference(resolution,[],[f89466,f453])).
fof(f91103,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),'T_0')) ) | $spl598),
  inference(cnf_transformation,[],[f91103_D])).
fof(f91103_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),'T_0')) ) <=> ~$spl598),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl598])])).
fof(f91085,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl572)),
  inference(resolution,[],[f90162,f28027])).
fof(f91084,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl572)),
  inference(resolution,[],[f90162,f30061])).
fof(f91083,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl572)),
  inference(resolution,[],[f90162,f30058])).
fof(f91102,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),X11) | ~p(s(bool,X11))) ) | $spl596),
  inference(cnf_transformation,[],[f91102_D])).
fof(f91102_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl596),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl596])])).
fof(f91080,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl572)),
  inference(resolution,[],[f90162,f28884])).
fof(f90923,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | $spl585),
  inference(cnf_transformation,[],[f90923_D])).
fof(f90923_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) <=> ~$spl585),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl585])])).
fof(f90991,plain,(
  ( ! [X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl19 | $spl22 | $spl198)),
  inference(resolution,[],[f89558,f73680])).
fof(f89558,plain,(
  ( ! [X253,X252] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),X252,X253)) ) | $spl198),
  inference(subsumption_resolution,[],[f89323,f595])).
fof(f89323,plain,(
  ( ! [X253,X252] : (~p(s(bool,'T_0')) | sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),X252,X253)) ) | $spl198),
  inference(superposition,[],[f463,f33478])).
fof(f90997,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl592),
  inference(cnf_transformation,[],[f90997_D])).
fof(f90997_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) <=> ~$spl592),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl592])])).
fof(f90985,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl198)),
  inference(resolution,[],[f89558,f29451])).
fof(f90926,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl590),
  inference(cnf_transformation,[],[f90926_D])).
fof(f90926_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) <=> ~$spl590),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl590])])).
fof(f90908,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl198)),
  inference(resolution,[],[f89552,f28678])).
fof(f89552,plain,(
  ( ! [X225,X224] : (sP8(X224,X225,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl198),
  inference(subsumption_resolution,[],[f89309,f595])).
fof(f89309,plain,(
  ( ! [X225,X224] : (~p(s(bool,'T_0')) | sP8(X224,X225,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl198),
  inference(superposition,[],[f453,f33478])).
fof(f90925,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),'T_0')) ) | $spl588),
  inference(cnf_transformation,[],[f90925_D])).
fof(f90925_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),'T_0')) ) <=> ~$spl588),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl588])])).
fof(f90907,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl198)),
  inference(resolution,[],[f89552,f28027])).
fof(f90906,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl198)),
  inference(resolution,[],[f89552,f30061])).
fof(f90905,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl198)),
  inference(resolution,[],[f89552,f30058])).
fof(f90924,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),X11) | ~p(s(bool,X11))) ) | $spl586),
  inference(cnf_transformation,[],[f90924_D])).
fof(f90924_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl586),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl586])])).
fof(f90902,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl198)),
  inference(resolution,[],[f89552,f28884])).
fof(f90092,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | $spl582),
  inference(cnf_transformation,[],[f90092_D])).
fof(f90092_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) <=> ~$spl582),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl582])])).
fof(f90675,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X2))) ) | $spl198),
  inference(subsumption_resolution,[],[f90668,f543])).
fof(f543,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f328])).
fof(f328,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X16] : ((p(s(bool,X16)) | p(s(bool,'F_0'))) <=> p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.OR_CLAUSES|split|3')).
fof(f90668,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X2))) ) | $spl198),
  inference(resolution,[],[f89572,f498])).
fof(f498,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f306])).
fof(f306,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f89572,plain,(
  ( ! [X297] : (sP11(happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),X297)) ) | $spl198),
  inference(subsumption_resolution,[],[f89355,f595])).
fof(f89355,plain,(
  ( ! [X297] : (~p(s(bool,'T_0')) | sP11(happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))),X297)) ) | $spl198),
  inference(superposition,[],[f500,f33478])).
fof(f500,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f307])).
fof(f90674,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X1))) ) | $spl198),
  inference(subsumption_resolution,[],[f90667,f543])).
fof(f90667,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X1))) ) | $spl198),
  inference(resolution,[],[f89572,f502])).
fof(f502,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f309])).
fof(f309,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(flattening,[],[f308])).
fof(f308,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP11(X0,X1)))),
  inference(nnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : (sP11(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f192,f232])).
fof(f192,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f191])).
fof(f191,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X12,X6] : ((~(p(s(bool,X6)) | p(s(bool,X12))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X12)) => p(s(bool,'F_0')))))),
  file('thm.list.EXISTS_CONG','thm.sat.sat_9_prove_ltb')).
fof(f90666,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X0))) ) | ($spl198 | $spl200)),
  inference(resolution,[],[f89572,f33481])).
fof(f90649,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(subsumption_resolution,[],[f90642,f543])).
fof(f90642,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(resolution,[],[f89571,f498])).
fof(f89571,plain,(
  ( ! [X296] : (sP11(X296,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl198),
  inference(subsumption_resolution,[],[f89354,f595])).
fof(f89354,plain,(
  ( ! [X296] : (~p(s(bool,'T_0')) | sP11(X296,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl198),
  inference(superposition,[],[f499,f33478])).
fof(f499,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f307])).
fof(f90648,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(subsumption_resolution,[],[f90641,f543])).
fof(f90641,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(resolution,[],[f89571,f502])).
fof(f90640,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | ($spl198 | $spl200)),
  inference(resolution,[],[f89571,f33481])).
fof(f90630,plain,(
  $false | ($spl26 | $spl199)),
  inference(subsumption_resolution,[],[f90602,f582])).
fof(f90602,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl199)),
  inference(superposition,[],[f90482,f4759])).
fof(f4759,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | $spl26),
  inference(resolution,[],[f4439,f562])).
fof(f90482,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))))) ) | $spl199),
  inference(subsumption_resolution,[],[f90371,f595])).
fof(f90371,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))))) ) | $spl199),
  inference(resolution,[],[f89655,f364])).
fof(f89655,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | $spl199),
  inference(trivial_inequality_removal,[],[f89646])).
fof(f89646,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | $spl199),
  inference(superposition,[],[f33477,f564])).
fof(f33477,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | $spl199),
  inference(cnf_transformation,[],[f33477_D])).
fof(f33477_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f90629,plain,(
  $false | ($spl24 | $spl199)),
  inference(subsumption_resolution,[],[f90601,f582])).
fof(f90601,plain,(
  ~p(s(bool,'T_0')) | ($spl24 | $spl199)),
  inference(superposition,[],[f90482,f4630])).
fof(f4630,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | $spl24),
  inference(resolution,[],[f4421,f562])).
fof(f90628,plain,(
  $false | ($spl199 | $spl568)),
  inference(subsumption_resolution,[],[f90600,f582])).
fof(f90600,plain,(
  ~p(s(bool,'T_0')) | ($spl199 | $spl568)),
  inference(superposition,[],[f90482,f83900])).
fof(f90627,plain,(
  $false | ($spl199 | $spl566)),
  inference(subsumption_resolution,[],[f90599,f582])).
fof(f90599,plain,(
  ~p(s(bool,'T_0')) | ($spl199 | $spl566)),
  inference(superposition,[],[f90482,f83868])).
fof(f90093,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | $spl583),
  inference(cnf_transformation,[],[f90093_D])).
fof(f90093_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) <=> ~$spl583),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl583])])).
fof(f89659,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | ~p(s(bool,X10))) ) | $spl199),
  inference(subsumption_resolution,[],[f89641,f562])).
fof(f89641,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | ~p(s(bool,X10))) ) | $spl199),
  inference(superposition,[],[f33477,f518])).
fof(f518,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f202])).
fof(f89661,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl199),
  inference(subsumption_resolution,[],[f89642,f562])).
fof(f89642,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl199),
  inference(superposition,[],[f33477,f518])).
fof(f90076,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl573)),
  inference(resolution,[],[f89774,f29192])).
fof(f89966,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X20))) ) | ($spl19 | $spl573)),
  inference(resolution,[],[f89770,f73790])).
fof(f89965,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl573)),
  inference(resolution,[],[f89770,f30061])).
fof(f89960,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) ) | ($spl22 | $spl573)),
  inference(resolution,[],[f89770,f28875])).
fof(f89465,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | $spl573),
  inference(cnf_transformation,[],[f89465_D])).
fof(f89465_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) <=> ~$spl573),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl573])])).
fof(f89662,plain,(
  ( ! [X11] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | ~p(s(bool,X11))) ) | ($spl46 | $spl199)),
  inference(forward_demodulation,[],[f89661,f10721])).
fof(f89660,plain,(
  ( ! [X10] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | ~p(s(bool,X10))) ) | ($spl46 | $spl199)),
  inference(forward_demodulation,[],[f89659,f10721])).
fof(f89466,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | $spl572),
  inference(cnf_transformation,[],[f89466_D])).
fof(f89466_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) <=> ~$spl572),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl572])])).
fof(f89605,plain,(
  ( ! [X326] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | s(bool,'T_0') != s(bool,X326)) ) | ($spl22 | $spl46 | $spl198)),
  inference(forward_demodulation,[],[f89604,f10721])).
fof(f89604,plain,(
  ( ! [X326] : (s(bool,'T_0') != s(bool,X326) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | ($spl22 | $spl198)),
  inference(subsumption_resolution,[],[f89407,f2929])).
fof(f89407,plain,(
  ( ! [X326] : (s(bool,'T_0') != s(bool,X326) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | s(bool,'F_0') = s(bool,X326)) ) | ($spl22 | $spl198)),
  inference(superposition,[],[f8572,f33478])).
fof(f89603,plain,(
  ( ! [X324] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | s(bool,'T_0') != s(bool,X324)) ) | ($spl22 | $spl46 | $spl198)),
  inference(forward_demodulation,[],[f89602,f10721])).
fof(f89602,plain,(
  ( ! [X324] : (s(bool,'T_0') != s(bool,X324) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | ($spl22 | $spl198)),
  inference(subsumption_resolution,[],[f89405,f561])).
fof(f89405,plain,(
  ( ! [X324] : (s(bool,'T_0') != s(bool,X324) | ~p(s(bool,X324)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | ($spl22 | $spl198)),
  inference(superposition,[],[f8469,f33478])).
fof(f89600,plain,(
  ( ! [X318,X319] : (s(bool,X319) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X318)) | p(s(bool,X319))) ) | ($spl22 | $spl46 | $spl198)),
  inference(forward_demodulation,[],[f89599,f10721])).
fof(f89599,plain,(
  ( ! [X318,X319] : (s(bool,X319) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X318)) | p(s(bool,X319))) ) | ($spl22 | $spl198)),
  inference(subsumption_resolution,[],[f89401,f562])).
fof(f89401,plain,(
  ( ! [X318,X319] : (s(bool,'T_0') != s(bool,X318) | s(bool,X319) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X318)) | p(s(bool,X319))) ) | ($spl22 | $spl198)),
  inference(superposition,[],[f8409,f33478])).
fof(f8409,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,X6) != s(bool,X7) | ~p(s(bool,X8)) | p(s(bool,X6))) ) | $spl22),
  inference(superposition,[],[f4474,f560])).
fof(f89581,plain,(
  ( ! [X303] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X303))) ) | ($spl46 | $spl198)),
  inference(forward_demodulation,[],[f89580,f10721])).
fof(f89580,plain,(
  ( ! [X303] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X303))) ) | $spl198),
  inference(subsumption_resolution,[],[f89361,f561])).
fof(f89361,plain,(
  ( ! [X303] : (s(bool,'T_0') = s(bool,X303) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X303))) ) | $spl198),
  inference(superposition,[],[f515,f33478])).
fof(f89579,plain,(
  ( ! [X302] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X302))) ) | ($spl46 | $spl198)),
  inference(forward_demodulation,[],[f89578,f10721])).
fof(f89578,plain,(
  ( ! [X302] : (p(s(bool,X302)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(subsumption_resolution,[],[f89360,f561])).
fof(f89360,plain,(
  ( ! [X302] : (s(bool,'T_0') = s(bool,X302) | p(s(bool,X302)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(superposition,[],[f515,f33478])).
fof(f89468,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X9))) ) | ($spl46 | $spl198)),
  inference(forward_demodulation,[],[f89467,f10721])).
fof(f89467,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X9))) ) | $spl198),
  inference(subsumption_resolution,[],[f89180,f561])).
fof(f89180,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,X9) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X9))) ) | $spl198),
  inference(superposition,[],[f33478,f515])).
fof(f89464,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK16,sK12))))) | p(s(bool,X8))) ) | ($spl46 | $spl198)),
  inference(forward_demodulation,[],[f89463,f10721])).
fof(f89463,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(subsumption_resolution,[],[f89179,f561])).
fof(f89179,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,X8) | p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))))) ) | $spl198),
  inference(superposition,[],[f33478,f515])).
fof(f25890,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X23,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12)))))) ) | $spl136),
  inference(cnf_transformation,[],[f25890_D])).
fof(f25890_D,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X23,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12)))))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f25873,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6))))) ) | $spl118),
  inference(cnf_transformation,[],[f25873_D])).
fof(f25873_D,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f88774,plain,(
  ( ! [X56,X55] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X55,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X55),s(sK12,sK26(sK16,X55,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88489,f10721])).
fof(f88489,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X55,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X55),s(sK12,sK26(sK16,X55,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56))))) ) | $spl46),
  inference(superposition,[],[f372,f10721])).
fof(f25888,plain,(
  ( ! [X21] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X21,sK12))))) | ~p(s(bool,X21))) ) | $spl134),
  inference(cnf_transformation,[],[f25888_D])).
fof(f25888_D,plain,(
  ( ! [X21] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X21,sK12))))) | ~p(s(bool,X21))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f88772,plain,(
  ( ! [X52,X51] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X52)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X51,sK12))))) | ~p(s(bool,X51))) ) | $spl46),
  inference(forward_demodulation,[],[f88487,f10721])).
fof(f88487,plain,(
  ( ! [X52,X51] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X51,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X52)))) | ~p(s(bool,X51))) ) | $spl46),
  inference(superposition,[],[f371,f10721])).
fof(f25886,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl132),
  inference(cnf_transformation,[],[f25886_D])).
fof(f25886_D,plain,(
  ( ! [X19] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f88771,plain,(
  ( ! [X50,X49] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X50)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X49,sK12))))) | ~p(s(bool,X49))) ) | $spl46),
  inference(forward_demodulation,[],[f88486,f10721])).
fof(f88486,plain,(
  ( ! [X50,X49] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X49,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X50)))) | ~p(s(bool,X49))) ) | $spl46),
  inference(superposition,[],[f369,f10721])).
fof(f25884,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X17,sK16,sK12))))) | ~p(s(bool,X17))) ) | $spl130),
  inference(cnf_transformation,[],[f25884_D])).
fof(f25884_D,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X17,sK16,sK12))))) | ~p(s(bool,X17))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f88770,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X48)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X47,sK16,sK12))))) | ~p(s(bool,X47))) ) | $spl46),
  inference(forward_demodulation,[],[f88485,f10721])).
fof(f88485,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X47,sK16,sK12))))) | ~p(s(bool,X47)) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X48))))) ) | $spl46),
  inference(superposition,[],[f366,f10721])).
fof(f25882,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) | $spl128),
  inference(cnf_transformation,[],[f25882_D])).
fof(f25882_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f88769,plain,(
  ( ! [X45,X46] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X45,sK16,sK12))))) | ~p(s(bool,X45))) ) | $spl46),
  inference(forward_demodulation,[],[f88484,f10721])).
fof(f88484,plain,(
  ( ! [X45,X46] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X45,sK16,sK12))))) | ~p(s(bool,X45)) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))))) ) | $spl46),
  inference(superposition,[],[f364,f10721])).
fof(f25878,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X11,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(X11,sK16,sK12)))))) ) | $spl124),
  inference(cnf_transformation,[],[f25878_D])).
fof(f25878_D,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X11,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(X11,sK16,sK12)))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f88768,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X41,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X41),s(sK12,sK20(X41,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88482,f10721])).
fof(f88482,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X41,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X41),s(sK12,sK20(X41,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42))))) ) | $spl46),
  inference(superposition,[],[f362,f10721])).
fof(f25880,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X13,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK20(sK16,X13,sK12)))))) ) | $spl126),
  inference(cnf_transformation,[],[f25880_D])).
fof(f25880_D,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X13,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK20(sK16,X13,sK12)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f88767,plain,(
  ( ! [X39,X40] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X40)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X39,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X39),s(sK12,sK20(sK16,X39,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88481,f10721])).
fof(f88481,plain,(
  ( ! [X39,X40] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X39,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X39),s(sK12,sK20(sK16,X39,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X40))))) ) | $spl46),
  inference(superposition,[],[f363,f10721])).
fof(f25876,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X9,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK18(sK16,X9,sK12)))))) ) | $spl122),
  inference(cnf_transformation,[],[f25876_D])).
fof(f25876_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X9,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK18(sK16,X9,sK12)))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f88766,plain,(
  ( ! [X35,X36] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X36)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X35,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X35),s(sK12,sK18(sK16,X35,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88479,f10721])).
fof(f88479,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X35,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X35),s(sK12,sK18(sK16,X35,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X36))))) ) | $spl46),
  inference(superposition,[],[f361,f10721])).
fof(f25874,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X7,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK19(X7,sK16,sK12)))))) ) | $spl120),
  inference(cnf_transformation,[],[f25874_D])).
fof(f25874_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X7,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK19(X7,sK16,sK12)))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f88765,plain,(
  ( ! [X30,X29] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X30)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X29,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X29),s(sK12,sK19(X29,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88476,f10721])).
fof(f88476,plain,(
  ( ! [X30,X29] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X29,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X29),s(sK12,sK19(X29,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X30))))) ) | $spl46),
  inference(superposition,[],[f360,f10721])).
fof(f88472,plain,(
  ( ! [X24,X25] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X24)))) | p(s(bool,happ(s(fun(sK12,bool),X25),s(sK12,sK26(sK16,X25,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X25),s(sK12,sK25(sK16,X25,sK12)))))) ) | $spl46),
  inference(superposition,[],[f373,f10721])).
fof(f88762,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22)))) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88471,f10721])).
fof(f88471,plain,(
  ( ! [X23,X22] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22)))) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl46),
  inference(superposition,[],[f372,f10721])).
fof(f88761,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X21,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X20)))) | ~p(s(bool,X21))) ) | $spl46),
  inference(forward_demodulation,[],[f88470,f10721])).
fof(f88470,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X20)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X21,sK12))))) | ~p(s(bool,X21))) ) | $spl46),
  inference(superposition,[],[f371,f10721])).
fof(f88760,plain,(
  ( ! [X19,X18] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X19,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18)))) | ~p(s(bool,X19))) ) | $spl46),
  inference(forward_demodulation,[],[f88469,f10721])).
fof(f88469,plain,(
  ( ! [X19,X18] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl46),
  inference(superposition,[],[f369,f10721])).
fof(f88759,plain,(
  ( ! [X17,X16] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X17,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X16)))) | ~p(s(bool,X17))) ) | $spl46),
  inference(forward_demodulation,[],[f88468,f10721])).
fof(f88468,plain,(
  ( ! [X17,X16] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X16)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X17,sK16,sK12))))) | ~p(s(bool,X17))) ) | $spl46),
  inference(superposition,[],[f366,f10721])).
fof(f88758,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X15,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14)))) | ~p(s(bool,X15))) ) | $spl46),
  inference(forward_demodulation,[],[f88467,f10721])).
fof(f88467,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) | $spl46),
  inference(superposition,[],[f364,f10721])).
fof(f88757,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X13,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X12)))) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK20(sK16,X13,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88466,f10721])).
fof(f88466,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X12)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X13,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK20(sK16,X13,sK12)))))) ) | $spl46),
  inference(superposition,[],[f363,f10721])).
fof(f88756,plain,(
  ( ! [X10,X11] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X11,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X10)))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(X11,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88465,f10721])).
fof(f88465,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X10)))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(X11,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X11,sK16,sK12)))))) ) | $spl46),
  inference(superposition,[],[f362,f10721])).
fof(f88755,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X9,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X8)))) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK18(sK16,X9,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88464,f10721])).
fof(f88464,plain,(
  ( ! [X8,X9] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X8)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X9,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK18(sK16,X9,sK12)))))) ) | $spl46),
  inference(superposition,[],[f361,f10721])).
fof(f88754,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X7,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6)))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK19(X7,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f88463,f10721])).
fof(f88463,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6)))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK19(X7,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X7,sK16,sK12)))))) ) | $spl46),
  inference(superposition,[],[f360,f10721])).
fof(f87840,plain,(
  ( ! [X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X31) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X32))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X32)))) ) | $spl7),
  inference(subsumption_resolution,[],[f87817,f864])).
fof(f87817,plain,(
  ( ! [X31,X32] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s(sK12,X32),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X31) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X32))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X32)))) ) | $spl7),
  inference(resolution,[],[f67192,f977])).
fof(f83900,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | $spl568),
  inference(cnf_transformation,[],[f83900_D])).
fof(f83900_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) <=> ~$spl568),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl568])])).
fof(f83915,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | p(s(bool,X0))) ) | $spl8),
  inference(forward_demodulation,[],[f83903,f73011])).
fof(f73011,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) | $spl8),
  inference(resolution,[],[f67587,f562])).
fof(f67587,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) | $spl8),
  inference(subsumption_resolution,[],[f67318,f595])).
fof(f67318,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) | $spl8),
  inference(superposition,[],[f355,f708])).
fof(f83903,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12))))) ) | $spl8),
  inference(resolution,[],[f75359,f1031])).
fof(f1031,plain,(
  ( ! [X198,X199] : (~sP11('const.bool.IN_2'(s(sK12,X198),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X199) | p(s(bool,X199)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X198))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X198)))) )),
  inference(subsumption_resolution,[],[f987,f543])).
fof(f987,plain,(
  ( ! [X198,X199] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X198))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X198))) | p(s(bool,'F_0')) | p(s(bool,X199)) | ~sP11('const.bool.IN_2'(s(sK12,X198),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X199)) )),
  inference(resolution,[],[f864,f498])).
fof(f75359,plain,(
  ( ! [X149] : (sP11('const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X149)) ) | $spl8),
  inference(resolution,[],[f67586,f500])).
fof(f83898,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | p(s(bool,X0))) ) | $spl8),
  inference(forward_demodulation,[],[f83886,f73011])).
fof(f83886,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12))))) ) | $spl8),
  inference(resolution,[],[f75358,f1030])).
fof(f1030,plain,(
  ( ! [X196,X197] : (~sP11(X197,'const.bool.IN_2'(s(sK12,X196),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | p(s(bool,X197)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X196))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X196)))) )),
  inference(subsumption_resolution,[],[f986,f543])).
fof(f986,plain,(
  ( ! [X196,X197] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X196))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X196))) | p(s(bool,'F_0')) | p(s(bool,X197)) | ~sP11(X197,'const.bool.IN_2'(s(sK12,X196),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f864,f498])).
fof(f75358,plain,(
  ( ! [X148] : (sP11(X148,'const.bool.IN_2'(s(sK12,sK17(sK13,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl8),
  inference(resolution,[],[f67586,f499])).
fof(f83868,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | $spl566),
  inference(cnf_transformation,[],[f83868_D])).
fof(f83868_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) <=> ~$spl566),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl566])])).
fof(f83883,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | p(s(bool,X0))) ) | $spl8),
  inference(forward_demodulation,[],[f83871,f72732])).
fof(f72732,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) | $spl8),
  inference(resolution,[],[f67583,f562])).
fof(f83871,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12))))) ) | $spl8),
  inference(resolution,[],[f75226,f1031])).
fof(f75226,plain,(
  ( ! [X149] : (sP11('const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X149)) ) | $spl8),
  inference(resolution,[],[f67584,f500])).
fof(f83866,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | p(s(bool,X0))) ) | $spl8),
  inference(forward_demodulation,[],[f83854,f72732])).
fof(f83854,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12))))) ) | $spl8),
  inference(resolution,[],[f75225,f1030])).
fof(f75225,plain,(
  ( ! [X148] : (sP11(X148,'const.bool.IN_2'(s(sK12,sK17(sK14,sK15,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl8),
  inference(resolution,[],[f67584,f499])).
fof(f81555,plain,(
  $false | ($spl140 | $spl461)),
  inference(subsumption_resolution,[],[f71282,f80445])).
fof(f80445,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | $spl140),
  inference(resolution,[],[f25909,f562])).
fof(f81554,plain,(
  $false | ($spl138 | $spl449)),
  inference(subsumption_resolution,[],[f70678,f78526])).
fof(f78526,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | $spl138),
  inference(resolution,[],[f25906,f562])).
fof(f70462,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | $spl441),
  inference(cnf_transformation,[],[f70462_D])).
fof(f70462_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) <=> ~$spl441),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl441])])).
fof(f80667,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl19 | $spl22 | $spl140)),
  inference(resolution,[],[f69498,f73680])).
fof(f69498,plain,(
  ( ! [X136,X137] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X136,X137)) ) | $spl140),
  inference(resolution,[],[f25909,f463])).
fof(f80115,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | $spl544),
  inference(cnf_transformation,[],[f80115_D])).
fof(f80115_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) <=> ~$spl544),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl544])])).
fof(f80660,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69498,f29451])).
fof(f70465,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | $spl446),
  inference(cnf_transformation,[],[f70465_D])).
fof(f70465_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f80576,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28678])).
fof(f69492,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | $spl140),
  inference(resolution,[],[f25909,f453])).
fof(f70464,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),'T_0')) ) | $spl444),
  inference(cnf_transformation,[],[f70464_D])).
fof(f70464_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),'T_0')) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f80575,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28027])).
fof(f80574,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f30061])).
fof(f80573,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f30058])).
fof(f70463,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X11) | ~p(s(bool,X11))) ) | $spl442),
  inference(cnf_transformation,[],[f70463_D])).
fof(f70463_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f80570,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28884])).
fof(f25909,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) | $spl140),
  inference(cnf_transformation,[],[f25909_D])).
fof(f25909_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f80323,plain,(
  ( ! [X4,X5] : (sP9(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X4,X5) | ~sP6(X5,X4,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | ~sP7(X4,X5,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | $spl564),
  inference(cnf_transformation,[],[f80323_D])).
fof(f80323_D,plain,(
  ( ! [X4,X5] : (sP9(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X4,X5) | ~sP6(X5,X4,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | ~sP7(X4,X5,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) <=> ~$spl564),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl564])])).
fof(f80295,plain,(
  ( ! [X4,X5] : (sP9(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X4,X5) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) | ~sP7(X4,X5,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | ~sP6(X5,X4,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | $spl141),
  inference(resolution,[],[f70517,f451])).
fof(f70517,plain,(
  ( ! [X101,X102] : (sP8(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X101,X102)) ) | $spl141),
  inference(resolution,[],[f25908,f455])).
fof(f25908,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) | $spl141),
  inference(cnf_transformation,[],[f25908_D])).
fof(f25908_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f71282,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | $spl461),
  inference(cnf_transformation,[],[f71282_D])).
fof(f71282_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) <=> ~$spl461),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl461])])).
fof(f80276,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | ~p(s(bool,X20))) ) | ($spl19 | $spl141)),
  inference(resolution,[],[f70516,f73790])).
fof(f70516,plain,(
  ( ! [X99,X100] : (sP8(X99,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X100)) ) | $spl141),
  inference(resolution,[],[f25908,f454])).
fof(f80275,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70516,f30061])).
fof(f71283,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X9))) ) | $spl462),
  inference(cnf_transformation,[],[f71283_D])).
fof(f71283_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X9))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f80270,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70516,f28875])).
fof(f80108,plain,(
  ( ! [X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl19 | $spl22 | $spl140)),
  inference(resolution,[],[f69498,f73680])).
fof(f80101,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69498,f29451])).
fof(f80046,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28678])).
fof(f80045,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28027])).
fof(f80044,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f30061])).
fof(f80043,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f30058])).
fof(f80040,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28884])).
fof(f79709,plain,(
  $false | ($spl141 | $spl441)),
  inference(subsumption_resolution,[],[f70462,f71476])).
fof(f71476,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | $spl141),
  inference(resolution,[],[f25908,f558])).
fof(f69957,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | $spl433),
  inference(cnf_transformation,[],[f69957_D])).
fof(f69957_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) <=> ~$spl433),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl433])])).
fof(f78784,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl19 | $spl22 | $spl138)),
  inference(resolution,[],[f69347,f73680])).
fof(f69347,plain,(
  ( ! [X136,X137] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X136,X137)) ) | $spl138),
  inference(resolution,[],[f25906,f463])).
fof(f78170,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | $spl520),
  inference(cnf_transformation,[],[f78170_D])).
fof(f78170_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f78778,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69347,f29451])).
fof(f69960,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | $spl438),
  inference(cnf_transformation,[],[f69960_D])).
fof(f69960_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f78672,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28678])).
fof(f69341,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | $spl138),
  inference(resolution,[],[f25906,f453])).
fof(f69959,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),'T_0')) ) | $spl436),
  inference(cnf_transformation,[],[f69959_D])).
fof(f69959_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),'T_0')) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f78671,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28027])).
fof(f78670,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f30061])).
fof(f78669,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f30058])).
fof(f69958,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X11) | ~p(s(bool,X11))) ) | $spl434),
  inference(cnf_transformation,[],[f69958_D])).
fof(f69958_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f78666,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28884])).
fof(f25906,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) | $spl138),
  inference(cnf_transformation,[],[f25906_D])).
fof(f25906_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f78401,plain,(
  ( ! [X4,X5] : (sP9(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X4,X5) | ~sP6(X5,X4,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | ~sP7(X4,X5,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | $spl540),
  inference(cnf_transformation,[],[f78401_D])).
fof(f78401_D,plain,(
  ( ! [X4,X5] : (sP9(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X4,X5) | ~sP6(X5,X4,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | ~sP7(X4,X5,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) <=> ~$spl540),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl540])])).
fof(f78374,plain,(
  ( ! [X4,X5] : (sP9(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X4,X5) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) | ~sP7(X4,X5,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | ~sP6(X5,X4,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | $spl139),
  inference(resolution,[],[f70012,f451])).
fof(f70012,plain,(
  ( ! [X101,X102] : (sP8(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X101,X102)) ) | $spl139),
  inference(resolution,[],[f25905,f455])).
fof(f25905,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) | $spl139),
  inference(cnf_transformation,[],[f25905_D])).
fof(f25905_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f70678,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | $spl449),
  inference(cnf_transformation,[],[f70678_D])).
fof(f70678_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) <=> ~$spl449),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl449])])).
fof(f78356,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | ~p(s(bool,X20))) ) | ($spl19 | $spl139)),
  inference(resolution,[],[f70011,f73790])).
fof(f70011,plain,(
  ( ! [X99,X100] : (sP8(X99,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X100)) ) | $spl139),
  inference(resolution,[],[f25905,f454])).
fof(f78355,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70011,f30061])).
fof(f70679,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X9))) ) | $spl450),
  inference(cnf_transformation,[],[f70679_D])).
fof(f70679_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X9))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f78350,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70011,f28875])).
fof(f78164,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl19 | $spl22 | $spl138)),
  inference(resolution,[],[f69347,f73680])).
fof(f78158,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69347,f29451])).
fof(f78107,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28678])).
fof(f78106,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28027])).
fof(f78105,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f30061])).
fof(f78104,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f30058])).
fof(f78101,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28884])).
fof(f77824,plain,(
  $false | ($spl139 | $spl433)),
  inference(subsumption_resolution,[],[f69957,f70963])).
fof(f70963,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | $spl139),
  inference(resolution,[],[f25905,f558])).
fof(f75271,plain,(
  ( ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6))))))) ) | $spl516),
  inference(cnf_transformation,[],[f75271_D])).
fof(f75271_D,plain,(
  ( ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6))))))) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f75402,plain,(
  s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | $spl518),
  inference(cnf_transformation,[],[f75402_D])).
fof(f75402_D,plain,(
  s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f75284,plain,(
  ( ! [X7] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X7)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X7))))))) ) | $spl8),
  inference(resolution,[],[f67586,f1008])).
fof(f1008,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),X0)))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X1))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X1))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X2))))))) )),
  inference(superposition,[],[f864,f583])).
fof(f75283,plain,(
  ( ! [X6] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6))))))) ) | $spl8),
  inference(resolution,[],[f67586,f1009])).
fof(f1009,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),X3)))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X4))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X4))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X3),s(fun(fun(sK12,bool),bool),X5))))))) )),
  inference(superposition,[],[f864,f583])).
fof(f75268,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4))))))) ) | $spl512),
  inference(cnf_transformation,[],[f75268_D])).
fof(f75268_D,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4))))))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f75282,plain,(
  ( ! [X5] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X5))))))) ) | $spl8),
  inference(resolution,[],[f67586,f1006])).
fof(f1006,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X1),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X0)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X1))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X1))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X0),s(fun('type.list.list'(sK12),bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X0),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X2))))))) )),
  inference(superposition,[],[f864,f583])).
fof(f75281,plain,(
  ( ! [X4] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4))))))) ) | $spl8),
  inference(resolution,[],[f67586,f1007])).
fof(f1007,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X4),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),X3)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X4))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X4))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X3),s(fun('type.list.list'(sK12),bool),X5))))))) )),
  inference(superposition,[],[f864,f583])).
fof(f75270,plain,(
  s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | $spl514),
  inference(cnf_transformation,[],[f75270_D])).
fof(f75270_D,plain,(
  s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f75151,plain,(
  ( ! [X7] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X7)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X7))))))) ) | $spl8),
  inference(resolution,[],[f67584,f1008])).
fof(f75150,plain,(
  ( ! [X6] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))),s(fun(fun(sK12,bool),bool),X6))))))) ) | $spl8),
  inference(resolution,[],[f67584,f1009])).
fof(f75149,plain,(
  ( ! [X5] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X5))))))) ) | $spl8),
  inference(resolution,[],[f67584,f1006])).
fof(f75148,plain,(
  ( ! [X4] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK15,sK12)))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X4))))))) ) | $spl8),
  inference(resolution,[],[f67584,f1007])).
fof(f65104,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X26)))))) ) | $spl374),
  inference(cnf_transformation,[],[f65104_D])).
fof(f65104_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X26)))))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f74194,plain,(
  ( ! [X39,X40] : (p(s(bool,X39)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X40)))))) ) | ($spl19 | $spl22 | $spl24)),
  inference(resolution,[],[f73680,f52898])).
fof(f52898,plain,(
  ( ! [X191,X189,X190] : (sP6('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X189)))),X190,X191)) ) | $spl24),
  inference(resolution,[],[f4534,f463])).
fof(f4534,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X0))))))) ) | $spl24),
  inference(resolution,[],[f4421,f352])).
fof(f352,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & ~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3)))))) & ((~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) & ~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) <=> (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X9,X17,X16] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X17),s('type.list.list'(X6),X16)))))) <=> (p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X17)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),X16))))))),
  file('thm.list.EXISTS_CONG','thm.list.EXISTS_DEF|split|1')).
fof(f66114,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X26)))))) ) | $spl380),
  inference(cnf_transformation,[],[f66114_D])).
fof(f66114_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X26)))))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f74193,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X38)))))) ) | ($spl19 | $spl22 | $spl26)),
  inference(resolution,[],[f73680,f55090])).
fof(f55090,plain,(
  ( ! [X191,X189,X190] : (sP6('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X189)))),X190,X191)) ) | $spl26),
  inference(resolution,[],[f4663,f463])).
fof(f4663,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X0))))))) ) | $spl26),
  inference(resolution,[],[f4439,f352])).
fof(f74067,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) | $spl503),
  inference(cnf_transformation,[],[f74067_D])).
fof(f74067_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) <=> ~$spl503),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl503])])).
fof(f74192,plain,(
  ( ! [X36] : (p(s(bool,X36)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f73680,f72989])).
fof(f72989,plain,(
  ( ! [X136,X137] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))),X136,X137)) ) | $spl8),
  inference(resolution,[],[f67587,f463])).
fof(f73932,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) | $spl493),
  inference(cnf_transformation,[],[f73932_D])).
fof(f73932_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) <=> ~$spl493),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl493])])).
fof(f74191,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f73680,f72710])).
fof(f72710,plain,(
  ( ! [X136,X137] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))),X136,X137)) ) | $spl8),
  inference(resolution,[],[f67583,f463])).
fof(f72191,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))) | $spl473),
  inference(cnf_transformation,[],[f72191_D])).
fof(f72191_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))) <=> ~$spl473),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl473])])).
fof(f74190,plain,(
  ( ! [X34] : (p(s(bool,X34)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl19 | $spl22 | $spl454)),
  inference(resolution,[],[f73680,f71883])).
fof(f71883,plain,(
  ( ! [X136,X137] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X136,X137)) ) | $spl454),
  inference(resolution,[],[f70843,f463])).
fof(f73144,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))) | $spl481),
  inference(cnf_transformation,[],[f73144_D])).
fof(f73144_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))) <=> ~$spl481),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl481])])).
fof(f74189,plain,(
  ( ! [X33] : (p(s(bool,X33)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | ($spl8 | $spl19 | $spl22)),
  inference(resolution,[],[f73680,f72840])).
fof(f72840,plain,(
  ( ! [X136,X137] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X136,X137)) ) | $spl8),
  inference(resolution,[],[f72649,f463])).
fof(f72649,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | $spl8),
  inference(resolution,[],[f67583,f375])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X0)))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f256,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK28(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK27,sK28])],[f255])).
fof(f255,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f254])).
fof(f254,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6] : ! [X9] : (~? [X11] : p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))) <=> ! [X11] : ~p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X11)))))),
  file('thm.list.EXISTS_CONG','thm.bool.NOT_EXISTS_THM')).
fof(f74158,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | ~p(s(bool,X11))) ) | ($spl19 | $spl141)),
  inference(resolution,[],[f73790,f70516])).
fof(f74157,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | ~p(s(bool,X10))) ) | ($spl19 | $spl139)),
  inference(resolution,[],[f73790,f70011])).
fof(f74149,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | $spl510),
  inference(cnf_transformation,[],[f74149_D])).
fof(f74149_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f74138,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72989,f29451])).
fof(f74070,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | $spl508),
  inference(cnf_transformation,[],[f74070_D])).
fof(f74070_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f74054,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72983,f28678])).
fof(f72983,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | $spl8),
  inference(resolution,[],[f67587,f453])).
fof(f74069,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))),'T_0')) ) | $spl506),
  inference(cnf_transformation,[],[f74069_D])).
fof(f74069_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))),'T_0')) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f74053,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72983,f28027])).
fof(f74052,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72983,f30061])).
fof(f74051,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72983,f30058])).
fof(f74068,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))),X11) | ~p(s(bool,X11))) ) | $spl504),
  inference(cnf_transformation,[],[f74068_D])).
fof(f74068_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f74048,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72983,f28884])).
fof(f73985,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | $spl500),
  inference(cnf_transformation,[],[f73985_D])).
fof(f73985_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f73974,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72710,f29451])).
fof(f73935,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | $spl498),
  inference(cnf_transformation,[],[f73935_D])).
fof(f73935_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f73919,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72704,f28678])).
fof(f72704,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | $spl8),
  inference(resolution,[],[f67583,f453])).
fof(f73934,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))),'T_0')) ) | $spl496),
  inference(cnf_transformation,[],[f73934_D])).
fof(f73934_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))),'T_0')) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f73918,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72704,f28027])).
fof(f73917,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72704,f30061])).
fof(f73916,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72704,f30058])).
fof(f73933,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))),X11) | ~p(s(bool,X11))) ) | $spl494),
  inference(cnf_transformation,[],[f73933_D])).
fof(f73933_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f73913,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72704,f28884])).
fof(f73369,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | $spl490),
  inference(cnf_transformation,[],[f73369_D])).
fof(f73369_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f73358,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71883,f29451])).
fof(f72194,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | $spl478),
  inference(cnf_transformation,[],[f72194_D])).
fof(f72194_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f73287,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f28678])).
fof(f71877,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | $spl454),
  inference(resolution,[],[f70843,f453])).
fof(f72193,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),'T_0')) ) | $spl476),
  inference(cnf_transformation,[],[f72193_D])).
fof(f72193_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),'T_0')) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f73286,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f28027])).
fof(f73285,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f30061])).
fof(f73284,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f30058])).
fof(f72192,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X11) | ~p(s(bool,X11))) ) | $spl474),
  inference(cnf_transformation,[],[f72192_D])).
fof(f72192_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f73281,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f28884])).
fof(f73197,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | $spl488),
  inference(cnf_transformation,[],[f73197_D])).
fof(f73197_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f73186,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72840,f29451])).
fof(f73147,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | $spl486),
  inference(cnf_transformation,[],[f73147_D])).
fof(f73147_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) <=> ~$spl486),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl486])])).
fof(f73131,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72834,f28678])).
fof(f72834,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | $spl8),
  inference(resolution,[],[f72649,f453])).
fof(f73146,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),'T_0')) ) | $spl484),
  inference(cnf_transformation,[],[f73146_D])).
fof(f73146_D,plain,(
  ( ! [X21] : (sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),'T_0')) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f73130,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72834,f28027])).
fof(f73129,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72834,f30061])).
fof(f73128,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72834,f30058])).
fof(f73145,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X11) | ~p(s(bool,X11))) ) | $spl482),
  inference(cnf_transformation,[],[f73145_D])).
fof(f73145_D,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X11) | ~p(s(bool,X11))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f73125,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f72834,f28884])).
fof(f71393,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | $spl466),
  inference(cnf_transformation,[],[f71393_D])).
fof(f71393_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f72500,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3)))) ) | $spl454),
  inference(resolution,[],[f72046,f1030])).
fof(f72046,plain,(
  ( ! [X151] : (sP11(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X151)) ) | $spl454),
  inference(resolution,[],[f71822,f500])).
fof(f71822,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | $spl454),
  inference(resolution,[],[f70843,f375])).
fof(f72506,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | p(s(bool,X2))) ) | $spl454),
  inference(subsumption_resolution,[],[f72499,f543])).
fof(f72499,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | p(s(bool,X2))) ) | $spl454),
  inference(resolution,[],[f72046,f498])).
fof(f72505,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | p(s(bool,X1))) ) | $spl454),
  inference(subsumption_resolution,[],[f72498,f543])).
fof(f72498,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | p(s(bool,X1))) ) | $spl454),
  inference(resolution,[],[f72046,f502])).
fof(f72497,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | p(s(bool,X0))) ) | ($spl200 | $spl454)),
  inference(resolution,[],[f72046,f33481])).
fof(f72490,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3)))) ) | $spl454),
  inference(resolution,[],[f72045,f1031])).
fof(f72045,plain,(
  ( ! [X150] : (sP11(X150,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | $spl454),
  inference(resolution,[],[f71822,f499])).
fof(f72496,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl454),
  inference(subsumption_resolution,[],[f72489,f543])).
fof(f72489,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl454),
  inference(resolution,[],[f72045,f498])).
fof(f72495,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl454),
  inference(subsumption_resolution,[],[f72488,f543])).
fof(f72488,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl454),
  inference(resolution,[],[f72045,f502])).
fof(f72487,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | ($spl200 | $spl454)),
  inference(resolution,[],[f72045,f33481])).
fof(f72466,plain,(
  $false | ($spl8 | $spl455)),
  inference(subsumption_resolution,[],[f67583,f72195])).
fof(f72195,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))))) ) | $spl455),
  inference(resolution,[],[f70842,f374])).
fof(f374,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK27(X1,X0)))))) )),
  inference(cnf_transformation,[],[f256])).
fof(f70842,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | $spl455),
  inference(cnf_transformation,[],[f70842_D])).
fof(f70842_D,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) <=> ~$spl455),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl455])])).
fof(f72178,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f28678])).
fof(f72177,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f28027])).
fof(f72176,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f30061])).
fof(f72175,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f30058])).
fof(f72172,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | ($spl22 | $spl454)),
  inference(resolution,[],[f71877,f28884])).
fof(f71800,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30))))))) ) | $spl470),
  inference(cnf_transformation,[],[f71800_D])).
fof(f71800_D,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30))))))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f71801,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl8),
  inference(subsumption_resolution,[],[f71783,f561])).
fof(f71783,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,X32)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl8),
  inference(superposition,[],[f67773,f518])).
fof(f67773,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X2),s(fun(fun(sK12,bool),bool),X3))))))) ) | $spl8),
  inference(superposition,[],[f67295,f583])).
fof(f71799,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13))))) ) | $spl8),
  inference(subsumption_resolution,[],[f71782,f561])).
fof(f71782,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X29))) ) | $spl8),
  inference(superposition,[],[f67773,f518])).
fof(f71339,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | $spl464),
  inference(cnf_transformation,[],[f71339_D])).
fof(f71339_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f71754,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70520,f29192])).
fof(f70520,plain,(
  ( ! [X107,X108] : (sP7(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X107,X108)) ) | $spl141),
  inference(resolution,[],[f25908,f459])).
fof(f71701,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70516,f30061])).
fof(f71696,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70516,f28875])).
fof(f71690,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30))))))) ) | $spl468),
  inference(cnf_transformation,[],[f71690_D])).
fof(f71690_D,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30))))))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f71691,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl8),
  inference(subsumption_resolution,[],[f71673,f561])).
fof(f71673,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,X32)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl8),
  inference(superposition,[],[f67772,f518])).
fof(f67772,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X1))))))) ) | $spl8),
  inference(superposition,[],[f67295,f583])).
fof(f71689,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13))))) ) | $spl8),
  inference(subsumption_resolution,[],[f71672,f561])).
fof(f71672,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X29))) ) | $spl8),
  inference(superposition,[],[f67772,f518])).
fof(f6301,plain,(
  ( ! [X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X53))))))) ) | $spl40),
  inference(cnf_transformation,[],[f6301_D])).
fof(f6301_D,plain,(
  ( ! [X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X53))))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f71571,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))))) | ~p(s(bool,X58))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f71554])).
fof(f71554,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X58) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))))) | ~p(s(bool,X58))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f71353,f562])).
fof(f71353,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X14),s(fun(bool,bool),X15))))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f71352,f711])).
fof(f71352,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X15)))) | s(bool,'T_0') != s(bool,X14) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X14),s(fun(bool,bool),X15))))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f67793,f711])).
fof(f67793,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X14),s(fun(bool,bool),X15))))))) ) | $spl8),
  inference(superposition,[],[f67295,f583])).
fof(f71572,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))))) | ~p(s(bool,X56))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f71553])).
fof(f71553,plain,(
  ( ! [X57,X56] : (s(bool,X56) != s(bool,X56) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))))) | ~p(s(bool,X56))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f71353,f564])).
fof(f6169,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,X52)) ) | $spl34),
  inference(cnf_transformation,[],[f6169_D])).
fof(f6169_D,plain,(
  ( ! [X52] : (s(bool,'F_0') = s(bool,X52)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f71573,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f71552])).
fof(f71552,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X54) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f71353,f587])).
fof(f71574,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f71551])).
fof(f71551,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X52) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f71353,f588])).
fof(f6299,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | ~p(s(bool,X22)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X22),s(fun(bool,bool),X24))))))) ) | $spl38),
  inference(cnf_transformation,[],[f6299_D])).
fof(f6299_D,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | ~p(s(bool,X22)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X22),s(fun(bool,bool),X24))))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f71582,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X25),s(fun(bool,bool),X27)))))) | ~p(s(bool,X25))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f71540,f561])).
fof(f71540,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X25),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f71353,f518])).
fof(f71581,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X22),s(fun(bool,bool),X24)))))) | ~p(s(bool,X22))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f71539,f561])).
fof(f71539,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X22),s(fun(bool,bool),X24)))))) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f71353,f518])).
fof(f71387,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3)))) ) | $spl140),
  inference(resolution,[],[f70316,f1031])).
fof(f70316,plain,(
  ( ! [X150] : (sP11(X150,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) ) | $spl140),
  inference(resolution,[],[f69437,f499])).
fof(f69437,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | $spl140),
  inference(resolution,[],[f25909,f375])).
fof(f71395,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl140),
  inference(subsumption_resolution,[],[f71386,f543])).
fof(f71386,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl140),
  inference(resolution,[],[f70316,f498])).
fof(f71394,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl140),
  inference(subsumption_resolution,[],[f71385,f543])).
fof(f71385,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | $spl140),
  inference(resolution,[],[f70316,f502])).
fof(f71384,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))))) ) | ($spl140 | $spl200)),
  inference(resolution,[],[f70316,f33481])).
fof(f70843,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | $spl454),
  inference(cnf_transformation,[],[f70843_D])).
fof(f70843_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f71377,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3)))) ) | $spl140),
  inference(resolution,[],[f70166,f1030])).
fof(f70166,plain,(
  ( ! [X151] : (sP11(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X151)) ) | $spl140),
  inference(resolution,[],[f69436,f500])).
fof(f69436,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | $spl140),
  inference(resolution,[],[f25909,f374])).
fof(f71383,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | p(s(bool,X2))) ) | $spl140),
  inference(subsumption_resolution,[],[f71376,f543])).
fof(f71376,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | p(s(bool,X2))) ) | $spl140),
  inference(resolution,[],[f70166,f498])).
fof(f71382,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | p(s(bool,X1))) ) | $spl140),
  inference(subsumption_resolution,[],[f71375,f543])).
fof(f71375,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | p(s(bool,X1))) ) | $spl140),
  inference(resolution,[],[f70166,f502])).
fof(f71374,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | p(s(bool,X0))) ) | ($spl140 | $spl200)),
  inference(resolution,[],[f70166,f33481])).
fof(f71367,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3)))) ) | $spl140),
  inference(resolution,[],[f70165,f1031])).
fof(f70165,plain,(
  ( ! [X150] : (sP11(X150,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | $spl140),
  inference(resolution,[],[f69436,f499])).
fof(f71373,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl140),
  inference(subsumption_resolution,[],[f71366,f543])).
fof(f71366,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl140),
  inference(resolution,[],[f70165,f498])).
fof(f71372,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl140),
  inference(subsumption_resolution,[],[f71365,f543])).
fof(f71365,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl140),
  inference(resolution,[],[f70165,f502])).
fof(f71364,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | ($spl140 | $spl200)),
  inference(resolution,[],[f70165,f33481])).
fof(f71326,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70520,f29192])).
fof(f71270,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70516,f30061])).
fof(f71265,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl141)),
  inference(resolution,[],[f70516,f28875])).
fof(f70811,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | $spl452),
  inference(cnf_transformation,[],[f70811_D])).
fof(f70811_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f71171,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70015,f29192])).
fof(f70015,plain,(
  ( ! [X107,X108] : (sP7(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X107,X108)) ) | $spl139),
  inference(resolution,[],[f25905,f459])).
fof(f71118,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70011,f30061])).
fof(f71113,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70011,f28875])).
fof(f71093,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29))))))) ) | $spl458),
  inference(cnf_transformation,[],[f71093_D])).
fof(f71093_D,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29))))))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f71094,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X30))))) ) | $spl8),
  inference(subsumption_resolution,[],[f71075,f561])).
fof(f71075,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,X31)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X30))))) ) | $spl8),
  inference(superposition,[],[f67775,f518])).
fof(f67775,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X2))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X2),s(fun('type.list.list'(sK12),bool),X3))))))) ) | $spl8),
  inference(superposition,[],[f67295,f583])).
fof(f71092,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27))))) ) | $spl8),
  inference(subsumption_resolution,[],[f71074,f561])).
fof(f71074,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27)))) | ~p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f67775,f518])).
fof(f70881,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29))))))) ) | $spl456),
  inference(cnf_transformation,[],[f70881_D])).
fof(f70881_D,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29))))))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f70882,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X30))))) ) | $spl8),
  inference(subsumption_resolution,[],[f70862,f561])).
fof(f70862,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,X31)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X30))))) ) | $spl8),
  inference(superposition,[],[f67774,f518])).
fof(f67774,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X0))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X0),s(fun('type.list.list'(sK12),bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X0),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X1))))))) ) | $spl8),
  inference(superposition,[],[f67295,f583])).
fof(f70880,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27))))) ) | $spl8),
  inference(subsumption_resolution,[],[f70861,f561])).
fof(f70861,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),X27)))) | ~p(s(bool,X28))) ) | $spl8),
  inference(superposition,[],[f67774,f518])).
fof(f70837,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3)))) ) | $spl138),
  inference(resolution,[],[f69660,f1031])).
fof(f69660,plain,(
  ( ! [X150] : (sP11(X150,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) ) | $spl138),
  inference(resolution,[],[f69285,f499])).
fof(f69285,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | $spl138),
  inference(resolution,[],[f25906,f374])).
fof(f70845,plain,(
  ( ! [X2] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl138),
  inference(subsumption_resolution,[],[f70836,f543])).
fof(f70836,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,X2)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl138),
  inference(resolution,[],[f69660,f498])).
fof(f70844,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl138),
  inference(subsumption_resolution,[],[f70835,f543])).
fof(f70835,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | $spl138),
  inference(resolution,[],[f69660,f502])).
fof(f70834,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))))) ) | ($spl138 | $spl200)),
  inference(resolution,[],[f69660,f33481])).
fof(f70798,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70015,f29192])).
fof(f6170,plain,(
  ( ! [X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X53))))))) ) | $spl36),
  inference(cnf_transformation,[],[f6170_D])).
fof(f6170_D,plain,(
  ( ! [X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X53))))))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f70763,plain,(
  ( ! [X59,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,X58))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f70746])).
fof(f70746,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X58) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,X58))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f70613,f562])).
fof(f70613,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X12) | ~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X13)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13))))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f67792,f711])).
fof(f67792,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),X13))))))) ) | $spl8),
  inference(superposition,[],[f67295,f583])).
fof(f70764,plain,(
  ( ! [X57,X56] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,X56))) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f70745])).
fof(f70745,plain,(
  ( ! [X57,X56] : (s(bool,X56) != s(bool,X56) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,X56))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f70613,f564])).
fof(f70765,plain,(
  ( ! [X54,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f70744])).
fof(f70744,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X54) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f70613,f587])).
fof(f70766,plain,(
  ( ! [X52,X53] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl8 | $spl10)),
  inference(trivial_inequality_removal,[],[f70743])).
fof(f70743,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X52) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f70613,f588])).
fof(f70774,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X27)))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f70732,f562])).
fof(f70732,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | ~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X27)))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f70613,f518])).
fof(f70773,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X24)))) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f70731,f562])).
fof(f70731,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X24)))) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f70613,f518])).
fof(f70666,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70011,f30061])).
fof(f70661,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl139)),
  inference(resolution,[],[f70011,f28875])).
fof(f70449,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28678])).
fof(f70448,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28027])).
fof(f70447,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f30061])).
fof(f70446,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f30058])).
fof(f70443,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl140)),
  inference(resolution,[],[f69492,f28884])).
fof(f69944,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | s(bool,X23) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') != s(bool,X23)) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28678])).
fof(f69943,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | sP9(X21,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),'T_0') | s(bool,X20) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28027])).
fof(f69942,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,X19) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f30061])).
fof(f69941,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f30058])).
fof(f69938,plain,(
  ( ! [X10,X11] : (sP9(X10,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))),X11) | ~p(s(bool,X11)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl138)),
  inference(resolution,[],[f69341,f28884])).
fof(f69239,plain,(
  ( ! [X34] : (~sP8(X34,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X34))) ) | $spl430),
  inference(cnf_transformation,[],[f69239_D])).
fof(f69239_D,plain,(
  ( ! [X34] : (~sP8(X34,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X34))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f69238,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X35) | ~p(s(bool,X34)) | ~sP8(X34,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl8 | $spl10 | $spl22)),
  inference(forward_demodulation,[],[f69237,f711])).
fof(f69237,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP8(X34,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl8 | $spl22)),
  inference(subsumption_resolution,[],[f69219,f457])).
fof(f69219,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP8(X34,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X34)) ) | ($spl8 | $spl22)),
  inference(trivial_inequality_removal,[],[f69203])).
fof(f69203,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP8(X34,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X34)) ) | ($spl8 | $spl22)),
  inference(resolution,[],[f67779,f28175])).
fof(f69127,plain,(
  $false | $spl44),
  inference(trivial_inequality_removal,[],[f69121])).
fof(f69121,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl44),
  inference(superposition,[],[f8475,f357])).
fof(f357,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X9] : s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X9),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.EXISTS_CONG','thm.list.EXISTS_DEF|split|0')).
fof(f69126,plain,(
  $false | ($spl10 | $spl44)),
  inference(trivial_inequality_removal,[],[f69124])).
fof(f69124,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl10 | $spl44)),
  inference(superposition,[],[f8475,f711])).
fof(f69125,plain,(
  $false | $spl44),
  inference(equality_resolution,[],[f8475])).
fof(f69072,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28)))) ) | $spl428),
  inference(cnf_transformation,[],[f69072_D])).
fof(f69072_D,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28)))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f69061,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,'T_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28))) | s(bool,'F_0') != s(bool,X29)) ) | ($spl22 | $spl118)),
  inference(resolution,[],[f68866,f30061])).
fof(f68866,plain,(
  ( ! [X182,X184,X183] : (sP8(X182,happ(s(fun(sK12,bool),sK15),s(sK12,X183)),X184)) ) | $spl118),
  inference(resolution,[],[f25873,f454])).
fof(f68791,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X26)) ) | $spl426),
  inference(cnf_transformation,[],[f68791_D])).
fof(f68791_D,plain,(
  ( ! [X26] : (s(bool,X26) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X26)) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f69031,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,X28) | s(bool,X28) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X29))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X29)))) ) | $spl8),
  inference(subsumption_resolution,[],[f69009,f864])).
fof(f69009,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X28) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X29))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X29)))) ) | $spl8),
  inference(resolution,[],[f67777,f977])).
fof(f67777,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f67295,f464])).
fof(f68790,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,X26) | s(bool,X26) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27)))) ) | $spl8),
  inference(subsumption_resolution,[],[f68767,f864])).
fof(f68767,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X26) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27)))) ) | $spl8),
  inference(resolution,[],[f67776,f977])).
fof(f67776,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl8),
  inference(superposition,[],[f67295,f464])).
fof(f68751,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f68632,f68539])).
fof(f68539,plain,(
  ( ! [X10] : (~p(s(bool,X10))) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f68538,f564])).
fof(f68538,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10))) ) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f705,f708])).
fof(f68632,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f579])).
fof(f579,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f343])).
fof(f343,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f165])).
fof(f165,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f12])).
fof(f12,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.EXISTS_CONG','thm.bool.NOT_CLAUSES|split|1')).
fof(f68750,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f68631,f68539])).
fof(f68631,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f577])).
fof(f577,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f342])).
fof(f342,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f20])).
fof(f20,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.list.EXISTS_CONG','thm.bool.NOT_CLAUSES|split|2')).
fof(f68749,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f68628,f68539])).
fof(f68628,plain,(
  ( ! [X216] : (p(s(bool,X216))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f536])).
fof(f536,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f325])).
fof(f325,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f324])).
fof(f324,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,X16))) <=> p(s(bool,'T_0')))),
  file('thm.list.EXISTS_CONG','thm.bool.IMP_CLAUSES|split|3')).
fof(f68748,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f68627,f68539])).
fof(f68627,plain,(
  ( ! [X215] : (p(s(bool,X215))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f533])).
fof(f533,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f323])).
fof(f323,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f322])).
fof(f322,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f205])).
fof(f205,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X16] : ((p(s(bool,X16)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.EXISTS_CONG','thm.bool.IMP_CLAUSES|split|1')).
fof(f68747,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f68626,f68539])).
fof(f68626,plain,(
  ( ! [X214] : (p(s(bool,X214))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f530])).
fof(f530,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f321])).
fof(f321,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f320])).
fof(f320,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f204])).
fof(f204,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X16] : ((p(s(bool,'T_0')) => p(s(bool,X16))) <=> p(s(bool,X16)))),
  file('thm.list.EXISTS_CONG','thm.bool.IMP_CLAUSES|split|0')).
fof(f68633,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f581])).
fof(f581,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f14])).
fof(f14,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.EXISTS_CONG','thm.bool.TRUTH')).
fof(f68634,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f582])).
fof(f68635,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f68539,f595])).
fof(f68671,plain,(
  $false | ($spl4 | $spl8 | $spl26)),
  inference(resolution,[],[f68539,f4663])).
fof(f68672,plain,(
  $false | ($spl4 | $spl8 | $spl24)),
  inference(resolution,[],[f68539,f4534])).
fof(f67828,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl424),
  inference(cnf_transformation,[],[f67828_D])).
fof(f67828_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f67815,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl1 | $spl22)),
  inference(resolution,[],[f769,f29192])).
fof(f769,plain,(
  ( ! [X109,X108] : (sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X108,X109)) ) | $spl1),
  inference(resolution,[],[f702,f459])).
fof(f712,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | $spl13),
  inference(cnf_transformation,[],[f712_D])).
fof(f712_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f67282,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X19)) ) | ($spl1 | $spl22)),
  inference(resolution,[],[f765,f30061])).
fof(f765,plain,(
  ( ! [X101,X100] : (sP8(X100,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X101)) ) | $spl1),
  inference(resolution,[],[f702,f454])).
fof(f67293,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X9))) ) | $spl422),
  inference(cnf_transformation,[],[f67293_D])).
fof(f67293_D,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X9))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f67277,plain,(
  ( ! [X8,X9] : (sP9(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X9)) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl1 | $spl22)),
  inference(resolution,[],[f765,f28875])).
fof(f1077,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | $spl14),
  inference(cnf_transformation,[],[f1077_D])).
fof(f1077_D,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f67215,plain,(
  ( ! [X9] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | p(s(bool,X9))) ) | $spl7),
  inference(subsumption_resolution,[],[f67197,f558])).
fof(f67197,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | p(s(bool,X9))) ) | $spl7),
  inference(superposition,[],[f707,f515])).
fof(f67214,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl7),
  inference(subsumption_resolution,[],[f67196,f558])).
fof(f67196,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl7),
  inference(superposition,[],[f707,f515])).
fof(f713,plain,(
  s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | $spl12),
  inference(cnf_transformation,[],[f713_D])).
fof(f713_D,plain,(
  s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f707,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl7),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f67051,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f588])).
fof(f711,plain,(
  s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | $spl10),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f709,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl9),
  inference(cnf_transformation,[],[f709_D])).
fof(f709_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f67050,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f588])).
fof(f67049,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f587])).
fof(f67048,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f587])).
fof(f702,plain,(
  ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | $spl1),
  inference(cnf_transformation,[],[f702_D])).
fof(f702_D,plain,(
  ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f67047,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f564])).
fof(f67046,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f562])).
fof(f703,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | $spl0),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f67045,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f560])).
fof(f67044,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f558])).
fof(f705,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X10))) ) | $spl4),
  inference(cnf_transformation,[],[f705_D])).
fof(f705_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X10))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f67043,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X11)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) )),
  inference(superposition,[],[f679,f518])).
fof(f67042,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f679,f518])).
fof(f704,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X8))) ) | $spl2),
  inference(cnf_transformation,[],[f704_D])).
fof(f704_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X8))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f67041,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | p(s(bool,X9))) )),
  inference(superposition,[],[f679,f515])).
fof(f67040,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) )),
  inference(superposition,[],[f679,f515])).
fof(f66995,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X22))) ) | $spl48),
  inference(subsumption_resolution,[],[f25018,f562])).
fof(f25018,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X22))) ) | $spl48),
  inference(superposition,[],[f10722,f562])).
fof(f66994,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X23))) ) | $spl48),
  inference(subsumption_resolution,[],[f25019,f562])).
fof(f25019,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl48),
  inference(superposition,[],[f10722,f564])).
fof(f66993,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X24)) ) | ($spl22 | $spl48)),
  inference(subsumption_resolution,[],[f25020,f2929])).
fof(f25020,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X24) | s(bool,'F_0') = s(bool,X24)) ) | $spl48),
  inference(superposition,[],[f10722,f587])).
fof(f66992,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X26)) ) | ($spl22 | $spl48)),
  inference(subsumption_resolution,[],[f25022,f2929])).
fof(f25022,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X26) | s(bool,'F_0') = s(bool,X26)) ) | $spl48),
  inference(superposition,[],[f10722,f588])).
fof(f66989,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X15))) ) | $spl420),
  inference(cnf_transformation,[],[f66989_D])).
fof(f66989_D,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X15))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f25286,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X13)) | p(s(bool,X12))) ) | $spl50),
  inference(subsumption_resolution,[],[f25226,f558])).
fof(f25226,plain,(
  ( ! [X12,X13] : (s(bool,X13) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X12) | p(s(bool,X13)) | p(s(bool,X12))) ) | $spl50),
  inference(superposition,[],[f10724,f515])).
fof(f25288,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X14)) | p(s(bool,X15))) ) | $spl50),
  inference(subsumption_resolution,[],[f25227,f558])).
fof(f25227,plain,(
  ( ! [X14,X15] : (s(bool,X15) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X14) | p(s(bool,X14)) | p(s(bool,X15))) ) | $spl50),
  inference(superposition,[],[f10724,f515])).
fof(f710,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | $spl11),
  inference(cnf_transformation,[],[f710_D])).
fof(f710_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f25290,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X20))) ) | $spl50),
  inference(subsumption_resolution,[],[f25230,f558])).
fof(f25230,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X20) | p(s(bool,X20))) ) | $spl50),
  inference(superposition,[],[f10724,f558])).
fof(f25292,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X21))) ) | $spl50),
  inference(subsumption_resolution,[],[f25231,f558])).
fof(f25231,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X21) | p(s(bool,X21))) ) | $spl50),
  inference(superposition,[],[f10724,f560])).
fof(f15214,plain,(
  ( ! [X56] : (s(bool,'T_0') = s(bool,X56)) ) | $spl62),
  inference(cnf_transformation,[],[f15214_D])).
fof(f15214_D,plain,(
  ( ! [X56] : (s(bool,'T_0') = s(bool,X56)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f25294,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X25)) ) | $spl50),
  inference(subsumption_resolution,[],[f25235,f587])).
fof(f25235,plain,(
  ( ! [X25] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | $spl50),
  inference(superposition,[],[f10724,f587])).
fof(f25296,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X27)) ) | $spl50),
  inference(subsumption_resolution,[],[f25237,f587])).
fof(f25237,plain,(
  ( ! [X27] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | $spl50),
  inference(superposition,[],[f10724,f588])).
fof(f25300,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42)))) ) | $spl50),
  inference(subsumption_resolution,[],[f25240,f998])).
fof(f25240,plain,(
  ( ! [X42] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X42),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42)))) ) | $spl50),
  inference(superposition,[],[f10724,f998])).
fof(f66986,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') != s(bool,X16)) ) | $spl418),
  inference(cnf_transformation,[],[f66986_D])).
fof(f66986_D,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') != s(bool,X16)) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f25323,plain,(
  ( ! [X14,X13] : (s(bool,X13) != s(bool,X14) | s(bool,'F_0') != s(bool,X14) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | ($spl22 | $spl50)),
  inference(subsumption_resolution,[],[f25262,f4474])).
fof(f25262,plain,(
  ( ! [X14,X13] : (s(bool,X13) != s(bool,X14) | s(bool,'F_0') != s(bool,X14) | p(s(bool,X13)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | $spl50),
  inference(superposition,[],[f10724,f515])).
fof(f25325,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') != s(bool,X16) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | ($spl22 | $spl50)),
  inference(subsumption_resolution,[],[f25263,f4474])).
fof(f25263,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'F_0') != s(bool,X16) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | p(s(bool,X15))) ) | $spl50),
  inference(superposition,[],[f10724,f515])).
fof(f25280,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | $spl50),
  inference(duplicate_literal_removal,[],[f25266])).
fof(f25266,plain,(
  ( ! [X21] : (s(bool,'F_0') != s(bool,X21) | s(bool,'F_0') != s(bool,X21) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | $spl50),
  inference(superposition,[],[f10724,f558])).
fof(f25279,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,X22) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | $spl50),
  inference(duplicate_literal_removal,[],[f25267])).
fof(f25267,plain,(
  ( ! [X22] : (s(bool,'F_0') != s(bool,X22) | s(bool,'F_0') != s(bool,X22) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | $spl50),
  inference(superposition,[],[f10724,f560])).
fof(f25278,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl50),
  inference(duplicate_literal_removal,[],[f25271])).
fof(f25271,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,X26) | s(bool,'F_0') != s(bool,X26) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl50),
  inference(superposition,[],[f10724,f587])).
fof(f25277,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,X28) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl50),
  inference(duplicate_literal_removal,[],[f25273])).
fof(f25273,plain,(
  ( ! [X28] : (s(bool,'F_0') != s(bool,X28) | s(bool,'F_0') != s(bool,X28) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl50),
  inference(superposition,[],[f10724,f588])).
fof(f28213,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl148),
  inference(duplicate_literal_removal,[],[f28199])).
fof(f28199,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X7))) ) | $spl148),
  inference(resolution,[],[f28180,f465])).
fof(f28180,plain,(
  ( ! [X3] : (~sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0') | ~p(s(bool,X3))) ) | $spl148),
  inference(resolution,[],[f28067,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X2,X1) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f284])).
fof(f31701,plain,(
  ( ! [X528] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X528))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X528)))) ) | $spl50),
  inference(subsumption_resolution,[],[f31532,f998])).
fof(f31532,plain,(
  ( ! [X528] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X528),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X528))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X528)))) ) | $spl50),
  inference(superposition,[],[f10724,f998])).
fof(f62353,plain,(
  ( ! [X7] : (p(s(bool,X7)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | $spl364),
  inference(resolution,[],[f62318,f459])).
fof(f62399,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(trivial_inequality_removal,[],[f62355])).
fof(f62355,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62318,f28845])).
fof(f62356,plain,(
  ( ! [X10] : (p(s(bool,X10)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62318,f29259])).
fof(f66977,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl416),
  inference(cnf_transformation,[],[f66977_D])).
fof(f66977_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f62472,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | ($spl22 | $spl364)),
  inference(subsumption_resolution,[],[f62419,f455])).
fof(f62419,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP8(X7,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62350,f27940])).
fof(f62350,plain,(
  ( ! [X4] : (~sP9(X4,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0') | p(s(bool,X4))) ) | $spl364),
  inference(resolution,[],[f62318,f446])).
fof(f62466,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl364),
  inference(duplicate_literal_removal,[],[f62424])).
fof(f62424,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X14))) ) | $spl364),
  inference(resolution,[],[f62350,f464])).
fof(f62464,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(trivial_inequality_removal,[],[f62463])).
fof(f62463,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,'T_0')) ) | ($spl22 | $spl364)),
  inference(duplicate_literal_removal,[],[f62427])).
fof(f62427,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X17))) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62350,f4491])).
fof(f62462,plain,(
  ( ! [X18] : (p(s(bool,X18)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) ) | ($spl22 | $spl364)),
  inference(trivial_inequality_removal,[],[f62461])).
fof(f62461,plain,(
  ( ! [X18] : (p(s(bool,X18)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | s(bool,'T_0') != s(bool,'T_0')) ) | ($spl22 | $spl364)),
  inference(duplicate_literal_removal,[],[f62428])).
fof(f62428,plain,(
  ( ! [X18] : (p(s(bool,X18)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X18))) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62350,f28891])).
fof(f62477,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(subsumption_resolution,[],[f62460,f595])).
fof(f62460,plain,(
  ( ! [X19] : (p(s(bool,X19)) | ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(duplicate_literal_removal,[],[f62429])).
fof(f62429,plain,(
  ( ! [X19] : (p(s(bool,X19)) | ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X19))) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62350,f28933])).
fof(f62479,plain,(
  ( ! [X23] : (p(s(bool,X23)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | ($spl22 | $spl364)),
  inference(subsumption_resolution,[],[f62456,f28166])).
fof(f62456,plain,(
  ( ! [X23] : (p(s(bool,X23)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP6('T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X23)) ) | ($spl22 | $spl364)),
  inference(duplicate_literal_removal,[],[f62433])).
fof(f62433,plain,(
  ( ! [X23] : (p(s(bool,X23)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X23)) | ~sP6('T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X23)) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62350,f29284])).
fof(f29284,plain,(
  ( ! [X10,X11,X9] : (sP9(X10,X9,X11) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X10)) | ~sP6(X11,X9,X10)) ) | $spl22),
  inference(subsumption_resolution,[],[f29279,f455])).
fof(f29279,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | sP9(X10,X9,X11) | p(s(bool,X10)) | ~sP8(X10,X11,X9) | ~sP6(X11,X9,X10)) ) | $spl22),
  inference(resolution,[],[f29259,f451])).
fof(f66975,plain,(
  ( ! [X26] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26))) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl364),
  inference(subsumption_resolution,[],[f62436,f864])).
fof(f62436,plain,(
  ( ! [X26] : (p(s(bool,'const.bool.IN_2'(s(sK12,X26),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26))) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl364),
  inference(resolution,[],[f62350,f976])).
fof(f1081,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl19),
  inference(cnf_transformation,[],[f1081_D])).
fof(f1081_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f66972,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X20))) ) | $spl22),
  inference(subsumption_resolution,[],[f4498,f558])).
fof(f4498,plain,(
  ( ! [X20] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X20) | p(s(bool,X20))) ) | $spl22),
  inference(superposition,[],[f2929,f558])).
fof(f66971,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X21))) ) | $spl22),
  inference(subsumption_resolution,[],[f4499,f558])).
fof(f4499,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X21) | p(s(bool,X21))) ) | $spl22),
  inference(superposition,[],[f2929,f560])).
fof(f66970,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X25)) ) | $spl22),
  inference(subsumption_resolution,[],[f4503,f2929])).
fof(f4503,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | $spl22),
  inference(superposition,[],[f2929,f587])).
fof(f66969,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X27)) ) | $spl22),
  inference(subsumption_resolution,[],[f4505,f2929])).
fof(f4505,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | $spl22),
  inference(superposition,[],[f2929,f588])).
fof(f4528,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X34)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f4508])).
fof(f4508,plain,(
  ( ! [X34] : (s(bool,X34) != s(bool,X34) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X34)) ) | $spl22),
  inference(superposition,[],[f2929,f588])).
fof(f4527,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X35)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f4509])).
fof(f4509,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,X35) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X35)) ) | $spl22),
  inference(superposition,[],[f2929,f587])).
fof(f4526,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X36))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f4510])).
fof(f4510,plain,(
  ( ! [X36] : (s(bool,X36) != s(bool,X36) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X36))) ) | $spl22),
  inference(superposition,[],[f2929,f564])).
fof(f4525,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X37))) ) | $spl22),
  inference(trivial_inequality_removal,[],[f4511])).
fof(f4511,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,X37) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X37))) ) | $spl22),
  inference(superposition,[],[f2929,f562])).
fof(f66967,plain,(
  ( ! [X483] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X483))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X483)))) )),
  inference(subsumption_resolution,[],[f7088,f864])).
fof(f7088,plain,(
  ( ! [X483] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X483),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X483))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X483)))) )),
  inference(superposition,[],[f561,f998])).
fof(f66966,plain,(
  ( ! [X487] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X487))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X487)))) )),
  inference(subsumption_resolution,[],[f7092,f864])).
fof(f7092,plain,(
  ( ! [X487] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X487),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X487))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X487)))) )),
  inference(superposition,[],[f563,f998])).
fof(f563,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f341])).
fof(f66965,plain,(
  ( ! [X506] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X506))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X506)))) ) | $spl22),
  inference(subsumption_resolution,[],[f7103,f998])).
fof(f7103,plain,(
  ( ! [X506] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X506),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X506))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X506)))) ) | $spl22),
  inference(superposition,[],[f2929,f998])).
fof(f7282,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0)))) )),
  inference(duplicate_literal_removal,[],[f7259])).
fof(f7259,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0)))) )),
  inference(superposition,[],[f1000,f998])).
fof(f66964,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21)))) )),
  inference(subsumption_resolution,[],[f7268,f864])).
fof(f7268,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21))) | p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1000,f558])).
fof(f66963,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X22)))) )),
  inference(subsumption_resolution,[],[f7269,f864])).
fof(f7269,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X22))) | p(s(bool,'const.bool.IN_2'(s(sK12,X22),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1000,f560])).
fof(f66962,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26)))) )),
  inference(subsumption_resolution,[],[f7273,f1000])).
fof(f7273,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X26),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1000,f587])).
fof(f66961,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X28)))) )),
  inference(subsumption_resolution,[],[f7275,f1000])).
fof(f7275,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X28))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X28),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1000,f588])).
fof(f66940,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl22),
  inference(subsumption_resolution,[],[f8439,f518])).
fof(f8439,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | ~p(s(bool,X33)) | ~p(s(bool,X32))) ) | $spl22),
  inference(superposition,[],[f4474,f562])).
fof(f66939,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl22),
  inference(subsumption_resolution,[],[f8440,f518])).
fof(f8440,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | ~p(s(bool,X35)) | ~p(s(bool,X34))) ) | $spl22),
  inference(superposition,[],[f4474,f564])).
fof(f66916,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | ~p(s(bool,X34))) ) | $spl406),
  inference(cnf_transformation,[],[f66916_D])).
fof(f66916_D,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,X35) | ~p(s(bool,X34))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f66938,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37))) ) | $spl22),
  inference(subsumption_resolution,[],[f8441,f4474])).
fof(f8441,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | ~p(s(bool,X37)) | s(bool,'F_0') = s(bool,X36)) ) | $spl22),
  inference(superposition,[],[f4474,f587])).
fof(f66937,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | ~p(s(bool,X41))) ) | $spl22),
  inference(subsumption_resolution,[],[f8443,f4474])).
fof(f8443,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | ~p(s(bool,X41)) | s(bool,'F_0') = s(bool,X40)) ) | $spl22),
  inference(superposition,[],[f4474,f588])).
fof(f25073,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl48),
  inference(subsumption_resolution,[],[f25014,f562])).
fof(f25014,plain,(
  ( ! [X17,X16] : (s(bool,X17) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X16) | ~p(s(bool,X16)) | ~p(s(bool,X17))) ) | $spl48),
  inference(superposition,[],[f10722,f518])).
fof(f25075,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X19)) | ~p(s(bool,X18))) ) | $spl48),
  inference(subsumption_resolution,[],[f25015,f562])).
fof(f25015,plain,(
  ( ! [X19,X18] : (s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X19)) | ~p(s(bool,X18))) ) | $spl48),
  inference(superposition,[],[f10722,f518])).
fof(f66934,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(X48,bool),X49),s('type.list.list'(X48),'const.list.NIL_0')))) ) | $spl414),
  inference(cnf_transformation,[],[f66934_D])).
fof(f66934_D,plain,(
  ( ! [X48,X49] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(X48,bool),X49),s('type.list.list'(X48),'const.list.NIL_0')))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f25036,plain,(
  ( ! [X48,X49] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(X48,bool),X49),s('type.list.list'(X48),'const.list.NIL_0')))) ) | $spl48),
  inference(superposition,[],[f10722,f357])).
fof(f66886,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') != s(bool,X24) | p(s(bool,X25))) ) | $spl388),
  inference(cnf_transformation,[],[f66886_D])).
fof(f66886_D,plain,(
  ( ! [X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,'T_0') != s(bool,X24) | p(s(bool,X25))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f25049,plain,(
  ( ! [X14,X13] : (s(bool,X13) != s(bool,X14) | s(bool,'T_0') != s(bool,X14) | p(s(bool,X13)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl48),
  inference(superposition,[],[f10722,f515])).
fof(f25050,plain,(
  ( ! [X15,X16] : (s(bool,X15) != s(bool,X16) | s(bool,'T_0') != s(bool,X16) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | p(s(bool,X15))) ) | $spl48),
  inference(superposition,[],[f10722,f515])).
fof(f1076,plain,(
  ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | $spl15),
  inference(cnf_transformation,[],[f1076_D])).
fof(f1076_D,plain,(
  ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f25086,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,X18) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X17))) ) | ($spl22 | $spl48)),
  inference(subsumption_resolution,[],[f25051,f8467])).
fof(f8467,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,X43) | ~p(s(bool,X43)) | s(bool,'T_0') = s(bool,X42)) ) | $spl22),
  inference(trivial_inequality_removal,[],[f8444])).
fof(f8444,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | ~p(s(bool,X43)) | s(bool,'T_0') = s(bool,X42)) ) | $spl22),
  inference(superposition,[],[f4474,f588])).
fof(f25051,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') != s(bool,X18) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X17))) ) | $spl48),
  inference(superposition,[],[f10722,f518])).
fof(f25088,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X20) | ~p(s(bool,X19)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | ($spl22 | $spl48)),
  inference(subsumption_resolution,[],[f25052,f8467])).
fof(f25052,plain,(
  ( ! [X19,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') != s(bool,X20) | ~p(s(bool,X19)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl48),
  inference(superposition,[],[f10722,f518])).
fof(f25067,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl48),
  inference(duplicate_literal_removal,[],[f25055])).
fof(f25055,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl48),
  inference(superposition,[],[f10722,f562])).
fof(f25066,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,X24) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl48),
  inference(duplicate_literal_removal,[],[f25056])).
fof(f25056,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,X24) | s(bool,'T_0') != s(bool,X24) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl48),
  inference(superposition,[],[f10722,f564])).
fof(f706,plain,(
  s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl6),
  inference(cnf_transformation,[],[f706_D])).
fof(f706_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f25065,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl48),
  inference(duplicate_literal_removal,[],[f25057])).
fof(f25057,plain,(
  ( ! [X25] : (s(bool,'T_0') != s(bool,X25) | s(bool,'T_0') != s(bool,X25) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl48),
  inference(superposition,[],[f10722,f587])).
fof(f25064,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl48),
  inference(duplicate_literal_removal,[],[f25059])).
fof(f25059,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,X27) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl48),
  inference(superposition,[],[f10722,f588])).
fof(f66930,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41)) ) | $spl412),
  inference(cnf_transformation,[],[f66930_D])).
fof(f66930_D,plain,(
  ( ! [X41,X40] : (s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41)) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f66933,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,X33)) ) | $spl22),
  inference(subsumption_resolution,[],[f25365,f8572])).
fof(f25365,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | s(bool,'F_0') = s(bool,X33) | ~p(s(bool,X32))) ) | $spl22),
  inference(superposition,[],[f4472,f562])).
fof(f4472,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'F_0') = s(bool,X0)) ) | $spl22),
  inference(superposition,[],[f2929,f588])).
fof(f66932,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35)) ) | $spl22),
  inference(subsumption_resolution,[],[f25366,f8572])).
fof(f25366,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | s(bool,'F_0') = s(bool,X35) | ~p(s(bool,X34))) ) | $spl22),
  inference(superposition,[],[f4472,f564])).
fof(f66931,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X37)) ) | $spl22),
  inference(subsumption_resolution,[],[f25367,f4472])).
fof(f25367,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl22),
  inference(superposition,[],[f4472,f587])).
fof(f66929,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41)) ) | $spl22),
  inference(subsumption_resolution,[],[f25369,f4472])).
fof(f25369,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X41) | s(bool,'F_0') = s(bool,X40)) ) | $spl22),
  inference(superposition,[],[f4472,f588])).
fof(f66928,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') = s(bool,X66)) ) | $spl410),
  inference(cnf_transformation,[],[f66928_D])).
fof(f66928_D,plain,(
  ( ! [X66] : (s(bool,X66) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') = s(bool,X66)) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f25378,plain,(
  ( ! [X66] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X66) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | s(bool,'F_0') = s(bool,X66)) ) | ($spl22 | $spl24)),
  inference(superposition,[],[f4472,f4630])).
fof(f66927,plain,(
  ( ! [X67] : (s(bool,X67) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') = s(bool,X67)) ) | $spl408),
  inference(cnf_transformation,[],[f66927_D])).
fof(f66927_D,plain,(
  ( ! [X67] : (s(bool,X67) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') = s(bool,X67)) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f25379,plain,(
  ( ! [X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X67) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | s(bool,'F_0') = s(bool,X67)) ) | ($spl22 | $spl26)),
  inference(superposition,[],[f4472,f4759])).
fof(f27960,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl144),
  inference(duplicate_literal_removal,[],[f27945])).
fof(f27945,plain,(
  ( ! [X3] : (p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl144),
  inference(resolution,[],[f27941,f467])).
fof(f28125,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'T_0') != s(bool,'F_0')) ) | $spl146),
  inference(duplicate_literal_removal,[],[f28102])).
fof(f28102,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X9))) ) | $spl146),
  inference(resolution,[],[f28086,f465])).
fof(f28086,plain,(
  ( ! [X3] : (~sP9(X3,'F_0','T_0') | ~p(s(bool,X3))) ) | $spl146),
  inference(resolution,[],[f28062,f447])).
fof(f28177,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl148),
  inference(resolution,[],[f28067,f453])).
fof(f66919,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') != s(bool,X18) | s(bool,X16) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl142),
  inference(subsumption_resolution,[],[f28439,f558])).
fof(f28439,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,'T_0') != s(bool,X18) | s(bool,X16) != s(bool,X18) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl142),
  inference(superposition,[],[f27910,f515])).
fof(f66918,plain,(
  ( ! [X21,X19,X20] : (s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl142),
  inference(subsumption_resolution,[],[f28440,f558])).
fof(f28440,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl142),
  inference(superposition,[],[f27910,f515])).
fof(f66917,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X32) != s(bool,X33) | ~p(s(bool,X32))) ) | ($spl22 | $spl142)),
  inference(subsumption_resolution,[],[f28445,f8467])).
fof(f28445,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X33) | s(bool,X32) != s(bool,X33) | ~p(s(bool,X32))) ) | $spl142),
  inference(superposition,[],[f27910,f562])).
fof(f66915,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X34) != s(bool,X35) | ~p(s(bool,X34))) ) | ($spl22 | $spl142)),
  inference(subsumption_resolution,[],[f28446,f8467])).
fof(f28446,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X35) | s(bool,X34) != s(bool,X35) | ~p(s(bool,X34))) ) | $spl142),
  inference(superposition,[],[f27910,f564])).
fof(f66913,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | s(bool,X40) != s(bool,X41)) ) | $spl404),
  inference(cnf_transformation,[],[f66913_D])).
fof(f66913_D,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,X41) | s(bool,X40) != s(bool,X41)) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f66914,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X37)) ) | $spl142),
  inference(subsumption_resolution,[],[f28447,f27910])).
fof(f28447,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X37) | s(bool,X36) != s(bool,X37) | s(bool,'F_0') = s(bool,X36)) ) | $spl142),
  inference(superposition,[],[f27910,f587])).
fof(f66912,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X41) | s(bool,X40) != s(bool,X41)) ) | $spl142),
  inference(subsumption_resolution,[],[f28449,f27910])).
fof(f28449,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X41) | s(bool,X40) != s(bool,X41) | s(bool,'F_0') = s(bool,X40)) ) | $spl142),
  inference(superposition,[],[f27910,f588])).
fof(f66911,plain,(
  ( ! [X66] : (s(bool,'T_0') != s(bool,X66) | s(bool,X66) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | $spl402),
  inference(cnf_transformation,[],[f66911_D])).
fof(f66911_D,plain,(
  ( ! [X66] : (s(bool,'T_0') != s(bool,X66) | s(bool,X66) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f28459,plain,(
  ( ! [X66] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X66) | s(bool,X66) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl24 | $spl142)),
  inference(superposition,[],[f27910,f4630])).
fof(f66910,plain,(
  ( ! [X67] : (s(bool,'T_0') != s(bool,X67) | s(bool,X67) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | $spl400),
  inference(cnf_transformation,[],[f66910_D])).
fof(f66910_D,plain,(
  ( ! [X67] : (s(bool,'T_0') != s(bool,X67) | s(bool,X67) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f28460,plain,(
  ( ! [X67] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X67) | s(bool,X67) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl26 | $spl142)),
  inference(superposition,[],[f27910,f4759])).
fof(f66906,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl398),
  inference(cnf_transformation,[],[f66906_D])).
fof(f66906_D,plain,(
  ( ! [X43,X42] : (s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f66909,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | ($spl22 | $spl142)),
  inference(subsumption_resolution,[],[f28511,f8467])).
fof(f28511,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X29) | p(s(bool,X28))) ) | $spl142),
  inference(superposition,[],[f28477,f558])).
fof(f28477,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,X43) | s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl142),
  inference(trivial_inequality_removal,[],[f28450])).
fof(f28450,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X43) | s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl142),
  inference(superposition,[],[f27910,f588])).
fof(f66908,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | ($spl22 | $spl142)),
  inference(subsumption_resolution,[],[f28512,f8467])).
fof(f28512,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | s(bool,'T_0') = s(bool,X31) | p(s(bool,X30))) ) | $spl142),
  inference(superposition,[],[f28477,f560])).
fof(f66907,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl142),
  inference(subsumption_resolution,[],[f28516,f28477])).
fof(f28516,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X38)) ) | $spl142),
  inference(superposition,[],[f28477,f587])).
fof(f66905,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl142),
  inference(subsumption_resolution,[],[f28518,f28477])).
fof(f28518,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X42) != s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | $spl142),
  inference(superposition,[],[f28477,f588])).
fof(f66904,plain,(
  ( ! [X83,X84,X82] : (s(bool,X84) != s(bool,'const.list.EXISTS_2'(s(fun(X82,bool),X83),s('type.list.list'(X82),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X84)) ) | $spl396),
  inference(cnf_transformation,[],[f66904_D])).
fof(f66904_D,plain,(
  ( ! [X83,X84,X82] : (s(bool,X84) != s(bool,'const.list.EXISTS_2'(s(fun(X82,bool),X83),s('type.list.list'(X82),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X84)) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f28538,plain,(
  ( ! [X83,X84,X82] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X84) != s(bool,'const.list.EXISTS_2'(s(fun(X82,bool),X83),s('type.list.list'(X82),'const.list.NIL_0'))) | s(bool,'T_0') = s(bool,X84)) ) | $spl142),
  inference(superposition,[],[f28477,f357])).
fof(f66903,plain,(
  ( ! [X33,X32] : (sP9(X32,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X33) | ~p(s(bool,X33))) ) | $spl394),
  inference(cnf_transformation,[],[f66903_D])).
fof(f66903_D,plain,(
  ( ! [X33,X32] : (sP9(X32,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X33) | ~p(s(bool,X33))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f28972,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | sP9(X32,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X33) | ~p(s(bool,X33))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f28943,f4630])).
fof(f28943,plain,(
  ( ! [X33,X32] : (sP9(X32,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X33) | ~p(s(bool,X33)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f28884,f4602])).
fof(f4602,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | $spl24),
  inference(resolution,[],[f4421,f453])).
fof(f66902,plain,(
  ( ! [X35,X34] : (sP9(X34,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X35) | ~p(s(bool,X35))) ) | $spl392),
  inference(cnf_transformation,[],[f66902_D])).
fof(f66902_D,plain,(
  ( ! [X35,X34] : (sP9(X34,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X35) | ~p(s(bool,X35))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f28973,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | sP9(X34,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X35) | ~p(s(bool,X35))) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f28944,f4759])).
fof(f28944,plain,(
  ( ! [X35,X34] : (sP9(X34,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X35) | ~p(s(bool,X35)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f28884,f4731])).
fof(f4731,plain,(
  ( ! [X125,X124] : (sP8(X124,X125,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | $spl26),
  inference(resolution,[],[f4439,f453])).
fof(f29026,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f28981])).
fof(f28981,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X3) != s(bool,X4)) ) | $spl22),
  inference(resolution,[],[f28891,f467])).
fof(f29146,plain,(
  ( ! [X58] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X58))) ) | ($spl22 | $spl148)),
  inference(subsumption_resolution,[],[f29145,f595])).
fof(f29145,plain,(
  ( ! [X58] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X58))) ) | ($spl22 | $spl148)),
  inference(subsumption_resolution,[],[f29114,f457])).
fof(f29114,plain,(
  ( ! [X58] : (~p(s(bool,'T_0')) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),'T_0',X58) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X58))) ) | ($spl22 | $spl148)),
  inference(resolution,[],[f29054,f28180])).
fof(f29054,plain,(
  ( ! [X2,X0,X1] : (sP9(X1,X2,X0) | ~p(s(bool,X0)) | ~sP7(X2,X0,X1) | ~p(s(bool,X2))) ) | $spl22),
  inference(resolution,[],[f29052,f453])).
fof(f29849,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,'T_0') != s(bool,X17)) ) | $spl22),
  inference(subsumption_resolution,[],[f29778,f558])).
fof(f29778,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,'T_0') != s(bool,X17) | s(bool,'F_0') != s(bool,X15)) ) | $spl22),
  inference(resolution,[],[f29451,f28844])).
fof(f29779,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,'T_0') != s(bool,X20)) ) | $spl22),
  inference(resolution,[],[f29451,f29028])).
fof(f30110,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X32)) | ~p(s(bool,X33))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f30083,f4630])).
fof(f30083,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f30058,f4602])).
fof(f30111,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X34)) | ~p(s(bool,X35))) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f30084,f4759])).
fof(f30084,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f30058,f4731])).
fof(f66898,plain,(
  ( ! [X483] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X483))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X483)))) )),
  inference(subsumption_resolution,[],[f31505,f864])).
fof(f31505,plain,(
  ( ! [X483] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X483),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X483))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X483)))) )),
  inference(superposition,[],[f561,f998])).
fof(f66897,plain,(
  ( ! [X487] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X487))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X487)))) )),
  inference(subsumption_resolution,[],[f31509,f864])).
fof(f31509,plain,(
  ( ! [X487] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK12,X487),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X487))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X487)))) )),
  inference(superposition,[],[f563,f998])).
fof(f66896,plain,(
  ( ! [X506] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X506))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X506)))) ) | $spl22),
  inference(subsumption_resolution,[],[f31520,f998])).
fof(f31520,plain,(
  ( ! [X506] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,X506),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X506))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X506)))) ) | $spl22),
  inference(superposition,[],[f2929,f998])).
fof(f66895,plain,(
  ( ! [X542,X541] : (s(bool,X542) != s(bool,'const.bool.IN_2'(s(sK12,X541),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X541))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X541))) | s(bool,'T_0') = s(bool,X542)) ) | $spl390),
  inference(cnf_transformation,[],[f66895_D])).
fof(f66895_D,plain,(
  ( ! [X542,X541] : (s(bool,X542) != s(bool,'const.bool.IN_2'(s(sK12,X541),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X541))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X541))) | s(bool,'T_0') = s(bool,X542)) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f31539,plain,(
  ( ! [X542,X541] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X542) != s(bool,'const.bool.IN_2'(s(sK12,X541),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,'T_0') = s(bool,X542) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X541))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X541)))) ) | $spl142),
  inference(superposition,[],[f28477,f998])).
fof(f31792,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0)))) )),
  inference(duplicate_literal_removal,[],[f31769])).
fof(f31769,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0)))) )),
  inference(superposition,[],[f1000,f998])).
fof(f66894,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21)))) )),
  inference(subsumption_resolution,[],[f31778,f864])).
fof(f31778,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21))) | p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1000,f558])).
fof(f66893,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X22)))) )),
  inference(subsumption_resolution,[],[f31779,f864])).
fof(f31779,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X22))) | p(s(bool,'const.bool.IN_2'(s(sK12,X22),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1000,f560])).
fof(f66892,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26)))) )),
  inference(subsumption_resolution,[],[f31783,f1000])).
fof(f31783,plain,(
  ( ! [X26] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X26),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1000,f587])).
fof(f66891,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X28)))) )),
  inference(subsumption_resolution,[],[f31785,f1000])).
fof(f31785,plain,(
  ( ! [X28] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X28))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X28))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X28),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1000,f588])).
fof(f32695,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') != s(bool,X43) | s(bool,X43) != s(bool,X44) | ~p(s(bool,X42)) | p(s(bool,X44))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f32661])).
fof(f32661,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') != s(bool,X43) | s(bool,X43) != s(bool,X44) | ~p(s(bool,X42)) | p(s(bool,X44)) | ~p(s(bool,X42))) ) | $spl22),
  inference(superposition,[],[f8409,f562])).
fof(f32694,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,X46) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X45)) | p(s(bool,X47))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f32662])).
fof(f32662,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') != s(bool,X46) | s(bool,X46) != s(bool,X47) | ~p(s(bool,X45)) | p(s(bool,X47)) | ~p(s(bool,X45))) ) | $spl22),
  inference(superposition,[],[f8409,f564])).
fof(f66890,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,X49) | s(bool,X49) != s(bool,X50) | ~p(s(bool,X48)) | p(s(bool,X50))) ) | $spl22),
  inference(subsumption_resolution,[],[f32663,f557])).
fof(f32663,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') != s(bool,X49) | s(bool,X49) != s(bool,X50) | ~p(s(bool,X48)) | p(s(bool,X50)) | s(bool,'F_0') = s(bool,X48)) ) | $spl22),
  inference(superposition,[],[f8409,f587])).
fof(f66889,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,X55) | s(bool,X55) != s(bool,X56) | ~p(s(bool,X54)) | p(s(bool,X56))) ) | $spl22),
  inference(subsumption_resolution,[],[f32665,f557])).
fof(f32665,plain,(
  ( ! [X54,X56,X55] : (s(bool,'T_0') != s(bool,X55) | s(bool,X55) != s(bool,X56) | ~p(s(bool,X54)) | p(s(bool,X56)) | s(bool,'F_0') = s(bool,X54)) ) | $spl22),
  inference(superposition,[],[f8409,f588])).
fof(f39604,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X32)) | s(bool,'T_0') != s(bool,X33)) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f39579,f4630])).
fof(f39579,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | s(bool,'T_0') != s(bool,X33) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f30061,f4602])).
fof(f39605,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X34)) | s(bool,'T_0') != s(bool,X35)) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f39580,f4759])).
fof(f39580,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | s(bool,'T_0') != s(bool,X35) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f30061,f4731])).
fof(f66885,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26))) | p(s(bool,X25)) | s(bool,'T_0') != s(bool,X24)) ) | $spl22),
  inference(subsumption_resolution,[],[f39802,f864])).
fof(f39802,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,X25) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X26))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X26))) | p(s(bool,X25)) | s(bool,'T_0') != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(sK12,X26),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl22),
  inference(resolution,[],[f977,f28891])).
fof(f66884,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X79))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X79)))) ) | $spl144),
  inference(subsumption_resolution,[],[f39823,f864])).
fof(f39823,plain,(
  ( ! [X79] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X79))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X79))) | p(s(bool,'const.bool.IN_2'(s(sK12,X79),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl144),
  inference(resolution,[],[f977,f27941])).
fof(f54656,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f54655,f4630])).
fof(f54655,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f54612,f561])).
fof(f54612,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f4603,f28845])).
fof(f4603,plain,(
  ( ! [X127,X126] : (~sP7(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X126,X127) | p(s(bool,X127)) | ~p(s(bool,X126))) ) | $spl24),
  inference(resolution,[],[f4421,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f290])).
fof(f54657,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f54613,f4630])).
fof(f54613,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f4603,f29259])).
fof(f54729,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9)) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f54728,f4630])).
fof(f54728,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f54684,f561])).
fof(f54684,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f4604,f28763])).
fof(f28763,plain,(
  ( ! [X10,X11,X9] : (sP7(X10,X9,X11) | s(bool,'T_0') != s(bool,X10) | s(bool,'F_0') != s(bool,X9)) ) | $spl22),
  inference(subsumption_resolution,[],[f28720,f457])).
fof(f28720,plain,(
  ( ! [X10,X11,X9] : (s(bool,'F_0') != s(bool,X9) | s(bool,'T_0') != s(bool,X10) | p(s(bool,X11)) | sP7(X10,X9,X11)) ) | $spl22),
  inference(resolution,[],[f4490,f446])).
fof(f4490,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') != s(bool,X1) | p(s(bool,X2))) ) | $spl22),
  inference(superposition,[],[f2929,f464])).
fof(f4604,plain,(
  ( ! [X128,X129] : (~sP7(X128,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X129) | p(s(bool,X129)) | ~p(s(bool,X128))) ) | $spl24),
  inference(resolution,[],[f4421,f456])).
fof(f57212,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11)) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f57211,f4759])).
fof(f57211,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f57168,f561])).
fof(f57168,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X11) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f4732,f28845])).
fof(f4732,plain,(
  ( ! [X127,X126] : (~sP7(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X126,X127) | p(s(bool,X127)) | ~p(s(bool,X126))) ) | $spl26),
  inference(resolution,[],[f4439,f456])).
fof(f57213,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X12)) | ~p(s(bool,X13))) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f57169,f4759])).
fof(f57169,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f4732,f29259])).
fof(f57285,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9)) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f57284,f4759])).
fof(f57284,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f57240,f561])).
fof(f57240,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f4733,f28763])).
fof(f4733,plain,(
  ( ! [X128,X129] : (~sP7(X128,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X129) | p(s(bool,X129)) | ~p(s(bool,X128))) ) | $spl26),
  inference(resolution,[],[f4439,f456])).
fof(f65635,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X14)) | p(s(bool,X15))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f65634,f4630])).
fof(f65634,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f65620,f562])).
fof(f65620,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | s(bool,'T_0') != s(bool,X14)) ) | ($spl22 | $spl24)),
  inference(duplicate_literal_removal,[],[f65551])).
fof(f65551,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | s(bool,'T_0') != s(bool,X14) | p(s(bool,X15))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f54607,f4491])).
fof(f54607,plain,(
  ( ! [X0,X1] : (~sP9(X0,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl24),
  inference(resolution,[],[f4603,f446])).
fof(f65638,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X18)) | p(s(bool,X19))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f65618,f4630])).
fof(f65618,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(duplicate_literal_removal,[],[f65553])).
fof(f65553,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | ~p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X19))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f54607,f28933])).
fof(f65641,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X26)) | p(s(bool,X27))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f65640,f4630])).
fof(f65640,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f65614,f463])).
fof(f65614,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | ~sP6(X26,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X27)) ) | ($spl22 | $spl24)),
  inference(duplicate_literal_removal,[],[f65557])).
fof(f65557,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X27)) | ~sP6(X26,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X27)) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f54607,f29284])).
fof(f65789,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X12)) | p(s(bool,X13))) ) | ($spl22 | $spl24)),
  inference(forward_demodulation,[],[f65788,f4630])).
fof(f65788,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | ($spl22 | $spl24)),
  inference(subsumption_resolution,[],[f65779,f562])).
fof(f65779,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | s(bool,'T_0') != s(bool,X12)) ) | ($spl22 | $spl24)),
  inference(duplicate_literal_removal,[],[f65695])).
fof(f65695,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | s(bool,'T_0') != s(bool,X12) | p(s(bool,X13))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f54680,f4490])).
fof(f54680,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl24),
  inference(resolution,[],[f4604,f446])).
fof(f66531,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X14)) | p(s(bool,X15))) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f66530,f4759])).
fof(f66530,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f66516,f562])).
fof(f66516,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | s(bool,'T_0') != s(bool,X14)) ) | ($spl22 | $spl26)),
  inference(duplicate_literal_removal,[],[f66447])).
fof(f66447,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | s(bool,'T_0') != s(bool,X14) | p(s(bool,X15))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f57163,f4491])).
fof(f57163,plain,(
  ( ! [X0,X1] : (~sP9(X0,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl26),
  inference(resolution,[],[f4732,f446])).
fof(f66534,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X18)) | p(s(bool,X19))) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f66514,f4759])).
fof(f66514,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(duplicate_literal_removal,[],[f66449])).
fof(f66449,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | ~p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X19))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f57163,f28933])).
fof(f66537,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X26)) | p(s(bool,X27))) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f66536,f4759])).
fof(f66536,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f66510,f463])).
fof(f66510,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | ~sP6(X26,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X27)) ) | ($spl22 | $spl26)),
  inference(duplicate_literal_removal,[],[f66453])).
fof(f66453,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X27)) | ~sP6(X26,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X27)) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f57163,f29284])).
fof(f66685,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X12)) | p(s(bool,X13))) ) | ($spl22 | $spl26)),
  inference(forward_demodulation,[],[f66684,f4759])).
fof(f66684,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | ($spl22 | $spl26)),
  inference(subsumption_resolution,[],[f66675,f562])).
fof(f66675,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | s(bool,'T_0') != s(bool,X12)) ) | ($spl22 | $spl26)),
  inference(duplicate_literal_removal,[],[f66591])).
fof(f66591,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | s(bool,'T_0') != s(bool,X12) | p(s(bool,X13))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f57236,f4490])).
fof(f57236,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl26),
  inference(resolution,[],[f4733,f446])).
fof(f66873,plain,(
  ( ! [X25] : (p(s(bool,X25)) | ~sP6('T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X25)) ) | $spl386),
  inference(cnf_transformation,[],[f66873_D])).
fof(f66873_D,plain,(
  ( ! [X25] : (p(s(bool,X25)) | ~sP6('T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X25)) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f62455,plain,(
  ( ! [X25] : (p(s(bool,X25)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP6('T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X25)) ) | ($spl22 | $spl364)),
  inference(duplicate_literal_removal,[],[f62435])).
fof(f62435,plain,(
  ( ! [X25] : (p(s(bool,X25)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP6('T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X25) | p(s(bool,X25))) ) | ($spl22 | $spl364)),
  inference(resolution,[],[f62350,f29288])).
fof(f66870,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f66869,f582])).
fof(f66869,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(forward_demodulation,[],[f66868,f15586])).
fof(f15586,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) | $spl12),
  inference(forward_demodulation,[],[f15469,f5506])).
fof(f5506,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK16,sK12)))) | $spl12),
  inference(resolution,[],[f4914,f562])).
fof(f4914,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK16,sK12))))) | $spl12),
  inference(subsumption_resolution,[],[f4883,f595])).
fof(f4883,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK16,sK12))))) | $spl12),
  inference(superposition,[],[f609,f713])).
fof(f15469,plain,(
  s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK16,sK12)))) | $spl12),
  inference(resolution,[],[f5094,f864])).
fof(f5094,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl12),
  inference(subsumption_resolution,[],[f5063,f595])).
fof(f5063,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl12),
  inference(superposition,[],[f655,f713])).
fof(f66868,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12))))) | $spl12),
  inference(subsumption_resolution,[],[f66839,f595])).
fof(f66839,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12))))) | $spl12),
  inference(superposition,[],[f3419,f15554])).
fof(f15554,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | $spl12),
  inference(resolution,[],[f5094,f562])).
fof(f3419,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X0),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))))) ) | $spl12),
  inference(resolution,[],[f3404,f356])).
fof(f3404,plain,(
  ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | $spl12),
  inference(trivial_inequality_removal,[],[f3395])).
fof(f3395,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | $spl12),
  inference(superposition,[],[f2942,f564])).
fof(f2942,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl12),
  inference(forward_demodulation,[],[f679,f713])).
fof(f66867,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f66866,f582])).
fof(f66866,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(forward_demodulation,[],[f66865,f11182])).
fof(f11182,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) | $spl12),
  inference(forward_demodulation,[],[f11065,f5337])).
fof(f5337,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK16,sK12)))) | $spl12),
  inference(resolution,[],[f2046,f562])).
fof(f2046,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK16,sK12))))) | $spl12),
  inference(subsumption_resolution,[],[f1800,f595])).
fof(f1800,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK16,sK12))))) | $spl12),
  inference(superposition,[],[f355,f713])).
fof(f11065,plain,(
  s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK16,sK12)))) | $spl12),
  inference(resolution,[],[f2045,f864])).
fof(f2045,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl12),
  inference(subsumption_resolution,[],[f1799,f595])).
fof(f1799,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | $spl12),
  inference(superposition,[],[f354,f713])).
fof(f66865,plain,(
  ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12))))) | $spl12),
  inference(subsumption_resolution,[],[f66838,f595])).
fof(f66838,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12))))) | $spl12),
  inference(superposition,[],[f3419,f11150])).
fof(f11150,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | $spl12),
  inference(resolution,[],[f2045,f562])).
fof(f66255,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X2)))))) ) | $spl384),
  inference(cnf_transformation,[],[f66255_D])).
fof(f66255_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X2)))))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f66233,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X2)))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f55090,f29451])).
fof(f66095,plain,(
  ( ! [X50,X51] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X50))))) | ~p(s(bool,X51))) ) | ($spl6 | $spl12 | $spl26)),
  inference(resolution,[],[f55084,f5695])).
fof(f5695,plain,(
  ( ! [X6,X7] : (~sP8(X6,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X7) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,X6))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f5694,f595])).
fof(f5694,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | ~sP8(X6,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X7)) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f5693,f713])).
fof(f5693,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP8(X6,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X7)) ) | $spl6),
  inference(subsumption_resolution,[],[f5692,f463])).
fof(f5692,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP8(X6,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X7) | ~sP6('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X7,X6)) ) | $spl6),
  inference(subsumption_resolution,[],[f5674,f457])).
fof(f5674,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'F_0') != s(bool,X7) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~sP8(X6,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X7) | ~sP7(X7,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X6) | ~sP6('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X7,X6)) ) | $spl6),
  inference(resolution,[],[f5617,f449])).
fof(f5617,plain,(
  ( ! [X4,X5] : (~sP9(X5,X4,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X5)) | s(bool,'F_0') != s(bool,X4)) ) | $spl6),
  inference(forward_demodulation,[],[f686,f706])).
fof(f686,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X5)) | ~sP9(X5,X4,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) )),
  inference(superposition,[],[f679,f466])).
fof(f55084,plain,(
  ( ! [X173,X171,X172] : (sP8(X171,X172,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X173)))))) ) | $spl26),
  inference(resolution,[],[f4663,f453])).
fof(f66115,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | $spl382),
  inference(cnf_transformation,[],[f66115_D])).
fof(f66115_D,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f66088,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f55084,f28678])).
fof(f66086,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X29)))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f55084,f30061])).
fof(f66085,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK16,sK12)),s('type.list.list'(sK12),X26)))))) ) | ($spl22 | $spl26)),
  inference(resolution,[],[f55084,f30058])).
fof(f65276,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X2)))))) ) | $spl378),
  inference(cnf_transformation,[],[f65276_D])).
fof(f65276_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X2)))))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f65254,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X2)))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f52898,f29451])).
fof(f60198,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X26)))))) ) | $spl304),
  inference(cnf_transformation,[],[f60198_D])).
fof(f60198_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X26)))))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f65178,plain,(
  ( ! [X78,X77] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X77))))) | ~p(s(bool,X78))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f41975])).
fof(f41975,plain,(
  ( ! [X173,X171,X172] : (sP8(X171,X172,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X173)))))) ) | $spl12),
  inference(resolution,[],[f11846,f453])).
fof(f11846,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X0))))))) ) | $spl12),
  inference(resolution,[],[f11546,f352])).
fof(f11546,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | $spl12),
  inference(subsumption_resolution,[],[f11276,f595])).
fof(f11276,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))))) | $spl12),
  inference(superposition,[],[f375,f11182])).
fof(f59391,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X26)))))) ) | $spl298),
  inference(cnf_transformation,[],[f59391_D])).
fof(f59391_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X26)))))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f65177,plain,(
  ( ! [X76,X75] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X75))))) | ~p(s(bool,X76))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f40231])).
fof(f40231,plain,(
  ( ! [X173,X171,X172] : (sP8(X171,X172,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X173)))))) ) | $spl12),
  inference(resolution,[],[f11672,f453])).
fof(f11672,plain,(
  ( ! [X0] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X0))))))) ) | $spl12),
  inference(resolution,[],[f11545,f352])).
fof(f11545,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | $spl12),
  inference(subsumption_resolution,[],[f11275,f595])).
fof(f11275,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))))) | $spl12),
  inference(superposition,[],[f374,f11182])).
fof(f65176,plain,(
  ( ! [X74,X73] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X73))))) | ~p(s(bool,X74))) ) | ($spl6 | $spl12 | $spl24)),
  inference(resolution,[],[f5695,f52892])).
fof(f52892,plain,(
  ( ! [X173,X171,X172] : (sP8(X171,X172,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X173)))))) ) | $spl24),
  inference(resolution,[],[f4534,f453])).
fof(f30121,plain,(
  ( ! [X60] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X60),s('type.list.list'(sK12),sK13)))))) ) | $spl196),
  inference(cnf_transformation,[],[f30121_D])).
fof(f30121_D,plain,(
  ( ! [X60] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X60),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f65175,plain,(
  ( ! [X72,X71] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X71),s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X72))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f19578])).
fof(f19578,plain,(
  ( ! [X173,X171,X172] : (sP8(X171,X172,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X173),s('type.list.list'(sK12),sK13)))))) ) | $spl12),
  inference(resolution,[],[f2047,f453])).
fof(f2047,plain,(
  ( ! [X2] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13))))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1803,f595])).
fof(f1803,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13))))))) ) | $spl12),
  inference(superposition,[],[f353,f713])).
fof(f353,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X3))))) )),
  inference(cnf_transformation,[],[f236])).
fof(f39905,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X26),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | $spl228),
  inference(cnf_transformation,[],[f39905_D])).
fof(f39905_D,plain,(
  ( ! [X26] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X26),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f65159,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X46),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | ~p(s(bool,X47))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f32513])).
fof(f32513,plain,(
  ( ! [X182,X180,X181] : (sP8(X180,X181,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X182),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f19522,f453])).
fof(f19522,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13)),sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f2047,f355])).
fof(f62197,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | $spl362),
  inference(cnf_transformation,[],[f62197_D])).
fof(f62197_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f65153,plain,(
  ( ! [X39,X38] : (s(bool,'F_0') != s(bool,X38) | ~p(s(bool,X39)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f50378])).
fof(f50378,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X147,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f454])).
fof(f20795,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X252,sK12))))) ) | $spl12),
  inference(resolution,[],[f11542,f562])).
fof(f11542,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X10,sK12))))) | ~p(s(bool,X10))) ) | $spl12),
  inference(subsumption_resolution,[],[f11272,f595])).
fof(f11272,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X10,sK12))))) | ~p(s(bool,X10))) ) | $spl12),
  inference(superposition,[],[f371,f11182])).
fof(f61851,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | $spl346),
  inference(cnf_transformation,[],[f61851_D])).
fof(f61851_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f65151,plain,(
  ( ! [X35,X34] : (s(bool,'F_0') != s(bool,X34) | ~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f48662])).
fof(f48662,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X147,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f454])).
fof(f20566,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X252,sK12))))) ) | $spl12),
  inference(resolution,[],[f11541,f562])).
fof(f11541,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X9,sK12))))) | ~p(s(bool,X9))) ) | $spl12),
  inference(subsumption_resolution,[],[f11271,f595])).
fof(f11271,plain,(
  ( ! [X9] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X9,sK12))))) | ~p(s(bool,X9))) ) | $spl12),
  inference(superposition,[],[f369,f11182])).
fof(f61500,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | $spl332),
  inference(cnf_transformation,[],[f61500_D])).
fof(f61500_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f65149,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X30) | ~p(s(bool,X31)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f46574])).
fof(f46574,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X147,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f454])).
fof(f20366,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X252,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f11540,f562])).
fof(f11540,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X8,sK15,sK12))))) | ~p(s(bool,X8))) ) | $spl12),
  inference(subsumption_resolution,[],[f11270,f595])).
fof(f11270,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X8,sK15,sK12))))) | ~p(s(bool,X8))) ) | $spl12),
  inference(superposition,[],[f366,f11182])).
fof(f61108,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | $spl318),
  inference(cnf_transformation,[],[f61108_D])).
fof(f61108_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f65147,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X26) | ~p(s(bool,X27)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f44436])).
fof(f44436,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X147,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f454])).
fof(f20172,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X252,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f11539,f562])).
fof(f11539,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X7,sK15,sK12))))) | ~p(s(bool,X7))) ) | $spl12),
  inference(subsumption_resolution,[],[f11269,f595])).
fof(f11269,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X7,sK15,sK12))))) | ~p(s(bool,X7))) ) | $spl12),
  inference(superposition,[],[f364,f11182])).
fof(f52682,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | $spl276),
  inference(cnf_transformation,[],[f52682_D])).
fof(f52682_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f65145,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f38230])).
fof(f38230,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X147,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f454])).
fof(f17382,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X252,sK12))))) ) | $spl12),
  inference(resolution,[],[f5250,f562])).
fof(f5250,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X9,sK12))))) | ~p(s(bool,X9))) ) | $spl12),
  inference(resolution,[],[f2046,f371])).
fof(f52362,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | $spl266),
  inference(cnf_transformation,[],[f52362_D])).
fof(f52362_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f65143,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X18) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f36415])).
fof(f36415,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X147,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f454])).
fof(f17107,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X252,sK12))))) ) | $spl12),
  inference(resolution,[],[f5249,f562])).
fof(f5249,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X8,sK12))))) | ~p(s(bool,X8))) ) | $spl12),
  inference(resolution,[],[f2046,f369])).
fof(f52193,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | $spl256),
  inference(cnf_transformation,[],[f52193_D])).
fof(f52193_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f65141,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f34803])).
fof(f34803,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X147,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f454])).
fof(f16832,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X252,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f5248,f562])).
fof(f5248,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X7,sK16,sK12))))) | ~p(s(bool,X7))) ) | $spl12),
  inference(resolution,[],[f2046,f366])).
fof(f52024,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | $spl246),
  inference(cnf_transformation,[],[f52024_D])).
fof(f52024_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f65139,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f5695,f33336])).
fof(f33336,plain,(
  ( ! [X149,X147,X148] : (sP8(X148,X147,X149) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X147,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f454])).
fof(f16638,plain,(
  ( ! [X252] : (~p(s(bool,X252)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X252,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f5247,f562])).
fof(f5247,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X6,sK16,sK12))))) | ~p(s(bool,X6))) ) | $spl12),
  inference(resolution,[],[f2046,f364])).
fof(f65105,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | $spl376),
  inference(cnf_transformation,[],[f65105_D])).
fof(f65105_D,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f65079,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f52892,f28678])).
fof(f65077,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X29)))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f52892,f30061])).
fof(f65076,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK16,sK12)),s('type.list.list'(sK12),X26)))))) ) | ($spl22 | $spl24)),
  inference(resolution,[],[f52892,f30058])).
fof(f62721,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl368),
  inference(cnf_transformation,[],[f62721_D])).
fof(f62721_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f64969,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X38) | ~p(s(bool,X39)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X38,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f50378])).
fof(f4463,plain,(
  ( ! [X8,X9] : (~sP8(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X9) | ~p(s(bool,X8))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f4462,f528])).
fof(f4462,plain,(
  ( ! [X8,X9] : (p(s(bool,'F_0')) | ~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | ~sP8(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f4461,f706])).
fof(f4461,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~sP8(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(subsumption_resolution,[],[f4460,f461])).
fof(f4460,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~sP8(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP6(X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X8)) ) | $spl12),
  inference(subsumption_resolution,[],[f4445,f459])).
fof(f4445,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~sP8(X8,X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X9,X8) | ~sP6(X9,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X8)) ) | $spl12),
  inference(resolution,[],[f3387,f450])).
fof(f3387,plain,(
  ( ! [X6,X7] : (~sP9(X7,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X6) | ~p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl12),
  inference(superposition,[],[f2942,f466])).
fof(f61961,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl348),
  inference(cnf_transformation,[],[f61961_D])).
fof(f61961_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f64967,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X34) | ~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X34,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f48662])).
fof(f61744,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X27,sK15,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl334),
  inference(cnf_transformation,[],[f61744_D])).
fof(f61744_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X27,sK15,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f64965,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,X30) | ~p(s(bool,X31)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X30,sK15,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f46574])).
fof(f61218,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X27,sK15,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl320),
  inference(cnf_transformation,[],[f61218_D])).
fof(f61218_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X27,sK15,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f64963,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') != s(bool,X26) | ~p(s(bool,X27)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X26,sK15,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f44436])).
fof(f52770,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl278),
  inference(cnf_transformation,[],[f52770_D])).
fof(f52770_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f64961,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X22,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f38230])).
fof(f52575,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl268),
  inference(cnf_transformation,[],[f52575_D])).
fof(f52575_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X27,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f64959,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X18) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X18,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f36415])).
fof(f52279,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X27,sK16,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl258),
  inference(cnf_transformation,[],[f52279_D])).
fof(f52279_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X27,sK16,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f64957,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,X14) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X14,sK16,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f34803])).
fof(f52110,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X27,sK16,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) | $spl248),
  inference(cnf_transformation,[],[f52110_D])).
fof(f52110_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X27,sK16,sK12)))) | s(bool,'T_0') != s(bool,X27)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f64955,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X10,sK16,sK12))))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f4463,f33336])).
fof(f63582,plain,(
  ( ! [X57] : (s(bool,'F_0') != s(bool,X57) | s(bool,X57) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl372),
  inference(cnf_transformation,[],[f63582_D])).
fof(f63582_D,plain,(
  ( ! [X57] : (s(bool,'F_0') != s(bool,X57) | s(bool,X57) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f64307,plain,(
  ( ! [X61,X60] : (s(bool,'F_0') != s(bool,X61) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X60))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X60))) | s(bool,X61) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl6),
  inference(subsumption_resolution,[],[f64233,f864])).
fof(f64233,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.bool.IN_2'(s(sK12,X60),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'F_0') != s(bool,X61) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X60))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X60))) | s(bool,X61) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl6),
  inference(resolution,[],[f3953,f976])).
fof(f3953,plain,(
  ( ! [X263,X264] : (~sP9(X264,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X263) | p(s(bool,X264)) | s(bool,'F_0') != s(bool,X263)) ) | $spl6),
  inference(superposition,[],[f467,f706])).
fof(f63581,plain,(
  ( ! [X57,X56] : (s(bool,'F_0') != s(bool,X57) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56))) | s(bool,X57) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl6),
  inference(subsumption_resolution,[],[f63485,f864])).
fof(f63485,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'F_0') != s(bool,X57) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56))) | s(bool,X57) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl6),
  inference(resolution,[],[f3952,f976])).
fof(f3952,plain,(
  ( ! [X261,X262] : (~sP9(X262,X261,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X262)) | s(bool,'F_0') != s(bool,X261)) ) | $spl6),
  inference(superposition,[],[f467,f706])).
fof(f62645,plain,(
  ( ! [X44] : (s(bool,'F_0') = s(bool,X44) | s(bool,X44) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl366),
  inference(cnf_transformation,[],[f62645_D])).
fof(f62645_D,plain,(
  ( ! [X44] : (s(bool,'F_0') = s(bool,X44) | s(bool,X44) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f62884,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') = s(bool,X46) | s(bool,X46) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X47))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X47)))) ) | $spl6),
  inference(subsumption_resolution,[],[f62833,f864])).
fof(f62833,plain,(
  ( ! [X47,X46] : (s(bool,'F_0') = s(bool,X46) | p(s(bool,'const.bool.IN_2'(s(sK12,X47),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X46) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X47))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X47)))) ) | $spl6),
  inference(resolution,[],[f3806,f977])).
fof(f3806,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X2) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f706,f464])).
fof(f62195,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | $spl360),
  inference(cnf_transformation,[],[f62195_D])).
fof(f62195_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f62798,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X42)) | ~p(s(bool,X43))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50382,f4732])).
fof(f50382,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X159,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f459])).
fof(f62193,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | $spl358),
  inference(cnf_transformation,[],[f62193_D])).
fof(f62193_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f62797,plain,(
  ( ! [X41,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X40)) | ~p(s(bool,X41))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50382,f4603])).
fof(f62191,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | $spl356),
  inference(cnf_transformation,[],[f62191_D])).
fof(f62191_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f62796,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X38)) | ~p(s(bool,X39))) ) | $spl12),
  inference(resolution,[],[f50382,f11915])).
fof(f11915,plain,(
  ( ! [X127,X126] : (~sP7(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X126,X127) | p(s(bool,X127)) | ~p(s(bool,X126))) ) | $spl12),
  inference(resolution,[],[f11546,f456])).
fof(f62189,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | $spl354),
  inference(cnf_transformation,[],[f62189_D])).
fof(f62189_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f62795,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X36)) | ~p(s(bool,X37))) ) | $spl12),
  inference(resolution,[],[f50382,f11741])).
fof(f11741,plain,(
  ( ! [X127,X126] : (~sP7(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X126,X127) | p(s(bool,X127)) | ~p(s(bool,X126))) ) | $spl12),
  inference(resolution,[],[f11545,f456])).
fof(f62794,plain,(
  ( ! [X35] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X35))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f50382,f62318])).
fof(f62793,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f50382,f2095])).
fof(f2095,plain,(
  ( ! [X222,X221] : (~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X221,X222) | p(s(bool,X222)) | ~p(s(bool,X221))) ) | $spl12),
  inference(subsumption_resolution,[],[f1911,f595])).
fof(f1911,plain,(
  ( ! [X222,X221] : (~p(s(bool,'T_0')) | ~p(s(bool,X221)) | p(s(bool,X222)) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X221,X222)) ) | $spl12),
  inference(superposition,[],[f456,f713])).
fof(f62799,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl370),
  inference(cnf_transformation,[],[f62799_D])).
fof(f62799_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f62780,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X0,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f50382,f29192])).
fof(f62772,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50381,f4733])).
fof(f50381,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X156,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f458])).
fof(f62771,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50381,f4604])).
fof(f62770,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f50381,f11916])).
fof(f11916,plain,(
  ( ! [X128,X129] : (~sP7(X128,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X129) | p(s(bool,X129)) | ~p(s(bool,X128))) ) | $spl12),
  inference(resolution,[],[f11546,f456])).
fof(f62769,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f50381,f11742])).
fof(f11742,plain,(
  ( ! [X128,X129] : (~sP7(X128,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X129) | p(s(bool,X129)) | ~p(s(bool,X128))) ) | $spl12),
  inference(resolution,[],[f11545,f456])).
fof(f62768,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f50381,f2096])).
fof(f2096,plain,(
  ( ! [X223,X224] : (~sP7(X223,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X224) | p(s(bool,X224)) | ~p(s(bool,X223))) ) | $spl12),
  inference(subsumption_resolution,[],[f1912,f595])).
fof(f1912,plain,(
  ( ! [X223,X224] : (~p(s(bool,'T_0')) | ~p(s(bool,X223)) | p(s(bool,X224)) | ~sP7(X223,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X224)) ) | $spl12),
  inference(superposition,[],[f456,f713])).
fof(f62756,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X81)) | ~p(s(bool,X82))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50379,f4729])).
fof(f4729,plain,(
  ( ! [X121,X120] : (~sP8(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X120,X121) | p(s(bool,X121)) | ~p(s(bool,X120))) ) | $spl26),
  inference(resolution,[],[f4439,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f287])).
fof(f50379,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X150,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f455])).
fof(f62755,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X79)) | ~p(s(bool,X80))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50379,f4600])).
fof(f4600,plain,(
  ( ! [X121,X120] : (~sP8(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X120,X121) | p(s(bool,X121)) | ~p(s(bool,X120))) ) | $spl24),
  inference(resolution,[],[f4421,f452])).
fof(f62754,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X77)) | ~p(s(bool,X78))) ) | $spl12),
  inference(resolution,[],[f50379,f11912])).
fof(f11912,plain,(
  ( ! [X121,X120] : (~sP8(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X120,X121) | p(s(bool,X121)) | ~p(s(bool,X120))) ) | $spl12),
  inference(resolution,[],[f11546,f452])).
fof(f62753,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X75)) | ~p(s(bool,X76))) ) | $spl12),
  inference(resolution,[],[f50379,f11738])).
fof(f11738,plain,(
  ( ! [X121,X120] : (~sP8(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X120,X121) | p(s(bool,X121)) | ~p(s(bool,X120))) ) | $spl12),
  inference(resolution,[],[f11545,f452])).
fof(f62752,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X73)) | ~p(s(bool,X74))) ) | $spl12),
  inference(resolution,[],[f50379,f2093])).
fof(f2093,plain,(
  ( ! [X210,X209] : (~sP8('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X209,X210) | p(s(bool,X210)) | ~p(s(bool,X209))) ) | $spl12),
  inference(subsumption_resolution,[],[f1905,f595])).
fof(f1905,plain,(
  ( ! [X210,X209] : (~p(s(bool,'T_0')) | ~p(s(bool,X209)) | p(s(bool,X210)) | ~sP8('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X209,X210)) ) | $spl12),
  inference(superposition,[],[f452,f713])).
fof(f62714,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X71)) | ~p(s(bool,X72))) ) | $spl12),
  inference(resolution,[],[f50378,f2094])).
fof(f2094,plain,(
  ( ! [X212,X211] : (~sP8(X211,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X212) | p(s(bool,X212)) | ~p(s(bool,X211))) ) | $spl12),
  inference(subsumption_resolution,[],[f1906,f595])).
fof(f1906,plain,(
  ( ! [X212,X211] : (~p(s(bool,'T_0')) | ~p(s(bool,X211)) | p(s(bool,X212)) | ~sP8(X211,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X212)) ) | $spl12),
  inference(superposition,[],[f452,f713])).
fof(f62713,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X69)) | ~p(s(bool,X70))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50378,f4730])).
fof(f4730,plain,(
  ( ! [X123,X122] : (~sP8(X122,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X123) | p(s(bool,X123)) | ~p(s(bool,X122))) ) | $spl26),
  inference(resolution,[],[f4439,f452])).
fof(f62712,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X67)) | ~p(s(bool,X68))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50378,f4601])).
fof(f4601,plain,(
  ( ! [X123,X122] : (~sP8(X122,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X123) | p(s(bool,X123)) | ~p(s(bool,X122))) ) | $spl24),
  inference(resolution,[],[f4421,f452])).
fof(f62711,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X65)) | ~p(s(bool,X66))) ) | $spl12),
  inference(resolution,[],[f50378,f11913])).
fof(f11913,plain,(
  ( ! [X123,X122] : (~sP8(X122,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X123) | p(s(bool,X123)) | ~p(s(bool,X122))) ) | $spl12),
  inference(resolution,[],[f11546,f452])).
fof(f62710,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X63)) | ~p(s(bool,X64))) ) | $spl12),
  inference(resolution,[],[f50378,f11739])).
fof(f11739,plain,(
  ( ! [X123,X122] : (~sP8(X122,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X123) | p(s(bool,X123)) | ~p(s(bool,X122))) ) | $spl12),
  inference(resolution,[],[f11545,f452])).
fof(f62696,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X32,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f50378,f3739])).
fof(f3739,plain,(
  ( ! [X12,X13] : (~sP8(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3738,f461])).
fof(f3738,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~sP8(X12,X13,'F_0') | ~sP6(X13,'F_0',X12)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3737,f528])).
fof(f3737,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | p(s(bool,'F_0')) | ~sP8(X12,X13,'F_0') | ~sP6(X13,'F_0',X12)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3727,f459])).
fof(f3727,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | p(s(bool,'F_0')) | ~sP8(X12,X13,'F_0') | ~sP7('F_0',X13,X12) | ~sP6(X13,'F_0',X12)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3361,f450])).
fof(f3361,plain,(
  ( ! [X10,X11] : (~sP9(X11,'F_0',X10) | ~p(s(bool,X11)) | s(bool,'T_0') != s(bool,X10)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1779,f466])).
fof(f1779,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl6 | $spl12)),
  inference(backward_demodulation,[],[f713,f1274])).
fof(f1274,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | $spl6),
  inference(backward_demodulation,[],[f706,f679])).
fof(f62694,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X27,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f50378,f30061])).
fof(f62644,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') = s(bool,X44) | s(bool,X44) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X45))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X45)))) ) | $spl6),
  inference(subsumption_resolution,[],[f62592,f864])).
fof(f62592,plain,(
  ( ! [X45,X44] : (s(bool,'F_0') = s(bool,X44) | p(s(bool,'const.bool.IN_2'(s(sK12,X45),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X44) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X45))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X45)))) ) | $spl6),
  inference(resolution,[],[f3805,f977])).
fof(f3805,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') = s(bool,X0) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f706,f464])).
fof(f62560,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X41)) | ~p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50366,f4716])).
fof(f4716,plain,(
  ( ! [X94,X95] : (~sP4(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X94,X95) | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl26),
  inference(resolution,[],[f4439,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f50366,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X111,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f279])).
fof(f62559,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50366,f4587])).
fof(f4587,plain,(
  ( ! [X94,X95] : (~sP4(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X94,X95) | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl24),
  inference(resolution,[],[f4421,f435])).
fof(f62558,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f50366,f11899])).
fof(f11899,plain,(
  ( ! [X94,X95] : (~sP4(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X94,X95) | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl12),
  inference(resolution,[],[f11546,f435])).
fof(f62557,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f50366,f11725])).
fof(f11725,plain,(
  ( ! [X94,X95] : (~sP4(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X94,X95) | p(s(bool,X95)) | ~p(s(bool,X94))) ) | $spl12),
  inference(resolution,[],[f11545,f435])).
fof(f62556,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f50366,f2081])).
fof(f2081,plain,(
  ( ! [X158,X157] : (~sP4('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X157,X158) | p(s(bool,X158)) | ~p(s(bool,X157))) ) | $spl12),
  inference(subsumption_resolution,[],[f1879,f595])).
fof(f1879,plain,(
  ( ! [X158,X157] : (~p(s(bool,'T_0')) | ~p(s(bool,X157)) | p(s(bool,X158)) | ~sP4('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X157,X158)) ) | $spl12),
  inference(superposition,[],[f435,f713])).
fof(f62535,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50365,f4717])).
fof(f4717,plain,(
  ( ! [X97,X96] : (~sP4(X96,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl26),
  inference(resolution,[],[f4439,f435])).
fof(f50365,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X108,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f279])).
fof(f62534,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50365,f4588])).
fof(f4588,plain,(
  ( ! [X97,X96] : (~sP4(X96,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl24),
  inference(resolution,[],[f4421,f435])).
fof(f62533,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f50365,f11900])).
fof(f11900,plain,(
  ( ! [X97,X96] : (~sP4(X96,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl12),
  inference(resolution,[],[f11546,f435])).
fof(f62532,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f50365,f11726])).
fof(f11726,plain,(
  ( ! [X97,X96] : (~sP4(X96,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X97) | p(s(bool,X97)) | ~p(s(bool,X96))) ) | $spl12),
  inference(resolution,[],[f11545,f435])).
fof(f62531,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f50365,f2082])).
fof(f2082,plain,(
  ( ! [X159,X160] : (~sP4(X159,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X160) | p(s(bool,X160)) | ~p(s(bool,X159))) ) | $spl12),
  inference(subsumption_resolution,[],[f1880,f595])).
fof(f1880,plain,(
  ( ! [X159,X160] : (~p(s(bool,'T_0')) | ~p(s(bool,X159)) | p(s(bool,X160)) | ~sP4(X159,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X160)) ) | $spl12),
  inference(superposition,[],[f435,f713])).
fof(f62522,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50352,f4700])).
fof(f4700,plain,(
  ( ! [X62,X63] : (~sP2(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X62,X63) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl26),
  inference(resolution,[],[f4439,f418])).
fof(f418,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f50352,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X69,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f272])).
fof(f62521,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X37)) | p(s(bool,X38))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50352,f4571])).
fof(f4571,plain,(
  ( ! [X62,X63] : (~sP2(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X62,X63) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl24),
  inference(resolution,[],[f4421,f418])).
fof(f62520,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f50352,f11883])).
fof(f11883,plain,(
  ( ! [X62,X63] : (~sP2(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X62,X63) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl12),
  inference(resolution,[],[f11546,f418])).
fof(f62519,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f50352,f11709])).
fof(f11709,plain,(
  ( ! [X62,X63] : (~sP2(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X62,X63) | p(s(bool,X63)) | p(s(bool,X62))) ) | $spl12),
  inference(resolution,[],[f11545,f418])).
fof(f62518,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f50352,f2069])).
fof(f2069,plain,(
  ( ! [X99,X100] : (~sP2('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X99,X100) | p(s(bool,X100)) | p(s(bool,X99))) ) | $spl12),
  inference(subsumption_resolution,[],[f1850,f595])).
fof(f1850,plain,(
  ( ! [X99,X100] : (~p(s(bool,'T_0')) | p(s(bool,X99)) | p(s(bool,X100)) | ~sP2('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X99,X100)) ) | $spl12),
  inference(superposition,[],[f418,f713])).
fof(f62370,plain,(
  ( ! [X24] : (p(s(bool,X24)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f62318,f48666])).
fof(f48666,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X159,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f459])).
fof(f62368,plain,(
  ( ! [X22] : (p(s(bool,X22)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12))))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f62318,f46578])).
fof(f46578,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X159,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f459])).
fof(f62366,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12))))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f62318,f44440])).
fof(f44440,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X159,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f459])).
fof(f62364,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f62318,f38234])).
fof(f38234,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X159,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f459])).
fof(f62362,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12))))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f62318,f36419])).
fof(f36419,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X159,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f459])).
fof(f62360,plain,(
  ( ! [X14] : (p(s(bool,X14)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f62318,f34807])).
fof(f34807,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X159,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f459])).
fof(f62358,plain,(
  ( ! [X12] : (p(s(bool,X12)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | ($spl12 | $spl364)),
  inference(resolution,[],[f62318,f33340])).
fof(f33340,plain,(
  ( ! [X161,X159,X160] : (sP7(X159,X160,X161) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X159,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f459])).
fof(f62131,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,X54) | s(bool,X54) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl350),
  inference(cnf_transformation,[],[f62131_D])).
fof(f62131_D,plain,(
  ( ! [X54] : (s(bool,'T_0') != s(bool,X54) | s(bool,X54) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f62332,plain,(
  ( ! [X57,X56] : (s(bool,'T_0') != s(bool,X57) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56))) | s(bool,X57) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(subsumption_resolution,[],[f62249,f864])).
fof(f62249,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'T_0') != s(bool,X57) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56))) | s(bool,X57) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(resolution,[],[f1932,f976])).
fof(f1932,plain,(
  ( ! [X263,X264] : (~sP9(X264,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X263) | p(s(bool,X264)) | s(bool,'T_0') != s(bool,X263)) ) | $spl12),
  inference(superposition,[],[f467,f713])).
fof(f62318,plain,(
  ( ! [X36] : (~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X36) | p(s(bool,X36))) ) | $spl364),
  inference(cnf_transformation,[],[f62318_D])).
fof(f62318_D,plain,(
  ( ! [X36] : (~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X36) | p(s(bool,X36))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f62320,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,X39) | p(s(bool,X38)) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl12 | $spl22)),
  inference(forward_demodulation,[],[f62319,f713])).
fof(f62319,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl12 | $spl22)),
  inference(subsumption_resolution,[],[f62282,f455])).
fof(f62282,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP8(X38,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl12 | $spl22)),
  inference(trivial_inequality_removal,[],[f62235])).
fof(f62235,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP8(X38,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X38)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f1932,f28175])).
fof(f62317,plain,(
  ( ! [X37,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl12 | $spl22)),
  inference(forward_demodulation,[],[f62284,f713])).
fof(f62284,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,X37) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl12 | $spl22)),
  inference(trivial_inequality_removal,[],[f62283])).
fof(f62283,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X37) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f62234])).
fof(f62234,plain,(
  ( ! [X37,X36] : (p(s(bool,X36)) | s(bool,'T_0') != s(bool,'T_0') | s(bool,X37) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X36)) | ~sP7('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),'T_0',X36)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f1932,f28176])).
fof(f62215,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | ~p(s(bool,X41)) | p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50336,f4686])).
fof(f4686,plain,(
  ( ! [X35,X34] : (~sP0(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X34,X35) | ~p(s(bool,X35)) | p(s(bool,X34))) ) | $spl26),
  inference(resolution,[],[f4439,f401])).
fof(f401,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f264])).
fof(f264,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f263])).
fof(f263,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f50336,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X21,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f404])).
fof(f404,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f264])).
fof(f62214,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50336,f4557])).
fof(f4557,plain,(
  ( ! [X35,X34] : (~sP0(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X34,X35) | ~p(s(bool,X35)) | p(s(bool,X34))) ) | $spl24),
  inference(resolution,[],[f4421,f401])).
fof(f62213,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | ~p(s(bool,X37)) | p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f50336,f11869])).
fof(f11869,plain,(
  ( ! [X35,X34] : (~sP0(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),X34,X35) | ~p(s(bool,X35)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f11546,f401])).
fof(f62212,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f50336,f11695])).
fof(f11695,plain,(
  ( ! [X35,X34] : (~sP0(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),X34,X35) | ~p(s(bool,X35)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f11545,f401])).
fof(f62211,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | ~p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f50336,f2056])).
fof(f2056,plain,(
  ( ! [X37,X38] : (~sP0('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X37,X38) | ~p(s(bool,X38)) | p(s(bool,X37))) ) | $spl12),
  inference(subsumption_resolution,[],[f1819,f595])).
fof(f1819,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X37)) | ~p(s(bool,X38)) | ~sP0('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X37,X38)) ) | $spl12),
  inference(superposition,[],[f401,f713])).
fof(f62181,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f50335,f2057])).
fof(f2057,plain,(
  ( ! [X39,X40] : (~sP0(X39,X40,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X40)) | ~p(s(bool,X39))) ) | $spl12),
  inference(subsumption_resolution,[],[f1820,f595])).
fof(f1820,plain,(
  ( ! [X39,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP0(X39,X40,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(superposition,[],[f401,f713])).
fof(f50335,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X18,sK12))))) ) | $spl12),
  inference(resolution,[],[f20795,f402])).
fof(f402,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f264])).
fof(f62180,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f50335,f4687])).
fof(f4687,plain,(
  ( ! [X37,X36] : (~sP0(X36,X37,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X36))) ) | $spl26),
  inference(resolution,[],[f4439,f401])).
fof(f62179,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f50335,f4558])).
fof(f4558,plain,(
  ( ! [X37,X36] : (~sP0(X36,X37,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X36))) ) | $spl24),
  inference(resolution,[],[f4421,f401])).
fof(f62178,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f50335,f11870])).
fof(f11870,plain,(
  ( ! [X37,X36] : (~sP0(X36,X37,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f11546,f401])).
fof(f62177,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f50335,f11696])).
fof(f11696,plain,(
  ( ! [X37,X36] : (~sP0(X36,X37,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f11545,f401])).
fof(f61849,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | $spl344),
  inference(cnf_transformation,[],[f61849_D])).
fof(f61849_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f62165,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X41)) | ~p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48666,f4732])).
fof(f61847,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | $spl342),
  inference(cnf_transformation,[],[f61847_D])).
fof(f61847_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f62164,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48666,f4603])).
fof(f61845,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | $spl340),
  inference(cnf_transformation,[],[f61845_D])).
fof(f61845_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f62163,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f48666,f11915])).
fof(f61843,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | $spl338),
  inference(cnf_transformation,[],[f61843_D])).
fof(f61843_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f62162,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f48666,f11741])).
fof(f62161,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f48666,f2095])).
fof(f62166,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl352),
  inference(cnf_transformation,[],[f62166_D])).
fof(f62166_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f62148,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X0,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f48666,f29192])).
fof(f62130,plain,(
  ( ! [X54,X53] : (s(bool,'T_0') != s(bool,X54) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X53))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X53))) | s(bool,X54) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(subsumption_resolution,[],[f62052,f864])).
fof(f62052,plain,(
  ( ! [X54,X53] : (p(s(bool,'const.bool.IN_2'(s(sK12,X53),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,'T_0') != s(bool,X54) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X53))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X53))) | s(bool,X54) = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(resolution,[],[f1931,f976])).
fof(f1931,plain,(
  ( ! [X261,X262] : (~sP9(X262,X261,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X262)) | s(bool,'T_0') != s(bool,X261)) ) | $spl12),
  inference(superposition,[],[f467,f713])).
fof(f62012,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48665,f4733])).
fof(f48665,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X156,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f458])).
fof(f62011,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48665,f4604])).
fof(f62010,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f48665,f11916])).
fof(f62009,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f48665,f11742])).
fof(f62008,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f48665,f2096])).
fof(f61996,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X81)) | ~p(s(bool,X82))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48663,f4729])).
fof(f48663,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X150,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f455])).
fof(f61995,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X79)) | ~p(s(bool,X80))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48663,f4600])).
fof(f61994,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X77)) | ~p(s(bool,X78))) ) | $spl12),
  inference(resolution,[],[f48663,f11912])).
fof(f61993,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X75)) | ~p(s(bool,X76))) ) | $spl12),
  inference(resolution,[],[f48663,f11738])).
fof(f61992,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X73)) | ~p(s(bool,X74))) ) | $spl12),
  inference(resolution,[],[f48663,f2093])).
fof(f61954,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X71)) | ~p(s(bool,X72))) ) | $spl12),
  inference(resolution,[],[f48662,f2094])).
fof(f61953,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X69)) | ~p(s(bool,X70))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48662,f4730])).
fof(f61952,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X67)) | ~p(s(bool,X68))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48662,f4601])).
fof(f61951,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X65)) | ~p(s(bool,X66))) ) | $spl12),
  inference(resolution,[],[f48662,f11913])).
fof(f61950,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X63)) | ~p(s(bool,X64))) ) | $spl12),
  inference(resolution,[],[f48662,f11739])).
fof(f61936,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X32,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f48662,f3739])).
fof(f61934,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X27,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f48662,f30061])).
fof(f61923,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X41)) | ~p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48650,f4716])).
fof(f48650,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X111,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f438])).
fof(f61922,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48650,f4587])).
fof(f61921,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f48650,f11899])).
fof(f61920,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f48650,f11725])).
fof(f61919,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f48650,f2081])).
fof(f61898,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48649,f4717])).
fof(f48649,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X108,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f437])).
fof(f61897,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48649,f4588])).
fof(f61896,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f48649,f11900])).
fof(f61895,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f48649,f11726])).
fof(f61894,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f48649,f2082])).
fof(f61885,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48636,f4700])).
fof(f48636,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X69,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f421])).
fof(f61884,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X37)) | p(s(bool,X38))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48636,f4571])).
fof(f61883,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f48636,f11883])).
fof(f61882,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f48636,f11709])).
fof(f61881,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f48636,f2069])).
fof(f61869,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | ~p(s(bool,X41)) | p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48620,f4686])).
fof(f48620,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X21,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f404])).
fof(f61868,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48620,f4557])).
fof(f61867,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | ~p(s(bool,X37)) | p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f48620,f11869])).
fof(f61866,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f48620,f11695])).
fof(f61865,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | ~p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f48620,f2056])).
fof(f61835,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f48619,f2057])).
fof(f48619,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X18,sK12))))) ) | $spl12),
  inference(resolution,[],[f20566,f402])).
fof(f61834,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f48619,f4687])).
fof(f61833,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f48619,f4558])).
fof(f61832,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f48619,f11870])).
fof(f61831,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f48619,f11696])).
fof(f61498,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | $spl330),
  inference(cnf_transformation,[],[f61498_D])).
fof(f61498_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f61819,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X41)) | ~p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46578,f4732])).
fof(f61496,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | $spl328),
  inference(cnf_transformation,[],[f61496_D])).
fof(f61496_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f61818,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46578,f4603])).
fof(f61494,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | $spl326),
  inference(cnf_transformation,[],[f61494_D])).
fof(f61494_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f61817,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f46578,f11915])).
fof(f61492,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | $spl324),
  inference(cnf_transformation,[],[f61492_D])).
fof(f61492_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f61816,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f46578,f11741])).
fof(f61815,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f46578,f2095])).
fof(f61820,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X0,sK15,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl336),
  inference(cnf_transformation,[],[f61820_D])).
fof(f61820_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X0,sK15,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f61802,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X0,sK15,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f46578,f29192])).
fof(f61795,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46577,f4733])).
fof(f46577,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X156,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f458])).
fof(f61794,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46577,f4604])).
fof(f61793,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f46577,f11916])).
fof(f61792,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f46577,f11742])).
fof(f61791,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f46577,f2096])).
fof(f61779,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X81)) | ~p(s(bool,X82))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46575,f4729])).
fof(f46575,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X150,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f455])).
fof(f61778,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X79)) | ~p(s(bool,X80))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46575,f4600])).
fof(f61777,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X77)) | ~p(s(bool,X78))) ) | $spl12),
  inference(resolution,[],[f46575,f11912])).
fof(f61776,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X75)) | ~p(s(bool,X76))) ) | $spl12),
  inference(resolution,[],[f46575,f11738])).
fof(f61775,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X73)) | ~p(s(bool,X74))) ) | $spl12),
  inference(resolution,[],[f46575,f2093])).
fof(f61737,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X71)) | ~p(s(bool,X72))) ) | $spl12),
  inference(resolution,[],[f46574,f2094])).
fof(f61736,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X69)) | ~p(s(bool,X70))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46574,f4730])).
fof(f61735,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X67)) | ~p(s(bool,X68))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46574,f4601])).
fof(f61734,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X65)) | ~p(s(bool,X66))) ) | $spl12),
  inference(resolution,[],[f46574,f11913])).
fof(f61733,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X63)) | ~p(s(bool,X64))) ) | $spl12),
  inference(resolution,[],[f46574,f11739])).
fof(f61719,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X32,sK15,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f46574,f3739])).
fof(f61717,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X27,sK15,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f46574,f30061])).
fof(f61706,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X41)) | ~p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46562,f4716])).
fof(f46562,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X111,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f438])).
fof(f61705,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46562,f4587])).
fof(f61704,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f46562,f11899])).
fof(f61703,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f46562,f11725])).
fof(f61702,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f46562,f2081])).
fof(f61547,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46561,f4717])).
fof(f46561,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X108,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f437])).
fof(f61546,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46561,f4588])).
fof(f61545,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f46561,f11900])).
fof(f61544,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f46561,f11726])).
fof(f61543,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f46561,f2082])).
fof(f61534,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46548,f4700])).
fof(f46548,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X69,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f421])).
fof(f61533,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X37)) | p(s(bool,X38))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46548,f4571])).
fof(f61532,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f46548,f11883])).
fof(f61531,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f46548,f11709])).
fof(f61530,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f46548,f2069])).
fof(f61518,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | ~p(s(bool,X41)) | p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46532,f4686])).
fof(f46532,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X21,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f404])).
fof(f61517,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46532,f4557])).
fof(f61516,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | ~p(s(bool,X37)) | p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f46532,f11869])).
fof(f61515,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f46532,f11695])).
fof(f61514,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | ~p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f46532,f2056])).
fof(f61484,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f46531,f2057])).
fof(f46531,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X18,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20366,f402])).
fof(f61483,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f46531,f4687])).
fof(f61482,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f46531,f4558])).
fof(f61481,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f46531,f11870])).
fof(f61480,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f46531,f11696])).
fof(f61106,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | $spl316),
  inference(cnf_transformation,[],[f61106_D])).
fof(f61106_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f61293,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X41)) | ~p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44440,f4732])).
fof(f61104,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | $spl314),
  inference(cnf_transformation,[],[f61104_D])).
fof(f61104_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f61292,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44440,f4603])).
fof(f61102,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | $spl312),
  inference(cnf_transformation,[],[f61102_D])).
fof(f61102_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f61291,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f44440,f11915])).
fof(f61100,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | $spl310),
  inference(cnf_transformation,[],[f61100_D])).
fof(f61100_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f61290,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f44440,f11741])).
fof(f61289,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f44440,f2095])).
fof(f61294,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X0,sK15,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl322),
  inference(cnf_transformation,[],[f61294_D])).
fof(f61294_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X0,sK15,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f61276,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X0,sK15,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f44440,f29192])).
fof(f61269,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44439,f4733])).
fof(f44439,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X156,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f458])).
fof(f61268,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44439,f4604])).
fof(f61267,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f44439,f11916])).
fof(f61266,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f44439,f11742])).
fof(f61265,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f44439,f2096])).
fof(f61253,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X81)) | ~p(s(bool,X82))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44437,f4729])).
fof(f44437,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X150,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f455])).
fof(f61252,plain,(
  ( ! [X80,X79] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X79)) | ~p(s(bool,X80))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44437,f4600])).
fof(f61251,plain,(
  ( ! [X78,X77] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X77)) | ~p(s(bool,X78))) ) | $spl12),
  inference(resolution,[],[f44437,f11912])).
fof(f61250,plain,(
  ( ! [X76,X75] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X75)) | ~p(s(bool,X76))) ) | $spl12),
  inference(resolution,[],[f44437,f11738])).
fof(f61249,plain,(
  ( ! [X74,X73] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X73)) | ~p(s(bool,X74))) ) | $spl12),
  inference(resolution,[],[f44437,f2093])).
fof(f61211,plain,(
  ( ! [X72,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X71)) | ~p(s(bool,X72))) ) | $spl12),
  inference(resolution,[],[f44436,f2094])).
fof(f61210,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X69)) | ~p(s(bool,X70))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44436,f4730])).
fof(f61209,plain,(
  ( ! [X68,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X67)) | ~p(s(bool,X68))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44436,f4601])).
fof(f61208,plain,(
  ( ! [X66,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X65)) | ~p(s(bool,X66))) ) | $spl12),
  inference(resolution,[],[f44436,f11913])).
fof(f61207,plain,(
  ( ! [X64,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X63)) | ~p(s(bool,X64))) ) | $spl12),
  inference(resolution,[],[f44436,f11739])).
fof(f61193,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X32,sK15,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f44436,f3739])).
fof(f61191,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X27,sK15,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f44436,f30061])).
fof(f61180,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X41)) | ~p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44424,f4716])).
fof(f44424,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X111,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f438])).
fof(f61179,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X39)) | ~p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44424,f4587])).
fof(f61178,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X37)) | ~p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f44424,f11899])).
fof(f61177,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X35)) | ~p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f44424,f11725])).
fof(f61176,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f44424,f2081])).
fof(f61155,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44423,f4717])).
fof(f44423,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X108,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f437])).
fof(f61154,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44423,f4588])).
fof(f61153,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f44423,f11900])).
fof(f61152,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f44423,f11726])).
fof(f61151,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f44423,f2082])).
fof(f61142,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44410,f4700])).
fof(f44410,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X69,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f421])).
fof(f61141,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X37)) | p(s(bool,X38))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44410,f4571])).
fof(f61140,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f44410,f11883])).
fof(f61139,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f44410,f11709])).
fof(f61138,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X31)) | p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f44410,f2069])).
fof(f61126,plain,(
  ( ! [X41,X42] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | ~p(s(bool,X41)) | p(s(bool,X42))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44394,f4686])).
fof(f44394,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X21,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f404])).
fof(f61125,plain,(
  ( ! [X39,X40] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | ~p(s(bool,X39)) | p(s(bool,X40))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44394,f4557])).
fof(f61124,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | ~p(s(bool,X37)) | p(s(bool,X38))) ) | $spl12),
  inference(resolution,[],[f44394,f11869])).
fof(f61123,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | ~p(s(bool,X35)) | p(s(bool,X36))) ) | $spl12),
  inference(resolution,[],[f44394,f11695])).
fof(f61122,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | ~p(s(bool,X33)) | p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f44394,f2056])).
fof(f61092,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK15,sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f44393,f2057])).
fof(f44393,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X18,sK15,sK12))))) ) | $spl12),
  inference(resolution,[],[f20172,f402])).
fof(f61091,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK15,sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f44393,f4687])).
fof(f61090,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK15,sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f44393,f4558])).
fof(f61089,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK15,sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f44393,f11870])).
fof(f61088,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK15,sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f44393,f11696])).
fof(f60405,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X2)))))) ) | $spl308),
  inference(cnf_transformation,[],[f60405_D])).
fof(f60405_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X2)))))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f60383,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X2)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f41981,f29451])).
fof(f41981,plain,(
  ( ! [X191,X189,X190] : (sP6('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X189)))),X190,X191)) ) | $spl12),
  inference(resolution,[],[f11846,f463])).
fof(f25430,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl116),
  inference(cnf_transformation,[],[f25430_D])).
fof(f25430_D,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f60315,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | s(bool,X42) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43)))) ) | $spl12),
  inference(subsumption_resolution,[],[f60277,f864])).
fof(f60277,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X42) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43)))) ) | $spl12),
  inference(resolution,[],[f1786,f977])).
fof(f1786,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X2) | s(bool,'T_0') = s(bool,X2) | p(s(bool,X3))) ) | $spl12),
  inference(superposition,[],[f713,f464])).
fof(f60199,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | $spl306),
  inference(cnf_transformation,[],[f60199_D])).
fof(f60199_D,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f60173,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f41975,f28678])).
fof(f60171,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X29)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f41975,f30061])).
fof(f60170,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK28(sK15,sK12)),s('type.list.list'(sK12),X26)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f41975,f30058])).
fof(f59614,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X2)))))) ) | $spl302),
  inference(cnf_transformation,[],[f59614_D])).
fof(f59614_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X2)))))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f59592,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X2)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f40237,f29451])).
fof(f40237,plain,(
  ( ! [X191,X189,X190] : (sP6('const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X189)))),X190,X191)) ) | $spl12),
  inference(resolution,[],[f11672,f463])).
fof(f59392,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | $spl300),
  inference(cnf_transformation,[],[f59392_D])).
fof(f59392_D,plain,(
  ( ! [X35,X34] : (s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f59366,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | s(bool,X34) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X35))))) | s(bool,'F_0') != s(bool,X34)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f40231,f28678])).
fof(f59364,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X29)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f40231,f30061])).
fof(f59363,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,sK27(sK15,sK12)),s('type.list.list'(sK12),X26)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f40231,f30058])).
fof(f58281,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X56)))) | ~p(s(bool,X54))) )),
  inference(subsumption_resolution,[],[f58203,f557])).
fof(f58203,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X56)))) | ~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f1027,f588])).
fof(f58280,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X50)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X49))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X49))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X50)))) | ~p(s(bool,X48))) )),
  inference(subsumption_resolution,[],[f58201,f557])).
fof(f58201,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X50)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X49))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X49))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X50)))) | ~p(s(bool,X48)) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f1027,f587])).
fof(f58263,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X47)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X47)))) | ~p(s(bool,X45))) )),
  inference(duplicate_literal_removal,[],[f58200])).
fof(f58200,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X47)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X47)))) | ~p(s(bool,X45)) | ~p(s(bool,X45))) )),
  inference(superposition,[],[f1027,f564])).
fof(f58264,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | ~p(s(bool,X42))) )),
  inference(duplicate_literal_removal,[],[f58199])).
fof(f58199,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | ~p(s(bool,X42)) | ~p(s(bool,X42))) )),
  inference(superposition,[],[f1027,f562])).
fof(f55253,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | $spl296),
  inference(cnf_transformation,[],[f55253_D])).
fof(f55253_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f57249,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4733,f38233])).
fof(f38233,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X156,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f458])).
fof(f55251,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) | $spl294),
  inference(cnf_transformation,[],[f55251_D])).
fof(f55251_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12)))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f57247,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4733,f36418])).
fof(f36418,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X156,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f458])).
fof(f55249,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12)))) | $spl292),
  inference(cnf_transformation,[],[f55249_D])).
fof(f55249_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12)))) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f57245,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4733,f34806])).
fof(f34806,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X156,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f458])).
fof(f55247,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12)))) | $spl290),
  inference(cnf_transformation,[],[f55247_D])).
fof(f55247_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12)))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f57243,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4733,f33339])).
fof(f33339,plain,(
  ( ! [X158,X156,X157] : (sP7(X157,X156,X158) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X156,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f458])).
fof(f57177,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4732,f38234])).
fof(f57175,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4732,f36419])).
fof(f57173,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4732,f34807])).
fof(f57171,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4732,f33340])).
fof(f57087,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4730,f38230])).
fof(f57085,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4730,f36415])).
fof(f57083,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4730,f34803])).
fof(f57081,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4730,f33336])).
fof(f57074,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X66)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X64))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X64))) | p(s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),X66)))) | ~p(s(bool,X65))) )),
  inference(subsumption_resolution,[],[f56983,f1000])).
fof(f56983,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X66)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X64))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X64))) | p(s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),X66)))) | ~p(s(bool,X65)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X64),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1026,f588])).
fof(f57073,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X60)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X58))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X58))) | p(s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),X60)))) | ~p(s(bool,X59))) )),
  inference(subsumption_resolution,[],[f56981,f1000])).
fof(f56981,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X60)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X58))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X58))) | p(s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),X60)))) | ~p(s(bool,X59)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X58),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1026,f587])).
fof(f57072,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X47),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X48)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,X47),s(fun(bool,bool),X48)))) | ~p(s(bool,X47))) )),
  inference(subsumption_resolution,[],[f56977,f864])).
fof(f56977,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X47),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X48)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,X47),s(fun(bool,bool),X48)))) | ~p(s(bool,X47)) | p(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f560])).
fof(f57071,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.bool.IN_2'(s(bool,X44),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X45)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,X44),s(fun(bool,bool),X45)))) | ~p(s(bool,X44))) )),
  inference(subsumption_resolution,[],[f56976,f864])).
fof(f56976,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.bool.IN_2'(s(bool,X44),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X45)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,X44),s(fun(bool,bool),X45)))) | ~p(s(bool,X44)) | p(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f558])).
fof(f57070,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X32),s(fun(bool,bool),X34)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X31))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X31))) | p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),X34)))) | ~p(s(bool,X33)) | p(s(bool,X32))) )),
  inference(subsumption_resolution,[],[f56973,f864])).
fof(f56973,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X32),s(fun(bool,bool),X34)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X31))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X31))) | p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),X34)))) | ~p(s(bool,X33)) | p(s(bool,'const.bool.IN_2'(s(sK12,X31),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | p(s(bool,X32))) )),
  inference(superposition,[],[f1026,f515])).
fof(f57069,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X28),s(fun(bool,bool),X30)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),X30)))) | ~p(s(bool,X29)) | p(s(bool,X28))) )),
  inference(subsumption_resolution,[],[f56972,f864])).
fof(f56972,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X28),s(fun(bool,bool),X30)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),X30)))) | ~p(s(bool,X29)) | p(s(bool,X28)) | p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f515])).
fof(f57010,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X4))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X4))) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6)))) | ~p(s(bool,X5))) )),
  inference(duplicate_literal_removal,[],[f56967])).
fof(f56967,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X4))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X4))) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X6)))) | ~p(s(bool,X5)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X4))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X4)))) )),
  inference(superposition,[],[f1026,f998])).
fof(f56778,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4729,f38231])).
fof(f38231,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X150,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f455])).
fof(f56776,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4729,f36416])).
fof(f36416,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X150,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f455])).
fof(f56774,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4729,f34804])).
fof(f34804,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X150,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f455])).
fof(f56772,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4729,f33337])).
fof(f33337,plain,(
  ( ! [X152,X151,X150] : (sP8(X150,X151,X152) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X150,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f455])).
fof(f56052,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4717,f38217])).
fof(f38217,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X108,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f437])).
fof(f56050,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4717,f36402])).
fof(f36402,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X108,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f437])).
fof(f56048,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4717,f34790])).
fof(f34790,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X108,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f437])).
fof(f56046,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4717,f33323])).
fof(f33323,plain,(
  ( ! [X109,X110,X108] : (sP4(X109,X108,X110) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X108,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f437])).
fof(f56000,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4716,f38218])).
fof(f38218,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X111,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f438])).
fof(f55998,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4716,f36403])).
fof(f36403,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X111,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f438])).
fof(f55996,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4716,f34791])).
fof(f34791,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X111,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f438])).
fof(f55994,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4716,f33324])).
fof(f33324,plain,(
  ( ! [X111,X112,X113] : (sP4(X111,X112,X113) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X111,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f438])).
fof(f55662,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4700,f38204])).
fof(f38204,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X69,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f421])).
fof(f55661,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4700,f36389])).
fof(f36389,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X69,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f421])).
fof(f55660,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4700,f34777])).
fof(f34777,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X69,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f421])).
fof(f55659,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4700,f33310])).
fof(f33310,plain,(
  ( ! [X70,X71,X69] : (sP2(X69,X70,X71) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X69,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f421])).
fof(f55285,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4687,f38187])).
fof(f38187,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X18,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f402])).
fof(f55283,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4687,f36372])).
fof(f36372,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X18,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f402])).
fof(f55281,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4687,f34760])).
fof(f34760,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X18,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f402])).
fof(f55279,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4687,f33293])).
fof(f33293,plain,(
  ( ! [X19,X20,X18] : (sP0(X19,X20,X18) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X18,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f402])).
fof(f55209,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4686,f38188])).
fof(f38188,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X21,sK12))))) ) | $spl12),
  inference(resolution,[],[f17382,f404])).
fof(f55207,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4686,f36373])).
fof(f36373,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X21,sK12))))) ) | $spl12),
  inference(resolution,[],[f17107,f404])).
fof(f55205,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4686,f34761])).
fof(f34761,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X21,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16832,f404])).
fof(f55203,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f4686,f33294])).
fof(f33294,plain,(
  ( ! [X23,X21,X22] : (sP0(X21,X22,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X21,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f16638,f404])).
fof(f53088,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | $spl288),
  inference(cnf_transformation,[],[f53088_D])).
fof(f53088_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f54693,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4604,f38233])).
fof(f53086,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) | $spl286),
  inference(cnf_transformation,[],[f53086_D])).
fof(f53086_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12)))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f54691,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4604,f36418])).
fof(f53084,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12)))) | $spl284),
  inference(cnf_transformation,[],[f53084_D])).
fof(f53084_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12)))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f54689,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4604,f34806])).
fof(f53082,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12)))) | $spl282),
  inference(cnf_transformation,[],[f53082_D])).
fof(f53082_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12)))) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f54687,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4604,f33339])).
fof(f54621,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4603,f38234])).
fof(f54619,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4603,f36419])).
fof(f54617,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4603,f34807])).
fof(f54615,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4603,f33340])).
fof(f54531,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4601,f38230])).
fof(f54529,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4601,f36415])).
fof(f54527,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4601,f34803])).
fof(f54525,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4601,f33336])).
fof(f54479,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4600,f38231])).
fof(f54477,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4600,f36416])).
fof(f54475,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4600,f34804])).
fof(f54473,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4600,f33337])).
fof(f53894,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4588,f38217])).
fof(f53892,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4588,f36402])).
fof(f53890,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4588,f34790])).
fof(f53888,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4588,f33323])).
fof(f53842,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4587,f38218])).
fof(f53840,plain,(
  ( ! [X19,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4587,f36403])).
fof(f53838,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4587,f34791])).
fof(f53836,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4587,f33324])).
fof(f53561,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4571,f38204])).
fof(f53560,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4571,f36389])).
fof(f53559,plain,(
  ( ! [X10,X11] : (p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4571,f34777])).
fof(f53558,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4571,f33310])).
fof(f53188,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4558,f38187])).
fof(f53186,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4558,f36372])).
fof(f53184,plain,(
  ( ! [X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4558,f34760])).
fof(f53182,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4558,f33293])).
fof(f53044,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,X23)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4557,f38188])).
fof(f53042,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4557,f36373])).
fof(f53040,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4557,f34761])).
fof(f53038,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,X11)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),sK16,sK12))))) ) | ($spl12 | $spl24)),
  inference(resolution,[],[f4557,f33294])).
fof(f52680,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | $spl274),
  inference(cnf_transformation,[],[f52680_D])).
fof(f52680_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f52830,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f38234,f11915])).
fof(f52678,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | $spl272),
  inference(cnf_transformation,[],[f52678_D])).
fof(f52678_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f52829,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f38234,f11741])).
fof(f52828,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f38234,f2095])).
fof(f52831,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl280),
  inference(cnf_transformation,[],[f52831_D])).
fof(f52831_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f52818,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X0,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f38234,f29192])).
fof(f52813,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f38233,f11916])).
fof(f52812,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f38233,f11742])).
fof(f52811,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f38233,f2096])).
fof(f52800,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X70)) | ~p(s(bool,X71))) ) | $spl12),
  inference(resolution,[],[f38231,f11912])).
fof(f52799,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X68)) | ~p(s(bool,X69))) ) | $spl12),
  inference(resolution,[],[f38231,f11738])).
fof(f52798,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X66)) | ~p(s(bool,X67))) ) | $spl12),
  inference(resolution,[],[f38231,f2093])).
fof(f52765,plain,(
  ( ! [X64,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X64)) | ~p(s(bool,X65))) ) | $spl12),
  inference(resolution,[],[f38230,f2094])).
fof(f52764,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X62)) | ~p(s(bool,X63))) ) | $spl12),
  inference(resolution,[],[f38230,f11913])).
fof(f52763,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X60)) | ~p(s(bool,X61))) ) | $spl12),
  inference(resolution,[],[f38230,f11739])).
fof(f52750,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X32,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f38230,f3739])).
fof(f52748,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X27,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f38230,f30061])).
fof(f52737,plain,(
  ( ! [X33,X34] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X33)) | ~p(s(bool,X34))) ) | $spl12),
  inference(resolution,[],[f38218,f11899])).
fof(f52736,plain,(
  ( ! [X31,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X31)) | ~p(s(bool,X32))) ) | $spl12),
  inference(resolution,[],[f38218,f11725])).
fof(f52735,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X29)) | ~p(s(bool,X30))) ) | $spl12),
  inference(resolution,[],[f38218,f2081])).
fof(f52718,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X27)) | ~p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f38217,f11900])).
fof(f52717,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X25)) | ~p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f38217,f11726])).
fof(f52716,plain,(
  ( ! [X24,X23] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X23)) | ~p(s(bool,X24))) ) | $spl12),
  inference(resolution,[],[f38217,f2082])).
fof(f52707,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl12),
  inference(resolution,[],[f38204,f11883])).
fof(f52706,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f38204,f11709])).
fof(f52705,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X25)) | p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f38204,f2069])).
fof(f52695,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f38188,f11869])).
fof(f52694,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | ~p(s(bool,X28)) | p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f38188,f11695])).
fof(f52693,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f38188,f2056])).
fof(f52672,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f38187,f2057])).
fof(f52671,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f38187,f11870])).
fof(f52670,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f38187,f11696])).
fof(f52360,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | $spl264),
  inference(cnf_transformation,[],[f52360_D])).
fof(f52360_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f52635,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f36419,f11915])).
fof(f52358,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | $spl262),
  inference(cnf_transformation,[],[f52358_D])).
fof(f52358_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f52634,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f36419,f11741])).
fof(f52633,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f36419,f2095])).
fof(f52636,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl270),
  inference(cnf_transformation,[],[f52636_D])).
fof(f52636_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X0,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f52623,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X0,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f36419,f29192])).
fof(f52618,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f36418,f11916])).
fof(f52617,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f36418,f11742])).
fof(f52616,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f36418,f2096])).
fof(f52605,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X70)) | ~p(s(bool,X71))) ) | $spl12),
  inference(resolution,[],[f36416,f11912])).
fof(f52604,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X68)) | ~p(s(bool,X69))) ) | $spl12),
  inference(resolution,[],[f36416,f11738])).
fof(f52603,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X66)) | ~p(s(bool,X67))) ) | $spl12),
  inference(resolution,[],[f36416,f2093])).
fof(f52570,plain,(
  ( ! [X64,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X64)) | ~p(s(bool,X65))) ) | $spl12),
  inference(resolution,[],[f36415,f2094])).
fof(f52569,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X62)) | ~p(s(bool,X63))) ) | $spl12),
  inference(resolution,[],[f36415,f11913])).
fof(f52568,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X60)) | ~p(s(bool,X61))) ) | $spl12),
  inference(resolution,[],[f36415,f11739])).
fof(f52555,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X32,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f36415,f3739])).
fof(f52553,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X27,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f36415,f30061])).
fof(f52542,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f36403,f11899])).
fof(f52541,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f36403,f11725])).
fof(f52540,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f36403,f2081])).
fof(f52524,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f36402,f11900])).
fof(f52523,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f36402,f11726])).
fof(f52522,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f36402,f2082])).
fof(f52514,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl12),
  inference(resolution,[],[f36389,f11883])).
fof(f52513,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f36389,f11709])).
fof(f52512,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X25)) | p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f36389,f2069])).
fof(f52502,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f36373,f11869])).
fof(f52501,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | ~p(s(bool,X28)) | p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f36373,f11695])).
fof(f52500,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f36373,f2056])).
fof(f52352,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f36372,f2057])).
fof(f52351,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f36372,f11870])).
fof(f52350,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f36372,f11696])).
fof(f52191,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | $spl254),
  inference(cnf_transformation,[],[f52191_D])).
fof(f52191_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f52339,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f34807,f11915])).
fof(f52189,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | $spl252),
  inference(cnf_transformation,[],[f52189_D])).
fof(f52189_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f52338,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f34807,f11741])).
fof(f52337,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f34807,f2095])).
fof(f52340,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X0,sK16,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl260),
  inference(cnf_transformation,[],[f52340_D])).
fof(f52340_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X0,sK16,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f52327,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X0,sK16,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f34807,f29192])).
fof(f52322,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f34806,f11916])).
fof(f52321,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f34806,f11742])).
fof(f52320,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f34806,f2096])).
fof(f52309,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X70)) | ~p(s(bool,X71))) ) | $spl12),
  inference(resolution,[],[f34804,f11912])).
fof(f52308,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X68)) | ~p(s(bool,X69))) ) | $spl12),
  inference(resolution,[],[f34804,f11738])).
fof(f52307,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X66)) | ~p(s(bool,X67))) ) | $spl12),
  inference(resolution,[],[f34804,f2093])).
fof(f52274,plain,(
  ( ! [X64,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X64)) | ~p(s(bool,X65))) ) | $spl12),
  inference(resolution,[],[f34803,f2094])).
fof(f52273,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X62)) | ~p(s(bool,X63))) ) | $spl12),
  inference(resolution,[],[f34803,f11913])).
fof(f52272,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X60)) | ~p(s(bool,X61))) ) | $spl12),
  inference(resolution,[],[f34803,f11739])).
fof(f52259,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X32,sK16,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f34803,f3739])).
fof(f52257,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X27,sK16,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f34803,f30061])).
fof(f52246,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f34791,f11899])).
fof(f52245,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f34791,f11725])).
fof(f52244,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f34791,f2081])).
fof(f52228,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f34790,f11900])).
fof(f52227,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f34790,f11726])).
fof(f52226,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f34790,f2082])).
fof(f52218,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl12),
  inference(resolution,[],[f34777,f11883])).
fof(f52217,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f34777,f11709])).
fof(f52216,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X25)) | p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f34777,f2069])).
fof(f52206,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f34761,f11869])).
fof(f52205,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | ~p(s(bool,X28)) | p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f34761,f11695])).
fof(f52204,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f34761,f2056])).
fof(f52183,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f34760,f2057])).
fof(f52182,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f34760,f11870])).
fof(f52181,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f34760,f11696])).
fof(f52022,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | $spl244),
  inference(cnf_transformation,[],[f52022_D])).
fof(f52022_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f52170,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f33340,f11915])).
fof(f52020,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | $spl242),
  inference(cnf_transformation,[],[f52020_D])).
fof(f52020_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f52169,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f33340,f11741])).
fof(f52168,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f33340,f2095])).
fof(f52171,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X0,sK16,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) | $spl250),
  inference(cnf_transformation,[],[f52171_D])).
fof(f52171_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X0,sK16,sK12)))) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X1))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f52158,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X0,sK16,sK12)))) | ~p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f33340,f29192])).
fof(f52153,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f33339,f11916])).
fof(f52152,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f33339,f11742])).
fof(f52151,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f33339,f2096])).
fof(f52140,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X70)) | ~p(s(bool,X71))) ) | $spl12),
  inference(resolution,[],[f33337,f11912])).
fof(f52139,plain,(
  ( ! [X68,X69] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X68)) | ~p(s(bool,X69))) ) | $spl12),
  inference(resolution,[],[f33337,f11738])).
fof(f52138,plain,(
  ( ! [X66,X67] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X66)) | ~p(s(bool,X67))) ) | $spl12),
  inference(resolution,[],[f33337,f2093])).
fof(f52105,plain,(
  ( ! [X64,X65] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X64)) | ~p(s(bool,X65))) ) | $spl12),
  inference(resolution,[],[f33336,f2094])).
fof(f52104,plain,(
  ( ! [X62,X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X62)) | ~p(s(bool,X63))) ) | $spl12),
  inference(resolution,[],[f33336,f11913])).
fof(f52103,plain,(
  ( ! [X61,X60] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X60)) | ~p(s(bool,X61))) ) | $spl12),
  inference(resolution,[],[f33336,f11739])).
fof(f52090,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X32,sK16,sK12)))) | s(bool,'T_0') != s(bool,X32) | ~p(s(bool,X33))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f33336,f3739])).
fof(f52088,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X27,sK16,sK12)))) | ~p(s(bool,X28)) | s(bool,'T_0') != s(bool,X27) | s(bool,'F_0') != s(bool,X29)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f33336,f30061])).
fof(f52077,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X30)) | ~p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f33324,f11899])).
fof(f52076,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X28)) | ~p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f33324,f11725])).
fof(f52075,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X26)) | ~p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f33324,f2081])).
fof(f52059,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f33323,f11900])).
fof(f52058,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f33323,f11726])).
fof(f52057,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f33323,f2082])).
fof(f52049,plain,(
  ( ! [X30,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X29)) | p(s(bool,X30))) ) | $spl12),
  inference(resolution,[],[f33310,f11883])).
fof(f52048,plain,(
  ( ! [X28,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X27)) | p(s(bool,X28))) ) | $spl12),
  inference(resolution,[],[f33310,f11709])).
fof(f52047,plain,(
  ( ! [X26,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X25)) | p(s(bool,X26))) ) | $spl12),
  inference(resolution,[],[f33310,f2069])).
fof(f52037,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | ~p(s(bool,X30)) | p(s(bool,X31))) ) | $spl12),
  inference(resolution,[],[f33294,f11869])).
fof(f52036,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | ~p(s(bool,X28)) | p(s(bool,X29))) ) | $spl12),
  inference(resolution,[],[f33294,f11695])).
fof(f52035,plain,(
  ( ! [X26,X27] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | ~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl12),
  inference(resolution,[],[f33294,f2056])).
fof(f52014,plain,(
  ( ! [X24,X25] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | p(s(bool,X24)) | ~p(s(bool,X25))) ) | $spl12),
  inference(resolution,[],[f33293,f2057])).
fof(f52013,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12))),sK16,sK12)))) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl12),
  inference(resolution,[],[f33293,f11870])).
fof(f52012,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12))),sK16,sK12)))) | p(s(bool,X20)) | ~p(s(bool,X21))) ) | $spl12),
  inference(resolution,[],[f33293,f11696])).
fof(f50530,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | $spl240),
  inference(cnf_transformation,[],[f50530_D])).
fof(f50530_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f50521,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f50493])).
fof(f50493,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f20795,f562])).
fof(f50522,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f50492])).
fof(f50492,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f20795,f564])).
fof(f50538,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f50491,f557])).
fof(f50491,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f20795,f587])).
fof(f50537,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f50490,f557])).
fof(f50490,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f20795,f588])).
fof(f50419,plain,(
  ( ! [X225] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | ~p(s(bool,X225))) ) | $spl12),
  inference(resolution,[],[f20795,f546])).
fof(f546,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f331])).
fof(f331,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f330])).
fof(f330,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : ((p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X16] : ((p(s(bool,'T_0')) | p(s(bool,X16))) <=> p(s(bool,'T_0')))),
  file('thm.list.EXISTS_CONG','thm.bool.OR_CLAUSES|split|0')).
fof(f50418,plain,(
  ( ! [X224] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | ~p(s(bool,X224))) ) | $spl12),
  inference(resolution,[],[f20795,f537])).
fof(f537,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f325])).
fof(f50417,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f20795,f536])).
fof(f50416,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f20795,f533])).
fof(f50415,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,'T_0',sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f20795,f530])).
fof(f48813,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | $spl238),
  inference(cnf_transformation,[],[f48813_D])).
fof(f48813_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f48804,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f48777])).
fof(f48777,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f20566,f562])).
fof(f48805,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f48776])).
fof(f48776,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f20566,f564])).
fof(f48821,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f48775,f557])).
fof(f48775,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f20566,f587])).
fof(f48820,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f48774,f557])).
fof(f48774,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f20566,f588])).
fof(f48703,plain,(
  ( ! [X225] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | ~p(s(bool,X225))) ) | $spl12),
  inference(resolution,[],[f20566,f546])).
fof(f48702,plain,(
  ( ! [X224] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | ~p(s(bool,X224))) ) | $spl12),
  inference(resolution,[],[f20566,f537])).
fof(f48701,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f20566,f536])).
fof(f48700,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f20566,f533])).
fof(f48699,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,'T_0',sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f20566,f530])).
fof(f48321,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f48278])).
fof(f48278,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X17) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f1013,f465])).
fof(f48020,plain,(
  ( ! [X54,X56,X55] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | ~p(s(bool,X55)) | s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK12,X54),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56)))) )),
  inference(subsumption_resolution,[],[f47988,f864])).
fof(f47988,plain,(
  ( ! [X54,X56,X55] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | ~p(s(bool,X55)) | p(s(bool,'const.bool.IN_2'(s(sK12,X56),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK12,X54),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56)))) )),
  inference(resolution,[],[f1011,f977])).
fof(f48016,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f47972])).
fof(f47972,plain,(
  ( ! [X17,X15,X16] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X15))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X15))) | ~p(s(bool,X16)) | p(s(bool,X17)) | p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(sK12,X15),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1011,f467])).
fof(f47678,plain,(
  ( ! [X54,X55,X53] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X53))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X53))) | ~p(s(bool,X54)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK12,X53),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55)))) )),
  inference(subsumption_resolution,[],[f47643,f864])).
fof(f47643,plain,(
  ( ! [X54,X55,X53] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X53))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X53))) | ~p(s(bool,X54)) | p(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK12,X53),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55)))) )),
  inference(resolution,[],[f1010,f977])).
fof(f47672,plain,(
  ( ! [X19,X20,X18] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X18))) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK12,X18),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f47628])).
fof(f47628,plain,(
  ( ! [X19,X20,X18] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X18))) | ~p(s(bool,X19)) | p(s(bool,X20)) | p(s(bool,X20)) | s(bool,X19) != s(bool,'const.bool.IN_2'(s(sK12,X18),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1010,f467])).
fof(f46724,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | $spl236),
  inference(cnf_transformation,[],[f46724_D])).
fof(f46724_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f46715,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f46689])).
fof(f46689,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f20366,f562])).
fof(f46716,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f46688])).
fof(f46688,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f20366,f564])).
fof(f46732,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f46687,f557])).
fof(f46687,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f20366,f587])).
fof(f46731,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f46686,f557])).
fof(f46686,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f20366,f588])).
fof(f46615,plain,(
  ( ! [X225] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | ~p(s(bool,X225))) ) | $spl12),
  inference(resolution,[],[f20366,f546])).
fof(f46614,plain,(
  ( ! [X224] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | ~p(s(bool,X224))) ) | $spl12),
  inference(resolution,[],[f20366,f537])).
fof(f46613,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f20366,f536])).
fof(f46612,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f20366,f533])).
fof(f46611,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22('T_0',sK15,sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f20366,f530])).
fof(f46076,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f46026])).
fof(f46026,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,X13) != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X12))) )),
  inference(resolution,[],[f964,f465])).
fof(f44585,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | $spl234),
  inference(cnf_transformation,[],[f44585_D])).
fof(f44585_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f44576,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f44551])).
fof(f44551,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f20172,f562])).
fof(f44577,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f44550])).
fof(f44550,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f20172,f564])).
fof(f44593,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f44549,f557])).
fof(f44549,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f20172,f587])).
fof(f44592,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f44548,f557])).
fof(f44548,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f20172,f588])).
fof(f44477,plain,(
  ( ! [X225] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | ~p(s(bool,X225))) ) | $spl12),
  inference(resolution,[],[f20172,f546])).
fof(f44476,plain,(
  ( ! [X224] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | ~p(s(bool,X224))) ) | $spl12),
  inference(resolution,[],[f20172,f537])).
fof(f44475,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f20172,f536])).
fof(f44474,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f20172,f533])).
fof(f44473,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21('T_0',sK15,sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f20172,f530])).
fof(f39974,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | $spl232),
  inference(cnf_transformation,[],[f39974_D])).
fof(f39974_D,plain,(
  ( ! [X2,X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f39957,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X2),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f32519,f29451])).
fof(f32519,plain,(
  ( ! [X198,X200,X199] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X198),s('type.list.list'(sK12),sK13)),sK16,sK12))),X199,X200)) ) | $spl12),
  inference(resolution,[],[f19522,f463])).
fof(f39906,plain,(
  ( ! [X33,X34] : (s(bool,X33) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | s(bool,'F_0') != s(bool,X33)) ) | $spl230),
  inference(cnf_transformation,[],[f39906_D])).
fof(f39906_D,plain,(
  ( ! [X33,X34] : (s(bool,X33) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | s(bool,'F_0') != s(bool,X33)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f39892,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | s(bool,X33) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X34),s('type.list.list'(sK12),sK13)),sK16,sK12)))) | s(bool,'F_0') != s(bool,X33)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f32513,f28678])).
fof(f39890,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | s(bool,'T_0') != s(bool,X28) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X29),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f32513,f30061])).
fof(f39889,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17('const.list.CONS_2'(s(sK12,X26),s('type.list.list'(sK12),sK13)),sK16,sK12))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f32513,f30058])).
fof(f39859,plain,(
  ( ! [X66,X67] : (s(bool,X66) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X67))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X67))) | s(bool,'T_0') = s(bool,X66)) ) | $spl12),
  inference(subsumption_resolution,[],[f39817,f864])).
fof(f39817,plain,(
  ( ! [X66,X67] : (s(bool,X66) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X67))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X67))) | s(bool,'T_0') = s(bool,X66) | p(s(bool,'const.bool.IN_2'(s(sK12,X67),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl12),
  inference(resolution,[],[f977,f1785])).
fof(f1785,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) ) | $spl12),
  inference(superposition,[],[f713,f464])).
fof(f39590,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,X58)) | s(bool,'T_0') != s(bool,X59) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X60),s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30061,f19578])).
fof(f39586,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | s(bool,'T_0') != s(bool,X50) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X51,sK12)))) | ~p(s(bool,X51))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30061,f20767])).
fof(f20767,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f11542,f453])).
fof(f39585,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,X46)) | s(bool,'T_0') != s(bool,X47) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X48,sK12)))) | ~p(s(bool,X48))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30061,f20538])).
fof(f20538,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f11541,f453])).
fof(f39584,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | s(bool,'T_0') != s(bool,X44) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X45,sK15,sK12)))) | ~p(s(bool,X45))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30061,f20338])).
fof(f20338,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X192,sK15,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f11540,f453])).
fof(f39583,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | s(bool,'T_0') != s(bool,X41) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X42,sK15,sK12)))) | ~p(s(bool,X42))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30061,f20144])).
fof(f20144,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X192,sK15,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f11539,f453])).
fof(f39565,plain,(
  ( ! [X59,X60] : (s(bool,X59) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X60),s('type.list.list'(sK12),sK13))))) | s(bool,'F_0') != s(bool,X59)) ) | $spl226),
  inference(cnf_transformation,[],[f39565_D])).
fof(f39565_D,plain,(
  ( ! [X59,X60] : (s(bool,X59) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X60),s('type.list.list'(sK12),sK13))))) | s(bool,'F_0') != s(bool,X59)) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f39539,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,X58)) | s(bool,X59) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X60),s('type.list.list'(sK12),sK13))))) | s(bool,'F_0') != s(bool,X59)) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f19578])).
fof(f39535,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | s(bool,X50) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X51,sK12)))) | s(bool,'F_0') != s(bool,X50) | ~p(s(bool,X51))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f20767])).
fof(f39534,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,X46)) | s(bool,X47) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X48,sK12)))) | s(bool,'F_0') != s(bool,X47) | ~p(s(bool,X48))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f20538])).
fof(f39533,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | s(bool,X44) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X45,sK15,sK12)))) | s(bool,'F_0') != s(bool,X44) | ~p(s(bool,X45))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f20338])).
fof(f39532,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | s(bool,X41) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X42,sK15,sK12)))) | s(bool,'F_0') != s(bool,X41) | ~p(s(bool,X42))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f20144])).
fof(f39527,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | s(bool,X30) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X31,sK12)))) | s(bool,'F_0') != s(bool,X30) | ~p(s(bool,X31))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f17354])).
fof(f17354,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f5250,f453])).
fof(f39526,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X26)) | s(bool,X27) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X28,sK12)))) | s(bool,'F_0') != s(bool,X27) | ~p(s(bool,X28))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f17079])).
fof(f17079,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X192,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f5249,f453])).
fof(f39525,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | s(bool,X24) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X25,sK16,sK12)))) | s(bool,'F_0') != s(bool,X24) | ~p(s(bool,X25))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f16804])).
fof(f16804,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X192,sK16,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f5248,f453])).
fof(f39524,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | s(bool,X21) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X22,sK16,sK12)))) | s(bool,'F_0') != s(bool,X21) | ~p(s(bool,X22))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28678,f16610])).
fof(f16610,plain,(
  ( ! [X194,X192,X193] : (sP8(X193,X194,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X192,sK16,sK12)))) | ~p(s(bool,X192))) ) | $spl12),
  inference(resolution,[],[f5247,f453])).
fof(f38375,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | $spl208),
  inference(cnf_transformation,[],[f38375_D])).
fof(f38375_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f38366,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f38342])).
fof(f38342,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f17382,f562])).
fof(f38367,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f38341])).
fof(f38341,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f17382,f564])).
fof(f38383,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f38340,f557])).
fof(f38340,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f17382,f587])).
fof(f38382,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f38339,f557])).
fof(f38339,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f17382,f588])).
fof(f38270,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | ~p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f17382,f546])).
fof(f38269,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | ~p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f17382,f537])).
fof(f38268,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f17382,f536])).
fof(f38267,plain,(
  ( ! [X220] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | p(s(bool,X220))) ) | $spl12),
  inference(resolution,[],[f17382,f533])).
fof(f38266,plain,(
  ( ! [X219] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,'T_0',sK12)))) | p(s(bool,X219))) ) | $spl12),
  inference(resolution,[],[f17382,f530])).
fof(f36559,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | $spl206),
  inference(cnf_transformation,[],[f36559_D])).
fof(f36559_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f36550,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f36527])).
fof(f36527,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f17107,f562])).
fof(f36551,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f36526])).
fof(f36526,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f17107,f564])).
fof(f36567,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f36525,f557])).
fof(f36525,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f17107,f587])).
fof(f36566,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f36524,f557])).
fof(f36524,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f17107,f588])).
fof(f36455,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | ~p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f17107,f546])).
fof(f36454,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | ~p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f17107,f537])).
fof(f36453,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f17107,f536])).
fof(f36452,plain,(
  ( ! [X220] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | p(s(bool,X220))) ) | $spl12),
  inference(resolution,[],[f17107,f533])).
fof(f36451,plain,(
  ( ! [X219] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,'T_0',sK12)))) | p(s(bool,X219))) ) | $spl12),
  inference(resolution,[],[f17107,f530])).
fof(f34946,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | $spl204),
  inference(cnf_transformation,[],[f34946_D])).
fof(f34946_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f34937,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f34915])).
fof(f34915,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f16832,f562])).
fof(f34938,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f34914])).
fof(f34914,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f16832,f564])).
fof(f34954,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f34913,f557])).
fof(f34913,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f16832,f587])).
fof(f34953,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f34912,f557])).
fof(f34912,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f16832,f588])).
fof(f34843,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | ~p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f16832,f546])).
fof(f34842,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | ~p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f16832,f537])).
fof(f34841,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f16832,f536])).
fof(f34840,plain,(
  ( ! [X220] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | p(s(bool,X220))) ) | $spl12),
  inference(resolution,[],[f16832,f533])).
fof(f34839,plain,(
  ( ! [X219] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22('T_0',sK16,sK12)))) | p(s(bool,X219))) ) | $spl12),
  inference(resolution,[],[f16832,f530])).
fof(f33478,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | $spl198),
  inference(cnf_transformation,[],[f33478_D])).
fof(f33478_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f33469,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f33448])).
fof(f33448,plain,(
  ( ! [X37] : (~p(s(bool,X37)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X37))) ) | $spl12),
  inference(superposition,[],[f16638,f562])).
fof(f33470,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f33447])).
fof(f33447,plain,(
  ( ! [X36] : (~p(s(bool,X36)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X36))) ) | $spl12),
  inference(superposition,[],[f16638,f564])).
fof(f33489,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f33446,f557])).
fof(f33446,plain,(
  ( ! [X35] : (~p(s(bool,X35)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | s(bool,'F_0') = s(bool,X35)) ) | $spl12),
  inference(superposition,[],[f16638,f587])).
fof(f33488,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12))))) ) | $spl12),
  inference(subsumption_resolution,[],[f33445,f557])).
fof(f33445,plain,(
  ( ! [X34] : (~p(s(bool,X34)) | s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | s(bool,'F_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f16638,f588])).
fof(f33481,plain,(
  ( ! [X225,X224] : (~sP11(X224,X225) | p(s(bool,X224)) | p(s(bool,X225))) ) | $spl200),
  inference(cnf_transformation,[],[f33481_D])).
fof(f33481_D,plain,(
  ( ! [X225,X224] : (~sP11(X224,X225) | p(s(bool,X224)) | p(s(bool,X225))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f33483,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('F_0',sK16,sK12)))) | $spl202),
  inference(cnf_transformation,[],[f33483_D])).
fof(f33483_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('F_0',sK16,sK12)))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f33386,plain,(
  ( ! [X229,X228] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('F_0',sK16,sK12)))) | p(s(bool,X228)) | p(s(bool,X229)) | ~sP11(X228,X229)) ) | $spl12),
  inference(resolution,[],[f16638,f498])).
fof(f33384,plain,(
  ( ! [X225,X224] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('F_0',sK16,sK12)))) | p(s(bool,X224)) | p(s(bool,X225)) | ~sP11(X224,X225)) ) | $spl12),
  inference(resolution,[],[f16638,f502])).
fof(f33376,plain,(
  ( ! [X223] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X223))) ) | $spl12),
  inference(resolution,[],[f16638,f546])).
fof(f33375,plain,(
  ( ! [X222] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | ~p(s(bool,X222))) ) | $spl12),
  inference(resolution,[],[f16638,f537])).
fof(f33374,plain,(
  ( ! [X221] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X221))) ) | $spl12),
  inference(resolution,[],[f16638,f536])).
fof(f33373,plain,(
  ( ! [X220] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X220))) ) | $spl12),
  inference(resolution,[],[f16638,f533])).
fof(f33372,plain,(
  ( ! [X219] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21('T_0',sK16,sK12)))) | p(s(bool,X219))) ) | $spl12),
  inference(resolution,[],[f16638,f530])).
fof(f30094,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X60),s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f19578])).
fof(f30090,plain,(
  ( ! [X50,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X51,sK12)))) | ~p(s(bool,X51))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f20767])).
fof(f30089,plain,(
  ( ! [X47,X48,X46] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X48,sK12)))) | ~p(s(bool,X48))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f20538])).
fof(f30088,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X45,sK15,sK12)))) | ~p(s(bool,X45))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f20338])).
fof(f30087,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X42,sK15,sK12)))) | ~p(s(bool,X42))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f20144])).
fof(f30082,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X31,sK12)))) | ~p(s(bool,X31))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f17354])).
fof(f30081,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X28,sK12)))) | ~p(s(bool,X28))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f17079])).
fof(f30080,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X25,sK16,sK12)))) | ~p(s(bool,X25))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f16804])).
fof(f30079,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,'F_0') != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X22,sK16,sK12)))) | ~p(s(bool,X22))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f16610])).
fof(f28968,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | $spl155),
  inference(cnf_transformation,[],[f28968_D])).
fof(f28968_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f30076,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f15526])).
fof(f15526,plain,(
  ( ! [X114,X115] : (sP8(X114,X115,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl12),
  inference(resolution,[],[f5094,f453])).
fof(f28965,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | $spl151),
  inference(cnf_transformation,[],[f28965_D])).
fof(f28965_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f30075,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f30058,f11122])).
fof(f11122,plain,(
  ( ! [X114,X115] : (sP8(X114,X115,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl12),
  inference(resolution,[],[f2045,f453])).
fof(f29888,plain,(
  ( ! [X187,X185] : (p(s(bool,X185)) | s(bool,X185) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X187),s('type.list.list'(sK12),sK13)))))) ) | $spl178),
  inference(cnf_transformation,[],[f29888_D])).
fof(f29888_D,plain,(
  ( ! [X187,X185] : (p(s(bool,X185)) | s(bool,X185) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X187),s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f29848,plain,(
  ( ! [X187,X185,X186] : (p(s(bool,X185)) | p(s(bool,X186)) | s(bool,X185) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X187),s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f19584])).
fof(f19584,plain,(
  ( ! [X191,X189,X190] : (sP6('const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),'const.list.CONS_2'(s(sK12,X189),s('type.list.list'(sK12),sK13)))),X190,X191)) ) | $spl12),
  inference(resolution,[],[f2047,f463])).
fof(f29845,plain,(
  ( ! [X180,X178,X179] : (p(s(bool,X178)) | p(s(bool,X179)) | s(bool,X178) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X180,sK12)))) | ~p(s(bool,X180))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f20773])).
fof(f20773,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f11542,f463])).
fof(f29844,plain,(
  ( ! [X177,X175,X176] : (p(s(bool,X175)) | p(s(bool,X176)) | s(bool,X175) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X177,sK12)))) | ~p(s(bool,X177))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f20544])).
fof(f20544,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f11541,f463])).
fof(f29843,plain,(
  ( ! [X173,X174,X172] : (p(s(bool,X172)) | p(s(bool,X173)) | s(bool,X172) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X174,sK15,sK12)))) | ~p(s(bool,X174))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f20344])).
fof(f20344,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X210,sK15,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f11540,f463])).
fof(f29842,plain,(
  ( ! [X171,X169,X170] : (p(s(bool,X169)) | p(s(bool,X170)) | s(bool,X169) != s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X171,sK15,sK12)))) | ~p(s(bool,X171))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f20150])).
fof(f20150,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X210,sK15,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f11539,f463])).
fof(f29841,plain,(
  ( ! [X167,X166,X168] : (p(s(bool,X166)) | p(s(bool,X167)) | s(bool,X166) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X168,sK12)))) | ~p(s(bool,X168))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f17360])).
fof(f17360,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f5250,f463])).
fof(f29840,plain,(
  ( ! [X165,X163,X164] : (p(s(bool,X163)) | p(s(bool,X164)) | s(bool,X163) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X165,sK12)))) | ~p(s(bool,X165))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f17085])).
fof(f17085,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X210,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f5249,f463])).
fof(f29839,plain,(
  ( ! [X161,X162,X160] : (p(s(bool,X160)) | p(s(bool,X161)) | s(bool,X160) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X162,sK16,sK12)))) | ~p(s(bool,X162))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f16810])).
fof(f16810,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X210,sK16,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f5248,f463])).
fof(f29838,plain,(
  ( ! [X158,X159,X157] : (p(s(bool,X157)) | p(s(bool,X158)) | s(bool,X157) != s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X159,sK16,sK12)))) | ~p(s(bool,X159))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f16616])).
fof(f16616,plain,(
  ( ! [X212,X210,X211] : (sP6(happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X210,sK16,sK12))),X211,X212) | ~p(s(bool,X210))) ) | $spl12),
  inference(resolution,[],[f5247,f463])).
fof(f29870,plain,(
  ( ! [X139] : (p(s(bool,X139)) | s(bool,X139) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl160),
  inference(cnf_transformation,[],[f29870_D])).
fof(f29870_D,plain,(
  ( ! [X139] : (p(s(bool,X139)) | s(bool,X139) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f29829,plain,(
  ( ! [X140,X139] : (p(s(bool,X139)) | p(s(bool,X140)) | s(bool,X139) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f15532])).
fof(f15532,plain,(
  ( ! [X127,X126] : (sP6('const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X126,X127)) ) | $spl12),
  inference(resolution,[],[f5094,f463])).
fof(f29869,plain,(
  ( ! [X137] : (p(s(bool,X137)) | s(bool,X137) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl158),
  inference(cnf_transformation,[],[f29869_D])).
fof(f29869_D,plain,(
  ( ! [X137] : (p(s(bool,X137)) | s(bool,X137) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f29828,plain,(
  ( ! [X138,X137] : (p(s(bool,X137)) | p(s(bool,X138)) | s(bool,X137) != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f29451,f11128])).
fof(f11128,plain,(
  ( ! [X127,X126] : (sP6('const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X126,X127)) ) | $spl12),
  inference(resolution,[],[f2045,f463])).
fof(f29756,plain,(
  ( ! [X94,X95,X93] : (~p(s(bool,X93)) | p(s(bool,X94)) | s(bool,X93) != s(bool,'const.bool.IN_2'(s(sK12,X95),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X95))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X95)))) ) | $spl22),
  inference(resolution,[],[f29192,f972])).
fof(f28969,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X15) | ~p(s(bool,X15))) ) | $spl156),
  inference(cnf_transformation,[],[f28969_D])).
fof(f28969_D,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X15) | ~p(s(bool,X15))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f28936,plain,(
  ( ! [X14,X15] : (sP9(X14,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X15) | ~p(s(bool,X15)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28884,f15526])).
fof(f28966,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13))) ) | $spl152),
  inference(cnf_transformation,[],[f28966_D])).
fof(f28966_D,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f28935,plain,(
  ( ! [X12,X13] : (sP9(X12,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X13) | ~p(s(bool,X13)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f28884,f11122])).
fof(f28067,plain,(
  ( ! [X27] : (~sP8(X27,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27))) ) | $spl148),
  inference(cnf_transformation,[],[f28067_D])).
fof(f28067_D,plain,(
  ( ! [X27] : (~sP8(X27,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f28066,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X26) | ~sP8(X27,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27))) ) | ($spl6 | $spl12 | $spl22)),
  inference(subsumption_resolution,[],[f28065,f528])).
fof(f28065,plain,(
  ( ! [X26,X27] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X26) | ~sP8(X27,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27))) ) | ($spl6 | $spl12 | $spl22)),
  inference(forward_demodulation,[],[f28064,f706])).
fof(f28064,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~sP8(X27,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27))) ) | ($spl6 | $spl12 | $spl22)),
  inference(forward_demodulation,[],[f28057,f706])).
fof(f28057,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~sP8(X27,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27))) ) | ($spl12 | $spl22)),
  inference(trivial_inequality_removal,[],[f28041])).
fof(f28041,plain,(
  ( ! [X26,X27] : (s(bool,X26) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~sP8(X27,'T_0','const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X27)) | s(bool,'T_0') != s(bool,'T_0')) ) | ($spl12 | $spl22)),
  inference(resolution,[],[f27940,f3387])).
fof(f28062,plain,(
  ( ! [X21] : (~sP8(X21,'T_0','F_0') | ~p(s(bool,X21))) ) | $spl146),
  inference(cnf_transformation,[],[f28062_D])).
fof(f28062_D,plain,(
  ( ! [X21] : (~sP8(X21,'T_0','F_0') | ~p(s(bool,X21))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f28061,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | ~sP8(X21,'T_0','F_0') | ~p(s(bool,X21))) ) | ($spl6 | $spl12 | $spl22)),
  inference(subsumption_resolution,[],[f28058,f528])).
fof(f28058,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | p(s(bool,'F_0')) | ~sP8(X21,'T_0','F_0') | ~p(s(bool,X21))) ) | ($spl6 | $spl12 | $spl22)),
  inference(trivial_inequality_removal,[],[f28038])).
fof(f28038,plain,(
  ( ! [X21,X20] : (s(bool,'F_0') != s(bool,X20) | p(s(bool,'F_0')) | ~sP8(X21,'T_0','F_0') | ~p(s(bool,X21)) | s(bool,'T_0') != s(bool,'T_0')) ) | ($spl6 | $spl12 | $spl22)),
  inference(resolution,[],[f27940,f3361])).
fof(f27941,plain,(
  ( ! [X17] : (sP9(X17,'F_0','T_0') | p(s(bool,X17))) ) | $spl144),
  inference(cnf_transformation,[],[f27941_D])).
fof(f27941_D,plain,(
  ( ! [X17] : (sP9(X17,'F_0','T_0') | p(s(bool,X17))) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f27936,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | p(s(bool,X17)) | sP9(X17,'F_0','T_0')) ) | ($spl6 | $spl12 | $spl22)),
  inference(duplicate_literal_removal,[],[f27934])).
fof(f27934,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X16) | s(bool,'F_0') != s(bool,X16) | p(s(bool,X17)) | sP9(X17,'F_0','T_0')) ) | ($spl6 | $spl12 | $spl22)),
  inference(resolution,[],[f27911,f3676])).
fof(f3676,plain,(
  ( ! [X6,X7] : (~sP6(X7,'F_0',X6) | p(s(bool,X6)) | sP9(X6,'F_0',X7)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3674,f455])).
fof(f3674,plain,(
  ( ! [X6,X7] : (sP9(X6,'F_0',X7) | p(s(bool,X6)) | ~sP8(X6,X7,'F_0') | ~sP6(X7,'F_0',X6)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3670,f451])).
fof(f3670,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3669,f458])).
fof(f3669,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | sP7('F_0',X6,X7)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3665,f457])).
fof(f3665,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3660,f446])).
fof(f3660,plain,(
  ( ! [X4,X5] : (sP9(X5,'F_0',X4) | ~p(s(bool,X4)) | p(s(bool,X5))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3655,f455])).
fof(f3655,plain,(
  ( ! [X2,X3] : (~sP8(X3,X2,'F_0') | ~p(s(bool,X2)) | sP9(X3,'F_0',X2)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3654,f562])).
fof(f3654,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3651,f3641])).
fof(f3641,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3636,f461])).
fof(f3636,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3359,f445])).
fof(f3359,plain,(
  ( ! [X6,X7] : (sP9(X7,'F_0',X6) | s(bool,'T_0') != s(bool,X6) | p(s(bool,X7))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1779,f464])).
fof(f3651,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3642,f449])).
fof(f3642,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f3637,f457])).
fof(f3637,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3359,f446])).
fof(f27910,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | s(bool,X3) != s(bool,X4)) ) | $spl142),
  inference(cnf_transformation,[],[f27910_D])).
fof(f27910_D,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') != s(bool,X4) | s(bool,X3) != s(bool,X4)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f27909,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f27882])).
fof(f27882,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,X3) != s(bool,X4) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl22),
  inference(resolution,[],[f4476,f467])).
fof(f27731,plain,(
  ( ! [X14,X15] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X15))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f968,f3739])).
fof(f26386,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f373])).
fof(f26385,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f372])).
fof(f26380,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f363])).
fof(f26379,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f362])).
fof(f26378,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f361])).
fof(f26377,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f360])).
fof(f26254,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f373])).
fof(f26253,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f372])).
fof(f26252,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f5249,f371])).
fof(f26248,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f363])).
fof(f26247,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f362])).
fof(f26246,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f361])).
fof(f26245,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f360])).
fof(f26122,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f373])).
fof(f26121,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f372])).
fof(f26120,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f5248,f371])).
fof(f26119,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl12),
  inference(resolution,[],[f5248,f369])).
fof(f26116,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f363])).
fof(f26115,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f362])).
fof(f26114,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f361])).
fof(f26113,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f360])).
fof(f25990,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f373])).
fof(f25989,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f372])).
fof(f25988,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f5247,f371])).
fof(f25987,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl12),
  inference(resolution,[],[f5247,f369])).
fof(f16671,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) | $spl82),
  inference(cnf_transformation,[],[f16671_D])).
fof(f16671_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f25986,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) | $spl12),
  inference(resolution,[],[f5247,f366])).
fof(f25984,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f363])).
fof(f25983,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f362])).
fof(f25982,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f361])).
fof(f25981,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f360])).
fof(f25907,plain,(
  ( ! [X62] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X62)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25614,f10721])).
fof(f25614,plain,(
  ( ! [X62] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X62))))) ) | $spl46),
  inference(superposition,[],[f375,f10721])).
fof(f25904,plain,(
  ( ! [X61] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X61)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25613,f10721])).
fof(f25613,plain,(
  ( ! [X61] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X61))))) ) | $spl46),
  inference(superposition,[],[f374,f10721])).
fof(f25902,plain,(
  ( ! [X56,X55] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X56)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X55,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X55),s(sK12,sK26(sK16,X55,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25610,f10721])).
fof(f25610,plain,(
  ( ! [X56,X55] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X55,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X55),s(sK12,sK26(sK16,X55,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X56))))) ) | $spl46),
  inference(superposition,[],[f372,f10721])).
fof(f25900,plain,(
  ( ! [X52,X51] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X52)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X51,sK12))))) | ~p(s(bool,X51))) ) | $spl46),
  inference(forward_demodulation,[],[f25608,f10721])).
fof(f25608,plain,(
  ( ! [X52,X51] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X51,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X52)))) | ~p(s(bool,X51))) ) | $spl46),
  inference(superposition,[],[f371,f10721])).
fof(f25899,plain,(
  ( ! [X50,X49] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X50)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X49,sK12))))) | ~p(s(bool,X49))) ) | $spl46),
  inference(forward_demodulation,[],[f25607,f10721])).
fof(f25607,plain,(
  ( ! [X50,X49] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X49,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X50)))) | ~p(s(bool,X49))) ) | $spl46),
  inference(superposition,[],[f369,f10721])).
fof(f25898,plain,(
  ( ! [X47,X48] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X48)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X47,sK16,sK12))))) | ~p(s(bool,X47))) ) | $spl46),
  inference(forward_demodulation,[],[f25606,f10721])).
fof(f25606,plain,(
  ( ! [X47,X48] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X47,sK16,sK12))))) | ~p(s(bool,X47)) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X48))))) ) | $spl46),
  inference(superposition,[],[f366,f10721])).
fof(f25897,plain,(
  ( ! [X45,X46] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X45,sK16,sK12))))) | ~p(s(bool,X45))) ) | $spl46),
  inference(forward_demodulation,[],[f25605,f10721])).
fof(f25605,plain,(
  ( ! [X45,X46] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X45,sK16,sK12))))) | ~p(s(bool,X45)) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))))) ) | $spl46),
  inference(superposition,[],[f364,f10721])).
fof(f25896,plain,(
  ( ! [X41,X42] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X41,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X41),s(sK12,sK20(X41,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25603,f10721])).
fof(f25603,plain,(
  ( ! [X41,X42] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X41,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X41),s(sK12,sK20(X41,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42))))) ) | $spl46),
  inference(superposition,[],[f362,f10721])).
fof(f25895,plain,(
  ( ! [X39,X40] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X40)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X39,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X39),s(sK12,sK20(sK16,X39,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25602,f10721])).
fof(f25602,plain,(
  ( ! [X39,X40] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X39,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X39),s(sK12,sK20(sK16,X39,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X40))))) ) | $spl46),
  inference(superposition,[],[f363,f10721])).
fof(f25894,plain,(
  ( ! [X35,X36] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X36)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X35,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X35),s(sK12,sK18(sK16,X35,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25600,f10721])).
fof(f25600,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X35,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X35),s(sK12,sK18(sK16,X35,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X36))))) ) | $spl46),
  inference(superposition,[],[f361,f10721])).
fof(f25893,plain,(
  ( ! [X30,X29] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X30)))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X29,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X29),s(sK12,sK19(X29,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25597,f10721])).
fof(f25597,plain,(
  ( ! [X30,X29] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X29,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X29),s(sK12,sK19(X29,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X30))))) ) | $spl46),
  inference(superposition,[],[f360,f10721])).
fof(f25593,plain,(
  ( ! [X24,X25] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X24)))) | p(s(bool,happ(s(fun(sK12,bool),X25),s(sK12,sK26(sK16,X25,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X25),s(sK12,sK25(sK16,X25,sK12)))))) ) | $spl46),
  inference(superposition,[],[f373,f10721])).
fof(f25889,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22)))) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25592,f10721])).
fof(f25592,plain,(
  ( ! [X23,X22] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X22)))) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl46),
  inference(superposition,[],[f372,f10721])).
fof(f25887,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK16,X21,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X20)))) | ~p(s(bool,X21))) ) | $spl46),
  inference(forward_demodulation,[],[f25591,f10721])).
fof(f25591,plain,(
  ( ! [X21,X20] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X20)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X21,sK12))))) | ~p(s(bool,X21))) ) | $spl46),
  inference(superposition,[],[f371,f10721])).
fof(f25885,plain,(
  ( ! [X19,X18] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK16,X19,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18)))) | ~p(s(bool,X19))) ) | $spl46),
  inference(forward_demodulation,[],[f25590,f10721])).
fof(f25590,plain,(
  ( ! [X19,X18] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X18)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl46),
  inference(superposition,[],[f369,f10721])).
fof(f25883,plain,(
  ( ! [X17,X16] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X17,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X16)))) | ~p(s(bool,X17))) ) | $spl46),
  inference(forward_demodulation,[],[f25589,f10721])).
fof(f25589,plain,(
  ( ! [X17,X16] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X16)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X17,sK16,sK12))))) | ~p(s(bool,X17))) ) | $spl46),
  inference(superposition,[],[f366,f10721])).
fof(f25881,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK21(X15,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14)))) | ~p(s(bool,X15))) ) | $spl46),
  inference(forward_demodulation,[],[f25588,f10721])).
fof(f25588,plain,(
  ( ! [X14,X15] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK21(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) | $spl46),
  inference(superposition,[],[f364,f10721])).
fof(f25879,plain,(
  ( ! [X12,X13] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK16,X13,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X12)))) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK20(sK16,X13,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25587,f10721])).
fof(f25587,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X12)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X13,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X13),s(sK12,sK20(sK16,X13,sK12)))))) ) | $spl46),
  inference(superposition,[],[f363,f10721])).
fof(f25877,plain,(
  ( ! [X10,X11] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X11,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X10)))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(X11,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25586,f10721])).
fof(f25586,plain,(
  ( ! [X10,X11] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X10)))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(X11,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X11,sK16,sK12)))))) ) | $spl46),
  inference(superposition,[],[f362,f10721])).
fof(f25875,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK16,X9,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X8)))) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK18(sK16,X9,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25585,f10721])).
fof(f25585,plain,(
  ( ! [X8,X9] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X8)))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X9,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK18(sK16,X9,sK12)))))) ) | $spl46),
  inference(superposition,[],[f361,f10721])).
fof(f25872,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X7,sK16,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6)))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK19(X7,sK16,sK12)))))) ) | $spl46),
  inference(forward_demodulation,[],[f25584,f10721])).
fof(f25584,plain,(
  ( ! [X6,X7] : (~p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6)))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK19(X7,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X7,sK16,sK12)))))) ) | $spl46),
  inference(superposition,[],[f360,f10721])).
fof(f25429,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23)))) ) | $spl12),
  inference(subsumption_resolution,[],[f25417,f864])).
fof(f25417,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') = s(bool,X22) | p(s(bool,'const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23)))) ) | $spl12),
  inference(resolution,[],[f1785,f977])).
fof(f20710,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11542,f373])).
fof(f20709,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11542,f372])).
fof(f20704,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11542,f363])).
fof(f20703,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11542,f362])).
fof(f20702,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11542,f361])).
fof(f20701,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11542,f360])).
fof(f20481,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11541,f373])).
fof(f20480,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11541,f372])).
fof(f20479,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f11541,f371])).
fof(f20475,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11541,f363])).
fof(f20474,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11541,f362])).
fof(f20473,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11541,f361])).
fof(f20472,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11541,f360])).
fof(f20281,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11540,f373])).
fof(f20280,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11540,f372])).
fof(f20279,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f11540,f371])).
fof(f20278,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl12),
  inference(resolution,[],[f11540,f369])).
fof(f20275,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11540,f363])).
fof(f20274,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11540,f362])).
fof(f20273,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11540,f361])).
fof(f20272,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11540,f360])).
fof(f20087,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK15,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK15,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11539,f373])).
fof(f20086,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK15,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK25(sK15,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11539,f372])).
fof(f20085,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK24(sK15,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f11539,f371])).
fof(f20084,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK23(sK15,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl12),
  inference(resolution,[],[f11539,f369])).
fof(f20083,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK22(X15,sK15,sK12))))) | ~p(s(bool,X15))) ) | $spl12),
  inference(resolution,[],[f11539,f366])).
fof(f20081,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(sK15,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK15,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11539,f363])).
fof(f20080,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK20(X9,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11539,f362])).
fof(f20079,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK19(sK15,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK15,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11539,f361])).
fof(f20078,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK15,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK18(X5,sK15,sK12)))))) ) | $spl12),
  inference(resolution,[],[f11539,f360])).
fof(f18685,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X56)))) | ~p(s(bool,X54))) )),
  inference(subsumption_resolution,[],[f18618,f557])).
fof(f18618,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X55))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X55))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X55),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X56)))) | ~p(s(bool,X54)) | s(bool,'F_0') = s(bool,X54)) )),
  inference(superposition,[],[f1027,f588])).
fof(f18684,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X50)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X49))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X49))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X50)))) | ~p(s(bool,X48))) )),
  inference(subsumption_resolution,[],[f18616,f557])).
fof(f18616,plain,(
  ( ! [X50,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X50)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X49))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X49))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X49),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X50)))) | ~p(s(bool,X48)) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f1027,f587])).
fof(f18668,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X47)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X47)))) | ~p(s(bool,X45))) )),
  inference(duplicate_literal_removal,[],[f18615])).
fof(f18615,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X47)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X46))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X46))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X46),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X47)))) | ~p(s(bool,X45)) | ~p(s(bool,X45))) )),
  inference(superposition,[],[f1027,f564])).
fof(f18669,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | ~p(s(bool,X42))) )),
  inference(duplicate_literal_removal,[],[f18614])).
fof(f18614,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X43))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X43))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK12,X43),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))),s(fun(bool,bool),X44)))) | ~p(s(bool,X42)) | ~p(s(bool,X42))) )),
  inference(superposition,[],[f1027,f562])).
fof(f17659,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X62)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X60))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X60))) | p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),X62)))) | ~p(s(bool,X61))) )),
  inference(subsumption_resolution,[],[f17589,f1000])).
fof(f17589,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X62)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X60))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X60))) | p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),X62)))) | ~p(s(bool,X61)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X60),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1026,f588])).
fof(f17658,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),X56)))) | ~p(s(bool,X55))) )),
  inference(subsumption_resolution,[],[f17587,f1000])).
fof(f17587,plain,(
  ( ! [X54,X56,X55] : (~p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X56)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X54))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X54))) | p(s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),X56)))) | ~p(s(bool,X55)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,X54),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(superposition,[],[f1026,f587])).
fof(f17657,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42))) | p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),X44)))) | ~p(s(bool,X43))) )),
  inference(subsumption_resolution,[],[f17583,f864])).
fof(f17583,plain,(
  ( ! [X43,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X44)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X42))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X42))) | p(s(bool,'const.bool.IN_2'(s(bool,X43),s(fun(bool,bool),X44)))) | ~p(s(bool,X43)) | p(s(bool,'const.bool.IN_2'(s(sK12,X42),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f560])).
fof(f17656,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X41)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X39))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X39))) | p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),X41)))) | ~p(s(bool,X40))) )),
  inference(subsumption_resolution,[],[f17582,f864])).
fof(f17582,plain,(
  ( ! [X39,X41,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X41)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X39))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X39))) | p(s(bool,'const.bool.IN_2'(s(bool,X40),s(fun(bool,bool),X41)))) | ~p(s(bool,X40)) | p(s(bool,'const.bool.IN_2'(s(sK12,X39),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f558])).
fof(f17655,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X28),s(fun(bool,bool),X30)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),X30)))) | ~p(s(bool,X29)) | p(s(bool,X28))) )),
  inference(subsumption_resolution,[],[f17579,f864])).
fof(f17579,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X28),s(fun(bool,bool),X30)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),X30)))) | ~p(s(bool,X29)) | p(s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | p(s(bool,X28))) )),
  inference(superposition,[],[f1026,f515])).
fof(f17653,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X24),s(fun(bool,bool),X26)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23))) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X26)))) | ~p(s(bool,X25)) | p(s(bool,X24))) )),
  inference(subsumption_resolution,[],[f17578,f864])).
fof(f17578,plain,(
  ( ! [X26,X24,X23,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X24),s(fun(bool,bool),X26)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23))) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X26)))) | ~p(s(bool,X25)) | p(s(bool,X24)) | p(s(bool,'const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) )),
  inference(superposition,[],[f1026,f515])).
fof(f17616,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | ~p(s(bool,X1))) )),
  inference(duplicate_literal_removal,[],[f17573])).
fof(f17573,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0))) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | ~p(s(bool,X1)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X0))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X0)))) )),
  inference(superposition,[],[f1026,f998])).
fof(f17297,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f373])).
fof(f17296,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f372])).
fof(f17291,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f363])).
fof(f17290,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f362])).
fof(f17289,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f361])).
fof(f17288,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5250,f360])).
fof(f17022,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f373])).
fof(f17021,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f372])).
fof(f17020,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f5249,f371])).
fof(f17016,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f363])).
fof(f17015,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f362])).
fof(f17014,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f361])).
fof(f17013,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5249,f360])).
fof(f16747,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f373])).
fof(f16746,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f372])).
fof(f16745,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f5248,f371])).
fof(f16744,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl12),
  inference(resolution,[],[f5248,f369])).
fof(f16741,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f363])).
fof(f16740,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f362])).
fof(f16739,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f361])).
fof(f16738,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5248,f360])).
fof(f16734,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24))))))) ) | $spl92),
  inference(cnf_transformation,[],[f16734_D])).
fof(f16734_D,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24))))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f16735,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X25),s(fun(fun(sK12,bool),bool),X27)))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X25),s('type.list.list'(sK12),sK13)))) | p(s(bool,X26))) ) | $spl6),
  inference(subsumption_resolution,[],[f16715,f558])).
fof(f16715,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X25),s(fun(fun(sK12,bool),bool),X27)))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X25),s('type.list.list'(sK12),sK13)))) | p(s(bool,X26))) ) | $spl6),
  inference(superposition,[],[f16541,f515])).
fof(f16541,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X2),s(fun(fun(sK12,bool),bool),X3))))))) ) | $spl6),
  inference(forward_demodulation,[],[f681,f706])).
fof(f681,plain,(
  ( ! [X2,X3] : (s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X2),s(fun(fun(sK12,bool),bool),X3))))))) )),
  inference(superposition,[],[f679,f583])).
fof(f16733,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24)))))) | p(s(bool,X23)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13))))) ) | $spl6),
  inference(subsumption_resolution,[],[f16714,f558])).
fof(f16714,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24)))))) | p(s(bool,X23)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13))))) ) | $spl6),
  inference(superposition,[],[f16541,f515])).
fof(f16553,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK26(sK16,X23,sK12))))) | ~p(s(bool,happ(s(fun(sK12,bool),X23),s(sK12,sK25(sK16,X23,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f373])).
fof(f16552,plain,(
  ( ! [X21,X20] : (~p(s(bool,X20)) | p(s(bool,happ(s(fun(sK12,bool),X21),s(sK12,sK26(sK16,X21,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK25(sK16,X21,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f372])).
fof(f16551,plain,(
  ( ! [X19,X18] : (~p(s(bool,X18)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK24(sK16,X19,sK12))))) | ~p(s(bool,X19))) ) | $spl12),
  inference(resolution,[],[f5247,f371])).
fof(f16550,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK23(sK16,X17,sK12))))) | ~p(s(bool,X17))) ) | $spl12),
  inference(resolution,[],[f5247,f369])).
fof(f16549,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK22(X15,sK16,sK12))))) | ~p(s(bool,X15))) ) | $spl12),
  inference(resolution,[],[f5247,f366])).
fof(f16547,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(sK16,X11,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X11),s(sK12,sK20(sK16,X11,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f363])).
fof(f16546,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK12,bool),X9),s(sK12,sK20(X9,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK20(X9,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f362])).
fof(f16545,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK19(sK16,X7,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),X7),s(sK12,sK18(sK16,X7,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f361])).
fof(f16544,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK12,bool),X5),s(sK12,sK19(X5,sK16,sK12))))) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK18(X5,sK16,sK12)))))) ) | $spl12),
  inference(resolution,[],[f5247,f360])).
fof(f16469,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24))))))) ) | $spl72),
  inference(cnf_transformation,[],[f16469_D])).
fof(f16469_D,plain,(
  ( ! [X24,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24)))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24))))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f16470,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X25),s(fun(fun(sK12,bool),bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X25),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X27)))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X25),s('type.list.list'(sK12),sK13)))) | p(s(bool,X26))) ) | $spl6),
  inference(subsumption_resolution,[],[f16450,f558])).
fof(f16450,plain,(
  ( ! [X26,X27,X25] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X25),s(fun(fun(sK12,bool),bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X25),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X27)))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X25),s('type.list.list'(sK12),sK13)))) | p(s(bool,X26))) ) | $spl6),
  inference(superposition,[],[f16343,f515])).
fof(f16343,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X1))))))) ) | $spl6),
  inference(forward_demodulation,[],[f680,f706])).
fof(f680,plain,(
  ( ! [X0,X1] : (s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X1))))))) )),
  inference(superposition,[],[f679,f583])).
fof(f16468,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24)))))) | p(s(bool,X23)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13))))) ) | $spl6),
  inference(subsumption_resolution,[],[f16449,f558])).
fof(f16449,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X22),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK16),s(fun(fun(sK12,bool),bool),X24)))))) | p(s(bool,X23)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X22),s('type.list.list'(sK12),sK13))))) ) | $spl6),
  inference(superposition,[],[f16343,f515])).
fof(f16243,plain,(
  ( ! [X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X63)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X63)))))) | p(s(bool,X62))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f16224])).
fof(f16224,plain,(
  ( ! [X62,X63] : (s(bool,X62) != s(bool,X62) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X63)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X63)))))) | p(s(bool,X62))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f16118,f558])).
fof(f16118,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X14),s(fun(bool,bool),X15))))))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f16117,f713])).
fof(f16117,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X15)))) | s(bool,'F_0') != s(bool,X14) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X14),s(fun(bool,bool),X15))))))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f16116,f713])).
fof(f16116,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X14),s(fun(bool,bool),X15))))))) ) | $spl6),
  inference(forward_demodulation,[],[f701,f706])).
fof(f701,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),X15)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X14),s(fun(bool,bool),X15))))))) )),
  inference(superposition,[],[f679,f583])).
fof(f16244,plain,(
  ( ! [X61,X60] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X61)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X61)))))) | p(s(bool,X60))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f16223])).
fof(f16223,plain,(
  ( ! [X61,X60] : (s(bool,X60) != s(bool,X60) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X61)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X61)))))) | p(s(bool,X60))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f16118,f560])).
fof(f16245,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f16222])).
fof(f16222,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X58) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f16118,f587])).
fof(f16246,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f16221])).
fof(f16221,plain,(
  ( ! [X57,X56] : (s(bool,X56) != s(bool,X56) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f16118,f588])).
fof(f16256,plain,(
  ( ! [X18,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))) | p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X16),s(fun(bool,bool),X18))))))) ) | $spl70),
  inference(cnf_transformation,[],[f16256_D])).
fof(f16256_D,plain,(
  ( ! [X18,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))) | p(s(bool,X16)) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X16),s(fun(bool,bool),X18))))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f16257,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X21)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X19),s(fun(bool,bool),X21)))))) | p(s(bool,X19)) | p(s(bool,X20))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f16208,f558])).
fof(f16208,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X21)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X19),s(fun(bool,bool),X21)))))) | p(s(bool,X19)) | p(s(bool,X20))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f16118,f515])).
fof(f16255,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X16),s(fun(bool,bool),X18)))))) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f16207,f558])).
fof(f16207,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X16),s(fun(bool,bool),X18)))))) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f16118,f515])).
fof(f15607,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) | $spl68),
  inference(cnf_transformation,[],[f15607_D])).
fof(f15607_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f15621,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) | p(s(bool,X0))) ) | $spl12),
  inference(forward_demodulation,[],[f15610,f5506])).
fof(f15610,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f15544,f1031])).
fof(f15544,plain,(
  ( ! [X141] : (sP11('const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X141)) ) | $spl12),
  inference(resolution,[],[f5094,f500])).
fof(f15605,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) | p(s(bool,X0))) ) | $spl12),
  inference(forward_demodulation,[],[f15594,f5506])).
fof(f15594,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK14,sK16,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK14,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f15543,f1030])).
fof(f15543,plain,(
  ( ! [X140] : (sP11(X140,'const.bool.IN_2'(s(sK12,sK17(sK14,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl12),
  inference(resolution,[],[f5094,f499])).
fof(f15356,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30))))))) ) | $spl66),
  inference(cnf_transformation,[],[f15356_D])).
fof(f15356_D,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30))))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f15357,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(subsumption_resolution,[],[f15339,f561])).
fof(f15339,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,X32)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(superposition,[],[f3381,f518])).
fof(f3381,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X2),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X2),s(fun(fun(sK12,bool),bool),X3))))))) ) | $spl12),
  inference(superposition,[],[f2942,f583])).
fof(f15355,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(subsumption_resolution,[],[f15338,f561])).
fof(f15338,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X29))) ) | $spl12),
  inference(superposition,[],[f3381,f518])).
fof(f15323,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30))))))) ) | $spl64),
  inference(cnf_transformation,[],[f15323_D])).
fof(f15323_D,plain,(
  ( ! [X30,X28] : (p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30))))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f15324,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(subsumption_resolution,[],[f15306,f561])).
fof(f15306,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),X33)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X31),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X33)))))) | ~p(s(bool,X32)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X31),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(superposition,[],[f3380,f518])).
fof(f3380,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X0),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X0),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X1))))))) ) | $spl12),
  inference(superposition,[],[f2942,f583])).
fof(f15322,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(subsumption_resolution,[],[f15305,f561])).
fof(f15305,plain,(
  ( ! [X30,X28,X29] : (s(bool,'T_0') != s(bool,X29) | p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),X30)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK12,bool),X28),s(fun(fun(sK12,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK12,bool),sK15),s(fun(fun(sK12,bool),bool),X30)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),X28),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X29))) ) | $spl12),
  inference(superposition,[],[f3380,f518])).
fof(f15199,plain,(
  ( ! [X62,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X63)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X63)))) | p(s(bool,X62))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f15181])).
fof(f15181,plain,(
  ( ! [X62,X63] : (s(bool,X62) != s(bool,X62) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X63)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X63)))) | p(s(bool,X62))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f15137,f558])).
fof(f15137,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X12) | ~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X13)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13))))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f15136,f713])).
fof(f15136,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X12) | p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),X13))))))) ) | $spl6),
  inference(forward_demodulation,[],[f700,f706])).
fof(f700,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),X13)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X12),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))),s(fun(bool,bool),X13))))))) )),
  inference(superposition,[],[f679,f583])).
fof(f15200,plain,(
  ( ! [X61,X60] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X61)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X61)))) | p(s(bool,X60))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f15180])).
fof(f15180,plain,(
  ( ! [X61,X60] : (s(bool,X60) != s(bool,X60) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X61)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X61)))) | p(s(bool,X60))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f15137,f560])).
fof(f15201,plain,(
  ( ! [X59,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f15179])).
fof(f15179,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X58) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f15137,f587])).
fof(f15202,plain,(
  ( ! [X57,X56] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f15178])).
fof(f15178,plain,(
  ( ! [X57,X56] : (s(bool,X56) != s(bool,X56) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f15137,f588])).
fof(f15212,plain,(
  ( ! [X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))))) | p(s(bool,X16)) | p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),X18))))) ) | $spl60),
  inference(cnf_transformation,[],[f15212_D])).
fof(f15212_D,plain,(
  ( ! [X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))))) | p(s(bool,X16)) | p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),X18))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f15213,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | ~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | p(s(bool,X19))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f15165,f557])).
fof(f15165,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | ~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | p(s(bool,X19)) | p(s(bool,X20))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f15137,f515])).
fof(f15211,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | ~p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),X18)))) | p(s(bool,X16))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f15164,f557])).
fof(f15164,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | ~p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),X18)))) | p(s(bool,X17)) | p(s(bool,X16))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f15137,f515])).
fof(f14840,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29))))))) ) | $spl58),
  inference(cnf_transformation,[],[f14840_D])).
fof(f14840_D,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29))))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f14841,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X30))))) ) | $spl12),
  inference(subsumption_resolution,[],[f14822,f561])).
fof(f14822,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,X31)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X30))))) ) | $spl12),
  inference(superposition,[],[f3383,f518])).
fof(f3383,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X2))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X2),s(fun('type.list.list'(sK12),bool),X3))))))) ) | $spl12),
  inference(superposition,[],[f2942,f583])).
fof(f14839,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27))))) ) | $spl12),
  inference(subsumption_resolution,[],[f14821,f561])).
fof(f14821,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27)))) | ~p(s(bool,X28))) ) | $spl12),
  inference(superposition,[],[f3383,f518])).
fof(f14704,plain,(
  $false | ($spl24 | $spl32)),
  inference(trivial_inequality_removal,[],[f14684])).
fof(f14684,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl32)),
  inference(superposition,[],[f6167,f4630])).
fof(f14703,plain,(
  $false | ($spl26 | $spl32)),
  inference(trivial_inequality_removal,[],[f14685])).
fof(f14685,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl26 | $spl32)),
  inference(superposition,[],[f6167,f4759])).
fof(f14702,plain,(
  $false | ($spl12 | $spl32)),
  inference(trivial_inequality_removal,[],[f14686])).
fof(f14686,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl32)),
  inference(superposition,[],[f6167,f5337])).
fof(f14701,plain,(
  $false | ($spl12 | $spl32)),
  inference(trivial_inequality_removal,[],[f14687])).
fof(f14687,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl32)),
  inference(superposition,[],[f6167,f5506])).
fof(f14700,plain,(
  $false | ($spl12 | $spl32)),
  inference(trivial_inequality_removal,[],[f14688])).
fof(f14688,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl32)),
  inference(superposition,[],[f6167,f11182])).
fof(f14699,plain,(
  $false | ($spl12 | $spl32)),
  inference(trivial_inequality_removal,[],[f14689])).
fof(f14689,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl32)),
  inference(superposition,[],[f6167,f11768])).
fof(f11768,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK27(sK15,sK12)))) | $spl12),
  inference(resolution,[],[f11545,f562])).
fof(f14698,plain,(
  $false | ($spl12 | $spl32)),
  inference(trivial_inequality_removal,[],[f14690])).
fof(f14690,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl32)),
  inference(superposition,[],[f6167,f11942])).
fof(f11942,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK28(sK15,sK12)))) | $spl12),
  inference(resolution,[],[f11546,f562])).
fof(f14697,plain,(
  $false | ($spl12 | $spl32)),
  inference(trivial_inequality_removal,[],[f14694])).
fof(f14694,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl12 | $spl32)),
  inference(superposition,[],[f6167,f713])).
fof(f14696,plain,(
  $false | $spl32),
  inference(equality_resolution,[],[f6167])).
fof(f14637,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29))))))) ) | $spl56),
  inference(cnf_transformation,[],[f14637_D])).
fof(f14637_D,plain,(
  ( ! [X29,X27] : (p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29))))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f14638,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X30))))) ) | $spl12),
  inference(subsumption_resolution,[],[f14618,f561])).
fof(f14618,plain,(
  ( ! [X30,X31,X32] : (s(bool,'T_0') != s(bool,X31) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),X32)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X30),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X32)))))) | ~p(s(bool,X31)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X30))))) ) | $spl12),
  inference(superposition,[],[f3382,f518])).
fof(f3382,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X0))) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X0),s(fun('type.list.list'(sK12),bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X0),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X1))))))) ) | $spl12),
  inference(superposition,[],[f2942,f583])).
fof(f14636,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27))))) ) | $spl12),
  inference(subsumption_resolution,[],[f14617,f561])).
fof(f14617,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') != s(bool,X28) | p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s('type.list.list'(sK12),X27),s(fun('type.list.list'(sK12),bool),'const.pred_set.INSERT_2'(s('type.list.list'(sK12),sK13),s(fun('type.list.list'(sK12),bool),X29)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),X27)))) | ~p(s(bool,X28))) ) | $spl12),
  inference(superposition,[],[f3382,f518])).
fof(f14137,plain,(
  ( ! [X6,X8,X7] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X6))) | ~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f14121])).
fof(f14121,plain,(
  ( ! [X6,X8,X7] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X6))) | ~p(s(bool,X7)) | ~p(s(bool,X8)) | s(bool,X8) != s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X7))) )),
  inference(resolution,[],[f1013,f465])).
fof(f13922,plain,(
  ( ! [X28,X29,X27] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | ~p(s(bool,X28)) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X29))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X29)))) )),
  inference(subsumption_resolution,[],[f13914,f864])).
fof(f13914,plain,(
  ( ! [X28,X29,X27] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X27))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X27))) | ~p(s(bool,X28)) | p(s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X28) != s(bool,'const.bool.IN_2'(s(sK12,X27),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X29))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X29)))) )),
  inference(resolution,[],[f1011,f977])).
fof(f13921,plain,(
  ( ! [X4,X5,X3] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3))) | ~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,X3),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f13905])).
fof(f13905,plain,(
  ( ! [X4,X5,X3] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X3))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X3))) | ~p(s(bool,X4)) | p(s(bool,X5)) | p(s(bool,X5)) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,X3),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1011,f467])).
fof(f13860,plain,(
  ( ! [X30,X31,X29] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X29))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X29))) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X31))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X31)))) )),
  inference(subsumption_resolution,[],[f13850,f864])).
fof(f13850,plain,(
  ( ! [X30,X31,X29] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X29))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X29))) | ~p(s(bool,X30)) | p(s(bool,'const.bool.IN_2'(s(sK12,X31),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(sK12,X29),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X31))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X31)))) )),
  inference(resolution,[],[f1010,f977])).
fof(f13858,plain,(
  ( ! [X6,X8,X7] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X6))) | ~p(s(bool,X7)) | p(s(bool,X8)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f13841])).
fof(f13841,plain,(
  ( ! [X6,X8,X7] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X6))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X6))) | ~p(s(bool,X7)) | p(s(bool,X8)) | p(s(bool,X8)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(sK12,X6),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(resolution,[],[f1010,f467])).
fof(f12602,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X5))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X5))) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) )),
  inference(duplicate_literal_removal,[],[f12581])).
fof(f12581,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | ~p(s(bool,X4)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X5))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X5))) | s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK12,X5),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X3))) )),
  inference(resolution,[],[f964,f465])).
fof(f11203,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) | $spl52),
  inference(cnf_transformation,[],[f11203_D])).
fof(f11203_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f11217,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) | p(s(bool,X0))) ) | $spl12),
  inference(forward_demodulation,[],[f11206,f5337])).
fof(f11206,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f11140,f1031])).
fof(f11140,plain,(
  ( ! [X141] : (sP11('const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))),X141)) ) | $spl12),
  inference(resolution,[],[f2045,f500])).
fof(f11201,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) | p(s(bool,X0))) ) | $spl12),
  inference(forward_demodulation,[],[f11190,f5337])).
fof(f11190,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,sK17(sK13,sK16,sK12)))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK16,sK12))))) ) | $spl12),
  inference(resolution,[],[f11139,f1030])).
fof(f11139,plain,(
  ( ! [X140] : (sP11(X140,'const.bool.IN_2'(s(sK12,sK17(sK13,sK16,sK12)),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13)))))) ) | $spl12),
  inference(resolution,[],[f2045,f499])).
fof(f10724,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X22)) ) | $spl50),
  inference(cnf_transformation,[],[f10724_D])).
fof(f10724_D,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X22)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10733,plain,(
  ( ! [X35,X36] : (s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X36))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X36))) | s(bool,'F_0') != s(bool,X35)) ) | $spl6),
  inference(subsumption_resolution,[],[f10705,f864])).
fof(f10705,plain,(
  ( ! [X35,X36] : (s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X36))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X36))) | s(bool,'F_0') != s(bool,X35) | p(s(bool,'const.bool.IN_2'(s(sK12,X36),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl6),
  inference(resolution,[],[f977,f5540])).
fof(f5540,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl6),
  inference(forward_demodulation,[],[f685,f706])).
fof(f685,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)),X2) | p(s(bool,X3))) )),
  inference(superposition,[],[f679,f464])).
fof(f10722,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X20)) ) | $spl48),
  inference(cnf_transformation,[],[f10722_D])).
fof(f10722_D,plain,(
  ( ! [X20] : (s(bool,X20) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X20)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f10732,plain,(
  ( ! [X33,X34] : (s(bool,X33) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X34))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X34))) | s(bool,'T_0') != s(bool,X33)) ) | $spl12),
  inference(subsumption_resolution,[],[f10704,f864])).
fof(f10704,plain,(
  ( ! [X33,X34] : (s(bool,X33) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X34))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X34))) | s(bool,'T_0') != s(bool,X33) | p(s(bool,'const.bool.IN_2'(s(sK12,X34),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl12),
  inference(resolution,[],[f977,f3385])).
fof(f3385,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)),X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl12),
  inference(superposition,[],[f2942,f464])).
fof(f10723,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23))) | s(bool,'F_0') != s(bool,X22)) ) | $spl6),
  inference(subsumption_resolution,[],[f10697,f864])).
fof(f10697,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X23))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X23))) | s(bool,'F_0') != s(bool,X22) | p(s(bool,'const.bool.IN_2'(s(sK12,X23),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl6),
  inference(resolution,[],[f977,f5205])).
fof(f5205,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl6),
  inference(forward_demodulation,[],[f684,f706])).
fof(f684,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | p(s(bool,X1))) )),
  inference(superposition,[],[f679,f464])).
fof(f10720,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21))) | s(bool,'T_0') != s(bool,X20)) ) | $spl12),
  inference(subsumption_resolution,[],[f10696,f864])).
fof(f10696,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X21))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X21))) | s(bool,'T_0') != s(bool,X20) | p(s(bool,'const.bool.IN_2'(s(sK12,X21),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))))) ) | $spl12),
  inference(resolution,[],[f977,f3384])).
fof(f3384,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl12),
  inference(superposition,[],[f2942,f464])).
fof(f8474,plain,(
  ( ! [X18,X16] : (s(bool,X16) != s(bool,X18) | p(s(bool,X16)) | ~p(s(bool,X18))) ) | $spl42),
  inference(cnf_transformation,[],[f8474_D])).
fof(f8474_D,plain,(
  ( ! [X18,X16] : (s(bool,X16) != s(bool,X18) | p(s(bool,X16)) | ~p(s(bool,X18))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f8476,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X21)) | p(s(bool,X19))) ) | $spl22),
  inference(subsumption_resolution,[],[f8434,f557])).
fof(f8434,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,X19) != s(bool,X21) | ~p(s(bool,X21)) | p(s(bool,X19)) | p(s(bool,X20))) ) | $spl22),
  inference(superposition,[],[f4474,f515])).
fof(f8473,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X18)) | p(s(bool,X16))) ) | $spl22),
  inference(subsumption_resolution,[],[f8433,f557])).
fof(f8433,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,X16) != s(bool,X18) | ~p(s(bool,X18)) | p(s(bool,X17)) | p(s(bool,X16))) ) | $spl22),
  inference(superposition,[],[f4474,f515])).
fof(f6288,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))))) | ~p(s(bool,X58))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6276])).
fof(f6276,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X58) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))))) | ~p(s(bool,X58))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3375,f562])).
fof(f3375,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X18) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X19)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X18),s(fun(bool,bool),X19))))))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1779,f583])).
fof(f6289,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))))) | ~p(s(bool,X56))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6275])).
fof(f6275,plain,(
  ( ! [X57,X56] : (s(bool,X56) != s(bool,X56) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))))) | ~p(s(bool,X56))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3375,f564])).
fof(f6290,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6274])).
fof(f6274,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X54) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3375,f587])).
fof(f6291,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6273])).
fof(f6273,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X52) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3375,f588])).
fof(f6300,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X25),s(fun(bool,bool),X27)))))) | ~p(s(bool,X25))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6262,f561])).
fof(f6262,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X25),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3375,f518])).
fof(f6298,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X22),s(fun(bool,bool),X24)))))) | ~p(s(bool,X22))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6261,f561])).
fof(f6261,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X22),s(fun(bool,bool),X24)))))) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3375,f518])).
fof(f6155,plain,(
  ( ! [X59,X58] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))))) | ~p(s(bool,X58))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6143])).
fof(f6143,plain,(
  ( ! [X59,X58] : (s(bool,X58) != s(bool,X58) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X59)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X59)))))) | ~p(s(bool,X58))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3374,f562])).
fof(f3374,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X16) | p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X17))))))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1779,f583])).
fof(f6156,plain,(
  ( ! [X57,X56] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))))) | ~p(s(bool,X56))) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6142])).
fof(f6142,plain,(
  ( ! [X57,X56] : (s(bool,X56) != s(bool,X56) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X57)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X57)))))) | ~p(s(bool,X56))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3374,f564])).
fof(f6157,plain,(
  ( ! [X54,X55] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6141])).
fof(f6141,plain,(
  ( ! [X54,X55] : (s(bool,X54) != s(bool,X54) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X55)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X55)))))) | s(bool,'F_0') = s(bool,X54)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3374,f587])).
fof(f6158,plain,(
  ( ! [X52,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl6 | $spl12)),
  inference(trivial_inequality_removal,[],[f6140])).
fof(f6140,plain,(
  ( ! [X52,X53] : (s(bool,X52) != s(bool,X52) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X53)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X53)))))) | s(bool,'F_0') = s(bool,X52)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3374,f588])).
fof(f6168,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X25))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6129,f561])).
fof(f6129,plain,(
  ( ! [X26,X27,X25] : (s(bool,'T_0') != s(bool,X26) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),X27)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X27)))))) | ~p(s(bool,X26)) | ~p(s(bool,X25))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3374,f518])).
fof(f6165,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))))) | ~p(s(bool,X22))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f6128,f561])).
fof(f6128,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') != s(bool,X23) | p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X24)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X24)))))) | ~p(s(bool,X22)) | ~p(s(bool,X23))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f3374,f518])).
fof(f6046,plain,(
  ( ! [X14,X15] : (s(bool,happ(s(fun(sK12,bool),sK15),s(sK12,X14))) = s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,X14))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK12,X14),s(fun(sK12,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK12),sK13))))) | ~p(s(bool,X15))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f968,f3739])).
fof(f4439,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) | $spl26),
  inference(cnf_transformation,[],[f4439_D])).
fof(f4439_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4471,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) | p(s(bool,X1))) ) | $spl0),
  inference(subsumption_resolution,[],[f4465,f543])).
fof(f4465,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) | p(s(bool,X1))) ) | $spl0),
  inference(resolution,[],[f2857,f498])).
fof(f2857,plain,(
  ( ! [X151] : (sP11(happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))),X151)) ) | $spl0),
  inference(resolution,[],[f2318,f500])).
fof(f2318,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) | $spl0),
  inference(resolution,[],[f1140,f375])).
fof(f1140,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK17(sK13,sK16,sK12))))) | $spl0),
  inference(resolution,[],[f703,f355])).
fof(f4470,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) | p(s(bool,X0))) ) | $spl0),
  inference(subsumption_resolution,[],[f4464,f543])).
fof(f4464,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) | p(s(bool,X0))) ) | $spl0),
  inference(resolution,[],[f2857,f502])).
fof(f2929,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl22),
  inference(cnf_transformation,[],[f2929_D])).
fof(f2929_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4456,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5)) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f4454,f706])).
fof(f4454,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f4443])).
fof(f4443,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X4))) ) | $spl12),
  inference(resolution,[],[f3387,f465])).
fof(f4440,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4432,f543])).
fof(f4432,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))))) ) | $spl0),
  inference(resolution,[],[f2856,f498])).
fof(f2856,plain,(
  ( ! [X150] : (sP11(X150,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12))))) ) | $spl0),
  inference(resolution,[],[f2318,f499])).
fof(f4437,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4431,f543])).
fof(f4431,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK28(sK16,sK12)))))) ) | $spl0),
  inference(resolution,[],[f2856,f502])).
fof(f4421,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) | $spl24),
  inference(cnf_transformation,[],[f4421_D])).
fof(f4421_D,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4430,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) | p(s(bool,X1))) ) | $spl0),
  inference(subsumption_resolution,[],[f4424,f543])).
fof(f4424,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) | p(s(bool,X1))) ) | $spl0),
  inference(resolution,[],[f2521,f498])).
fof(f2521,plain,(
  ( ! [X151] : (sP11(happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))),X151)) ) | $spl0),
  inference(resolution,[],[f2317,f500])).
fof(f2317,plain,(
  p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) | $spl0),
  inference(resolution,[],[f1140,f374])).
fof(f4429,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) | p(s(bool,X0))) ) | $spl0),
  inference(subsumption_resolution,[],[f4423,f543])).
fof(f4423,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) | p(s(bool,X0))) ) | $spl0),
  inference(resolution,[],[f2521,f502])).
fof(f4422,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4414,f543])).
fof(f4414,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))))) ) | $spl0),
  inference(resolution,[],[f2520,f498])).
fof(f2520,plain,(
  ( ! [X150] : (sP11(X150,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12))))) ) | $spl0),
  inference(resolution,[],[f2317,f499])).
fof(f4419,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))))) ) | $spl0),
  inference(subsumption_resolution,[],[f4413,f543])).
fof(f4413,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK12,bool),sK16),s(sK12,sK27(sK16,sK12)))))) ) | $spl0),
  inference(resolution,[],[f2520,f502])).
fof(f4406,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5)) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f4404,f706])).
fof(f4404,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f4390])).
fof(f4390,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X4))) ) | $spl12),
  inference(resolution,[],[f3386,f465])).
fof(f3386,plain,(
  ( ! [X4,X5] : (~sP9(X5,X4,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl12),
  inference(superposition,[],[f2942,f466])).
fof(f4386,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f4385,f706])).
fof(f4385,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f4375])).
fof(f4375,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(resolution,[],[f3385,f467])).
fof(f4372,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f4371,f706])).
fof(f4371,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f4360])).
fof(f4360,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) ) | $spl12),
  inference(resolution,[],[f3384,f467])).
fof(f3731,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X9)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f3725])).
fof(f3725,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | s(bool,'F_0') != s(bool,X9) | ~p(s(bool,X8))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3361,f465])).
fof(f3686,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f3679])).
fof(f3679,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3360,f465])).
fof(f3360,plain,(
  ( ! [X8,X9] : (~sP9(X9,X8,'F_0') | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1779,f466])).
fof(f3640,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f3635])).
fof(f3635,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3359,f467])).
fof(f3580,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f3575])).
fof(f3575,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f3358,f467])).
fof(f3358,plain,(
  ( ! [X4,X5] : (sP9(X5,X4,'F_0') | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f1779,f464])).
fof(f3408,plain,(
  $false | ($spl6 | $spl12 | $spl18)),
  inference(subsumption_resolution,[],[f1080,f1779])).
fof(f1080,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl18),
  inference(cnf_transformation,[],[f1080_D])).
fof(f1080_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f3407,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(subsumption_resolution,[],[f3391,f562])).
fof(f3391,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl12),
  inference(superposition,[],[f2942,f518])).
fof(f3406,plain,(
  ( ! [X10] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X10))) ) | $spl12),
  inference(subsumption_resolution,[],[f3390,f562])).
fof(f3390,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X10))) ) | $spl12),
  inference(superposition,[],[f2942,f518])).
fof(f3378,plain,(
  $false | ($spl6 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f2943,f1779])).
fof(f2943,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl10 | $spl12)),
  inference(forward_demodulation,[],[f711,f713])).
fof(f3349,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3236,f1135])).
fof(f3236,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(resolution,[],[f1135,f579])).
fof(f3348,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3235,f1135])).
fof(f3235,plain,(
  p(s(bool,'F_0')) | $spl20),
  inference(resolution,[],[f1135,f577])).
fof(f3347,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3232,f1135])).
fof(f3232,plain,(
  ( ! [X214] : (p(s(bool,X214))) ) | $spl20),
  inference(resolution,[],[f1135,f536])).
fof(f3346,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3231,f1135])).
fof(f3231,plain,(
  ( ! [X213] : (p(s(bool,X213))) ) | $spl20),
  inference(resolution,[],[f1135,f533])).
fof(f3345,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f3230,f1135])).
fof(f3230,plain,(
  ( ! [X212] : (p(s(bool,X212))) ) | $spl20),
  inference(resolution,[],[f1135,f530])).
fof(f3237,plain,(
  $false | $spl20),
  inference(resolution,[],[f1135,f581])).
fof(f3238,plain,(
  $false | $spl20),
  inference(resolution,[],[f1135,f582])).
fof(f3239,plain,(
  $false | $spl20),
  inference(resolution,[],[f1135,f595])).
fof(f3144,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f3058,f1078])).
fof(f3058,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f1078,f580])).
fof(f580,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f343])).
fof(f3143,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f3057,f1078])).
fof(f3057,plain,(
  ~p(s(bool,'T_0')) | $spl16),
  inference(resolution,[],[f1078,f578])).
fof(f578,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f342])).
fof(f3142,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f3054,f1078])).
fof(f3054,plain,(
  ( ! [X239] : (~p(s(bool,X239))) ) | $spl16),
  inference(resolution,[],[f1078,f572])).
fof(f3141,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f3050,f1078])).
fof(f3050,plain,(
  ( ! [X235] : (~p(s(bool,X235))) ) | $spl16),
  inference(resolution,[],[f1078,f528])).
fof(f3056,plain,(
  $false | $spl16),
  inference(resolution,[],[f1078,f576])).
fof(f576,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.EXISTS_CONG',aHL_FALSITY)).
fof(f3061,plain,(
  $false | $spl16),
  inference(resolution,[],[f1078,f586])).
fof(f586,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.list.EXISTS_CONG','thm.pred_set.NOT_IN_EMPTY')).
fof(f2946,plain,(
  $false | ($spl1 | $spl12)),
  inference(subsumption_resolution,[],[f2945,f582])).
fof(f2945,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl12)),
  inference(forward_demodulation,[],[f702,f713])).
fof(f2944,plain,(
  $false | ($spl10 | $spl12 | $spl19)),
  inference(subsumption_resolution,[],[f2943,f1081])).
fof(f2928,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f2927,f713])).
fof(f2927,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f2918])).
fof(f2918,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) ) | $spl6),
  inference(resolution,[],[f1710,f467])).
fof(f1710,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f1274,f464])).
fof(f2753,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2666,f2566])).
fof(f2566,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2565,f560])).
fof(f2565,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f704,f706])).
fof(f2666,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl6)),
  inference(resolution,[],[f2566,f580])).
fof(f2752,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2665,f2566])).
fof(f2665,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl6)),
  inference(resolution,[],[f2566,f578])).
fof(f2751,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2662,f2566])).
fof(f2662,plain,(
  ( ! [X239] : (~p(s(bool,X239))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f2566,f572])).
fof(f2750,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f2658,f2566])).
fof(f2658,plain,(
  ( ! [X235] : (~p(s(bool,X235))) ) | ($spl2 | $spl6)),
  inference(resolution,[],[f2566,f528])).
fof(f2664,plain,(
  $false | ($spl2 | $spl6)),
  inference(resolution,[],[f2566,f576])).
fof(f2669,plain,(
  $false | ($spl2 | $spl6)),
  inference(resolution,[],[f2566,f586])).
fof(f1138,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1130,f707])).
fof(f1130,plain,(
  s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl9),
  inference(trivial_inequality_removal,[],[f1126])).
fof(f1126,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl9),
  inference(superposition,[],[f709,f588])).
fof(f1137,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1131,f707])).
fof(f1131,plain,(
  s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl9),
  inference(trivial_inequality_removal,[],[f1124])).
fof(f1124,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | $spl9),
  inference(superposition,[],[f709,f587])).
fof(f1136,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl9),
  inference(subsumption_resolution,[],[f1119,f562])).
fof(f1119,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl9),
  inference(superposition,[],[f709,f518])).
fof(f1134,plain,(
  ( ! [X10] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X10))) ) | $spl9),
  inference(subsumption_resolution,[],[f1118,f562])).
fof(f1118,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X10))) ) | $spl9),
  inference(superposition,[],[f709,f518])).
fof(f1064,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | $spl7),
  inference(superposition,[],[f707,f564])).
fof(f1063,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | $spl7),
  inference(superposition,[],[f707,f562])).
fof(f1079,plain,(
  ( ! [X9] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | p(s(bool,X9))) ) | $spl7),
  inference(subsumption_resolution,[],[f1058,f558])).
fof(f1058,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13)))) | p(s(bool,X9))) ) | $spl7),
  inference(superposition,[],[f707,f515])).
fof(f1075,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl7),
  inference(subsumption_resolution,[],[f1057,f558])).
fof(f1057,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))))) ) | $spl7),
  inference(superposition,[],[f707,f515])).
fof(f699,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f588])).
fof(f698,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f588])).
fof(f697,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'T_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f587])).
fof(f696,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))),
  inference(superposition,[],[f679,f587])).
fof(f695,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f564])).
fof(f694,plain,(
  s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f562])).
fof(f693,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f560])).
fof(f692,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))),
  inference(superposition,[],[f679,f558])).
fof(f691,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,X11)) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) )),
  inference(superposition,[],[f679,f518])).
fof(f690,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f679,f518])).
fof(f689,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13)))) | p(s(bool,X9))) )),
  inference(superposition,[],[f679,f515])).
fof(f688,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK15),s('type.list.list'(sK12),sK13))) | p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK12,bool),sK16),s('type.list.list'(sK12),sK13))))) )),
  inference(superposition,[],[f679,f515])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 2550
% Passive clauses: 31549
% Generated clauses: 121130
% Final active clauses: 1242
% Final passive clauses: 14295
% Input formulas: 69
% Initial clauses: 245
% 
% Selected by SInE selection: 67
% SInE iterations: 1
% 
% Duplicate literals: 4683
% Trivial inequalities: 777
% Fw subsumption resolutions: 14389
% Fw demodulations: 19707
% Bw demodulations: 43
% 
% Simple tautologies: 8112
% Equational tautologies: 50
% Forward subsumptions: 42323
% Bw demodulations to eq. taut.: 142
% 
% Binary resolution: 51119
% Forward superposition: 16275
% Backward superposition: 12965
% Equality resolution: 23
% 
% Split clauses: 1516
% Split components: 3555
% SAT solver clauses: 1668
% SAT solver unit clauses: 19
% SAT solver binary clauses: 128
% 
% Sat splits: 1516
% Sat splitting refutations: 63
% 
% Memory used [KB]: 30958
% Time elapsed: 2.334 s
% ------------------------------
% ------------------------------
% Success in time 14.051 s
