Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quotient
% (106009)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 155
% Passive clauses: 2232
% Generated clauses: 22212
% Final active clauses: 142
% Final passive clauses: 1889
% Input formulas: 22
% Initial clauses: 44
% 
% Duplicate literals: 1489
% Trivial inequalities: 109
% Fw subsumption resolutions: 2953
% Fw demodulations: 3610
% Bw demodulations: 12
% 
% Simple tautologies: 352
% Equational tautologies: 20
% Forward subsumptions: 9968
% 
% Binary resolution: 429
% Forward superposition: 8122
% Backward superposition: 5099
% Self superposition: 116
% Equality factoring: 34
% Equality resolution: 25
% 
% Split clauses: 843
% Split components: 1692
% SAT solver clauses: 1559
% SAT solver unit clauses: 5
% SAT solver binary clauses: 49
% 
% Sat splits: 843
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1558
% TWLsolver calls for satisfiability: 200
% 
% Memory used [KB]: 12153
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% (106027)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 98
% Passive clauses: 1035
% Generated clauses: 8919
% Final active clauses: 79
% Final passive clauses: 781
% Input formulas: 22
% Initial clauses: 44
% 
% Duplicate literals: 500
% Trivial inequalities: 149
% Fw subsumption resolutions: 448
% Fw demodulations: 1196
% Bw demodulations: 163
% Condensations: 724
% Global subsumptions: 4
% 
% Simple tautologies: 179
% Equational tautologies: 18
% Forward subsumptions: 3305
% Bw demodulations to eq. taut.: 11
% 
% Binary resolution: 92
% Forward superposition: 2156
% Backward superposition: 2273
% Self superposition: 83
% Equality factoring: 54
% Equality resolution: 4
% 
% Split clauses: 36
% Split components: 77
% SAT solver clauses: 1395
% SAT solver unit clauses: 54
% SAT solver binary clauses: 30
% SAT solver learnt clauses: 8
% 
% Sat splits: 36
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1379
% TWLsolver calls for satisfiability: 322
% 
% Memory used [KB]: 11513
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quotient
% (106050)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 429
% Passive clauses: 9812
% Generated clauses: 102771
% Final active clauses: 23
% Final passive clauses: 369
% Input formulas: 22
% Initial clauses: 44
% Discarded non-redundant clauses: 4372
% 
% Duplicate literals: 6223
% Trivial inequalities: 1043
% Fw subsumption resolutions: 6277
% Fw demodulations: 12899
% Bw demodulations: 4880
% Condensations: 9365
% Global subsumptions: 186
% 
% Simple tautologies: 1193
% Equational tautologies: 660
% Forward subsumptions: 53588
% Bw demodulations to eq. taut.: 496
% 
% Unit resulting resolution: 2311
% Forward superposition: 21444
% Backward superposition: 20703
% Self superposition: 417
% Equality factoring: 236
% Equality resolution: 33
% 
% Split clauses: 268
% Split components: 541
% SAT solver clauses: 10603
% SAT solver unit clauses: 109
% SAT solver binary clauses: 728
% 
% Sat splits: 268
% Sat splitting refutations: 66
% 
% Memory used [KB]: 62045
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 183 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f10361,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1922,f1936,f1934,f1944,f10319,f8497,f1973,f1933,f10345,f10347,f10356,f10359])).
fof(f10359,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f10358,f261])).
fof(f261,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f260,f131])).
fof(f131,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK1,sK0),sK4),s(fun(sK0,sK1),sK5)))),
  inference(resolution,[],[f81,f113])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.bool.EQ_CLAUSES|split|1')).
fof(f81,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK1,sK0),sK4),s(fun(sK0,sK1),sK5))))),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK1,sK0),sK4),s(fun(sK0,sK1),sK5)))) & (p(s(bool,'const.quotient.===>_4'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK2,fun(sK2,bool)),sK6),s(fun(sK1,sK2),sK7),s(fun(sK1,sK2),sK8)))) & ~p(s(bool,'const.quotient.===>_4'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK2,fun(sK2,bool)),sK6),s(fun(sK1,sK2),sK7),s(fun(sK1,sK2),'const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,X0),X4),s(fun(X0,X1),X5)))) & ? [X6,X7,X8] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X2,fun(X2,bool)),X6),s(fun(X1,X2),X7),s(fun(X1,X2),X8)))) & ~p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X2,fun(X2,bool)),X6),s(fun(X1,X2),X7),s(fun(X1,X2),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),'const.quotient.respects_1'(s(fun(X1,fun(X1,bool)),X3))),s(fun(X1,X2),X8))))))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,X0),X4),s(fun(X0,X1),X5)))) => ! [X6,X7,X8] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X2,fun(X2,bool)),X6),s(fun(X1,X2),X7),s(fun(X1,X2),X8)))) => p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X2,fun(X2,bool)),X6),s(fun(X1,X2),X7),s(fun(X1,X2),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),'const.quotient.respects_1'(s(fun(X1,fun(X1,bool)),X3))),s(fun(X1,X2),X8))))))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1,X2] : ! [X3,X4,X5] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X3),s(fun(X1,X0),X4),s(fun(X0,X1),X5)))) => ! [X6,X7,X8] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X2,fun(X2,bool)),X6),s(fun(X1,X2),X7),s(fun(X1,X2),X8)))) => p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X3),s(fun(X2,fun(X2,bool)),X6),s(fun(X1,X2),X7),s(fun(X1,X2),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),'const.quotient.respects_1'(s(fun(X1,fun(X1,bool)),X3))),s(fun(X1,X2),X8))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : ! [X9,X10,X11] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,X6),X10),s(fun(X6,X7),X11)))) => ! [X12,X13,X14] : (p(s(bool,'const.quotient.===>_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X8,fun(X8,bool)),X12),s(fun(X7,X8),X13),s(fun(X7,X8),X14)))) => p(s(bool,'const.quotient.===>_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X8,fun(X8,bool)),X12),s(fun(X7,X8),X13),s(fun(X7,X8),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),'const.quotient.respects_1'(s(fun(X7,fun(X7,bool)),X9))),s(fun(X7,X8),X14))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : ! [X9,X10,X11] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,X6),X10),s(fun(X6,X7),X11)))) => ! [X12,X13,X14] : (p(s(bool,'const.quotient.===>_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X8,fun(X8,bool)),X12),s(fun(X7,X8),X13),s(fun(X7,X8),X14)))) => p(s(bool,'const.quotient.===>_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X8,fun(X8,bool)),X12),s(fun(X7,X8),X13),s(fun(X7,X8),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),'const.quotient.respects_1'(s(fun(X7,fun(X7,bool)),X9))),s(fun(X7,X8),X14))))))))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.quotient.ABSTRACT_RES_ABSTRACT')).
fof(f260,plain,(
  s(bool,'F_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK1,sK0),sK4),s(fun(sK0,sK1),sK5)))),
  inference(subsumption_resolution,[],[f191,f115])).
fof(f115,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT',aHL_TRUTH)).
fof(f191,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK1,sK0),sK4),s(fun(sK0,sK1),sK5)))),
  inference(superposition,[],[f106,f131])).
fof(f106,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.bool.EQ_CLAUSES|split|3')).
fof(f10358,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl1),
  inference(forward_demodulation,[],[f10357,f1414])).
fof(f1414,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2))))),
  inference(resolution,[],[f1388,f113])).
fof(f1388,plain,(
  p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))),
  inference(resolution,[],[f375,f146])).
fof(f146,plain,(
  ( ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X2),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3)))))) | p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))))) )),
  inference(forward_demodulation,[],[f145,f123])).
fof(f123,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X7] : ! [X16,X17] : s(bool,'const.bool.IN_2'(s(X7,X17),s(fun(X7,bool),X16))) = s(bool,happ(s(fun(X7,bool),X16),s(X7,X17)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.pred_set.SPECIFICATION')).
fof(f145,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,X2))))) )),
  inference(forward_demodulation,[],[f126,f122])).
fof(f122,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.quotient.respects_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2] : s(bool,'const.quotient.respects_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.quotient.respects_2'(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X7] : ! [X18,X17] : s(bool,'const.quotient.respects_2'(s(fun(X7,fun(X7,bool)),X18),s(X7,X17))) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X18),s(X7,X17))),s(X7,X17)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.quotient.RESPECTS')).
fof(f126,plain,(
  ( ! [X2,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X2))),s(sK1,X2)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X3))),s(sK1,X2))))) )),
  inference(resolution,[],[f81,f117])).
fof(f117,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))) & ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(ennf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) <=> (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) & s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X8,X7] : ! [X18,X19,X20] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X18),s(fun(X7,X8),X19),s(fun(X8,X7),X20)))) => ! [X21,X22] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X18),s(X7,X21))),s(X7,X22)))) <=> (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X18),s(X7,X21))),s(X7,X21)))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X18),s(X7,X22))),s(X7,X22)))) & s(X8,happ(s(fun(X7,X8),X19),s(X7,X21))) = s(X8,happ(s(fun(X7,X8),X19),s(X7,X22))))))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.quotient.QUOTIENT_REL')).
fof(f375,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))))),
  inference(forward_demodulation,[],[f361,f123])).
fof(f361,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))),
  inference(resolution,[],[f83,f89])).
fof(f89,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,sK11(X5,X4,X3,X2,X1,X0)))),s(X1,sK12(X5,X4,X3,X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) | ! [X8,X9] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X9)))))))) & ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,sK11(X5,X4,X3,X2,X1,X0)))),s(X1,sK12(X5,X4,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,sK11(X5,X4,X3,X2,X1,X0)))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK12(X5,X4,X3,X2,X1,X0)))))))) | p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12])],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) | ! [X8,X9] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X9)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X9)))))))) & (? [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))) | p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5))))))),
  inference(rectify,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) | ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7)))))))) & (? [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) & ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))) | p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5))))))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) <=> ! [X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) <=> ! [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))))),
  inference(flattening,[],[f32])).
fof(f32,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) <=> ! [X6,X7] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) => p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X6))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X7))))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8,X7] : ! [X9,X12,X13,X14] : (p(s(bool,'const.quotient.===>_4'(s(fun(X7,fun(X7,bool)),X9),s(fun(X8,fun(X8,bool)),X12),s(fun(X7,X8),X13),s(fun(X7,X8),X14)))) <=> ! [X17,X23] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X9),s(X7,X17))),s(X7,X23)))) => p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X12),s(X8,happ(s(fun(X7,X8),X13),s(X7,X17))))),s(X8,happ(s(fun(X7,X8),X14),s(X7,X23))))))))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.quotient.FUN_REL')).
fof(f83,plain,(
  ~p(s(bool,'const.quotient.===>_4'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK2,fun(sK2,bool)),sK6),s(fun(sK1,sK2),sK7),s(fun(sK1,sK2),'const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8))))))),
  inference(cnf_transformation,[],[f60])).
fof(f10357,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f10342,f115])).
fof(f10342,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(duplicate_literal_removal,[],[f10335])).
fof(f10335,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(superposition,[],[f1934,f6241])).
fof(f6241,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X11),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),X10))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X11))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),X10),s(sK1,X11)))) )),
  inference(superposition,[],[f6140,f87])).
fof(f87,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.quotient.respects_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f31])).
fof(f31,plain,(
  ! [X0,X1,X2,X3] : s(X0,'const.quotient.respects_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(flattening,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X0,'const.quotient.respects_2'(s(fun(X1,fun(X1,X0)),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),'const.quotient.respects_1'(s(fun(X1,fun(X1,X0)),X2))),s(X1,X3)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8,X7] : ! [X28,X29] : s(X8,'const.quotient.respects_2'(s(fun(X7,fun(X7,X8)),X28),s(X7,X29))) = s(X8,happ(s(fun(X7,X8),'const.quotient.respects_1'(s(fun(X7,fun(X7,X8)),X28))),s(X7,X29)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','ARITY_1_2_const.quotient.respects')).
fof(f6140,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X88),s(sK1,X87))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X87))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X87),s(fun(sK1,bool),X88)))) )),
  inference(duplicate_literal_removal,[],[f6106])).
fof(f6106,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),X88),s(sK1,X87))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X87))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X87))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X87),s(fun(sK1,bool),X88)))) )),
  inference(superposition,[],[f1691,f3982])).
fof(f3982,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') = s(bool,'const.bool.RES_ABSTRACT_3'(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(fun(sK1,bool),X7),s(sK1,X6))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X6),s(fun(sK1,bool),X7)))) )),
  inference(forward_demodulation,[],[f3776,f123])).
fof(f3776,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),X7),s(sK1,X6))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))) | s(bool,'F_0') = s(bool,'const.bool.RES_ABSTRACT_3'(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(fun(sK1,bool),X7),s(sK1,X6)))) )),
  inference(superposition,[],[f1691,f120])).
fof(f120,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT',aHL_BOOL_CASES)).
fof(f1691,plain,(
  ( ! [X6,X8,X7] : (s(X7,happ(s(fun(sK1,X7),X8),s(sK1,X6))) = s(X7,'const.bool.RES_ABSTRACT_3'(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(fun(sK1,X7),X8),s(sK1,X6))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6)))) )),
  inference(resolution,[],[f1643,f121])).
fof(f121,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2))))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),X2)))) => s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) = s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8,X7] : ! [X26,X27,X17] : (p(s(bool,'const.bool.IN_2'(s(X7,X17),s(fun(X7,bool),X26)))) => s(X8,happ(s(fun(X7,X8),X27),s(X7,X17))) = s(X8,'const.bool.RES_ABSTRACT_3'(s(fun(X7,bool),X26),s(fun(X7,X8),X27),s(X7,X17))))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.bool.RES_ABSTRACT_DEF|split|0')).
fof(f1643,plain,(
  ( ! [X10] : (p(s(bool,'const.bool.IN_2'(s(sK1,X10),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X10)))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X10)))) )),
  inference(resolution,[],[f1619,f109])).
fof(f109,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X15] : (s(bool,'F_0') = s(bool,X15) <=> ~p(s(bool,X15)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.bool.EQ_CLAUSES|split|2')).
fof(f1619,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0)))) | p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))))))) )),
  inference(duplicate_literal_removal,[],[f1618])).
fof(f1618,plain,(
  ( ! [X0] : (~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0)))) | p(s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0)))))) | ~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))))) )),
  inference(equality_resolution,[],[f150])).
fof(f150,plain,(
  ( ! [X6,X7] : (s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X6))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X7))) | ~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X7)))) | p(s(bool,'const.bool.IN_2'(s(sK1,X7),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6)))))) | ~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))))) )),
  inference(forward_demodulation,[],[f149,f122])).
fof(f149,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X7)))) | p(s(bool,'const.bool.IN_2'(s(sK1,X7),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6)))))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X6))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X6))))) )),
  inference(forward_demodulation,[],[f148,f122])).
fof(f148,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.bool.IN_2'(s(sK1,X7),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6)))))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X6))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X7))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X6))))) )),
  inference(forward_demodulation,[],[f128,f123])).
fof(f128,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X7)))) | s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X6))) != s(sK0,happ(s(fun(sK1,sK0),sK4),s(sK1,X7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X7))),s(sK1,X7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X6))),s(sK1,X6))))) )),
  inference(resolution,[],[f81,f119])).
fof(f119,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | s(X0,happ(s(fun(X1,X0),X3),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X6))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f80])).
fof(f10356,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f10355,f261])).
fof(f10355,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl1),
  inference(forward_demodulation,[],[f10354,f1414])).
fof(f10354,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(subsumption_resolution,[],[f10334,f115])).
fof(f10334,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(superposition,[],[f1934,f9260])).
fof(f9260,plain,(
  ( ! [X62] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X62),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X62)))) )),
  inference(trivial_inequality_removal,[],[f9182])).
fof(f9182,plain,(
  ( ! [X62] : (s(bool,'F_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X62))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X62),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) )),
  inference(superposition,[],[f1694,f9125])).
fof(f9125,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) )),
  inference(subsumption_resolution,[],[f9092,f717])).
fof(f717,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X4))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X4),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5)))))) )),
  inference(resolution,[],[f572,f108])).
fof(f108,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f572,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X4)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X4),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X5)))))) )),
  inference(resolution,[],[f146,f109])).
fof(f9092,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK1,X0),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X1)))))) )),
  inference(resolution,[],[f8895,f572])).
fof(f8895,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),X8),s(sK1,X9)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X9),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),X8))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X9)))) )),
  inference(trivial_inequality_removal,[],[f8892])).
fof(f8892,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X9))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X9),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),X8))))) | ~p(s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),X8),s(sK1,X9))))) )),
  inference(superposition,[],[f6723,f111])).
fof(f111,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X15] : (s(bool,'T_0') = s(bool,X15) <=> p(s(bool,X15)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.bool.EQ_CLAUSES|split|0')).
fof(f6723,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),X10),s(sK1,X11))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X11))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X11),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),X10)))))) )),
  inference(superposition,[],[f6483,f87])).
fof(f6483,plain,(
  ( ! [X144,X145] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),X144),s(sK1,X145))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X145))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X145),s(fun(sK1,bool),X144)))) )),
  inference(subsumption_resolution,[],[f6289,f114])).
fof(f114,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f45])).
fof(f45,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT',aHL_FALSITY)).
fof(f6289,plain,(
  ( ! [X144,X145] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),X144),s(sK1,X145))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X145))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,X145),s(fun(sK1,bool),X144)))) )),
  inference(superposition,[],[f110,f6140])).
fof(f110,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f1694,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK1,X11),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X11))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X11)))) )),
  inference(resolution,[],[f1643,f108])).
fof(f10347,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f10346,f261])).
fof(f10346,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl1),
  inference(forward_demodulation,[],[f10343,f1414])).
fof(f10343,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(duplicate_literal_removal,[],[f10324])).
fof(f10324,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(resolution,[],[f1934,f8752])).
fof(f8752,plain,(
  ( ! [X70,X71] : (p(s(bool,'const.bool.IN_2'(s(sK1,X70),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),X71)))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X70))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),X71),s(sK1,X70)))) )),
  inference(trivial_inequality_removal,[],[f8710])).
fof(f8710,plain,(
  ( ! [X70,X71] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.IN_2'(s(sK1,X70),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),X71)))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X70))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),X71),s(sK1,X70)))) )),
  inference(superposition,[],[f112,f6241])).
fof(f112,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f78])).
fof(f10345,plain,(
  $false | $spl1),
  inference(subsumption_resolution,[],[f10344,f261])).
fof(f10344,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl1),
  inference(forward_demodulation,[],[f10323,f1414])).
fof(f10323,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))) | $spl1),
  inference(resolution,[],[f1934,f9407])).
fof(f9407,plain,(
  ( ! [X48] : (p(s(bool,'const.bool.IN_2'(s(sK1,X48),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X48)))) )),
  inference(trivial_inequality_removal,[],[f9365])).
fof(f9365,plain,(
  ( ! [X48] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.bool.IN_2'(s(sK1,X48),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) | s(bool,'F_0') = s(bool,'const.quotient.respects_2'(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X48)))) )),
  inference(superposition,[],[f112,f9260])).
fof(f1933,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) | $spl0),
  inference(cnf_transformation,[],[f1933_D])).
fof(f1933_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1973,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0)))),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(sK1,X0),X3),s(sK1,sK9(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0)))))))))) ) | $spl4),
  inference(cnf_transformation,[],[f1973_D])).
fof(f1973_D,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0)))),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(sK1,X0),X3),s(sK1,sK9(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0)))))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f8497,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0)))),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(sK1,X0),X3),s(sK1,sK9(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))))))) )),
  inference(subsumption_resolution,[],[f8390,f92])).
fof(f92,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : (((~p(s(bool,'T_0')) | ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'T_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : ((p(s(bool,'T_0')) & p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X15] : ((p(s(bool,'T_0')) & p(s(bool,X15))) <=> p(s(bool,X15)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.bool.AND_CLAUSES|split|0')).
fof(f8390,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0)))),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(sK1,X0),X3),s(sK1,sK9(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))))))) )),
  inference(superposition,[],[f154,f1391])).
fof(f1391,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2))))))),
  inference(resolution,[],[f375,f113])).
fof(f154,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X14),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X19)))))) | ~p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(sK1,X15),X18),s(sK1,sK10(sK3,X16,sK4,X17,X18,sK5,X19,X14,sK1,X15,sK0)))),s(fun(X15,bool),happ(s(fun(X15,fun(X15,bool)),X16),s(X15,happ(s(fun(sK1,X15),X17),s(sK1,sK9(sK3,X16,sK4,X17,X18,sK5,X19,X14,sK1,X15,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,X14),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))))))) )),
  inference(forward_demodulation,[],[f153,f123])).
fof(f153,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(X15,happ(s(fun(sK1,X15),X18),s(sK1,sK10(sK3,X16,sK4,X17,X18,sK5,X19,X14,sK1,X15,sK0)))),s(fun(X15,bool),happ(s(fun(X15,fun(X15,bool)),X16),s(X15,happ(s(fun(sK1,X15),X17),s(sK1,sK9(sK3,X16,sK4,X17,X18,sK5,X19,X14,sK1,X15,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,X14),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X19))),s(sK1,X14))))) )),
  inference(forward_demodulation,[],[f130,f123])).
fof(f130,plain,(
  ( ! [X14,X19,X17,X15,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(sK1,X14),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) | ~p(s(bool,happ(s(fun(X15,bool),happ(s(fun(X15,fun(X15,bool)),X16),s(X15,happ(s(fun(sK1,X15),X17),s(sK1,sK9(sK3,X16,sK4,X17,X18,sK5,X19,X14,sK1,X15,sK0)))))),s(X15,happ(s(fun(sK1,X15),X18),s(sK1,sK10(sK3,X16,sK4,X17,X18,sK5,X19,X14,sK1,X15,sK0))))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X19))),s(sK1,X14))))) )),
  inference(resolution,[],[f81,f86])).
fof(f86,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (p(s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,sK9(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,sK10(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0))))))) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) | (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,sK9(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))),s(X2,sK10(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,sK9(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,sK10(X10,X9,X8,X7,X6,X5,X4,X3,X2,X1,X0)))))))) | p(s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) | ~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) | ? [X11,X12] : (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) | p(s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) | (~p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) | (? [X11,X12] : (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) & ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) | p(s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(ennf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) => (! [X11,X12] : (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) => p(s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5,X6,X7,X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X2,fun(X2,bool)),X10),s(fun(X2,X0),X8),s(fun(X0,X2),X5)))) => (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X4))),s(X2,X3)))) => (! [X11,X12] : (p(s(bool,happ(s(fun(X2,bool),happ(s(fun(X2,fun(X2,bool)),X10),s(X2,X11))),s(X2,X12)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X9),s(X1,happ(s(fun(X2,X1),X7),s(X2,X11))))),s(X1,happ(s(fun(X2,X1),X6),s(X2,X12))))))) => p(s(bool,'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.quotient.respects_1'(s(fun(X2,fun(X2,bool)),X10)))))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6,X8,X7] : ! [X23,X17,X11,X14,X13,X10,X12,X9] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X9),s(fun(X7,X6),X10),s(fun(X6,X7),X11)))) => (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X9),s(X7,X17))),s(X7,X23)))) => (! [X24,X25] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X9),s(X7,X24))),s(X7,X25)))) => p(s(bool,happ(s(fun(X8,bool),happ(s(fun(X8,fun(X8,bool)),X12),s(X8,happ(s(fun(X7,X8),X13),s(X7,X24))))),s(X8,happ(s(fun(X7,X8),X14),s(X7,X25))))))) => p(s(bool,'const.bool.IN_2'(s(X7,X23),s(fun(X7,bool),'const.quotient.respects_1'(s(fun(X7,fun(X7,bool)),X9)))))))))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','thm.quotient.quotient_87_prove_ltb')).
fof(f10319,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f1936,f1384])).
fof(f1384,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))))))),
  inference(resolution,[],[f375,f172])).
fof(f172,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK1,X1),s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0)))))) | p(s(bool,'const.bool.IN_2'(s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,X1))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,X0))))))))) )),
  inference(forward_demodulation,[],[f171,f123])).
fof(f171,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,X1))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,X0)))))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))))) )),
  inference(forward_demodulation,[],[f158,f123])).
fof(f158,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,X0))))),s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,X1)))))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK3),s(sK1,X0))),s(sK1,X1))))) )),
  inference(resolution,[],[f82,f88])).
fof(f88,plain,(
  ( ! [X4,X2,X0,X8,X5,X3,X1,X9] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,X8))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,X9)))))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X8))),s(X1,X9)))) | ~p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f82,plain,(
  p(s(bool,'const.quotient.===>_4'(s(fun(sK1,fun(sK1,bool)),sK3),s(fun(sK2,fun(sK2,bool)),sK6),s(fun(sK1,sK2),sK7),s(fun(sK1,sK2),sK8))))),
  inference(cnf_transformation,[],[f60])).
fof(f1944,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,happ(s(fun(sK1,X0),X1),s(sK1,sK10(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0)))),s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X2),s(X0,happ(s(fun(sK1,X0),X3),s(sK1,sK9(sK3,X2,sK4,X3,X1,sK5,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2),sK1,X0,sK0))))))))) | p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))))))) )),
  inference(resolution,[],[f154,f375])).
fof(f1934,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) | $spl1),
  inference(cnf_transformation,[],[f1934_D])).
fof(f1934_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1936,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2))))))))) | $spl3),
  inference(cnf_transformation,[],[f1936_D])).
fof(f1936_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2))))))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1922,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,happ(s(fun(sK1,sK2),sK8),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2))))))))) | ~p(s(bool,'const.bool.IN_2'(s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)),s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))))))),
  inference(superposition,[],[f377,f121])).
fof(f377,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,'const.bool.RES_ABSTRACT_3'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))))))),
  inference(forward_demodulation,[],[f376,f84])).
fof(f84,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(X0,'const.bool.RES_ABSTRACT_3'(s(fun(X1,bool),X2),s(fun(X1,X0),X3),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),'const.bool.RES_ABSTRACT_2'(s(fun(X1,bool),X2),s(fun(X1,X0),X3))),s(X1,X4)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8,X7] : ! [X28,X29,X30] : s(X8,'const.bool.RES_ABSTRACT_3'(s(fun(X7,bool),X28),s(fun(X7,X8),X29),s(X7,X30))) = s(X8,happ(s(fun(X7,X8),'const.bool.RES_ABSTRACT_2'(s(fun(X7,bool),X28),s(fun(X7,X8),X29))),s(X7,X30)))),
  file('thm.quotient.ABSTRACT_RES_ABSTRACT','ARITY_2_3_const.bool.RES_ABSTRACT')).
fof(f376,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK2,happ(s(fun(sK1,sK2),'const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8))),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))),s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))))))),
  inference(forward_demodulation,[],[f362,f123])).
fof(f362,plain,(
  ~p(s(bool,happ(s(fun(sK2,bool),happ(s(fun(sK2,fun(sK2,bool)),sK6),s(sK2,happ(s(fun(sK1,sK2),sK7),s(sK1,sK11('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))),s(sK2,happ(s(fun(sK1,sK2),'const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8))),s(sK1,sK12('const.bool.RES_ABSTRACT_2'(s(fun(sK1,bool),'const.quotient.respects_1'(s(fun(sK1,fun(sK1,bool)),sK3))),s(fun(sK1,sK2),sK8)),sK7,sK6,sK3,sK1,sK2)))))))),
  inference(resolution,[],[f83,f90])).
fof(f90,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.quotient.===>_4'(s(fun(X1,fun(X1,bool)),X2),s(fun(X0,fun(X0,bool)),X3),s(fun(X1,X0),X4),s(fun(X1,X0),X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,happ(s(fun(X1,X0),X4),s(X1,sK11(X5,X4,X3,X2,X1,X0)))))),s(X0,happ(s(fun(X1,X0),X5),s(X1,sK12(X5,X4,X3,X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f64])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 161
% Passive clauses: 2510
% Generated clauses: 10221
% Final active clauses: 153
% Final passive clauses: 2079
% Input formulas: 22
% Initial clauses: 43
% 
% Selected by SInE selection: 21
% SInE iterations: 1
% 
% Duplicate literals: 112
% Trivial inequalities: 270
% Fw subsumption resolutions: 455
% Fw demodulations: 3668
% Bw demodulations: 7
% 
% Simple tautologies: 98
% Equational tautologies: 5
% Forward subsumptions: 3378
% 
% Binary resolution: 598
% Forward superposition: 2747
% Backward superposition: 2344
% Self superposition: 6
% Equality factoring: 4
% Equality resolution: 1
% 
% Split clauses: 3
% Split components: 6
% SAT solver clauses: 9
% SAT solver unit clauses: 5
% SAT solver binary clauses: 3
% 
% Sat splits: 3
% Sat splitting refutations: 5
% 
% Memory used [KB]: 7931
% Time elapsed: 0.414 s
% ------------------------------
% ------------------------------
% Success in time 12.11 s
