Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (62960)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 113
% Passive clauses: 4210
% Generated clauses: 17190
% Final active clauses: 86
% Final passive clauses: 3461
% Input formulas: 14
% Initial clauses: 29
% 
% Duplicate literals: 841
% Trivial inequalities: 165
% Fw subsumption resolutions: 2418
% Fw demodulations: 1060
% Bw demodulations: 14
% 
% Simple tautologies: 294
% Equational tautologies: 79
% Forward subsumptions: 6871
% 
% Binary resolution: 457
% Factoring: 2
% Forward superposition: 7250
% Backward superposition: 4283
% Self superposition: 224
% Equality factoring: 174
% Equality resolution: 19
% 
% Split clauses: 571
% Split components: 1179
% SAT solver clauses: 1396
% SAT solver unit clauses: 19
% SAT solver binary clauses: 215
% 
% Sat splits: 571
% Sat splitting refutations: 11
% 
% TWLsolver clauses: 1380
% TWLsolver calls for satisfiability: 196
% 
% Memory used [KB]: 9978
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (62979)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 43
% Passive clauses: 442
% Generated clauses: 5727
% Final active clauses: 40
% Final passive clauses: 385
% Input formulas: 14
% Initial clauses: 29
% 
% Duplicate literals: 643
% Trivial inequalities: 38
% Fw subsumption resolutions: 231
% Fw demodulations: 36
% Bw demodulations: 16
% Condensations: 1015
% Global subsumptions: 1
% 
% Simple tautologies: 116
% Equational tautologies: 72
% Forward subsumptions: 1813
% 
% Binary resolution: 114
% Factoring: 2
% Forward superposition: 1248
% Backward superposition: 1635
% Self superposition: 190
% Equality factoring: 166
% Equality resolution: 10
% 
% Split clauses: 15
% Split components: 30
% SAT solver clauses: 693
% SAT solver unit clauses: 24
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 2
% 
% Sat splits: 15
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 689
% TWLsolver calls for satisfiability: 146
% 
% Memory used [KB]: 5245
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (63011)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 597
% Passive clauses: 10681
% Generated clauses: 155941
% Final active clauses: 127
% Final passive clauses: 1053
% Input formulas: 14
% Initial clauses: 29
% 
% Duplicate literals: 6517
% Trivial inequalities: 1170
% Fw subsumption resolutions: 14386
% Fw demodulations: 19887
% Bw demodulations: 8591
% Condensations: 17277
% Global subsumptions: 218
% 
% Simple tautologies: 1430
% Equational tautologies: 1003
% Forward subsumptions: 80899
% Bw demodulations to eq. taut.: 109
% 
% Unit resulting resolution: 6868
% Factoring: 1
% Forward superposition: 39783
% Backward superposition: 28175
% Self superposition: 701
% Equality factoring: 375
% Equality resolution: 123
% 
% Split clauses: 466
% Split components: 947
% SAT solver clauses: 12903
% SAT solver unit clauses: 378
% SAT solver binary clauses: 1119
% 
% Sat splits: 466
% Sat splitting refutations: 1070
% 
% Memory used [KB]: 70489
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (63258)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 406
% Passive clauses: 31469
% Generated clauses: 80414
% Final active clauses: 304
% Final passive clauses: 29299
% Input formulas: 14
% Initial clauses: 27
% 
% Selected by SInE selection: 12
% SInE iterations: 1
% 
% Duplicate literals: 2169
% Trivial inequalities: 821
% Fw subsumption resolutions: 11455
% Fw demodulations: 4432
% Bw demodulations: 76
% 
% Simple tautologies: 1486
% Equational tautologies: 79
% Forward subsumptions: 27072
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 10687
% Factoring: 8
% Forward superposition: 31869
% Backward superposition: 17590
% Self superposition: 350
% Equality factoring: 170
% Equality resolution: 86
% 
% Split clauses: 2191
% Split components: 4446
% SAT solver clauses: 2264
% SAT solver unit clauses: 16
% SAT solver binary clauses: 360
% 
% Sat splits: 2191
% Sat splitting refutations: 36
% 
% Memory used [KB]: 36715
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (63394)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 87
% Passive clauses: 1676
% Generated clauses: 11415
% Final active clauses: 65
% Final passive clauses: 1328
% Input formulas: 14
% Initial clauses: 29
% 
% Duplicate literals: 541
% Trivial inequalities: 267
% Fw subsumption resolutions: 1894
% Fw demodulations: 3
% Bw demodulations: 193
% Global subsumptions: 3
% 
% Simple tautologies: 639
% Equational tautologies: 55
% Forward subsumptions: 6040
% 
% Unit resulting resolution: 145
% Factoring: 1
% Forward superposition: 4260
% Backward superposition: 3658
% Self superposition: 99
% Equality factoring: 46
% Equality resolution: 13
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 1752
% SAT solver unit clauses: 30
% SAT solver binary clauses: 82
% 
% Sat splits: 19
% Sat splitting refutations: 52
% 
% Memory used [KB]: 7803
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (63430)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 176
% Passive clauses: 1789
% Generated clauses: 11191
% Final active clauses: 110
% Final passive clauses: 1183
% Input formulas: 14
% Initial clauses: 29
% 
% Duplicate literals: 561
% Trivial inequalities: 233
% Fw subsumption resolutions: 701
% Bw subsumption resolutions: 34
% Fw demodulations: 750
% Bw demodulations: 216
% Condensations: 1237
% Global subsumptions: 3
% 
% Simple tautologies: 464
% Equational tautologies: 43
% Forward subsumptions: 5385
% Backward subsumptions: 132
% 
% Binary resolution: 1214
% Factoring: 5
% Forward superposition: 3551
% Backward superposition: 1765
% Self superposition: 86
% Equality factoring: 58
% Equality resolution: 43
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 1830
% SAT solver unit clauses: 26
% SAT solver binary clauses: 142
% 
% Sat splits: 4
% Sat splitting refutations: 32
% 
% Memory used [KB]: 9594
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f10981,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f60,f90,f94,f96,f93,f98,f92,f100,f91,f102,f103,f107,f108,f112,f115,f122,f129,f128,f133,f127,f131,f135,f136,f138,f139,f141,f179,f134,f237,f236,f241,f235,f239,f243,f244,f246,f296,f242,f354,f353,f358,f352,f356,f360,f361,f363,f397,f359,f371,f399,f422,f447,f445,f468,f467,f472,f466,f470,f474,f475,f477,f511,f473,f531,f556,f554,f583,f573,f577,f581,f586,f485,f589,f658,f656,f666,f655,f668,f653,f661,f670,f671,f116,f676,f675,f680,f674,f678,f117,f687,f398,f742,f730,f732,f736,f740,f746,f761,f759,f758,f667,f756,f763,f768,f767,f775,f773,f783,f772,f785,f770,f778,f787,f788,f123,f793,f792,f797,f791,f795,f124,f798,f799,f800,f806,f801,f808,f802,f810,f803,f812,f804,f814,f815,f817,f825,f588,f881,f879,f878,f784,f875,f883,f889,f887,f911,f909,f919,f917,f920,f1031,f1029,f1028,f1025,f1024,f446,f1034,f669,f1035,f1040,f1042,f1045,f119,f1112,f1110,f1108,f1121,f1122,f1118,f1259,f1257,f1256,f1253,f1271,f1252,f1262,f786,f1263,f1268,f1274,f1373,f1343,f444,f1426,f1396,f553,f2332,f2333,f2334,f2335,f2336,f2388,f2389,f2390,f2391,f2392,f2405,f2406,f2440,f2441,f2452,f2453,f2454,f2486,f2487,f2548,f2779,f2549,f2780,f2610,f2611,f2645,f2770,f2646,f2657,f2658,f2659,f2774,f2783,f104,f2788,f2787,f2792,f2786,f2790,f105,f3717,f3718,f3719,f3720,f3779,f3780,f3781,f3782,f3798,f3800,f3840,f3841,f3853,f3854,f4021,f2775,f4023,f2776,f4063,f4195,f4064,f2771,f4076,f2772,f4077,f2773,f4200,f4201,f4203,f4212,f4204,f4205,f4206,f813,f4207,f4208,f4220,f4209,f4222,f4210,f4224,f4225,f816,f4226,f807,f4227,f4229,f4230,f4232,f4233,f805,f4234,f4238,f126,f4379,f4377,f4376,f4375,f4214,f4373,f4381,f4382,f4385,f4218,f4392,f4390,f4389,f4388,f4213,f4386,f4395,f4215,f4396,f4409,f4407,f4406,f4411,f4417,f4415,f4435,f5095,f5093,f5092,f5091,f5090,f5105,f5089,f5107,f5099,f5110,f5103,f5720,f5746,f5727,f5748,f5728,f5750,f5787,f5817,f5794,f5819,f5795,f5821,f5956,f5958,f6030,f6032,f6056,f6057,f6121,f6123,f6339,f6545,f6340,f6546,f6404,f6544,f6406,f6543,f6573,f6625,f6672,f6723,f7217,f7218,f7222,f7225,f7226,f7227,f7316,f7921,f7317,f7922,f7321,f7324,f7325,f7926,f7326,f7927,f7424,f7426,f7747,f7925,f7749,f7923,f7977,f8017,f7981,f8018,f7986,f8019,f7987,f8020,f8075,f8114,f8077,f8115,f8085,f8116,f8086,f8117,f8144,f8145,f8233,f8302,f8303,f8354,f8559,f8560,f8081,f8561,f5723,f4401,f8563,f8564,f4231,f8582,f8573,f8574,f4221,f8576,f8584,f8577,f8586,f8578,f8588,f8589,f109,f8571,f4216,f8718,f8716,f8714,f8713,f8712,f8710,f8709,f8708,f8732,f8706,f8704,f8701,f8730,f8733,f8734,f110,f8748,f8746,f8745,f8769,f8770,f8742,f8772,f8774,f8775,f8755,f8778,f8779,f8768,f8789,f8792,f8819,f8822,f8854,f9196,f9199,f9227,f9230,f9263,f9290,f9293,f9298,f9365,f9301,f9367,f9327,f9330,f9363,f9428,f9445,f9418,f9091,f9088,f9059,f9453,f9054,f9051,f9451,f8607,f9454,f9842,f9845,f9850,f9916,f9879,f9882,f140,f9915,f8886,f8605,f449,f811,f9954,f9952,f9951,f9950,f9973,f9975,f9977,f9979,f9981,f9983,f9984,f9941,f9987,f9961,f9989,f9990,f9991,f4219,f10115,f10113,f10112,f10137,f10138,f10139,f10140,f10141,f10142,f10105,f10144,f10146,f10148,f10150,f10152,f10154,f10155,f10097,f10157,f10095,f10159,f10094,f10120,f10161,f10123,f10163,f10164,f10136,f10305,f10867,f10865,f10864,f10888,f10889,f10890,f10891,f10892,f10893,f10857,f10895,f10897,f10899,f10901,f10903,f10905,f10906,f10874,f10909,f10910,f10887,f10936,f10934,f10933,f10957,f10958,f10959,f10960,f10961,f10962,f10926,f10964,f10966,f10968,f10970,f10972,f10974,f10975,f10977,f401,f10943,f10979,f10980,f10956])).
fof(f10956,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | $spl262),
  inference(cnf_transformation,[],[f10956_D])).
fof(f10956_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f10980,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83)),
  inference(global_subsumption,[],[f10949,f60,f786,f10936,f10934,f10917,f10944])).
fof(f10944,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl83),
  inference(resolution,[],[f786,f76])).
fof(f76,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.compl_insert','thm.bool.IMP_ANTISYM_AX')).
fof(f10917,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,X0))))))) ) | $spl83),
  inference(unit_resulting_resolution,[],[f786,f70])).
fof(f70,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X7,X8,X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X7),s(X6,X13)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X7)))) & ~s(X6,X8) = s(X6,X13)))),
  file('thm.pred_set.compl_insert','thm.pred_set.IN_DELETE')).
fof(f10949,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl83)),
  inference(superposition,[],[f786,f4901])).
fof(f4901,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl30),
  inference(resolution,[],[f4331,f4304])).
fof(f4304,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl30),
  inference(forward_demodulation,[],[f4296,f4292])).
fof(f4292,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0')))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f62,f2817,f76])).
fof(f2817,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))))) ) | $spl30),
  inference(forward_demodulation,[],[f2794,f137])).
fof(f137,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl30),
  inference(cnf_transformation,[],[f137_D])).
fof(f137_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2794,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f137,f71])).
fof(f71,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f62,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f18])).
fof(f18,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.compl_insert',aHL_FALSITY)).
fof(f4296,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'F_0')))))) ) | $spl30),
  inference(resolution,[],[f2817,f76])).
fof(f4331,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl30),
  inference(resolution,[],[f4304,f79])).
fof(f79,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f42])).
fof(f10979,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83)),
  inference(global_subsumption,[],[f10948,f60,f786,f10936,f10934,f10917,f10944])).
fof(f10948,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl30 | $spl83)),
  inference(superposition,[],[f786,f4901])).
fof(f10943,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83)),
  inference(resolution,[],[f786,f4304])).
fof(f401,plain,(
  ~p(s(bool,'T_0')) | $spl49),
  inference(cnf_transformation,[],[f401_D])).
fof(f401_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f10977,plain,(
  ~p(s(bool,'T_0')) | ($spl65 | $spl83 | $spl104)),
  inference(forward_demodulation,[],[f10916,f1118])).
fof(f10916,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | ($spl65 | $spl83)),
  inference(unit_resulting_resolution,[],[f555,f786,f73])).
fof(f73,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f28])).
fof(f28,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : ! [X8,X13,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X13),s(fun(X6,bool),X7)))))) <=> (s(X6,X8) = s(X6,X13) | p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X7))))))),
  file('thm.pred_set.compl_insert','thm.pred_set.IN_INSERT')).
fof(f555,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) | $spl65),
  inference(cnf_transformation,[],[f555_D])).
fof(f555_D,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f10975,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl33 | $spl83)),
  inference(forward_demodulation,[],[f10919,f4292])).
fof(f10919,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl33 | $spl83)),
  inference(unit_resulting_resolution,[],[f141,f786,f4589])).
fof(f4589,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X6),s(bool,X4)))))) ) | $spl30),
  inference(resolution,[],[f4298,f76])).
fof(f4298,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X1),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f2817,f63])).
fof(f63,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.compl_insert',aHL_BOOL_CASES)).
fof(f10974,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl83)),
  inference(forward_demodulation,[],[f10973,f4292])).
fof(f10973,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl11 | $spl30 | $spl83)),
  inference(forward_demodulation,[],[f10920,f137])).
fof(f10920,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))))))) ) | ($spl11 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f104,f786,f4589])).
fof(f10972,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl30 | $spl38 | $spl83)),
  inference(forward_demodulation,[],[f10971,f4292])).
fof(f10971,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl30 | $spl38 | $spl83)),
  inference(forward_demodulation,[],[f10921,f245])).
fof(f245,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl38),
  inference(cnf_transformation,[],[f245_D])).
fof(f245_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f10921,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))))))) ) | ($spl15 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f109,f786,f4589])).
fof(f10970,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl19 | $spl30 | $spl44 | $spl83)),
  inference(forward_demodulation,[],[f10969,f4292])).
fof(f10969,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl19 | $spl30 | $spl44 | $spl83)),
  inference(forward_demodulation,[],[f10922,f362])).
fof(f362,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl44),
  inference(cnf_transformation,[],[f362_D])).
fof(f362_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f10922,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl19 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f116,f786,f4589])).
fof(f10968,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl23 | $spl30 | $spl60 | $spl83)),
  inference(forward_demodulation,[],[f10967,f4292])).
fof(f10967,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl23 | $spl30 | $spl60 | $spl83)),
  inference(forward_demodulation,[],[f10923,f476])).
fof(f476,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl60),
  inference(cnf_transformation,[],[f476_D])).
fof(f476_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f10923,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl23 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f123,f786,f4589])).
fof(f10966,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83 | $spl87 | $spl96)),
  inference(forward_demodulation,[],[f10965,f4292])).
fof(f10965,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl83 | $spl87 | $spl96)),
  inference(forward_demodulation,[],[f10924,f817])).
fof(f10924,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl30 | $spl83 | $spl87)),
  inference(unit_resulting_resolution,[],[f806,f786,f4589])).
fof(f10964,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10963,f4292])).
fof(f10963,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl83 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10925,f8768])).
fof(f10925,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl83 | $spl155)),
  inference(unit_resulting_resolution,[],[f4414,f786,f4589])).
fof(f4414,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl155),
  inference(cnf_transformation,[],[f4414_D])).
fof(f4414_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f10926,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl33 | $spl83)),
  inference(unit_resulting_resolution,[],[f141,f786,f4353])).
fof(f4353,plain,(
  ( ! [X2,X1] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,X2)) ) | $spl30),
  inference(resolution,[],[f4321,f76])).
fof(f4321,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl30),
  inference(forward_demodulation,[],[f4314,f4311])).
fof(f4311,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f61,f2818,f79])).
fof(f2818,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl30),
  inference(forward_demodulation,[],[f2793,f137])).
fof(f2793,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),X0))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f137,f74])).
fof(f74,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f61,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.compl_insert',aHL_TRUTH)).
fof(f4314,plain,(
  ( ! [X2,X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))))) ) | $spl30),
  inference(resolution,[],[f2818,f79])).
fof(f10962,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl83)),
  inference(forward_demodulation,[],[f10927,f137])).
fof(f10927,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f104,f786,f4353])).
fof(f10961,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl30 | $spl38 | $spl83)),
  inference(forward_demodulation,[],[f10928,f245])).
fof(f10928,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f109,f786,f4353])).
fof(f10960,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl19 | $spl30 | $spl44 | $spl83)),
  inference(forward_demodulation,[],[f10929,f362])).
fof(f10929,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl19 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f116,f786,f4353])).
fof(f10959,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl23 | $spl30 | $spl60 | $spl83)),
  inference(forward_demodulation,[],[f10930,f476])).
fof(f10930,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl23 | $spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f123,f786,f4353])).
fof(f10958,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83 | $spl87 | $spl96)),
  inference(forward_demodulation,[],[f10931,f817])).
fof(f10931,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83 | $spl87)),
  inference(unit_resulting_resolution,[],[f806,f786,f4353])).
fof(f10957,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10932,f8768])).
fof(f10932,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83 | $spl155)),
  inference(unit_resulting_resolution,[],[f4414,f786,f4353])).
fof(f10933,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl83)),
  inference(unit_resulting_resolution,[],[f786,f4304])).
fof(f10934,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | $spl83),
  inference(unit_resulting_resolution,[],[f62,f786,f76])).
fof(f10936,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1))) | $spl83),
  inference(unit_resulting_resolution,[],[f62,f786,f76])).
fof(f10887,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl260),
  inference(cnf_transformation,[],[f10887_D])).
fof(f10887_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f10910,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73)),
  inference(global_subsumption,[],[f10880,f60,f669,f10867,f10865,f10848,f10875])).
fof(f10875,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | $spl73),
  inference(resolution,[],[f669,f76])).
fof(f10848,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,X0))))))) ) | $spl73),
  inference(unit_resulting_resolution,[],[f669,f70])).
fof(f10880,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl73)),
  inference(superposition,[],[f669,f4901])).
fof(f10909,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73)),
  inference(global_subsumption,[],[f10879,f60,f669,f10867,f10865,f10848,f10875])).
fof(f10879,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl30 | $spl73)),
  inference(superposition,[],[f669,f4901])).
fof(f10874,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73)),
  inference(resolution,[],[f669,f4304])).
fof(f10906,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl33 | $spl73)),
  inference(forward_demodulation,[],[f10850,f4292])).
fof(f10850,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl33 | $spl73)),
  inference(unit_resulting_resolution,[],[f141,f669,f4589])).
fof(f10905,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl73)),
  inference(forward_demodulation,[],[f10904,f4292])).
fof(f10904,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl11 | $spl30 | $spl73)),
  inference(forward_demodulation,[],[f10851,f137])).
fof(f10851,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))))))) ) | ($spl11 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f104,f669,f4589])).
fof(f10903,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl30 | $spl38 | $spl73)),
  inference(forward_demodulation,[],[f10902,f4292])).
fof(f10902,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl15 | $spl30 | $spl38 | $spl73)),
  inference(forward_demodulation,[],[f10852,f245])).
fof(f10852,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))))))) ) | ($spl15 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f109,f669,f4589])).
fof(f10901,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl19 | $spl30 | $spl44 | $spl73)),
  inference(forward_demodulation,[],[f10900,f4292])).
fof(f10900,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl19 | $spl30 | $spl44 | $spl73)),
  inference(forward_demodulation,[],[f10853,f362])).
fof(f10853,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl19 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f116,f669,f4589])).
fof(f10899,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl23 | $spl30 | $spl60 | $spl73)),
  inference(forward_demodulation,[],[f10898,f4292])).
fof(f10898,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl23 | $spl30 | $spl60 | $spl73)),
  inference(forward_demodulation,[],[f10854,f476])).
fof(f10854,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl23 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f123,f669,f4589])).
fof(f10897,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73 | $spl87 | $spl96)),
  inference(forward_demodulation,[],[f10896,f4292])).
fof(f10896,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl73 | $spl87 | $spl96)),
  inference(forward_demodulation,[],[f10855,f817])).
fof(f10855,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl30 | $spl73 | $spl87)),
  inference(unit_resulting_resolution,[],[f806,f669,f4589])).
fof(f10895,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10894,f4292])).
fof(f10894,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl73 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10856,f8768])).
fof(f10856,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl73 | $spl155)),
  inference(unit_resulting_resolution,[],[f4414,f669,f4589])).
fof(f10857,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl33 | $spl73)),
  inference(unit_resulting_resolution,[],[f141,f669,f4353])).
fof(f10893,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl73)),
  inference(forward_demodulation,[],[f10858,f137])).
fof(f10858,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f104,f669,f4353])).
fof(f10892,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl30 | $spl38 | $spl73)),
  inference(forward_demodulation,[],[f10859,f245])).
fof(f10859,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f109,f669,f4353])).
fof(f10891,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl19 | $spl30 | $spl44 | $spl73)),
  inference(forward_demodulation,[],[f10860,f362])).
fof(f10860,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl19 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f116,f669,f4353])).
fof(f10890,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl23 | $spl30 | $spl60 | $spl73)),
  inference(forward_demodulation,[],[f10861,f476])).
fof(f10861,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl23 | $spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f123,f669,f4353])).
fof(f10889,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73 | $spl87 | $spl96)),
  inference(forward_demodulation,[],[f10862,f817])).
fof(f10862,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73 | $spl87)),
  inference(unit_resulting_resolution,[],[f806,f669,f4353])).
fof(f10888,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10863,f8768])).
fof(f10863,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73 | $spl155)),
  inference(unit_resulting_resolution,[],[f4414,f669,f4353])).
fof(f10864,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | ($spl30 | $spl73)),
  inference(unit_resulting_resolution,[],[f669,f4304])).
fof(f10865,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl73),
  inference(unit_resulting_resolution,[],[f62,f669,f76])).
fof(f10867,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1))) | $spl73),
  inference(unit_resulting_resolution,[],[f62,f669,f76])).
fof(f10305,plain,(
  $false | $spl144),
  inference(trivial_inequality_removal,[],[f10270])).
fof(f10270,plain,(
  s(sK0,sK2) != s(sK0,sK2) | $spl144),
  inference(unit_resulting_resolution,[],[f4222,f71])).
fof(f10136,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl252),
  inference(cnf_transformation,[],[f10136_D])).
fof(f10136_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f10164,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl145)),
  inference(global_subsumption,[],[f10129,f60,f4221,f10115,f10113,f10097,f10096,f10124])).
fof(f10124,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | $spl145),
  inference(resolution,[],[f4221,f76])).
fof(f10096,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))),s(sK0,X0))))))) ) | $spl145),
  inference(unit_resulting_resolution,[],[f4221,f70])).
fof(f10129,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl30 | $spl145)),
  inference(superposition,[],[f4221,f4901])).
fof(f10163,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl145)),
  inference(global_subsumption,[],[f10128,f60,f4221,f10115,f10113,f10097,f10096,f10124])).
fof(f10128,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl30 | $spl145)),
  inference(superposition,[],[f4221,f4901])).
fof(f10123,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl145)),
  inference(resolution,[],[f4221,f4304])).
fof(f10161,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | $spl258),
  inference(cnf_transformation,[],[f10161_D])).
fof(f10161_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f10120,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | ($spl30 | $spl145)),
  inference(resolution,[],[f4221,f4356])).
fof(f4356,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),X7)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),'const.pred_set.COMPL_1'(s(fun(X5,bool),X7)))))) ) | $spl30),
  inference(resolution,[],[f4321,f69])).
fof(f69,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X8,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.compl_insert','thm.pred_set.IN_COMPL')).
fof(f10094,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | ($spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f4221,f4356])).
fof(f10159,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) | $spl257),
  inference(cnf_transformation,[],[f10159_D])).
fof(f10159_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) <=> ~$spl257),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl257])])).
fof(f10095,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) | ($spl65 | $spl145)),
  inference(unit_resulting_resolution,[],[f555,f4221,f73])).
fof(f10157,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) | $spl254),
  inference(cnf_transformation,[],[f10157_D])).
fof(f10157_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f10097,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) | $spl145),
  inference(unit_resulting_resolution,[],[f4221,f69])).
fof(f10155,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl33 | $spl145)),
  inference(forward_demodulation,[],[f10098,f4292])).
fof(f10098,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl30 | $spl33 | $spl145)),
  inference(unit_resulting_resolution,[],[f141,f4221,f4589])).
fof(f10154,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl11 | $spl30 | $spl145)),
  inference(forward_demodulation,[],[f10153,f4292])).
fof(f10153,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl11 | $spl30 | $spl145)),
  inference(forward_demodulation,[],[f10099,f137])).
fof(f10099,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))))))) ) | ($spl11 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f104,f4221,f4589])).
fof(f10152,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl15 | $spl30 | $spl38 | $spl145)),
  inference(forward_demodulation,[],[f10151,f4292])).
fof(f10151,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl15 | $spl30 | $spl38 | $spl145)),
  inference(forward_demodulation,[],[f10100,f245])).
fof(f10100,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))))))) ) | ($spl15 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f109,f4221,f4589])).
fof(f10150,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl19 | $spl30 | $spl44 | $spl145)),
  inference(forward_demodulation,[],[f10149,f4292])).
fof(f10149,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl19 | $spl30 | $spl44 | $spl145)),
  inference(forward_demodulation,[],[f10101,f362])).
fof(f10101,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl19 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f116,f4221,f4589])).
fof(f10148,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl23 | $spl30 | $spl60 | $spl145)),
  inference(forward_demodulation,[],[f10147,f4292])).
fof(f10147,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl23 | $spl30 | $spl60 | $spl145)),
  inference(forward_demodulation,[],[f10102,f476])).
fof(f10102,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl23 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f123,f4221,f4589])).
fof(f10146,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl87 | $spl96 | $spl145)),
  inference(forward_demodulation,[],[f10145,f4292])).
fof(f10145,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl30 | $spl87 | $spl96 | $spl145)),
  inference(forward_demodulation,[],[f10103,f817])).
fof(f10103,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl30 | $spl87 | $spl145)),
  inference(unit_resulting_resolution,[],[f806,f4221,f4589])).
fof(f10144,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl145 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10143,f4292])).
fof(f10143,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl30 | $spl145 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10104,f8768])).
fof(f10104,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) ) | ($spl30 | $spl145 | $spl155)),
  inference(unit_resulting_resolution,[],[f4414,f4221,f4589])).
fof(f10105,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl33 | $spl145)),
  inference(unit_resulting_resolution,[],[f141,f4221,f4353])).
fof(f10142,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl11 | $spl30 | $spl145)),
  inference(forward_demodulation,[],[f10106,f137])).
fof(f10106,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | ($spl11 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f104,f4221,f4353])).
fof(f10141,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl15 | $spl30 | $spl38 | $spl145)),
  inference(forward_demodulation,[],[f10107,f245])).
fof(f10107,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl15 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f109,f4221,f4353])).
fof(f10140,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl19 | $spl30 | $spl44 | $spl145)),
  inference(forward_demodulation,[],[f10108,f362])).
fof(f10108,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl19 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f116,f4221,f4353])).
fof(f10139,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl23 | $spl30 | $spl60 | $spl145)),
  inference(forward_demodulation,[],[f10109,f476])).
fof(f10109,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl23 | $spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f123,f4221,f4353])).
fof(f10138,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl87 | $spl96 | $spl145)),
  inference(forward_demodulation,[],[f10110,f817])).
fof(f10110,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl87 | $spl145)),
  inference(unit_resulting_resolution,[],[f806,f4221,f4353])).
fof(f10137,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl145 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f10111,f8768])).
fof(f10111,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl145 | $spl155)),
  inference(unit_resulting_resolution,[],[f4414,f4221,f4353])).
fof(f10112,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl30 | $spl145)),
  inference(unit_resulting_resolution,[],[f4221,f4304])).
fof(f10113,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl145),
  inference(unit_resulting_resolution,[],[f62,f4221,f76])).
fof(f10115,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl145),
  inference(unit_resulting_resolution,[],[f62,f4221,f76])).
fof(f4219,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl142),
  inference(cnf_transformation,[],[f4219_D])).
fof(f4219_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f9991,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl92)),
  inference(global_subsumption,[],[f9970,f60,f811,f9954,f9952,f9940,f9962,f9963,f9989])).
fof(f9963,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl92),
  inference(superposition,[],[f811,f63])).
fof(f9962,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) ) | $spl92),
  inference(resolution,[],[f811,f79])).
fof(f9940,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl92),
  inference(unit_resulting_resolution,[],[f811,f75])).
fof(f75,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f9970,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl30 | $spl92)),
  inference(superposition,[],[f811,f4941])).
fof(f4941,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl30),
  inference(resolution,[],[f4353,f4321])).
fof(f9990,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl92)),
  inference(global_subsumption,[],[f9969,f60,f811,f9954,f9952,f9940,f9962,f9963,f9989])).
fof(f9969,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl92)),
  inference(superposition,[],[f811,f4941])).
fof(f9989,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl92),
  inference(global_subsumption,[],[f60,f811,f9954,f9952,f9940,f9962,f9963])).
fof(f9961,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl92)),
  inference(resolution,[],[f811,f4321])).
fof(f9987,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0))))))) | $spl250),
  inference(cnf_transformation,[],[f9987_D])).
fof(f9987_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0))))))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f9941,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0))))))) | ($spl65 | $spl92)),
  inference(unit_resulting_resolution,[],[f555,f811,f72])).
fof(f72,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f52])).
fof(f9984,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl33 | $spl92)),
  inference(forward_demodulation,[],[f9943,f4311])).
fof(f9943,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | ($spl30 | $spl33 | $spl92)),
  inference(unit_resulting_resolution,[],[f141,f811,f4627])).
fof(f4627,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,X5)) | s(bool,'T_0') = s(bool,X4) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X6)))))) ) | $spl30),
  inference(resolution,[],[f4315,f79])).
fof(f4315,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f2818,f63])).
fof(f9983,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl11 | $spl30 | $spl92)),
  inference(forward_demodulation,[],[f9982,f4311])).
fof(f9982,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | ($spl11 | $spl30 | $spl92)),
  inference(forward_demodulation,[],[f9944,f137])).
fof(f9944,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),X0)))))) ) | ($spl11 | $spl30 | $spl92)),
  inference(unit_resulting_resolution,[],[f104,f811,f4627])).
fof(f9981,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl15 | $spl30 | $spl38 | $spl92)),
  inference(forward_demodulation,[],[f9980,f4311])).
fof(f9980,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | ($spl15 | $spl30 | $spl38 | $spl92)),
  inference(forward_demodulation,[],[f9945,f245])).
fof(f9945,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))),s(fun(bool,bool),X0)))))) ) | ($spl15 | $spl30 | $spl92)),
  inference(unit_resulting_resolution,[],[f109,f811,f4627])).
fof(f9979,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl19 | $spl30 | $spl44 | $spl92)),
  inference(forward_demodulation,[],[f9978,f4311])).
fof(f9978,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | ($spl19 | $spl30 | $spl44 | $spl92)),
  inference(forward_demodulation,[],[f9946,f362])).
fof(f9946,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),X0)))))) ) | ($spl19 | $spl30 | $spl92)),
  inference(unit_resulting_resolution,[],[f116,f811,f4627])).
fof(f9977,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl23 | $spl30 | $spl60 | $spl92)),
  inference(forward_demodulation,[],[f9976,f4311])).
fof(f9976,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | ($spl23 | $spl30 | $spl60 | $spl92)),
  inference(forward_demodulation,[],[f9947,f476])).
fof(f9947,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),X0)))))) ) | ($spl23 | $spl30 | $spl92)),
  inference(unit_resulting_resolution,[],[f123,f811,f4627])).
fof(f9975,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl87 | $spl92 | $spl96)),
  inference(forward_demodulation,[],[f9974,f4311])).
fof(f9974,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | ($spl30 | $spl87 | $spl92 | $spl96)),
  inference(forward_demodulation,[],[f9948,f817])).
fof(f9948,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),X0)))))) ) | ($spl30 | $spl87 | $spl92)),
  inference(unit_resulting_resolution,[],[f806,f811,f4627])).
fof(f9973,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl92 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f9972,f4311])).
fof(f9972,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0)))))) ) | ($spl30 | $spl92 | $spl155 | $spl228)),
  inference(forward_demodulation,[],[f9949,f8768])).
fof(f9949,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(fun(bool,bool),X0)))))) ) | ($spl30 | $spl92 | $spl155)),
  inference(unit_resulting_resolution,[],[f4414,f811,f4627])).
fof(f9950,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl33 | $spl92)),
  inference(unit_resulting_resolution,[],[f141,f811,f4331])).
fof(f9951,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl30 | $spl92)),
  inference(unit_resulting_resolution,[],[f811,f4321])).
fof(f9952,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl92),
  inference(unit_resulting_resolution,[],[f61,f811,f79])).
fof(f9954,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl92),
  inference(unit_resulting_resolution,[],[f61,f811,f79])).
fof(f811,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl92),
  inference(cnf_transformation,[],[f811_D])).
fof(f811_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f449,plain,(
  p(s(bool,'F_0')) | $spl54),
  inference(cnf_transformation,[],[f449_D])).
fof(f449_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8605,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl96),
  inference(superposition,[],[f69,f817])).
fof(f8886,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | ($spl30 | $spl87)),
  inference(unit_resulting_resolution,[],[f806,f4356])).
fof(f9915,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl246),
  inference(cnf_transformation,[],[f9915_D])).
fof(f9915_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f140,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl32),
  inference(cnf_transformation,[],[f140_D])).
fof(f140_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f9882,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl142)),
  inference(superposition,[],[f4901,f4219])).
fof(f9879,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X30)) ) | ($spl30 | $spl142)),
  inference(superposition,[],[f4901,f4219])).
fof(f9916,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) ) | $spl248),
  inference(cnf_transformation,[],[f9916_D])).
fof(f9916_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f9850,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) ) | ($spl30 | $spl142)),
  inference(superposition,[],[f4358,f4219])).
fof(f4358,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.bool.IN_2'(s(X12,X13),s(fun(X12,bool),X15)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X12,X13),s(fun(X12,bool),'const.pred_set.INSERT_2'(s(X12,X14),s(fun(X12,bool),X15)))))) ) | $spl30),
  inference(resolution,[],[f4321,f75])).
fof(f9845,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X3) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl142)),
  inference(superposition,[],[f4219,f4901])).
fof(f9842,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl30 | $spl142)),
  inference(superposition,[],[f4219,f4901])).
fof(f9454,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X1)))))) ) | $spl244),
  inference(cnf_transformation,[],[f9454_D])).
fof(f9454_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X1)))))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f8607,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X1)))))) ) | ($spl30 | $spl96)),
  inference(superposition,[],[f4335,f817])).
fof(f4335,plain,(
  ( ! [X6,X8,X7,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X9)))))) ) | $spl30),
  inference(resolution,[],[f4304,f70])).
fof(f9451,plain,(
  ( ! [X33] : (s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | $spl240),
  inference(cnf_transformation,[],[f9451_D])).
fof(f9451_D,plain,(
  ( ! [X33] : (s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f9051,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl30 | $spl88)),
  inference(superposition,[],[f808,f4901])).
fof(f9054,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X3) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl30 | $spl88)),
  inference(superposition,[],[f808,f4901])).
fof(f9453,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) | $spl242),
  inference(cnf_transformation,[],[f9453_D])).
fof(f9453_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f9059,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) | ($spl30 | $spl88)),
  inference(superposition,[],[f4358,f808])).
fof(f9088,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X30)) ) | ($spl30 | $spl88)),
  inference(superposition,[],[f4901,f808])).
fof(f9091,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl30 | $spl88)),
  inference(superposition,[],[f4901,f808])).
fof(f9418,plain,(
  ~p(s(bool,'T_0')) | ($spl64 | $spl88)),
  inference(superposition,[],[f9168,f808])).
fof(f9168,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK2))))))) ) | $spl64),
  inference(forward_demodulation,[],[f9148,f556])).
fof(f9148,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK2))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f556,f71])).
fof(f9445,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl30 | $spl64 | $spl222)),
  inference(backward_demodulation,[],[f9391,f8586])).
fof(f9391,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK2)))))))) ) | ($spl30 | $spl64)),
  inference(unit_resulting_resolution,[],[f9168,f4356])).
fof(f9428,plain,(
  p(s(bool,'F_0')) | ($spl64 | $spl220)),
  inference(backward_demodulation,[],[f9407,f8584])).
fof(f9407,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK2)))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f62,f9168,f76])).
fof(f9363,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X0)) ) | $spl234),
  inference(cnf_transformation,[],[f9363_D])).
fof(f9363_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f9330,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl140)),
  inference(superposition,[],[f4901,f4218])).
fof(f9327,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X30)) ) | ($spl30 | $spl140)),
  inference(superposition,[],[f4901,f4218])).
fof(f9367,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | $spl238),
  inference(cnf_transformation,[],[f9367_D])).
fof(f9367_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f9301,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | ($spl30 | $spl140)),
  inference(superposition,[],[f4334,f4218])).
fof(f4334,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),X5)))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X3,X4),s(fun(X3,bool),'const.pred_set.COMPL_1'(s(fun(X3,bool),X5)))))) ) | $spl30),
  inference(resolution,[],[f4304,f68])).
fof(f68,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f50])).
fof(f9365,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) ) | $spl236),
  inference(cnf_transformation,[],[f9365_D])).
fof(f9365_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f9298,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) ) | ($spl30 | $spl140)),
  inference(superposition,[],[f4358,f4218])).
fof(f9293,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X3) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl140)),
  inference(superposition,[],[f4218,f4901])).
fof(f9290,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl30 | $spl140)),
  inference(superposition,[],[f4218,f4901])).
fof(f9263,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X0)) ) | $spl232),
  inference(cnf_transformation,[],[f9263_D])).
fof(f9263_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f9230,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X31) | s(bool,X31) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl30 | $spl36)),
  inference(superposition,[],[f4901,f243])).
fof(f9227,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X28) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X28)) ) | ($spl30 | $spl36)),
  inference(superposition,[],[f4901,f243])).
fof(f9199,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X3) | s(bool,X3) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl30 | $spl36)),
  inference(superposition,[],[f243,f4901])).
fof(f9196,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X0) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X0)) ) | ($spl30 | $spl36)),
  inference(superposition,[],[f243,f4901])).
fof(f8854,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl230),
  inference(cnf_transformation,[],[f8854_D])).
fof(f8854_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f8822,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X33) | s(bool,X33) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl30 | $spl38)),
  inference(superposition,[],[f4941,f245])).
fof(f8819,plain,(
  ( ! [X30] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X30) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X30)) ) | ($spl30 | $spl38)),
  inference(superposition,[],[f4941,f245])).
fof(f8792,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X5) | s(bool,X5) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl30 | $spl38)),
  inference(superposition,[],[f245,f4941])).
fof(f8789,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X2) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | ($spl30 | $spl38)),
  inference(superposition,[],[f245,f4941])).
fof(f8768,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl228),
  inference(cnf_transformation,[],[f8768_D])).
fof(f8768_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f8779,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl139)),
  inference(global_subsumption,[],[f8761,f60,f4216,f8748,f8746,f8737,f8756])).
fof(f8756,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,X4) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) ) | $spl139),
  inference(resolution,[],[f4216,f76])).
fof(f8737,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK1),s(sK0,X0))))))) ) | $spl139),
  inference(unit_resulting_resolution,[],[f4216,f70])).
fof(f8761,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X2) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl139)),
  inference(superposition,[],[f4216,f4901])).
fof(f8778,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl139)),
  inference(global_subsumption,[],[f8760,f60,f4216,f8748,f8746,f8737,f8756])).
fof(f8760,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,X1)) ) | ($spl30 | $spl139)),
  inference(superposition,[],[f4216,f4901])).
fof(f8755,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl139)),
  inference(resolution,[],[f4216,f4304])).
fof(f8775,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl33 | $spl139)),
  inference(forward_demodulation,[],[f8739,f4292])).
fof(f8739,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0')))))) ) | ($spl30 | $spl33 | $spl139)),
  inference(unit_resulting_resolution,[],[f141,f4216,f4589])).
fof(f8774,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl139)),
  inference(forward_demodulation,[],[f8773,f4292])).
fof(f8773,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0')))))) ) | ($spl11 | $spl30 | $spl139)),
  inference(forward_demodulation,[],[f8740,f137])).
fof(f8740,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))))))) ) | ($spl11 | $spl30 | $spl139)),
  inference(unit_resulting_resolution,[],[f104,f4216,f4589])).
fof(f8772,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl19 | $spl30 | $spl44 | $spl139)),
  inference(forward_demodulation,[],[f8771,f4292])).
fof(f8771,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0')))))) ) | ($spl19 | $spl30 | $spl44 | $spl139)),
  inference(forward_demodulation,[],[f8741,f362])).
fof(f8741,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))))))) ) | ($spl19 | $spl30 | $spl139)),
  inference(unit_resulting_resolution,[],[f116,f4216,f4589])).
fof(f8742,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl33 | $spl139)),
  inference(unit_resulting_resolution,[],[f141,f4216,f4353])).
fof(f8770,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl11 | $spl30 | $spl139)),
  inference(forward_demodulation,[],[f8743,f137])).
fof(f8743,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | ($spl11 | $spl30 | $spl139)),
  inference(unit_resulting_resolution,[],[f104,f4216,f4353])).
fof(f8769,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl19 | $spl30 | $spl44 | $spl139)),
  inference(forward_demodulation,[],[f8744,f362])).
fof(f8744,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl19 | $spl30 | $spl139)),
  inference(unit_resulting_resolution,[],[f116,f4216,f4353])).
fof(f8745,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl139)),
  inference(unit_resulting_resolution,[],[f4216,f4304])).
fof(f8746,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl139),
  inference(unit_resulting_resolution,[],[f62,f4216,f76])).
fof(f8748,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl139),
  inference(unit_resulting_resolution,[],[f62,f4216,f76])).
fof(f110,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl14),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f8734,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl39)),
  inference(global_subsumption,[],[f8726,f60,f246,f8718,f8713,f8709,f8730])).
fof(f8726,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X6)) ) | ($spl30 | $spl39)),
  inference(superposition,[],[f246,f4941])).
fof(f8733,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl39)),
  inference(global_subsumption,[],[f8725,f60,f246,f8718,f8713,f8709,f8730])).
fof(f8725,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl30 | $spl39)),
  inference(superposition,[],[f246,f4941])).
fof(f8730,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl39),
  inference(trivial_inequality_removal,[],[f8719])).
fof(f8719,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl39),
  inference(superposition,[],[f246,f63])).
fof(f8701,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl33 | $spl39)),
  inference(unit_resulting_resolution,[],[f141,f246,f4901])).
fof(f8704,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl33 | $spl39)),
  inference(unit_resulting_resolution,[],[f141,f246,f4901])).
fof(f8706,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl39)),
  inference(unit_resulting_resolution,[],[f61,f246,f4331])).
fof(f8732,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | $spl226),
  inference(cnf_transformation,[],[f8732_D])).
fof(f8732_D,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f8708,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | ($spl30 | $spl39)),
  inference(unit_resulting_resolution,[],[f246,f4304])).
fof(f8709,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl39),
  inference(unit_resulting_resolution,[],[f246,f63])).
fof(f8710,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl33 | $spl39)),
  inference(unit_resulting_resolution,[],[f141,f246,f4941])).
fof(f8712,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | ($spl30 | $spl33 | $spl39)),
  inference(unit_resulting_resolution,[],[f141,f246,f4353])).
fof(f8713,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | $spl39),
  inference(unit_resulting_resolution,[],[f62,f246,f76])).
fof(f8714,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl33 | $spl39)),
  inference(unit_resulting_resolution,[],[f141,f246,f4941])).
fof(f8716,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl30 | $spl39)),
  inference(unit_resulting_resolution,[],[f62,f246,f4353])).
fof(f8718,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | $spl39),
  inference(unit_resulting_resolution,[],[f62,f246,f76])).
fof(f4216,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | $spl139),
  inference(cnf_transformation,[],[f4216_D])).
fof(f4216_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f8571,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | ($spl64 | $spl83)),
  inference(backward_demodulation,[],[f556,f786])).
fof(f109,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f8589,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl7 | $spl36)),
  inference(backward_demodulation,[],[f243,f100])).
fof(f8588,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) | $spl224),
  inference(cnf_transformation,[],[f8588_D])).
fof(f8588_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f8578,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) | ($spl64 | $spl172)),
  inference(backward_demodulation,[],[f556,f5821])).
fof(f8586,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | $spl222),
  inference(cnf_transformation,[],[f8586_D])).
fof(f8586_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f8577,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | ($spl64 | $spl166)),
  inference(backward_demodulation,[],[f556,f5750])).
fof(f8584,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(sK0,sK2)))))) | $spl220),
  inference(cnf_transformation,[],[f8584_D])).
fof(f8584_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(sK0,sK2)))))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f8576,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(sK0,sK2)))))) | ($spl64 | $spl160)),
  inference(backward_demodulation,[],[f556,f5107])).
fof(f4221,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl145),
  inference(cnf_transformation,[],[f4221_D])).
fof(f4221_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f8574,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | ($spl64 | $spl107)),
  inference(backward_demodulation,[],[f556,f1120])).
fof(f1120,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl107),
  inference(cnf_transformation,[],[f1120_D])).
fof(f1120_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f8573,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl64 | $spl104)),
  inference(backward_demodulation,[],[f556,f1118])).
fof(f8582,plain,(
  ~p(s(bool,'T_0')) | ($spl64 | $spl83 | $spl154)),
  inference(forward_demodulation,[],[f8571,f4415])).
fof(f4231,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl151),
  inference(cnf_transformation,[],[f4231_D])).
fof(f4231_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f8564,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl96 | $spl135)),
  inference(backward_demodulation,[],[f817,f4212])).
fof(f8563,plain,(
  ~p(s(bool,'T_0')) | ($spl139 | $spl154)),
  inference(forward_demodulation,[],[f4216,f4415])).
fof(f4401,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl152),
  inference(cnf_transformation,[],[f4401_D])).
fof(f4401_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f5723,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | ($spl30 | $spl154)),
  inference(superposition,[],[f4334,f4415])).
fof(f8561,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X27),s(fun(sK0,bool),sK1)))))) ) | $spl218),
  inference(cnf_transformation,[],[f8561_D])).
fof(f8561_D,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X27),s(fun(sK0,bool),sK1)))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f8081,plain,(
  ( ! [X27] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X27),s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl154)),
  inference(superposition,[],[f4358,f4415])).
fof(f8560,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))),
  introduced(theory_axiom,[])).
fof(f8559,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f8354,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(bool,X41),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X42),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X41),s(fun(bool,bool),X43)))))))) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | $spl216),
  inference(cnf_transformation,[],[f8354_D])).
fof(f8354_D,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(bool,X41),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X42),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X41),s(fun(bool,bool),X43)))))))) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f8303,plain,(
  ( ! [X47,X45,X48,X46] : (p(s(bool,'const.bool.IN_2'(s(bool,X45),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X46),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X45),s(fun(bool,bool),X47)))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | $spl30),
  inference(superposition,[],[f4773,f4901])).
fof(f4773,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f4306,f63])).
fof(f4306,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2818,f75])).
fof(f8302,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,'const.bool.IN_2'(s(bool,X41),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X42),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X41),s(fun(bool,bool),X43)))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X44) | s(bool,X41) = s(bool,X44)) ) | $spl30),
  inference(superposition,[],[f4773,f4901])).
fof(f8233,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X41),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X42),s(bool,X41))),s(bool,X43)))))) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) | $spl214),
  inference(cnf_transformation,[],[f8233_D])).
fof(f8233_D,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X41),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X42),s(bool,X41))),s(bool,X43)))))) | s(bool,X41) = s(bool,X44) | s(bool,'F_0') = s(bool,X44)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f8145,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X45),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X46),s(bool,X45))),s(bool,X47)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X45) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) ) | $spl30),
  inference(superposition,[],[f4671,f4901])).
fof(f4671,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X1),s(bool,X0))),s(bool,X2)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f4288,f63])).
fof(f4288,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))),s(bool,X1))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2817,f70])).
fof(f8144,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X41),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X42),s(bool,X41))),s(bool,X43)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X44) | s(bool,X41) = s(bool,X44)) ) | $spl30),
  inference(superposition,[],[f4671,f4901])).
fof(f8117,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X32),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) | $spl212),
  inference(cnf_transformation,[],[f8117_D])).
fof(f8117_D,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X32),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f8086,plain,(
  ( ! [X32] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X32),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) | ($spl30 | $spl58)),
  inference(superposition,[],[f4358,f474])).
fof(f8116,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X31),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) | $spl210),
  inference(cnf_transformation,[],[f8116_D])).
fof(f8116_D,plain,(
  ( ! [X31] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X31),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f8085,plain,(
  ( ! [X31] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X31),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))))) ) | ($spl30 | $spl42)),
  inference(superposition,[],[f4358,f360])).
fof(f8115,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X20),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X18),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X19)))))))))) ) | $spl208),
  inference(cnf_transformation,[],[f8115_D])).
fof(f8115_D,plain,(
  ( ! [X19,X20,X18] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X20),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X18),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X19)))))))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f8077,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X20),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X18),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X19)))))))))) ) | $spl30),
  inference(superposition,[],[f4358,f4756])).
fof(f4756,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f61,f4306,f79])).
fof(f8114,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X15),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14))),s(bool,'T_0')))))))) ) | $spl206),
  inference(cnf_transformation,[],[f8114_D])).
fof(f8114_D,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X15),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14))),s(bool,'T_0')))))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f8075,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X15),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X14))),s(bool,'T_0')))))))) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f4358,f4827])).
fof(f4827,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))),s(bool,'T_0')))))) ) | ($spl30 | $spl33)),
  inference(unit_resulting_resolution,[],[f61,f4307,f79])).
fof(f4307,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))),s(bool,'T_0'))))))) ) | ($spl30 | $spl33)),
  inference(unit_resulting_resolution,[],[f141,f2818,f72])).
fof(f8020,plain,(
  ( ! [X30] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X30)))))) ) | $spl204),
  inference(cnf_transformation,[],[f8020_D])).
fof(f8020_D,plain,(
  ( ! [X30] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X30)))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7987,plain,(
  ( ! [X30] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X30)))))) ) | ($spl30 | $spl60)),
  inference(superposition,[],[f4335,f476])).
fof(f8019,plain,(
  ( ! [X29] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X29)))))) ) | $spl202),
  inference(cnf_transformation,[],[f8019_D])).
fof(f8019_D,plain,(
  ( ! [X29] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X29)))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f7986,plain,(
  ( ! [X29] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))),s(sK0,X29)))))) ) | ($spl30 | $spl44)),
  inference(superposition,[],[f4335,f362])).
fof(f8018,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X21),s(bool,'F_0'))))),s(bool,X22)))))) ) | $spl200),
  inference(cnf_transformation,[],[f8018_D])).
fof(f8018_D,plain,(
  ( ! [X21,X22] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X21),s(bool,'F_0'))))),s(bool,X22)))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f7981,plain,(
  ( ! [X21,X22] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X21),s(bool,'F_0'))))),s(bool,X22)))))) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f4335,f4802])).
fof(f4802,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0')))))))) ) | ($spl30 | $spl33)),
  inference(unit_resulting_resolution,[],[f62,f4287,f76])).
fof(f4287,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))))))) ) | ($spl30 | $spl33)),
  inference(unit_resulting_resolution,[],[f141,f2817,f73])).
fof(f8017,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0'))),s(bool,X12))),s(bool,X13)))))) ) | $spl198),
  inference(cnf_transformation,[],[f8017_D])).
fof(f8017_D,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0'))),s(bool,X12))),s(bool,X13)))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f7977,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X11),s(bool,'F_0'))),s(bool,X12))),s(bool,X13)))))) ) | $spl30),
  inference(superposition,[],[f4335,f4664])).
fof(f4664,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))),s(bool,X1)))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f62,f4288,f76])).
fof(f7923,plain,(
  ( ! [X45,X46,X44] : (s(bool,X46) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X44),s(bool,'F_0'))),s(bool,X45))))) | s(bool,'T_0') = s(bool,X46)) ) | $spl190),
  inference(cnf_transformation,[],[f7923_D])).
fof(f7923_D,plain,(
  ( ! [X45,X46,X44] : (s(bool,X46) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X44),s(bool,'F_0'))),s(bool,X45))))) | s(bool,'T_0') = s(bool,X46)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7749,plain,(
  ( ! [X161,X159,X160] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X161) | s(bool,X161) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X159),s(bool,'F_0'))),s(bool,X160)))))) ) | $spl30),
  inference(superposition,[],[f4664,f4941])).
fof(f7925,plain,(
  ( ! [X56] : (s(bool,X56) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X56)) ) | $spl192),
  inference(cnf_transformation,[],[f7925_D])).
fof(f7925_D,plain,(
  ( ! [X56] : (s(bool,X56) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X56)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f7747,plain,(
  ( ! [X155] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X155) | s(bool,X155) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl152)),
  inference(superposition,[],[f4401,f4941])).
fof(f7426,plain,(
  ( ! [X161,X159,X160] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X161) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X159),s(bool,'F_0'))),s(bool,X160))))) | s(bool,'T_0') = s(bool,X161)) ) | $spl30),
  inference(superposition,[],[f4664,f4941])).
fof(f7424,plain,(
  ( ! [X155] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X155) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X155)) ) | ($spl30 | $spl152)),
  inference(superposition,[],[f4401,f4941])).
fof(f7927,plain,(
  ( ! [X58] : (s(bool,X58) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | s(bool,'T_0') = s(bool,X58)) ) | $spl196),
  inference(cnf_transformation,[],[f7927_D])).
fof(f7927_D,plain,(
  ( ! [X58] : (s(bool,X58) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | s(bool,'T_0') = s(bool,X58)) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f7326,plain,(
  ( ! [X58] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X58) | s(bool,X58) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) ) | ($spl30 | $spl60)),
  inference(superposition,[],[f4941,f476])).
fof(f7926,plain,(
  ( ! [X57] : (s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | s(bool,'T_0') = s(bool,X57)) ) | $spl194),
  inference(cnf_transformation,[],[f7926_D])).
fof(f7926_D,plain,(
  ( ! [X57] : (s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | s(bool,'T_0') = s(bool,X57)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7325,plain,(
  ( ! [X57] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X57) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) ) | ($spl30 | $spl44)),
  inference(superposition,[],[f4941,f362])).
fof(f7324,plain,(
  ( ! [X56] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X56) | s(bool,X56) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | ($spl30 | $spl152)),
  inference(superposition,[],[f4941,f4401])).
fof(f7321,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X46) | s(bool,X46) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X44),s(bool,'F_0'))),s(bool,X45)))))) ) | $spl30),
  inference(superposition,[],[f4941,f4664])).
fof(f7922,plain,(
  ( ! [X33] : (s(bool,X33) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X33)) ) | $spl188),
  inference(cnf_transformation,[],[f7922_D])).
fof(f7922_D,plain,(
  ( ! [X33] : (s(bool,X33) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X33)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7317,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X33) | s(bool,X33) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) ) | $spl30),
  inference(superposition,[],[f4941,f137])).
fof(f7921,plain,(
  ( ! [X32] : (s(bool,X32) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X32)) ) | $spl186),
  inference(cnf_transformation,[],[f7921_D])).
fof(f7921_D,plain,(
  ( ! [X32] : (s(bool,X32) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X32)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7316,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X32) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl16 | $spl30)),
  inference(superposition,[],[f4941,f111])).
fof(f111,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl16),
  inference(cnf_transformation,[],[f111_D])).
fof(f111_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f7227,plain,(
  ( ! [X58] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X58) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | s(bool,'T_0') = s(bool,X58)) ) | ($spl30 | $spl60)),
  inference(superposition,[],[f4941,f476])).
fof(f7226,plain,(
  ( ! [X57] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X57) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | s(bool,'T_0') = s(bool,X57)) ) | ($spl30 | $spl44)),
  inference(superposition,[],[f4941,f362])).
fof(f7225,plain,(
  ( ! [X56] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X56) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X56)) ) | ($spl30 | $spl152)),
  inference(superposition,[],[f4941,f4401])).
fof(f7222,plain,(
  ( ! [X45,X46,X44] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X46) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X44),s(bool,'F_0'))),s(bool,X45))))) | s(bool,'T_0') = s(bool,X46)) ) | $spl30),
  inference(superposition,[],[f4941,f4664])).
fof(f7218,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X33) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X33)) ) | $spl30),
  inference(superposition,[],[f4941,f137])).
fof(f7217,plain,(
  ( ! [X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,X32)) ) | ($spl16 | $spl30)),
  inference(superposition,[],[f4941,f111])).
fof(f6723,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3))),s(bool,'T_0')))))))) ) | $spl184),
  inference(cnf_transformation,[],[f6723_D])).
fof(f6723_D,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3))),s(bool,'T_0')))))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f6672,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X3))),s(bool,'T_0')))))))) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f4334,f4827])).
fof(f6625,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'F_0')))))))))) ) | $spl182),
  inference(cnf_transformation,[],[f6625_D])).
fof(f6625_D,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'F_0')))))))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f6573,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'F_0')))))))))) ) | ($spl30 | $spl33)),
  inference(superposition,[],[f4356,f4802])).
fof(f6543,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28))))))) | s(bool,'F_0') = s(bool,X29)) ) | $spl174),
  inference(cnf_transformation,[],[f6543_D])).
fof(f6543_D,plain,(
  ( ! [X28,X29,X27] : (s(bool,X29) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28))))))) | s(bool,'F_0') = s(bool,X29)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f6406,plain,(
  ( ! [X127,X125,X126] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X127) | s(bool,X127) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X125),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X126)))))))) ) | $spl30),
  inference(superposition,[],[f4756,f4901])).
fof(f6544,plain,(
  ( ! [X33] : (s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,X33)) ) | $spl176),
  inference(cnf_transformation,[],[f6544_D])).
fof(f6544_D,plain,(
  ( ! [X33] : (s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,X33)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f6404,plain,(
  ( ! [X119] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X119) | s(bool,X119) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl154)),
  inference(superposition,[],[f4415,f4901])).
fof(f6546,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X64),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X64),s(fun(bool,bool),X65)))))))) | s(bool,'F_0') = s(bool,X66) | s(bool,X64) = s(bool,X66)) ) | $spl180),
  inference(cnf_transformation,[],[f6546_D])).
fof(f6546_D,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X64),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X64),s(fun(bool,bool),X65)))))))) | s(bool,'F_0') = s(bool,X66) | s(bool,X64) = s(bool,X66)) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f6340,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X64),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X64),s(fun(bool,bool),X65)))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X66) | s(bool,X64) = s(bool,X66)) ) | $spl30),
  inference(superposition,[],[f4572,f4901])).
fof(f4572,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f4308,f63])).
fof(f4308,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2818,f68])).
fof(f6545,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X62),s(bool,X61)))))))) | s(bool,'F_0') = s(bool,X63) | s(bool,X61) = s(bool,X63)) ) | $spl178),
  inference(cnf_transformation,[],[f6545_D])).
fof(f6545_D,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X62),s(bool,X61)))))))) | s(bool,'F_0') = s(bool,X63) | s(bool,X61) = s(bool,X63)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f6339,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X62),s(bool,X61)))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X63) | s(bool,X61) = s(bool,X63)) ) | $spl30),
  inference(superposition,[],[f4549,f4901])).
fof(f4549,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X1),s(bool,X0)))))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl30),
  inference(superposition,[],[f4289,f63])).
fof(f4289,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))))))) ) | $spl30),
  inference(unit_resulting_resolution,[],[f2817,f69])).
fof(f6123,plain,(
  ( ! [X127,X125,X126] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X127) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X125),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X126))))))) | s(bool,'F_0') = s(bool,X127)) ) | $spl30),
  inference(superposition,[],[f4756,f4901])).
fof(f6121,plain,(
  ( ! [X119] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X119) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,X119)) ) | ($spl30 | $spl154)),
  inference(superposition,[],[f4415,f4901])).
fof(f6057,plain,(
  ( ! [X66,X64,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,X64),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X64),s(fun(bool,bool),X65)))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X64) = s(bool,X66) | s(bool,'F_0') = s(bool,X66)) ) | $spl30),
  inference(superposition,[],[f4572,f4901])).
fof(f6056,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X62),s(bool,X61)))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X61) = s(bool,X63) | s(bool,'F_0') = s(bool,X63)) ) | $spl30),
  inference(superposition,[],[f4549,f4901])).
fof(f6032,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X33) | s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) ) | ($spl30 | $spl154)),
  inference(superposition,[],[f4901,f4415])).
fof(f6030,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X29) | s(bool,X29) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28)))))))) ) | $spl30),
  inference(superposition,[],[f4901,f4756])).
fof(f5958,plain,(
  ( ! [X33] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X33) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,X33)) ) | ($spl30 | $spl154)),
  inference(superposition,[],[f4901,f4415])).
fof(f5956,plain,(
  ( ! [X28,X29,X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X28))))))) | s(bool,'F_0') = s(bool,X29)) ) | $spl30),
  inference(superposition,[],[f4901,f4756])).
fof(f5821,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) | $spl172),
  inference(cnf_transformation,[],[f5821_D])).
fof(f5821_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f5795,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) | ($spl30 | $spl60)),
  inference(superposition,[],[f4356,f476])).
fof(f5819,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) | $spl170),
  inference(cnf_transformation,[],[f5819_D])).
fof(f5819_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f5794,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) | ($spl30 | $spl44)),
  inference(superposition,[],[f4356,f362])).
fof(f5817,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X7),s(bool,'F_0'))),s(bool,X8)))))))) ) | $spl168),
  inference(cnf_transformation,[],[f5817_D])).
fof(f5817_D,plain,(
  ( ! [X8,X7] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X7),s(bool,'F_0'))),s(bool,X8)))))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f5787,plain,(
  ( ! [X8,X7] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X7),s(bool,'F_0'))),s(bool,X8)))))))) ) | $spl30),
  inference(superposition,[],[f4356,f4664])).
fof(f5750,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | $spl166),
  inference(cnf_transformation,[],[f5750_D])).
fof(f5750_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f5728,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | ($spl30 | $spl58)),
  inference(superposition,[],[f4334,f474])).
fof(f5748,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | $spl164),
  inference(cnf_transformation,[],[f5748_D])).
fof(f5748_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f5727,plain,(
  ~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))) | ($spl30 | $spl42)),
  inference(superposition,[],[f4334,f360])).
fof(f5746,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))))))))) ) | $spl162),
  inference(cnf_transformation,[],[f5746_D])).
fof(f5746_D,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))))))))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f5720,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.COMPL_1'(s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X11)))))))))) ) | $spl30),
  inference(superposition,[],[f4334,f4756])).
fof(f5103,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | $spl156),
  inference(cnf_transformation,[],[f5103_D])).
fof(f5103_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f5110,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | $spl136),
  inference(global_subsumption,[],[f60,f4213,f5095,f5093,f5090,f5088,f5100,f5101])).
fof(f5101,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | $spl136),
  inference(superposition,[],[f4213,f63])).
fof(f5100,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) ) | $spl136),
  inference(resolution,[],[f4213,f79])).
fof(f5088,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))))) ) | $spl136),
  inference(unit_resulting_resolution,[],[f4213,f75])).
fof(f5099,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | ($spl30 | $spl136)),
  inference(resolution,[],[f4213,f4321])).
fof(f5107,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0))))))) | $spl160),
  inference(cnf_transformation,[],[f5107_D])).
fof(f5107_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0))))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f5089,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))),s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0))))))) | ($spl65 | $spl136)),
  inference(unit_resulting_resolution,[],[f555,f4213,f72])).
fof(f5105,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))))) | $spl159),
  inference(cnf_transformation,[],[f5105_D])).
fof(f5105_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f5090,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))))) | $spl136),
  inference(unit_resulting_resolution,[],[f4213,f68])).
fof(f5091,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | ($spl30 | $spl33 | $spl136)),
  inference(unit_resulting_resolution,[],[f141,f4213,f4331])).
fof(f5092,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | ($spl30 | $spl136)),
  inference(unit_resulting_resolution,[],[f4213,f4321])).
fof(f5093,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | $spl136),
  inference(unit_resulting_resolution,[],[f61,f4213,f79])).
fof(f5095,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))) | $spl136),
  inference(unit_resulting_resolution,[],[f61,f4213,f79])).
fof(f4435,plain,(
  $false | $spl93),
  inference(trivial_inequality_removal,[],[f4420])).
fof(f4420,plain,(
  s(sK0,sK2) != s(sK0,sK2) | $spl93),
  inference(unit_resulting_resolution,[],[f812,f74])).
fof(f4415,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl154),
  inference(cnf_transformation,[],[f4415_D])).
fof(f4415_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f4417,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl138),
  inference(global_subsumption,[],[f60,f4215,f4409,f4407,f4404,f4412,f4413])).
fof(f4413,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl138),
  inference(superposition,[],[f4215,f63])).
fof(f4412,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) ) | $spl138),
  inference(resolution,[],[f4215,f79])).
fof(f4404,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK1))))))) ) | $spl138),
  inference(unit_resulting_resolution,[],[f4215,f75])).
fof(f4411,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl138)),
  inference(resolution,[],[f4215,f4321])).
fof(f4406,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | ($spl30 | $spl138)),
  inference(unit_resulting_resolution,[],[f4215,f4321])).
fof(f4407,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl138),
  inference(unit_resulting_resolution,[],[f61,f4215,f79])).
fof(f4409,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))) | $spl138),
  inference(unit_resulting_resolution,[],[f61,f4215,f79])).
fof(f4396,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | ($spl30 | $spl91)),
  inference(resolution,[],[f809,f4304])).
fof(f809,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl91),
  inference(cnf_transformation,[],[f809_D])).
fof(f809_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f4215,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | $spl138),
  inference(cnf_transformation,[],[f4215_D])).
fof(f4215_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f4395,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | $spl91),
  inference(resolution,[],[f809,f69])).
fof(f4386,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | $spl91),
  inference(unit_resulting_resolution,[],[f809,f69])).
fof(f4213,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl136),
  inference(cnf_transformation,[],[f4213_D])).
fof(f4213_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f4388,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl91),
  inference(unit_resulting_resolution,[],[f809,f69])).
fof(f4389,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | ($spl30 | $spl91)),
  inference(unit_resulting_resolution,[],[f809,f4304])).
fof(f4390,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl91),
  inference(unit_resulting_resolution,[],[f62,f809,f76])).
fof(f4392,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl91),
  inference(unit_resulting_resolution,[],[f62,f809,f76])).
fof(f4218,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl140),
  inference(cnf_transformation,[],[f4218_D])).
fof(f4218_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f4385,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl90),
  inference(global_subsumption,[],[f60,f810,f4379,f4377,f4375,f4374,f4373,f4381,f4383,f4384])).
fof(f4384,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl90),
  inference(superposition,[],[f810,f63])).
fof(f4383,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl90),
  inference(resolution,[],[f810,f79])).
fof(f4374,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) ) | $spl90),
  inference(unit_resulting_resolution,[],[f810,f75])).
fof(f4382,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | ($spl30 | $spl90)),
  inference(resolution,[],[f810,f4321])).
fof(f4381,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | $spl90),
  inference(resolution,[],[f810,f68])).
fof(f4373,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | $spl90),
  inference(unit_resulting_resolution,[],[f810,f68])).
fof(f4214,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl137),
  inference(cnf_transformation,[],[f4214_D])).
fof(f4214_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f4375,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl90),
  inference(unit_resulting_resolution,[],[f810,f68])).
fof(f4376,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | ($spl30 | $spl90)),
  inference(unit_resulting_resolution,[],[f810,f4321])).
fof(f4377,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f61,f810,f79])).
fof(f4379,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl90),
  inference(unit_resulting_resolution,[],[f61,f810,f79])).
fof(f126,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl25),
  inference(cnf_transformation,[],[f126_D])).
fof(f126_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f4238,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl9 | $spl60)),
  inference(backward_demodulation,[],[f476,f102])).
fof(f4234,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | p(s(bool,'F_0')) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f805,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl86),
  inference(cnf_transformation,[],[f805_D])).
fof(f805_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f4233,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl22 | $spl64)),
  inference(forward_demodulation,[],[f124,f556])).
fof(f4232,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl150),
  inference(cnf_transformation,[],[f4232_D])).
fof(f4232_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f4230,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl24 | $spl64)),
  inference(forward_demodulation,[],[f125,f556])).
fof(f125,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl24),
  inference(cnf_transformation,[],[f125_D])).
fof(f125_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4229,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl149),
  inference(cnf_transformation,[],[f4229_D])).
fof(f4229_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f4227,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | ($spl57 | $spl64)),
  inference(forward_demodulation,[],[f471,f556])).
fof(f471,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl57),
  inference(cnf_transformation,[],[f471_D])).
fof(f471_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f807,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl89),
  inference(cnf_transformation,[],[f807_D])).
fof(f807_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f4226,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl59 | $spl64)),
  inference(forward_demodulation,[],[f473,f556])).
fof(f816,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl97),
  inference(cnf_transformation,[],[f816_D])).
fof(f816_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f4225,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl61 | $spl64)),
  inference(forward_demodulation,[],[f477,f556])).
fof(f4224,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl146),
  inference(cnf_transformation,[],[f4224_D])).
fof(f4224_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f4210,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl64 | $spl108)),
  inference(backward_demodulation,[],[f556,f1268])).
fof(f4222,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl144),
  inference(cnf_transformation,[],[f4222_D])).
fof(f4222_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f4209,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | ($spl64 | $spl106)),
  inference(backward_demodulation,[],[f556,f1121])).
fof(f4220,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl143),
  inference(cnf_transformation,[],[f4220_D])).
fof(f4220_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f4208,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl64 | $spl105)),
  inference(backward_demodulation,[],[f556,f1117])).
fof(f1117,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl105),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f4207,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | ($spl64 | $spl98)),
  inference(backward_demodulation,[],[f556,f887])).
fof(f813,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl94),
  inference(cnf_transformation,[],[f813_D])).
fof(f813_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4206,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | ($spl64 | $spl84)),
  inference(backward_demodulation,[],[f556,f797])).
fof(f4205,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))) | ($spl64 | $spl83)),
  inference(backward_demodulation,[],[f556,f786])).
fof(f4204,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | ($spl64 | $spl81)),
  inference(backward_demodulation,[],[f556,f784])).
fof(f4212,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl135),
  inference(cnf_transformation,[],[f4212_D])).
fof(f4212_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f4203,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl9 | $spl64)),
  inference(backward_demodulation,[],[f556,f102])).
fof(f4201,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),
  introduced(theory_axiom,[])).
fof(f4200,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,'F_0')) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f2773,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X15)) ) | $spl118),
  inference(cnf_transformation,[],[f2773_D])).
fof(f2773_D,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X15)) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4077,plain,(
  ( ! [X89] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X89) | s(bool,X89) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl58 | $spl60)),
  inference(superposition,[],[f474,f2984])).
fof(f2984,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl60),
  inference(resolution,[],[f2880,f2853])).
fof(f2853,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl60),
  inference(forward_demodulation,[],[f2845,f2841])).
fof(f2841,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0')))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f62,f1307,f76])).
fof(f1307,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))))) ) | $spl60),
  inference(forward_demodulation,[],[f1278,f476])).
fof(f1278,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f476,f71])).
fof(f2845,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'F_0')))))) ) | $spl60),
  inference(resolution,[],[f1307,f76])).
fof(f2880,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl60),
  inference(resolution,[],[f2853,f79])).
fof(f2772,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X14)) ) | $spl116),
  inference(cnf_transformation,[],[f2772_D])).
fof(f2772_D,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X14)) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4076,plain,(
  ( ! [X88] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X88) | s(bool,X88) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl42 | $spl60)),
  inference(superposition,[],[f360,f2984])).
fof(f2771,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X13)) ) | $spl114),
  inference(cnf_transformation,[],[f2771_D])).
fof(f2771_D,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4064,plain,(
  ( ! [X48] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X48) | s(bool,X48) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl14 | $spl60)),
  inference(superposition,[],[f110,f2984])).
fof(f4195,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X12)) ) | $spl132),
  inference(cnf_transformation,[],[f4195_D])).
fof(f4195_D,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f4063,plain,(
  ( ! [X47] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X47) | s(bool,X47) = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) ) | ($spl28 | $spl60)),
  inference(superposition,[],[f135,f2984])).
fof(f2776,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X32),s(fun(bool,bool),X33)))))) | s(bool,'F_0') = s(bool,X34) | s(bool,X32) = s(bool,X34)) ) | $spl124),
  inference(cnf_transformation,[],[f2776_D])).
fof(f2776_D,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X32),s(fun(bool,bool),X33)))))) | s(bool,'F_0') = s(bool,X34) | s(bool,X32) = s(bool,X34)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f4023,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(bool,X36),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X36),s(fun(bool,bool),X37)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X38) | s(bool,X36) = s(bool,X38)) ) | $spl60),
  inference(superposition,[],[f2864,f2984])).
fof(f2864,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl60),
  inference(superposition,[],[f1309,f63])).
fof(f1309,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl60),
  inference(forward_demodulation,[],[f1277,f476])).
fof(f1277,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))),s(fun(bool,bool),X0))))))) ) | $spl60),
  inference(unit_resulting_resolution,[],[f476,f74])).
fof(f2775,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,X29)))))) | s(bool,'F_0') = s(bool,X31) | s(bool,X29) = s(bool,X31)) ) | $spl122),
  inference(cnf_transformation,[],[f2775_D])).
fof(f2775_D,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,X29)))))) | s(bool,'F_0') = s(bool,X31) | s(bool,X29) = s(bool,X31)) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f4021,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X33),s(bool,X32)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X34) | s(bool,X32) = s(bool,X34)) ) | $spl60),
  inference(superposition,[],[f2847,f2984])).
fof(f2847,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X1),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl60),
  inference(superposition,[],[f1307,f63])).
fof(f3854,plain,(
  ( ! [X89] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X89) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X89)) ) | ($spl58 | $spl60)),
  inference(superposition,[],[f474,f2984])).
fof(f3853,plain,(
  ( ! [X88] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X88) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X88)) ) | ($spl42 | $spl60)),
  inference(superposition,[],[f360,f2984])).
fof(f3841,plain,(
  ( ! [X48] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X48) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X48)) ) | ($spl14 | $spl60)),
  inference(superposition,[],[f110,f2984])).
fof(f3840,plain,(
  ( ! [X47] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X47) = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X47)) ) | ($spl28 | $spl60)),
  inference(superposition,[],[f135,f2984])).
fof(f3800,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(bool,X36),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X36),s(fun(bool,bool),X37)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X36) = s(bool,X38) | s(bool,'F_0') = s(bool,X38)) ) | $spl60),
  inference(superposition,[],[f2864,f2984])).
fof(f3798,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X33),s(bool,X32)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | $spl60),
  inference(superposition,[],[f2847,f2984])).
fof(f3782,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X15) | s(bool,X15) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl58 | $spl60)),
  inference(superposition,[],[f2984,f474])).
fof(f3781,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X14) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl42 | $spl60)),
  inference(superposition,[],[f2984,f360])).
fof(f3780,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl14 | $spl60)),
  inference(superposition,[],[f2984,f110])).
fof(f3779,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X12) | s(bool,X12) = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) ) | ($spl28 | $spl60)),
  inference(superposition,[],[f2984,f135])).
fof(f3720,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X15)) ) | ($spl58 | $spl60)),
  inference(superposition,[],[f2984,f474])).
fof(f3719,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X14)) ) | ($spl42 | $spl60)),
  inference(superposition,[],[f2984,f360])).
fof(f3718,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X13) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl14 | $spl60)),
  inference(superposition,[],[f2984,f110])).
fof(f3717,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X12) = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X12)) ) | ($spl28 | $spl60)),
  inference(superposition,[],[f2984,f135])).
fof(f105,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl10),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2790,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl31),
  inference(trivial_inequality_removal,[],[f2789])).
fof(f2789,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl31),
  inference(superposition,[],[f138,f63])).
fof(f2786,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl31),
  inference(unit_resulting_resolution,[],[f138,f63])).
fof(f2792,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | $spl130),
  inference(cnf_transformation,[],[f2792_D])).
fof(f2792_D,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f2787,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | $spl31),
  inference(unit_resulting_resolution,[],[f62,f138,f76])).
fof(f2788,plain,(
  p(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | $spl31),
  inference(unit_resulting_resolution,[],[f62,f138,f76])).
fof(f104,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f104_D])).
fof(f104_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f2783,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | ($spl3 | $spl28)),
  inference(backward_demodulation,[],[f135,f96])).
fof(f2774,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X17))))) | s(bool,'F_0') = s(bool,X18)) ) | $spl120),
  inference(cnf_transformation,[],[f2774_D])).
fof(f2774_D,plain,(
  ( ! [X17,X18] : (s(bool,X18) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X17))))) | s(bool,'F_0') = s(bool,X18)) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f2659,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X82) | s(bool,X82) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X81)))))) ) | $spl10),
  inference(superposition,[],[f595,f1554])).
fof(f1554,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f893,f263])).
fof(f263,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) ) | $spl10),
  inference(forward_demodulation,[],[f257,f253])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0')))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f62,f174,f76])).
fof(f174,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0'))))))) ) | $spl10),
  inference(forward_demodulation,[],[f153,f105])).
fof(f153,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'T_0'))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f105,f71])).
fof(f257,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X2),s(bool,'T_0')))))) ) | $spl10),
  inference(resolution,[],[f174,f76])).
fof(f893,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl10),
  inference(resolution,[],[f263,f79])).
fof(f595,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0)))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f61,f175,f79])).
fof(f175,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))))) ) | $spl10),
  inference(forward_demodulation,[],[f152,f105])).
fof(f152,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X0))))))) ) | $spl10),
  inference(unit_resulting_resolution,[],[f105,f74])).
fof(f2658,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X79) | s(bool,X79) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl10 | $spl58)),
  inference(superposition,[],[f474,f1554])).
fof(f2657,plain,(
  ( ! [X78] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X78) | s(bool,X78) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl10 | $spl42)),
  inference(superposition,[],[f360,f1554])).
fof(f2646,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X39) | s(bool,X39) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f110,f1554])).
fof(f2770,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X12)) ) | $spl112),
  inference(cnf_transformation,[],[f2770_D])).
fof(f2770_D,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f2645,plain,(
  ( ! [X38] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X38) | s(bool,X38) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) ) | $spl10),
  inference(superposition,[],[f105,f1554])).
fof(f2611,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X32),s(fun(bool,bool),X33)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X34) | s(bool,X32) = s(bool,X34)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f632,f1554])).
fof(f632,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f216,f63])).
fof(f216,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl12),
  inference(forward_demodulation,[],[f196,f106])).
fof(f106,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl12),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f196,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X0))))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f106,f74])).
fof(f2610,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,X29)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X31) | s(bool,X29) = s(bool,X31)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f615,f1554])).
fof(f615,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X1),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f215,f63])).
fof(f215,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))))) ) | $spl12),
  inference(forward_demodulation,[],[f197,f106])).
fof(f197,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X0),s(bool,'F_0'))))))) ) | $spl12),
  inference(unit_resulting_resolution,[],[f106,f71])).
fof(f2780,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X35),s(fun(bool,bool),X36))))) | s(bool,'F_0') = s(bool,X35)) ) | $spl128),
  inference(cnf_transformation,[],[f2780_D])).
fof(f2780_D,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X35),s(fun(bool,bool),X36))))) | s(bool,'F_0') = s(bool,X35)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f2549,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X35),s(fun(bool,bool),X36))))) | s(bool,'F_0') = s(bool,X35) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl10),
  inference(superposition,[],[f595,f1554])).
fof(f2779,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X34),s(bool,X33))))) | s(bool,'F_0') = s(bool,X33)) ) | $spl126),
  inference(cnf_transformation,[],[f2779_D])).
fof(f2779_D,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X34),s(bool,X33))))) | s(bool,'F_0') = s(bool,X33)) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f2548,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X34),s(bool,X33))))) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl10),
  inference(superposition,[],[f253,f1554])).
fof(f2487,plain,(
  ( ! [X35,X36] : (s(bool,X35) = s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X35),s(fun(bool,bool),X36))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X35)) ) | $spl10),
  inference(superposition,[],[f595,f1554])).
fof(f2486,plain,(
  ( ! [X33,X34] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X34),s(bool,X33))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X33)) ) | $spl10),
  inference(superposition,[],[f253,f1554])).
fof(f2454,plain,(
  ( ! [X81,X82] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X82) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X81))))) | s(bool,'F_0') = s(bool,X82)) ) | $spl10),
  inference(superposition,[],[f595,f1554])).
fof(f2453,plain,(
  ( ! [X79] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X79) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X79)) ) | ($spl10 | $spl58)),
  inference(superposition,[],[f474,f1554])).
fof(f2452,plain,(
  ( ! [X78] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X78) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X78)) ) | ($spl10 | $spl42)),
  inference(superposition,[],[f360,f1554])).
fof(f2441,plain,(
  ( ! [X39] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X39) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X39)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f110,f1554])).
fof(f2440,plain,(
  ( ! [X38] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X38) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X38)) ) | $spl10),
  inference(superposition,[],[f105,f1554])).
fof(f2406,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(bool,X32),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X32),s(fun(bool,bool),X33)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f632,f1554])).
fof(f2405,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(bool,X29),s(fun(bool,bool),'const.pred_set.DELETE_2'(s(fun(bool,bool),X30),s(bool,X29)))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X29) = s(bool,X31) | s(bool,'F_0') = s(bool,X31)) ) | ($spl10 | $spl12)),
  inference(superposition,[],[f615,f1554])).
fof(f2392,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X18) | s(bool,X18) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X17)))))) ) | $spl10),
  inference(superposition,[],[f1554,f595])).
fof(f2391,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X15) | s(bool,X15) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl10 | $spl58)),
  inference(superposition,[],[f1554,f474])).
fof(f2390,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X14) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) ) | ($spl10 | $spl42)),
  inference(superposition,[],[f1554,f360])).
fof(f2389,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X13) | s(bool,X13) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f1554,f110])).
fof(f2388,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X12) | s(bool,X12) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) ) | $spl10),
  inference(superposition,[],[f1554,f105])).
fof(f2336,plain,(
  ( ! [X17,X18] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X18) = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X17))))) | s(bool,'F_0') = s(bool,X18)) ) | $spl10),
  inference(superposition,[],[f1554,f595])).
fof(f2335,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X15)) ) | ($spl10 | $spl58)),
  inference(superposition,[],[f1554,f474])).
fof(f2334,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'F_0') = s(bool,X14)) ) | ($spl10 | $spl42)),
  inference(superposition,[],[f1554,f360])).
fof(f2333,plain,(
  ( ! [X13] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X13) = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'F_0') = s(bool,X13)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f1554,f110])).
fof(f2332,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X12) = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'F_0') = s(bool,X12)) ) | $spl10),
  inference(superposition,[],[f1554,f105])).
fof(f553,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl62),
  inference(cnf_transformation,[],[f553_D])).
fof(f553_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1396,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl58),
  inference(superposition,[],[f70,f474])).
fof(f1426,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl58 | $spl78)),
  inference(forward_demodulation,[],[f1396,f783])).
fof(f444,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl50),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1343,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl42),
  inference(superposition,[],[f70,f360])).
fof(f1373,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'T_0')) | ($spl42 | $spl68)),
  inference(forward_demodulation,[],[f1343,f666])).
fof(f1274,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl9 | $spl60)),
  inference(backward_demodulation,[],[f476,f102])).
fof(f1268,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl108),
  inference(cnf_transformation,[],[f1268_D])).
fof(f1268_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1263,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl67)),
  inference(resolution,[],[f589,f263])).
fof(f786,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl83),
  inference(cnf_transformation,[],[f786_D])).
fof(f786_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f1262,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl67),
  inference(resolution,[],[f589,f75])).
fof(f1252,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl67),
  inference(unit_resulting_resolution,[],[f589,f75])).
fof(f1271,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl111),
  inference(cnf_transformation,[],[f1271_D])).
fof(f1271_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f1253,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | ($spl65 | $spl67)),
  inference(unit_resulting_resolution,[],[f555,f589,f73])).
fof(f1256,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl67)),
  inference(unit_resulting_resolution,[],[f589,f263])).
fof(f1257,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl67),
  inference(unit_resulting_resolution,[],[f62,f589,f76])).
fof(f1259,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl67),
  inference(unit_resulting_resolution,[],[f62,f589,f76])).
fof(f1118,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl104),
  inference(cnf_transformation,[],[f1118_D])).
fof(f1118_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1122,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl66),
  inference(global_subsumption,[],[f60,f588,f1112,f1110,f1107,f1115,f1116])).
fof(f1116,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl66),
  inference(superposition,[],[f588,f63])).
fof(f1115,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) ) | $spl66),
  inference(resolution,[],[f588,f79])).
fof(f1107,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl66),
  inference(unit_resulting_resolution,[],[f588,f75])).
fof(f1121,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl106),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f1108,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | ($spl65 | $spl66)),
  inference(unit_resulting_resolution,[],[f555,f588,f72])).
fof(f1110,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl66),
  inference(unit_resulting_resolution,[],[f61,f588,f79])).
fof(f1112,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl66),
  inference(unit_resulting_resolution,[],[f61,f588,f79])).
fof(f119,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl21),
  inference(cnf_transformation,[],[f119_D])).
fof(f119_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f1045,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl5 | $spl44)),
  inference(backward_demodulation,[],[f362,f98])).
fof(f1042,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))),
  introduced(theory_axiom,[])).
fof(f1040,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl102),
  inference(cnf_transformation,[],[f1040_D])).
fof(f1040_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1035,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl47)),
  inference(resolution,[],[f399,f263])).
fof(f669,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl73),
  inference(cnf_transformation,[],[f669_D])).
fof(f669_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f1034,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(resolution,[],[f399,f75])).
fof(f446,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)) | $spl53),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f1024,plain,(
  s(sK0,sK2) != s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)) | $spl47),
  inference(unit_resulting_resolution,[],[f399,f74])).
fof(f1025,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(unit_resulting_resolution,[],[f399,f75])).
fof(f1028,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | ($spl10 | $spl47)),
  inference(unit_resulting_resolution,[],[f399,f263])).
fof(f1029,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl47),
  inference(unit_resulting_resolution,[],[f62,f399,f76])).
fof(f1031,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl47),
  inference(unit_resulting_resolution,[],[f62,f399,f76])).
fof(f920,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))),
  introduced(theory_axiom,[])).
fof(f917,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl100),
  inference(cnf_transformation,[],[f917_D])).
fof(f917_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f919,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl46),
  inference(global_subsumption,[],[f60,f398,f911,f909,f907,f914,f915])).
fof(f915,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl46),
  inference(superposition,[],[f398,f63])).
fof(f914,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) ) | $spl46),
  inference(resolution,[],[f398,f79])).
fof(f907,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))) ) | $spl46),
  inference(unit_resulting_resolution,[],[f398,f75])).
fof(f909,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl46),
  inference(unit_resulting_resolution,[],[f61,f398,f79])).
fof(f911,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) | $spl46),
  inference(unit_resulting_resolution,[],[f61,f398,f79])).
fof(f887,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl98),
  inference(cnf_transformation,[],[f887_D])).
fof(f887_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f889,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl62),
  inference(global_subsumption,[],[f60,f553,f881,f879,f878,f876,f875,f883,f884,f885])).
fof(f885,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl62),
  inference(superposition,[],[f553,f63])).
fof(f884,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl62),
  inference(resolution,[],[f553,f79])).
fof(f876,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f553,f75])).
fof(f883,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl62),
  inference(resolution,[],[f553,f68])).
fof(f875,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl62),
  inference(unit_resulting_resolution,[],[f553,f68])).
fof(f784,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl81),
  inference(cnf_transformation,[],[f784_D])).
fof(f784_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f878,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl62),
  inference(unit_resulting_resolution,[],[f553,f68])).
fof(f879,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f61,f553,f79])).
fof(f881,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f61,f553,f79])).
fof(f588,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl66),
  inference(cnf_transformation,[],[f588_D])).
fof(f588_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f825,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl60),
  inference(superposition,[],[f69,f476])).
fof(f817,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl96),
  inference(cnf_transformation,[],[f817_D])).
fof(f817_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f815,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl60 | $spl64)),
  inference(forward_demodulation,[],[f476,f556])).
fof(f814,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl95),
  inference(cnf_transformation,[],[f814_D])).
fof(f814_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f804,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | ($spl64 | $spl85)),
  inference(backward_demodulation,[],[f556,f796])).
fof(f796,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl85),
  inference(cnf_transformation,[],[f796_D])).
fof(f796_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f812,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl93),
  inference(cnf_transformation,[],[f812_D])).
fof(f812_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) <=> ~$spl93),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl93])])).
fof(f803,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | ($spl64 | $spl67)),
  inference(backward_demodulation,[],[f556,f589])).
fof(f810,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl90),
  inference(cnf_transformation,[],[f810_D])).
fof(f810_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f802,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | ($spl62 | $spl64)),
  inference(backward_demodulation,[],[f556,f553])).
fof(f808,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl88),
  inference(cnf_transformation,[],[f808_D])).
fof(f808_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f801,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | ($spl58 | $spl64)),
  inference(backward_demodulation,[],[f556,f474])).
fof(f806,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl87),
  inference(cnf_transformation,[],[f806_D])).
fof(f806_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f800,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl23 | $spl64)),
  inference(backward_demodulation,[],[f556,f123])).
fof(f799,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))))),
  introduced(theory_axiom,[])).
fof(f798,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))))),
  introduced(theory_axiom,[])).
fof(f124,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl22),
  inference(cnf_transformation,[],[f124_D])).
fof(f124_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f795,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl61),
  inference(trivial_inequality_removal,[],[f794])).
fof(f794,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl61),
  inference(superposition,[],[f477,f63])).
fof(f791,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f477,f63])).
fof(f797,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl84),
  inference(cnf_transformation,[],[f797_D])).
fof(f797_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f792,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f62,f477,f76])).
fof(f793,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f62,f477,f76])).
fof(f123,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl23),
  inference(cnf_transformation,[],[f123_D])).
fof(f123_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f788,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl9 | $spl58)),
  inference(backward_demodulation,[],[f474,f102])).
fof(f787,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl82),
  inference(cnf_transformation,[],[f787_D])).
fof(f787_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f778,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl63),
  inference(resolution,[],[f554,f69])).
fof(f770,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),sK1)))) | $spl63),
  inference(unit_resulting_resolution,[],[f554,f69])).
fof(f785,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl80),
  inference(cnf_transformation,[],[f785_D])).
fof(f785_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f772,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl63),
  inference(unit_resulting_resolution,[],[f554,f69])).
fof(f783,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl78),
  inference(cnf_transformation,[],[f783_D])).
fof(f783_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f773,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f62,f554,f76])).
fof(f775,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f62,f554,f76])).
fof(f767,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl76),
  inference(cnf_transformation,[],[f767_D])).
fof(f767_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f768,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl50),
  inference(global_subsumption,[],[f60,f444,f761,f759,f758,f757,f756,f763,f764,f765])).
fof(f765,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl50),
  inference(superposition,[],[f444,f63])).
fof(f764,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) ) | $spl50),
  inference(resolution,[],[f444,f79])).
fof(f757,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))))))) ) | $spl50),
  inference(unit_resulting_resolution,[],[f444,f75])).
fof(f763,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl50),
  inference(resolution,[],[f444,f68])).
fof(f756,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl50),
  inference(unit_resulting_resolution,[],[f444,f68])).
fof(f667,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl71),
  inference(cnf_transformation,[],[f667_D])).
fof(f667_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f758,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl50),
  inference(unit_resulting_resolution,[],[f444,f68])).
fof(f759,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl50),
  inference(unit_resulting_resolution,[],[f61,f444,f79])).
fof(f761,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl50),
  inference(unit_resulting_resolution,[],[f61,f444,f79])).
fof(f746,plain,(
  $false | ($spl33 | $spl54)),
  inference(global_subsumption,[],[f60,f141,f449,f739,f737,f735])).
fof(f735,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f141,f449,f79])).
fof(f737,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl54),
  inference(unit_resulting_resolution,[],[f61,f449,f79])).
fof(f739,plain,(
  ~p(s(bool,'T_0')) | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f141,f449,f79])).
fof(f740,plain,(
  $false | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f61,f141,f449,f79])).
fof(f736,plain,(
  $false | ($spl33 | $spl54)),
  inference(unit_resulting_resolution,[],[f61,f141,f449,f79])).
fof(f732,plain,(
  $false | $spl54),
  inference(unit_resulting_resolution,[],[f62,f449,f64])).
fof(f64,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f19])).
fof(f19,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X11] : (p(s(bool,'F_0')) => p(s(bool,X11)))),
  file('thm.pred_set.compl_insert','thm.bool.FALSITY')).
fof(f730,plain,(
  $false | $spl54),
  inference(unit_resulting_resolution,[],[f62,f449])).
fof(f742,plain,(
  $false | $spl54),
  inference(resolution,[],[f449,f62])).
fof(f398,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl46),
  inference(cnf_transformation,[],[f398_D])).
fof(f398_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f687,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl44),
  inference(superposition,[],[f69,f362])).
fof(f117,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl18),
  inference(cnf_transformation,[],[f117_D])).
fof(f117_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f678,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl45),
  inference(trivial_inequality_removal,[],[f677])).
fof(f677,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl45),
  inference(superposition,[],[f363,f63])).
fof(f674,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f363,f63])).
fof(f680,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl74),
  inference(cnf_transformation,[],[f680_D])).
fof(f680_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f675,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f62,f363,f76])).
fof(f676,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))))) | $spl45),
  inference(unit_resulting_resolution,[],[f62,f363,f76])).
fof(f116,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl19),
  inference(cnf_transformation,[],[f116_D])).
fof(f116_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f671,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl5 | $spl42)),
  inference(backward_demodulation,[],[f360,f98])).
fof(f670,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl72),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f661,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl51),
  inference(resolution,[],[f445,f69])).
fof(f653,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),sK1)))) | $spl51),
  inference(unit_resulting_resolution,[],[f445,f69])).
fof(f668,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl70),
  inference(cnf_transformation,[],[f668_D])).
fof(f668_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f655,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))))) | $spl51),
  inference(unit_resulting_resolution,[],[f445,f69])).
fof(f666,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl68),
  inference(cnf_transformation,[],[f666_D])).
fof(f666_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f656,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl51),
  inference(unit_resulting_resolution,[],[f62,f445,f76])).
fof(f658,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))))) | $spl51),
  inference(unit_resulting_resolution,[],[f62,f445,f76])).
fof(f589,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl67),
  inference(cnf_transformation,[],[f589_D])).
fof(f589_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f485,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl22),
  inference(superposition,[],[f68,f124])).
fof(f586,plain,(
  $false | ($spl33 | $spl49)),
  inference(global_subsumption,[],[f60,f141,f401,f580,f578,f576])).
fof(f576,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f141,f401,f76])).
fof(f578,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl49),
  inference(unit_resulting_resolution,[],[f62,f401,f76])).
fof(f580,plain,(
  p(s(bool,'F_0')) | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f141,f401,f76])).
fof(f581,plain,(
  $false | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f62,f141,f401,f76])).
fof(f577,plain,(
  $false | ($spl33 | $spl49)),
  inference(unit_resulting_resolution,[],[f62,f141,f401,f76])).
fof(f573,plain,(
  $false | $spl49),
  inference(unit_resulting_resolution,[],[f61,f401])).
fof(f583,plain,(
  $false | $spl49),
  inference(resolution,[],[f401,f61])).
fof(f554,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl63),
  inference(cnf_transformation,[],[f554_D])).
fof(f554_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f556,plain,(
  s(sK0,sK2) = s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) | $spl64),
  inference(cnf_transformation,[],[f556_D])).
fof(f556_D,plain,(
  s(sK0,sK2) = s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f531,plain,(
  p(s(bool,'F_0')) | s(sK0,sK2) = s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl24),
  inference(superposition,[],[f72,f125])).
fof(f473,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl59),
  inference(cnf_transformation,[],[f473_D])).
fof(f473_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f511,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl22),
  inference(global_subsumption,[],[f60,f124,f484])).
fof(f484,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))) | $spl22),
  inference(superposition,[],[f67,f124])).
fof(f67,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X7,X11] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X11) <=> ! [X8] : s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X8),s(fun(X6,bool),X11))))),
  file('thm.pred_set.compl_insert','thm.pred_set.EXTENSION')).
fof(f477,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl61),
  inference(cnf_transformation,[],[f477_D])).
fof(f477_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f475,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl9 | $spl24)),
  inference(backward_demodulation,[],[f125,f102])).
fof(f474,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl58),
  inference(cnf_transformation,[],[f474_D])).
fof(f474_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f470,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl25),
  inference(trivial_inequality_removal,[],[f469])).
fof(f469,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl25),
  inference(superposition,[],[f126,f63])).
fof(f466,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl25),
  inference(unit_resulting_resolution,[],[f126,f63])).
fof(f472,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl56),
  inference(cnf_transformation,[],[f472_D])).
fof(f472_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f467,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f62,f126,f76])).
fof(f468,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl25),
  inference(unit_resulting_resolution,[],[f62,f126,f76])).
fof(f445,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl51),
  inference(cnf_transformation,[],[f445_D])).
fof(f445_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f447,plain,(
  s(sK0,sK2) = s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)) | $spl52),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  s(sK0,sK2) = s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f422,plain,(
  p(s(bool,'F_0')) | s(sK0,sK2) = s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl20),
  inference(superposition,[],[f72,f118])).
fof(f118,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl20),
  inference(cnf_transformation,[],[f118_D])).
fof(f118_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f399,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl47),
  inference(cnf_transformation,[],[f399_D])).
fof(f399_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f371,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))))) | $spl18),
  inference(superposition,[],[f68,f117])).
fof(f359,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl43),
  inference(cnf_transformation,[],[f359_D])).
fof(f359_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f397,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl18),
  inference(global_subsumption,[],[f60,f117,f370])).
fof(f370,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))) | $spl18),
  inference(superposition,[],[f67,f117])).
fof(f363,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl45),
  inference(cnf_transformation,[],[f363_D])).
fof(f363_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f361,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | ($spl5 | $spl20)),
  inference(backward_demodulation,[],[f118,f98])).
fof(f360,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl42),
  inference(cnf_transformation,[],[f360_D])).
fof(f360_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f356,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl21),
  inference(trivial_inequality_removal,[],[f355])).
fof(f355,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl21),
  inference(superposition,[],[f119,f63])).
fof(f352,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl21),
  inference(unit_resulting_resolution,[],[f119,f63])).
fof(f358,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl40),
  inference(cnf_transformation,[],[f358_D])).
fof(f358_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f353,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f62,f119,f76])).
fof(f354,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f62,f119,f76])).
fof(f242,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl37),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f296,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl14),
  inference(global_subsumption,[],[f60,f110,f274])).
fof(f274,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))) | $spl14),
  inference(superposition,[],[f88,f110])).
fof(f88,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.compl_insert',aHL_EXT)).
fof(f246,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl39),
  inference(cnf_transformation,[],[f246_D])).
fof(f246_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f244,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | ($spl7 | $spl16)),
  inference(backward_demodulation,[],[f111,f100])).
fof(f243,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f239,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl17),
  inference(trivial_inequality_removal,[],[f238])).
fof(f238,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl17),
  inference(superposition,[],[f112,f63])).
fof(f235,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl17),
  inference(unit_resulting_resolution,[],[f112,f63])).
fof(f241,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | $spl34),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f236,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f62,f112,f76])).
fof(f237,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0))))) | $spl17),
  inference(unit_resulting_resolution,[],[f62,f112,f76])).
fof(f134,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl29),
  inference(cnf_transformation,[],[f134_D])).
fof(f134_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f179,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl10),
  inference(global_subsumption,[],[f60,f105,f174,f175,f159])).
fof(f159,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))) | $spl10),
  inference(superposition,[],[f88,f105])).
fof(f141,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl33),
  inference(cnf_transformation,[],[f141_D])).
fof(f141_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f139,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl10 | $spl31)),
  inference(forward_demodulation,[],[f138,f105])).
fof(f138,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl31),
  inference(cnf_transformation,[],[f138_D])).
fof(f138_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f136,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | ($spl3 | $spl12)),
  inference(backward_demodulation,[],[f106,f96])).
fof(f135,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl28),
  inference(cnf_transformation,[],[f135_D])).
fof(f135_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f131,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl13),
  inference(trivial_inequality_removal,[],[f130])).
fof(f130,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl13),
  inference(superposition,[],[f107,f63])).
fof(f127,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl13),
  inference(unit_resulting_resolution,[],[f107,f63])).
fof(f133,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | $spl26),
  inference(cnf_transformation,[],[f133_D])).
fof(f133_D,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f128,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f62,f107,f76])).
fof(f129,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0))))) | $spl13),
  inference(unit_resulting_resolution,[],[f62,f107,f76])).
fof(f122,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl9),
  inference(superposition,[],[f102,f63])).
fof(f115,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) | $spl5),
  inference(superposition,[],[f98,f63])).
fof(f112,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f112_D])).
fof(f112_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f108,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl7),
  inference(superposition,[],[f100,f63])).
fof(f107,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl13),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f103,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl3),
  inference(superposition,[],[f96,f63])).
fof(f102,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl9),
  inference(cnf_transformation,[],[f102_D])).
fof(f102_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f91,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f90,f67])).
fof(f100,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f100_D])).
fof(f100_D,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f92,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f90,f88])).
fof(f98,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl5),
  inference(cnf_transformation,[],[f98_D])).
fof(f98_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f93,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))))) | $spl1),
  inference(unit_resulting_resolution,[],[f90,f67])).
fof(f96,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f96_D])).
fof(f96_D,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f94,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1)))),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f90,f88])).
fof(f90,plain,(
  s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))) | $spl1),
  inference(cnf_transformation,[],[f90_D])).
fof(f90_D,plain,(
  s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f60,plain,(
  s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK2),s(fun(sK0,bool),sK1))))) != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(sK0,sK2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f37])).
fof(f37,plain,(
  ? [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) != s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(X0,X2)))),
  inference(ennf_transformation,[],[f17])).
fof(f17,plain,(
  ~! [X0,X1,X2] : s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(X0,X2)))),
  inference(flattening,[],[f16])).
fof(f16,plain,(
  ~! [X0] : ! [X1,X2] : s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X1))))) = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(X0,X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X8),s(fun(X6,bool),X7))))) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(X6,X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X8),s(fun(X6,bool),X7))))) = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(X6,X8)))),
  file('thm.pred_set.compl_insert','thm.pred_set.compl_insert')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 206
% Passive clauses: 1114
% Generated clauses: 10825
% Final active clauses: 71
% Final passive clauses: 445
% Input formulas: 14
% Initial clauses: 29
% 
% Duplicate literals: 439
% Trivial inequalities: 30
% Fw demodulations: 2598
% Bw demodulations: 197
% Global subsumptions: 41
% 
% Simple tautologies: 90
% Equational tautologies: 196
% Forward subsumptions: 5869
% Backward subsumptions: 2
% 
% Binary resolution: 516
% Unit resulting resolution: 1593
% Forward superposition: 1926
% Backward superposition: 2962
% Self superposition: 87
% Equality factoring: 96
% 
% Split clauses: 138
% Split components: 278
% SAT solver clauses: 2022
% SAT solver unit clauses: 9
% SAT solver binary clauses: 602
% 
% Sat splits: 138
% Sat splitting refutations: 13
% 
% Memory used [KB]: 6908
% Time elapsed: 0.376 s
% ------------------------------
% ------------------------------
% Success in time 18.474 s
