Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (82323)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 127
% Passive clauses: 4590
% Generated clauses: 17024
% Final active clauses: 98
% Final passive clauses: 3747
% Input formulas: 29
% Initial clauses: 64
% 
% Duplicate literals: 741
% Trivial inequalities: 279
% Fw subsumption resolutions: 2688
% Fw demodulations: 761
% Bw demodulations: 17
% 
% Simple tautologies: 373
% Equational tautologies: 80
% Forward subsumptions: 6072
% 
% Binary resolution: 666
% Factoring: 3
% Forward superposition: 6911
% Backward superposition: 4286
% Self superposition: 207
% Equality factoring: 170
% Equality resolution: 27
% 
% Split clauses: 711
% Split components: 1456
% SAT solver clauses: 1573
% SAT solver unit clauses: 49
% SAT solver binary clauses: 507
% 
% Sat splits: 711
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1572
% TWLsolver calls for satisfiability: 186
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (82338)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 105
% Passive clauses: 993
% Generated clauses: 11786
% Final active clauses: 90
% Final passive clauses: 861
% Input formulas: 29
% Initial clauses: 64
% 
% Duplicate literals: 807
% Trivial inequalities: 146
% Fw subsumption resolutions: 1065
% Fw demodulations: 756
% Bw demodulations: 33
% Condensations: 1425
% Global subsumptions: 9
% 
% Simple tautologies: 234
% Equational tautologies: 45
% Forward subsumptions: 4687
% 
% Binary resolution: 477
% Factoring: 3
% Forward superposition: 3343
% Backward superposition: 2450
% Self superposition: 118
% Equality factoring: 102
% Equality resolution: 21
% 
% Split clauses: 125
% Split components: 267
% SAT solver clauses: 1593
% SAT solver unit clauses: 67
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 18
% 
% Sat splits: 125
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1557
% TWLsolver calls for satisfiability: 281
% 
% Memory used [KB]: 10362
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (82349)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 303
% Passive clauses: 10204
% Generated clauses: 107834
% Final active clauses: 37
% Final passive clauses: 112
% Input formulas: 29
% Initial clauses: 64
% Discarded non-redundant clauses: 555
% 
% Duplicate literals: 3437
% Trivial inequalities: 981
% Fw subsumption resolutions: 11217
% Fw demodulations: 11278
% Bw demodulations: 8921
% Condensations: 8937
% Global subsumptions: 1081
% 
% Simple tautologies: 965
% Equational tautologies: 1663
% Forward subsumptions: 51373
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 281
% 
% Unit resulting resolution: 16745
% Factoring: 1
% Forward superposition: 22349
% Backward superposition: 9587
% Self superposition: 304
% Equality factoring: 206
% Equality resolution: 70
% 
% Split clauses: 239
% Split components: 502
% SAT solver clauses: 18035
% SAT solver unit clauses: 265
% SAT solver binary clauses: 2688
% 
% Sat splits: 239
% Sat splitting refutations: 6269
% 
% Memory used [KB]: 63325
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (82682)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 534
% Passive clauses: 33044
% Generated clauses: 106214
% Final active clauses: 126
% Final passive clauses: 28825
% Input formulas: 29
% Initial clauses: 63
% 
% Selected by SInE selection: 28
% SInE iterations: 1
% 
% Duplicate literals: 3044
% Trivial inequalities: 2014
% Fw subsumption resolutions: 13804
% Fw demodulations: 10302
% Bw demodulations: 365
% 
% Simple tautologies: 2355
% Equational tautologies: 73
% Forward subsumptions: 38739
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 6905
% Forward superposition: 41525
% Backward superposition: 26478
% Self superposition: 504
% Equality factoring: 140
% Equality resolution: 121
% 
% Split clauses: 3515
% Split components: 7171
% SAT solver clauses: 3652
% SAT solver unit clauses: 27
% SAT solver binary clauses: 1965
% 
% Sat splits: 3515
% Sat splitting refutations: 87
% 
% Memory used [KB]: 35308
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (82818)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 1412
% Generated clauses: 14741
% Final active clauses: 83
% Final passive clauses: 1279
% Input formulas: 29
% Initial clauses: 64
% 
% Duplicate literals: 520
% Trivial inequalities: 355
% Fw subsumption resolutions: 1831
% Fw demodulations: 1669
% Bw demodulations: 39
% Global subsumptions: 1
% 
% Simple tautologies: 539
% Equational tautologies: 106
% Forward subsumptions: 7854
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 1068
% Factoring: 2
% Forward superposition: 5195
% Backward superposition: 3792
% Self superposition: 85
% Equality factoring: 44
% Equality resolution: 11
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 1423
% SAT solver unit clauses: 31
% SAT solver binary clauses: 83
% 
% Sat splits: 7
% Sat splitting refutations: 1
% 
% Memory used [KB]: 8059
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (82829)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 199
% Passive clauses: 4010
% Generated clauses: 17828
% Final active clauses: 82
% Final passive clauses: 1134
% Input formulas: 29
% Initial clauses: 64
% 
% Duplicate literals: 602
% Trivial inequalities: 522
% Fw subsumption resolutions: 873
% Bw subsumption resolutions: 3
% Fw demodulations: 2001
% Bw demodulations: 297
% Condensations: 1019
% Global subsumptions: 1
% 
% Simple tautologies: 555
% Equational tautologies: 22
% Forward subsumptions: 8179
% Backward subsumptions: 127
% 
% Binary resolution: 1534
% Forward superposition: 7756
% Backward superposition: 3042
% Self superposition: 26
% Equality factoring: 48
% Equality resolution: 46
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 4015
% SAT solver unit clauses: 12
% SAT solver binary clauses: 160
% 
% Sat splits: 2
% 
% Memory used [KB]: 16502
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f3005,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f110,f176,f182,f184,f181,f186,f180,f188,f179,f190,f178,f192,f177,f194,f209,f206,f203,f202,f201,f200,f221,f197,f223,f215,f216,f219,f228,f231,f229,f233,f247,f245,f242,f241,f240,f239,f258,f236,f260,f252,f253,f256,f269,f268,f267,f266,f265,f273,f271,f275,f276,f278,f296,f295,f294,f293,f292,f300,f298,f302,f305,f309,f307,f310,f314,f312,f379,f378,f377,f376,f375,f383,f381,f385,f387,f306,f419,f418,f417,f424,f416,f415,f422,f430,f429,f428,f427,f426,f434,f432,f436,f438,f311,f444,f443,f442,f449,f441,f440,f447,f450,f452,f448,f789,f791,f750,f793,f811,f859,f861,f820,f795,f863,f919,f917,f914,f913,f912,f911,f930,f908,f932,f924,f925,f928,f977,f975,f972,f971,f970,f969,f988,f966,f990,f982,f983,f986,f1329,f1328,f1326,f1325,f1322,f1321,f1341,f1320,f1319,f1342,f1344,f1335,f1336,f1347,f1339,f1350,f1354,f1351,f1356,f1352,f1358,f1359,f1361,f1362,f1364,f1365,f1458,f1457,f1455,f1454,f1451,f1450,f1470,f1449,f1448,f1471,f1346,f1473,f1464,f1465,f1474,f1468,f1847,f1845,f1842,f1841,f1840,f1839,f1858,f1836,f1860,f1852,f1853,f1856,f2129,f2127,f2124,f2123,f2122,f2121,f2140,f2118,f2142,f2134,f2135,f2138,f2987,f2989])).
fof(f2989,plain,(
  $false | ($spl90 | $spl98)),
  inference(unit_resulting_resolution,[],[f1358,f2376,f123])).
fof(f123,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.SING_DELETE','thm.bool.EQ_CLAUSES|split|0')).
fof(f2376,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK1))))))) ) | $spl90),
  inference(forward_demodulation,[],[f2360,f1344])).
fof(f2360,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),X0),s(sK0,sK1))))))) ) | $spl90),
  inference(unit_resulting_resolution,[],[f1344,f152])).
fof(f152,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(flattening,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(nnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X12,X7,X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X12),s(X6,X13)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X12)))) & ~s(X6,X7) = s(X6,X13)))),
  file('thm.pred_set.SING_DELETE','thm.pred_set.IN_DELETE')).
fof(f2987,plain,(
  $false | ($spl90 | $spl98)),
  inference(unit_resulting_resolution,[],[f1358,f2376,f121])).
fof(f121,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.SING_DELETE','thm.bool.EQ_CLAUSES|split|1')).
fof(f2138,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl116),
  inference(cnf_transformation,[],[f2138_D])).
fof(f2138_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f2135,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(resolution,[],[f990,f126])).
fof(f126,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.SING_DELETE','thm.bool.EQ_CLAUSES|split|2')).
fof(f2134,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(resolution,[],[f990,f128])).
fof(f128,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.SING_DELETE','thm.bool.EQ_CLAUSES|split|3')).
fof(f2142,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))))) | $spl121),
  inference(cnf_transformation,[],[f2142_D])).
fof(f2142_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))))) <=> ~$spl121),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl121])])).
fof(f2118,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))))) | ($spl1 | $spl85)),
  inference(unit_resulting_resolution,[],[f176,f990,f154])).
fof(f154,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X7,X13,X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X13),s(fun(X6,bool),X12)))))) <=> (s(X6,X7) = s(X6,X13) | p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X12))))))),
  file('thm.pred_set.SING_DELETE','thm.pred_set.IN_INSERT')).
fof(f2140,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl119),
  inference(cnf_transformation,[],[f2140_D])).
fof(f2140_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f2121,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f990,f121])).
fof(f2122,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f990,f123])).
fof(f2123,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f990,f126])).
fof(f2124,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f990,f128])).
fof(f2127,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f117,f990,f157])).
fof(f157,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.SING_DELETE','thm.bool.IMP_ANTISYM_AX')).
fof(f117,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f35])).
fof(f35,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.SING_DELETE',aHL_FALSITY)).
fof(f2129,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl85),
  inference(unit_resulting_resolution,[],[f117,f990,f157])).
fof(f1856,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl110),
  inference(cnf_transformation,[],[f1856_D])).
fof(f1856_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f1853,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(resolution,[],[f932,f126])).
fof(f1852,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(resolution,[],[f932,f128])).
fof(f1860,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))))) | $spl115),
  inference(cnf_transformation,[],[f1860_D])).
fof(f1860_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1836,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))))) | ($spl1 | $spl79)),
  inference(unit_resulting_resolution,[],[f176,f932,f154])).
fof(f1858,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl113),
  inference(cnf_transformation,[],[f1858_D])).
fof(f1858_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1839,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(unit_resulting_resolution,[],[f932,f121])).
fof(f1840,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(unit_resulting_resolution,[],[f932,f123])).
fof(f1841,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(unit_resulting_resolution,[],[f932,f126])).
fof(f1842,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(unit_resulting_resolution,[],[f932,f128])).
fof(f1845,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(unit_resulting_resolution,[],[f117,f932,f157])).
fof(f1847,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))))) | $spl79),
  inference(unit_resulting_resolution,[],[f117,f932,f157])).
fof(f1468,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl104),
  inference(cnf_transformation,[],[f1468_D])).
fof(f1468_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f1474,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(global_subsumption,[],[f110,f863,f1458,f1457,f1455,f1454,f1451,f1450,f1449,f1448,f1447,f1461,f1464,f1465,f1466])).
fof(f1466,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(superposition,[],[f863,f119])).
fof(f119,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.SING_DELETE',aHL_BOOL_CASES)).
fof(f1461,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | $spl72),
  inference(resolution,[],[f863,f160])).
fof(f160,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f72])).
fof(f1447,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) ) | $spl72),
  inference(unit_resulting_resolution,[],[f863,f156])).
fof(f156,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f102])).
fof(f1465,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(resolution,[],[f863,f122])).
fof(f122,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f79])).
fof(f1464,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(resolution,[],[f863,f124])).
fof(f124,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f80])).
fof(f1473,plain,(
  s(sK0,sK1) = s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)) | $spl108),
  inference(cnf_transformation,[],[f1473_D])).
fof(f1473_D,plain,(
  s(sK0,sK1) = s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f1346,plain,(
  p(s(bool,'F_0')) | $spl92),
  inference(cnf_transformation,[],[f1346_D])).
fof(f1346_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1471,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)) | ($spl3 | $spl14 | $spl72)),
  inference(forward_demodulation,[],[f1460,f225])).
fof(f225,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) ) | ($spl3 | $spl14)),
  inference(backward_demodulation,[],[f219,f210])).
fof(f210,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))) = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f148,f184,f157])).
fof(f148,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X7] : ~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.SING_DELETE','thm.pred_set.NOT_IN_EMPTY')).
fof(f1460,plain,(
  s(sK0,sK1) = s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl72),
  inference(resolution,[],[f863,f154])).
fof(f1448,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f863,f122])).
fof(f1449,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f863,f124])).
fof(f1470,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl107),
  inference(cnf_transformation,[],[f1470_D])).
fof(f1470_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f1450,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f863,f125])).
fof(f125,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f1451,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f863,f127])).
fof(f127,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f1454,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f112,f863,f160])).
fof(f112,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SING_DELETE',aHL_TRUTH)).
fof(f1455,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f111,f863,f160])).
fof(f111,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.SING_DELETE','thm.bool.TRUTH')).
fof(f1457,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f112,f863,f160])).
fof(f1458,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl72),
  inference(unit_resulting_resolution,[],[f111,f863,f160])).
fof(f1365,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))),
  introduced(theory_axiom,[])).
fof(f1364,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl102),
  inference(cnf_transformation,[],[f1364_D])).
fof(f1364_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1362,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl86 | $spl90)),
  inference(forward_demodulation,[],[f1339,f1344])).
fof(f1361,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl101),
  inference(cnf_transformation,[],[f1361_D])).
fof(f1361_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f1359,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | ($spl89 | $spl90)),
  inference(forward_demodulation,[],[f1341,f1344])).
fof(f1358,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl98),
  inference(cnf_transformation,[],[f1358_D])).
fof(f1358_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f1352,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | ($spl32 | $spl90)),
  inference(backward_demodulation,[],[f1344,f275])).
fof(f1356,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl96),
  inference(cnf_transformation,[],[f1356_D])).
fof(f1356_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f1351,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | ($spl66 | $spl90)),
  inference(backward_demodulation,[],[f1344,f793])).
fof(f1354,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl95),
  inference(cnf_transformation,[],[f1354_D])).
fof(f1354_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f1350,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ($spl65 | $spl90)),
  inference(backward_demodulation,[],[f1344,f791])).
fof(f1339,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl86),
  inference(cnf_transformation,[],[f1339_D])).
fof(f1339_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1347,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(global_subsumption,[],[f110,f793,f1329,f1328,f1326,f1325,f1322,f1321,f1320,f1319,f1318,f1332,f1335,f1336,f1337])).
fof(f1337,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(superposition,[],[f793,f119])).
fof(f1332,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) ) | $spl66),
  inference(resolution,[],[f793,f160])).
fof(f1318,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))))))) ) | $spl66),
  inference(unit_resulting_resolution,[],[f793,f156])).
fof(f1336,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(resolution,[],[f793,f122])).
fof(f1335,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(resolution,[],[f793,f124])).
fof(f1344,plain,(
  s(sK0,sK1) = s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)) | $spl90),
  inference(cnf_transformation,[],[f1344_D])).
fof(f1344_D,plain,(
  s(sK0,sK1) = s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f1342,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)) | ($spl3 | $spl14 | $spl66)),
  inference(forward_demodulation,[],[f1331,f225])).
fof(f1331,plain,(
  s(sK0,sK1) = s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl66),
  inference(resolution,[],[f793,f154])).
fof(f1319,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f793,f122])).
fof(f1320,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f793,f124])).
fof(f1341,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl89),
  inference(cnf_transformation,[],[f1341_D])).
fof(f1341_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f1321,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f793,f125])).
fof(f1322,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f793,f127])).
fof(f1325,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f112,f793,f160])).
fof(f1326,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f111,f793,f160])).
fof(f1328,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f112,f793,f160])).
fof(f1329,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))))) | $spl66),
  inference(unit_resulting_resolution,[],[f111,f793,f160])).
fof(f986,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl80),
  inference(cnf_transformation,[],[f986_D])).
fof(f986_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f983,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(resolution,[],[f260,f126])).
fof(f982,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(resolution,[],[f260,f128])).
fof(f990,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))) | $spl85),
  inference(cnf_transformation,[],[f990_D])).
fof(f990_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f966,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))) | ($spl1 | $spl29)),
  inference(unit_resulting_resolution,[],[f176,f260,f154])).
fof(f988,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl83),
  inference(cnf_transformation,[],[f988_D])).
fof(f988_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f969,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f260,f121])).
fof(f970,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f260,f123])).
fof(f971,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f260,f126])).
fof(f972,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f260,f128])).
fof(f975,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f117,f260,f157])).
fof(f977,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl29),
  inference(unit_resulting_resolution,[],[f117,f260,f157])).
fof(f928,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl74),
  inference(cnf_transformation,[],[f928_D])).
fof(f928_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f925,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(resolution,[],[f223,f126])).
fof(f924,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(resolution,[],[f223,f128])).
fof(f932,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))) | $spl79),
  inference(cnf_transformation,[],[f932_D])).
fof(f932_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f908,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))))) | ($spl1 | $spl19)),
  inference(unit_resulting_resolution,[],[f176,f223,f154])).
fof(f930,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl77),
  inference(cnf_transformation,[],[f930_D])).
fof(f930_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f911,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f223,f121])).
fof(f912,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f223,f123])).
fof(f913,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f223,f126])).
fof(f914,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f223,f128])).
fof(f917,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f117,f223,f157])).
fof(f919,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))) | $spl19),
  inference(unit_resulting_resolution,[],[f117,f223,f157])).
fof(f863,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl72),
  inference(cnf_transformation,[],[f863_D])).
fof(f863_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f795,plain,(
  ~p(s(bool,'T_0')) | $spl69),
  inference(cnf_transformation,[],[f795_D])).
fof(f795_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f820,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl38),
  inference(superposition,[],[f151,f302])).
fof(f151,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f861,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl71),
  inference(cnf_transformation,[],[f861_D])).
fof(f861_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f859,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl38),
  inference(global_subsumption,[],[f110,f302,f819])).
fof(f819,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))) | $spl38),
  inference(superposition,[],[f150,f302])).
fof(f150,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : ! [X12,X10] : (s(fun(X6,bool),X10) = s(fun(X6,bool),X12) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X10))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X12))))),
  file('thm.pred_set.SING_DELETE','thm.pred_set.EXTENSION')).
fof(f811,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f793,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl66),
  inference(cnf_transformation,[],[f793_D])).
fof(f793_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f750,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))))) | $spl32),
  inference(superposition,[],[f151,f275])).
fof(f791,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl65),
  inference(cnf_transformation,[],[f791_D])).
fof(f791_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f789,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl32),
  inference(global_subsumption,[],[f110,f275,f749])).
fof(f749,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(fun(sK0,bool),'const.pred_set.EMPTY_0') = s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))) | $spl32),
  inference(superposition,[],[f150,f275])).
fof(f448,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl63),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f452,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | ($spl11 | $spl46)),
  inference(backward_demodulation,[],[f313,f192])).
fof(f313,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl46),
  inference(cnf_transformation,[],[f313_D])).
fof(f313_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f450,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | ~p(s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))))),
  introduced(theory_axiom,[])).
fof(f447,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl61),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f440,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f311,f124])).
fof(f441,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f311,f122])).
fof(f449,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl62),
  inference(cnf_transformation,[],[f449_D])).
fof(f449_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f442,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl45),
  inference(unit_resulting_resolution,[],[f311,f119])).
fof(f443,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f112,f311,f160])).
fof(f444,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl45),
  inference(unit_resulting_resolution,[],[f112,f311,f160])).
fof(f311,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl45),
  inference(cnf_transformation,[],[f311_D])).
fof(f311_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f438,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | ($spl11 | $spl58)),
  inference(backward_demodulation,[],[f436,f192])).
fof(f436,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl58),
  inference(cnf_transformation,[],[f436_D])).
fof(f436_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f432,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl47),
  inference(trivial_inequality_removal,[],[f431])).
fof(f431,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl47),
  inference(superposition,[],[f314,f119])).
fof(f434,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl56),
  inference(cnf_transformation,[],[f434_D])).
fof(f434_D,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f426,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl47),
  inference(unit_resulting_resolution,[],[f314,f128])).
fof(f427,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl47),
  inference(unit_resulting_resolution,[],[f314,f126])).
fof(f428,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl47),
  inference(unit_resulting_resolution,[],[f314,f119])).
fof(f429,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl47),
  inference(unit_resulting_resolution,[],[f117,f314,f157])).
fof(f430,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0))))) | $spl47),
  inference(unit_resulting_resolution,[],[f117,f314,f157])).
fof(f422,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl53),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f415,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl41),
  inference(unit_resulting_resolution,[],[f306,f124])).
fof(f416,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl41),
  inference(unit_resulting_resolution,[],[f306,f122])).
fof(f424,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl54),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f417,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl41),
  inference(unit_resulting_resolution,[],[f306,f119])).
fof(f418,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl41),
  inference(unit_resulting_resolution,[],[f112,f306,f160])).
fof(f419,plain,(
  ~p(s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl41),
  inference(unit_resulting_resolution,[],[f112,f306,f160])).
fof(f306,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl41),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f387,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | ($spl5 | $spl50)),
  inference(backward_demodulation,[],[f385,f186])).
fof(f385,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl50),
  inference(cnf_transformation,[],[f385_D])).
fof(f385_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f381,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl43),
  inference(trivial_inequality_removal,[],[f380])).
fof(f380,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl43),
  inference(superposition,[],[f309,f119])).
fof(f383,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl48),
  inference(cnf_transformation,[],[f383_D])).
fof(f383_D,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f375,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl43),
  inference(unit_resulting_resolution,[],[f309,f128])).
fof(f376,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl43),
  inference(unit_resulting_resolution,[],[f309,f126])).
fof(f377,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl43),
  inference(unit_resulting_resolution,[],[f309,f119])).
fof(f378,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl43),
  inference(unit_resulting_resolution,[],[f117,f309,f157])).
fof(f379,plain,(
  p(s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0))))) | $spl43),
  inference(unit_resulting_resolution,[],[f117,f309,f157])).
fof(f312,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl44),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f314,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl47),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f310,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl11),
  inference(superposition,[],[f192,f119])).
fof(f307,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl40),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f309,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl43),
  inference(cnf_transformation,[],[f309_D])).
fof(f309_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f305,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl5),
  inference(superposition,[],[f186,f119])).
fof(f302,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl38),
  inference(cnf_transformation,[],[f302_D])).
fof(f302_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f298,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl23),
  inference(trivial_inequality_removal,[],[f297])).
fof(f297,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl23),
  inference(superposition,[],[f233,f119])).
fof(f300,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl36),
  inference(cnf_transformation,[],[f300_D])).
fof(f300_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f292,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f233,f128])).
fof(f293,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f233,f126])).
fof(f294,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl23),
  inference(unit_resulting_resolution,[],[f233,f119])).
fof(f295,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f117,f233,f157])).
fof(f296,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f117,f233,f157])).
fof(f278,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl35),
  inference(cnf_transformation,[],[f278_D])).
fof(f278_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f276,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl21 | $spl32)),
  inference(backward_demodulation,[],[f275,f231])).
fof(f275,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl32),
  inference(cnf_transformation,[],[f275_D])).
fof(f275_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f271,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl21),
  inference(trivial_inequality_removal,[],[f270])).
fof(f270,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl21),
  inference(superposition,[],[f231,f119])).
fof(f273,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl30),
  inference(cnf_transformation,[],[f273_D])).
fof(f273_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f265,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f231,f128])).
fof(f266,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f231,f126])).
fof(f267,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl21),
  inference(unit_resulting_resolution,[],[f231,f119])).
fof(f268,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f117,f231,f157])).
fof(f269,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))))) | $spl21),
  inference(unit_resulting_resolution,[],[f117,f231,f157])).
fof(f256,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl24),
  inference(cnf_transformation,[],[f256_D])).
fof(f256_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f253,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(resolution,[],[f190,f126])).
fof(f252,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(resolution,[],[f190,f128])).
fof(f260,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))) | $spl29),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f236,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))) | ($spl1 | $spl9)),
  inference(unit_resulting_resolution,[],[f176,f190,f154])).
fof(f258,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl27),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f239,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(unit_resulting_resolution,[],[f190,f121])).
fof(f240,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(unit_resulting_resolution,[],[f190,f123])).
fof(f241,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(unit_resulting_resolution,[],[f190,f126])).
fof(f242,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(unit_resulting_resolution,[],[f190,f128])).
fof(f245,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(unit_resulting_resolution,[],[f117,f190,f157])).
fof(f247,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl9),
  inference(unit_resulting_resolution,[],[f117,f190,f157])).
fof(f233,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl23),
  inference(cnf_transformation,[],[f233_D])).
fof(f233_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f229,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | ($spl3 | $spl7 | $spl14)),
  inference(backward_demodulation,[],[f225,f188])).
fof(f231,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl21),
  inference(cnf_transformation,[],[f231_D])).
fof(f231_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f228,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | ($spl3 | $spl13 | $spl14)),
  inference(backward_demodulation,[],[f225,f194])).
fof(f219,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl14),
  inference(cnf_transformation,[],[f219_D])).
fof(f219_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f216,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(resolution,[],[f184,f126])).
fof(f215,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(resolution,[],[f184,f128])).
fof(f223,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))) | $spl19),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f197,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))))) | ($spl1 | $spl3)),
  inference(unit_resulting_resolution,[],[f176,f184,f154])).
fof(f221,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl17),
  inference(cnf_transformation,[],[f221_D])).
fof(f221_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f200,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f184,f121])).
fof(f201,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f184,f123])).
fof(f202,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f184,f126])).
fof(f203,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f184,f128])).
fof(f206,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f117,f184,f157])).
fof(f209,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0'))))) | $spl3),
  inference(unit_resulting_resolution,[],[f117,f184,f157])).
fof(f194,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl13),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f177,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f176,f150])).
fof(f192,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f178,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),'const.pred_set.EMPTY_0',bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f176,f173])).
fof(f173,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.SING_DELETE',aHL_EXT)).
fof(f190,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))) | $spl9),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f179,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f148,f176,f154])).
fof(f188,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl7),
  inference(cnf_transformation,[],[f188_D])).
fof(f188_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f180,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))))) | $spl1),
  inference(unit_resulting_resolution,[],[f176,f150])).
fof(f186,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl5),
  inference(cnf_transformation,[],[f186_D])).
fof(f186_D,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f181,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.EMPTY_0','const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f176,f173])).
fof(f184,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))) | $spl3),
  inference(cnf_transformation,[],[f184_D])).
fof(f184_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f182,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))),s(fun(fun(sK0,bool),bool),'const.pred_set.INSERT_2'(s(fun(sK0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK0,bool),bool),'const.pred_set.EMPTY_0')))))) | $spl1),
  inference(unit_resulting_resolution,[],[f148,f176,f154])).
fof(f176,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))) | $spl1),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f110,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))),s(sK0,sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f69])).
fof(f69,plain,(
  ? [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(X0,X1)))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(X0,X1)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))),s(X0,X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'))),s(X6,X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.EMPTY_0'))),s(X6,X7)))),
  file('thm.pred_set.SING_DELETE','thm.pred_set.SING_DELETE')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 107
% Passive clauses: 185
% Generated clauses: 2759
% Final active clauses: 70
% Final passive clauses: 77
% Input formulas: 29
% Initial clauses: 64
% 
% Duplicate literals: 4
% Trivial inequalities: 44
% Fw demodulations: 895
% Bw demodulations: 69
% Global subsumptions: 4
% 
% Simple tautologies: 57
% Equational tautologies: 27
% Forward subsumptions: 1422
% 
% Binary resolution: 148
% Unit resulting resolution: 694
% Forward superposition: 228
% Backward superposition: 606
% Self superposition: 2
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 482
% SAT solver unit clauses: 2
% SAT solver binary clauses: 339
% 
% Sat splits: 6
% Sat splitting refutations: 2
% 
% Memory used [KB]: 3198
% Time elapsed: 0.070 s
% ------------------------------
% ------------------------------
% Success in time 18.088 s
