Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (116970)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 142
% Passive clauses: 4697
% Generated clauses: 17656
% Final active clauses: 138
% Final passive clauses: 4308
% Input formulas: 33
% Initial clauses: 77
% 
% Duplicate literals: 749
% Trivial inequalities: 285
% Fw subsumption resolutions: 3099
% Fw demodulations: 1021
% Bw demodulations: 3
% 
% Simple tautologies: 420
% Equational tautologies: 79
% Forward subsumptions: 5825
% 
% Binary resolution: 831
% Factoring: 4
% Forward superposition: 5921
% Backward superposition: 5025
% Self superposition: 242
% Equality factoring: 192
% Equality resolution: 36
% 
% Split clauses: 672
% Split components: 1387
% SAT solver clauses: 1677
% SAT solver unit clauses: 35
% SAT solver binary clauses: 397
% 
% Sat splits: 672
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1671
% TWLsolver calls for satisfiability: 228
% 
% Memory used [KB]: 8443
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (116979)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 924
% Generated clauses: 14775
% Final active clauses: 132
% Final passive clauses: 769
% Input formulas: 33
% Initial clauses: 77
% 
% Duplicate literals: 294
% Trivial inequalities: 364
% Fw subsumption resolutions: 1441
% Fw demodulations: 3119
% Bw demodulations: 16
% Condensations: 490
% Global subsumptions: 9
% 
% Simple tautologies: 180
% Equational tautologies: 27
% Forward subsumptions: 6212
% 
% Binary resolution: 132
% Factoring: 1
% Forward superposition: 4237
% Backward superposition: 3448
% Self superposition: 117
% Equality factoring: 48
% Equality resolution: 35
% 
% Split clauses: 41
% Split components: 82
% SAT solver clauses: 1736
% SAT solver unit clauses: 89
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 19
% SAT solver learnt literals: 1
% 
% Sat splits: 41
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1699
% TWLsolver calls for satisfiability: 442
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f12916,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f134,f230,f228,f133,f229,f227,f716,f719,f717,f817,f819,f821,f823,f822,f826,f830,f834,f836,f838,f840,f844,f846,f848,f850,f852,f854,f915,f916,f921,f922,f923,f924,f925,f926,f927,f928,f1179,f1181,f1180,f1314,f1316,f1315,f1898,f1900,f1899,f1941,f1943,f1942,f2772,f2774,f3786,f3788,f3787,f4150,f4255,f4259,f5587,f5588,f5591,f5592,f5593,f5594,f5595,f5596,f5607,f5608,f5609,f5610,f5611,f5612,f5613,f5590,f5614,f5893,f5895,f5894,f6156,f6162,f6161,f6842,f6844,f6843,f7037,f7039,f7038,f7510,f7512,f7511,f8955,f8956,f8963,f8960,f8957,f8958,f8966,f8959,f8968,f8961,f8965,f10003,f10005,f10004,f10257,f10258,f10290,f10291,f10292,f10293,f10299,f10300,f10511,f10512,f10524,f10522,f11908,f11910,f718,f11909,f12081,f12082,f12104,f12103,f12114,f12128,f12139,f12160,f12171,f12813,f12815,f12818,f12820,f12822,f12475,f12826,f12714,f12825,f12824,f12915])).
fof(f12915,plain,(
  $false | ($spl0 | $spl21)),
  inference(subsumption_resolution,[],[f12914,f12898])).
fof(f12898,plain,(
  ( ! [X0] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X0)))) ) | $spl0),
  inference(forward_demodulation,[],[f143,f228])).
fof(f143,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X6)))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.rat.RAT_MUL_LZERO')).
fof(f12914,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl0 | $spl21)),
  inference(forward_demodulation,[],[f2773,f12898])).
fof(f2773,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | $spl21),
  inference(cnf_transformation,[],[f2773_D])).
fof(f2773_D,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f12824,plain,(
  ( ! [X850,X849] : (s(fun('type.rat.rat',bool),X849) = s(fun('type.rat.rat',bool),X850) | s(bool,'T_0') = s(bool,g(s('type.rat.rat',sK4(X849,X850,bool,'type.rat.rat'))))) ) | $spl86),
  inference(cnf_transformation,[],[f12824_D])).
fof(f12824_D,plain,(
  ( ! [X850,X849] : (s(fun('type.rat.rat',bool),X849) = s(fun('type.rat.rat',bool),X850) | s(bool,'T_0') = s(bool,g(s('type.rat.rat',sK4(X849,X850,bool,'type.rat.rat'))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f12825,plain,(
  ( ! [X851] : (s(bool,'F_0') != s(bool,f(s('type.rat.rat',X851)))) ) | $spl88),
  inference(cnf_transformation,[],[f12825_D])).
fof(f12825_D,plain,(
  ( ! [X851] : (s(bool,'F_0') != s(bool,f(s('type.rat.rat',X851)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f12714,plain,(
  ( ! [X850,X851,X849] : (s(bool,'F_0') != s(bool,f(s('type.rat.rat',X851))) | s(fun('type.rat.rat',bool),X849) = s(fun('type.rat.rat',bool),X850) | s(bool,'T_0') = s(bool,g(s('type.rat.rat',sK4(X849,X850,bool,'type.rat.rat'))))) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f2202,f12113])).
fof(f12113,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f12096,f4164])).
fof(f4164,plain,(
  ( ! [X0] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X0)))) ) | ($spl2 | $spl20)),
  inference(backward_demodulation,[],[f4162,f143])).
fof(f4162,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl2 | $spl20)),
  inference(global_subsumption,[],[f4144,f134,f133,f135,f144,f147,f146,f178,f179,f183,f180,f191,f221,f232,f246,f282,f422,f486,f684,f685,f716,f719,f821,f823,f702,f696,f692,f690,f177,f176,f175,f190,f189,f188,f698,f231,f151,f150,f1189,f1190,f1061,f1212,f1288,f1787,f1851,f1853,f1855,f1861,f1902,f1904,f1906,f1907,f1909,f1913,f1935,f2192,f2194,f2198,f2200,f2291,f2311,f2783,f2202,f2196,f229,f2847,f2949,f3588,f3589,f3771,f3786,f3788,f3789,f2583,f2581,f2571,f2562,f2553,f2541,f2433,f2425,f2418,f2411,f2310,f2303,f2298,f2293,f2250,f2246,f2241,f2236,f715,f714,f369,f195,f194,f193,f192,f143,f142])).
fof(f142,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f30])).
fof(f30,axiom,(
  ~s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.rat.RAT_|49|_NOT_|48|')).
fof(f192,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : (s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ((s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))) & (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : (s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | (s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : (s('type.rat.rat',X1) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : (~s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6,X23,X24] : (~s('type.rat.rat',X23) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat',X24) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X23))) <=> s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24)))))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.rat.RAT_LDIV_EQ')).
fof(f193,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f127])).
fof(f194,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ((s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2)))) & (s('type.rat.rat',X1) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1,X2] : (s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <=> s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2)))))),
  inference(ennf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2] : (s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <=> s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : (~s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) <=> s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X23,X24] : (~s('type.rat.rat',X24) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => (s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) <=> s('type.rat.rat',X23) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X24)))))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.rat.RAT_RDIV_EQ')).
fof(f195,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',X1) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f128])).
fof(f369,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f178,f143])).
fof(f714,plain,(
  ( ! [X28,X27] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X27)),X28,'type.rat.rat')) | p(s(bool,happ(s(fun('type.rat.rat',bool),X28),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f701,f143])).
fof(f701,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X28),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X27))) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X27)),X28,'type.rat.rat'))) )),
  inference(superposition,[],[f175,f143])).
fof(f715,plain,(
  ( ! [X35,X36] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X36,'type.rat.rat')) | p(s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) )),
  inference(forward_demodulation,[],[f704,f369])).
fof(f704,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),X36,'type.rat.rat'))) )),
  inference(superposition,[],[f175,f369])).
fof(f2236,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f192,f178])).
fof(f2241,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))),s('type.rat.rat',X7))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X6))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f192,f178])).
fof(f2246,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))),s('type.rat.rat',X6))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5)))) )),
  inference(superposition,[],[f192,f178])).
fof(f2250,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f192])).
fof(f2293,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))),s('type.rat.rat',X6))) | s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f193,f178])).
fof(f2298,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))),s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5)))) )),
  inference(superposition,[],[f193,f178])).
fof(f2303,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f193,f178])).
fof(f2310,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat',X1))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f193])).
fof(f2411,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))),s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f194,f178])).
fof(f2418,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))),s('type.rat.rat',X7))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f194,f178])).
fof(f2425,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5)))) )),
  inference(superposition,[],[f194,f178])).
fof(f2433,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X1))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f194])).
fof(f2541,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(unit_resulting_resolution,[],[f142,f143,f195])).
fof(f2553,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))),s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f195,f178])).
fof(f2562,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))),s('type.rat.rat',X7))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f195,f178])).
fof(f2571,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat',X6) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X4))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5)))) )),
  inference(superposition,[],[f195,f178])).
fof(f2581,plain,(
  ( ! [X6,X7] : (s('type.rat.rat',X7) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X7),s('type.rat.rat',X6))) | s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f195,f143])).
fof(f2583,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X1))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(equality_resolution,[],[f195])).
fof(f3789,plain,(
  ( ! [X8] : (s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) ) | $spl2),
  inference(subsumption_resolution,[],[f3781,f231])).
fof(f3781,plain,(
  ( ! [X8] : (p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'T_0')))) ) | $spl2),
  inference(superposition,[],[f2949,f144])).
fof(f3771,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X1,bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | $spl2),
  inference(superposition,[],[f2949,f246])).
fof(f3589,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (s(X3,happ(s(fun(X2,X3),X4),s(X2,X5))) != s(X3,happ(s(fun(X2,X3),X4),s(X2,X6))) | s(X2,X5) = s(X2,X6) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X2,X3),X4))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3583,f135])).
fof(f3583,plain,(
  ( ! [X6,X4,X2,X7,X5,X3] : (~p(s(bool,'T_0')) | s(X3,happ(s(fun(X2,X3),X4),s(X2,X5))) != s(X3,happ(s(fun(X2,X3),X4),s(X2,X6))) | s(X2,X5) = s(X2,X6) | s(bool,X7) = s(bool,'const.bool.ONE_ONE_1'(s(fun(X2,X3),X4))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f188,f246])).
fof(f3588,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X0))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f3582,f135])).
fof(f3582,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X0))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | $spl2),
  inference(superposition,[],[f188,f2847])).
fof(f2949,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) | s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X10,bool))) ) | $spl2),
  inference(forward_demodulation,[],[f2927,f2847])).
fof(f2927,plain,(
  ( ! [X10] : (p(s(bool,happ(s(fun(bool,bool),X10),s(bool,'T_0')))) | s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X10,bool))) ) | $spl2),
  inference(superposition,[],[f175,f2847])).
fof(f2847,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(unit_resulting_resolution,[],[f229,f147])).
fof(f2196,plain,(
  ( ! [X21,X19,X20] : (s(X21,f(s(bool,'F_0'))) != s(X21,g(s(bool,'F_0'))) | s(fun(bool,X21),X19) = s(fun(bool,X21),X20) | s(bool,'T_0') = s(bool,sK4(X19,X20,X21,bool))) )),
  inference(superposition,[],[f191,f144])).
fof(f2783,plain,(
  ( ! [X21,X22,X20] : (s(X21,happ(s(fun(X20,X21),X22),s(X20,sK2(X22,X20,X21)))) = s(X21,happ(s(fun(X20,X21),X22),s(X20,sK3(X22,X20,X21)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X20,X21),X22)))) )),
  inference(subsumption_resolution,[],[f2769,f231])).
fof(f2769,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'F_0')) | s(X21,happ(s(fun(X20,X21),X22),s(X20,sK2(X22,X20,X21)))) = s(X21,happ(s(fun(X20,X21),X22),s(X20,sK3(X22,X20,X21)))) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X20,X21),X22)))) )),
  inference(superposition,[],[f189,f144])).
fof(f2311,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(forward_demodulation,[],[f2289,f179])).
fof(f2289,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(unit_resulting_resolution,[],[f142,f179,f193])).
fof(f2291,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(unit_resulting_resolution,[],[f142,f178,f193])).
fof(f2200,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,X14) != s(bool,f(s(X11,sK4(X12,X13,bool,X11)))) | s(fun(X11,bool),X12) = s(fun(X11,bool),X13) | s(bool,'T_0') = s(bool,g(s(X11,sK4(X12,X13,bool,X11)))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f191,f246])).
fof(f2198,plain,(
  ( ! [X6,X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(X3,sK4(X4,X5,bool,X3)))) | s(fun(X3,bool),X4) = s(fun(X3,bool),X5) | s(bool,X6) = s(bool,g(s(X3,sK4(X4,X5,bool,X3)))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f191,f246])).
fof(f2194,plain,(
  ( ! [X14,X12,X13,X11] : (s(X13,f(s(bool,X14))) != s(X13,g(s(bool,X14))) | s(fun(bool,X13),X11) = s(fun(bool,X13),X12) | s(bool,'T_0') = s(bool,sK4(X11,X12,X13,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f191,f246])).
fof(f2192,plain,(
  ( ! [X6,X4,X5,X3] : (s(X5,f(s(bool,'T_0'))) != s(X5,g(s(bool,'T_0'))) | s(fun(bool,X5),X3) = s(fun(bool,X5),X4) | s(bool,X6) = s(bool,sK4(X3,X4,X5,bool)) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f191,f246])).
fof(f1935,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f1923,f135])).
fof(f1923,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | s(bool,X4) = s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f177,f246])).
fof(f1913,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun('type.rat.rat',bool),X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X38)))))) | s('type.rat.rat',X41) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X39))) | p(s(bool,happ(s(fun('type.rat.rat',bool),X40),s('type.rat.rat',X41))))) )),
  inference(superposition,[],[f177,f178])).
fof(f1909,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') != s(bool,X29) | ~p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X27)))) | p(s(bool,happ(s(fun(bool,bool),X28),s(bool,X29)))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f177,f144])).
fof(f1907,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X19)))) | s(bool,'T_0') != s(bool,X21) | p(s(bool,happ(s(fun(bool,bool),X20),s(bool,X21)))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f177,f246])).
fof(f1906,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,happ(s(fun(bool,bool),X17),s(bool,'F_0')))) | s(bool,X16) != s(bool,X18) | p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f177,f144])).
fof(f1904,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(bool,bool),X10),s(bool,X9)))) | s(bool,X8) != s(bool,X11) | p(s(bool,happ(s(fun(bool,bool),X10),s(bool,X11)))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f177,f246])).
fof(f1902,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | s(bool,X0) != s(bool,X2) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X2)))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f177,f246])).
fof(f1861,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK1(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f1856,f135])).
fof(f1856,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(bool,X3) = s(bool,happ(s(fun(X0,bool),X1),s(X0,sK1(X2,X1,X0)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f176,f246])).
fof(f1855,plain,(
  ( ! [X12,X13] : (~p(s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X13),s(bool,X12)))) | s(bool,'T_0') = s(bool,sK1(X12,X13,bool))) )),
  inference(superposition,[],[f176,f144])).
fof(f1853,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X8)))) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,X6)))) | s(bool,'T_0') = s(bool,sK1(X6,X7,bool)) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f176,f246])).
fof(f1851,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X1),s(bool,X0)))) | s(bool,X2) = s(bool,sK1(X0,X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f176,f246])).
fof(f1787,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X86),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85)))))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X84))),s('type.rat.rat',X85)),X86,'type.rat.rat'))) )),
  inference(forward_demodulation,[],[f1763,f179])).
fof(f1763,plain,(
  ( ! [X85,X86,X84] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X86),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85)))))) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X84))),s('type.rat.rat',X85))) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X84))),s('type.rat.rat',X85)),X86,'type.rat.rat'))) )),
  inference(superposition,[],[f175,f179])).
fof(f1288,plain,(
  ( ! [X7] : (s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X7,bool)) | p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0'))))) ) | $spl2),
  inference(forward_demodulation,[],[f1269,f1212])).
fof(f1269,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(bool,bool),X7),s(bool,'T_0')))) | s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X7,bool))) ) | $spl2),
  inference(superposition,[],[f175,f1212])).
fof(f1212,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(unit_resulting_resolution,[],[f229,f147])).
fof(f1061,plain,(
  ( ! [X14,X15] : (p(s(bool,happ(s(fun(bool,bool),X15),s(bool,'F_0')))) | s(bool,X14) = s(bool,sK1(X14,X15,bool)) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f175,f144])).
fof(f1190,plain,(
  ( ! [X24,X23,X22] : (s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,X24))) | s(X22,X24) = s(X22,sK1(X24,X23,X22))) )),
  inference(subsumption_resolution,[],[f1078,f231])).
fof(f1078,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,'F_0')) | s(X22,X24) = s(X22,sK1(X24,X23,X22)) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,X24)))) )),
  inference(superposition,[],[f175,f144])).
fof(f1189,plain,(
  ( ! [X21,X22,X20] : (s(X20,sK2(X22,X20,X21)) != s(X20,sK3(X22,X20,X21)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X20,X21),X22)))) )),
  inference(subsumption_resolution,[],[f1143,f231])).
fof(f1143,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,'F_0')) | s(X20,sK2(X22,X20,X21)) != s(X20,sK3(X22,X20,X21)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X20,X21),X22)))) )),
  inference(superposition,[],[f190,f144])).
fof(f150,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.EQ_CLAUSES|split|2')).
fof(f151,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f103])).
fof(f231,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f138,f135])).
fof(f138,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.NOT_CLAUSES|split|1')).
fof(f698,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(bool,bool),X23),s(bool,X22)))) | s(bool,'F_0') = s(bool,sK1('F_0',X23,bool)) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f175,f144])).
fof(f188,plain,(
  ( ! [X6,X2,X0,X5,X1] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6)) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & ((s(X0,happ(s(fun(X1,X0),X2),s(X1,sK2(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X2,X1,X0)))) & s(X1,sK2(X2,X1,X0)) != s(X1,sK3(X2,X1,X0))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2,sK3])],[f124])).
fof(f124,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X5,X6] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X6))) | s(X1,X5) = s(X1,X6))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(rectify,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4))) & (? [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) & s(X1,X3) != s(X1,X4)) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2))))))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) != s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) | s(X1,X3) = s(X1,X4)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0,X1] : ! [X2] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) <=> ! [X3,X4] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,X3))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) => s(X1,X3) = s(X1,X4)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7,X8] : ! [X5] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X8,X7),X5)))) <=> ! [X9,X10] : (s(X7,happ(s(fun(X8,X7),X5),s(X8,X9))) = s(X7,happ(s(fun(X8,X7),X5),s(X8,X10))) => s(X8,X9) = s(X8,X10)))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.ONE_ONE_DEF')).
fof(f189,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X0,happ(s(fun(X1,X0),X2),s(X1,sK2(X2,X1,X0)))) = s(X0,happ(s(fun(X1,X0),X2),s(X1,sK3(X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f190,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun(X1,X0),X2)))) | s(X1,sK2(X2,X1,X0)) != s(X1,sK3(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f125])).
fof(f175,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) = s(X0,sK1(X2,X1,X0))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0,X1,X2] : (((s(X0,X2) = s(X0,sK1(X2,X1,X0)) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK1(X2,X1,X0)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1])],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2] : ((? [X4] : (s(X0,X2) = s(X0,X4) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : ((? [X3] : (s(X0,X2) = s(X0,X3) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (s(X0,X2) = s(X0,X3) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) <=> p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X8] : ! [X19,X20] : (! [X14] : (s(X8,X14) = s(X8,X20) => p(s(bool,happ(s(fun(X8,bool),X19),s(X8,X14))))) <=> p(s(bool,happ(s(fun(X8,bool),X19),s(X8,X20)))))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.UNWIND_FORALL_THM2')).
fof(f176,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK1(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f177,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | s(X0,X2) != s(X0,X3) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f690,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X1),s(bool,'T_0')))) | s(bool,X0) = s(bool,sK1(X0,X1,bool)) | s(bool,X0) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f175,f246])).
fof(f692,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(bool,bool),X8),s(bool,X7)))) | s(bool,X6) = s(bool,sK1(X6,X8,bool)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f175,f246])).
fof(f696,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,happ(s(fun(bool,bool),X17),s(bool,X16)))) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool)) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f175,f246])).
fof(f702,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X29)))))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X29),s('type.rat.rat',X30))) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X29),s('type.rat.rat',X30)),X31,'type.rat.rat'))) )),
  inference(superposition,[],[f175,f178])).
fof(f685,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f504,f135])).
fof(f504,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f183,f246])).
fof(f684,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f500,f135])).
fof(f500,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f150,f246])).
fof(f486,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f146,f246])).
fof(f422,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f246,f246])).
fof(f282,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f277,f231])).
fof(f277,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f180,f144])).
fof(f246,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f144,f144])).
fof(f232,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f231,f146])).
fof(f221,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f220])).
fof(f220,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f209])).
fof(f209,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X14,X17,X15,X18] : ((s(bool,X14) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.IMP_CONG')).
fof(f191,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f90])).
fof(f90,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_MUL_ONE_ONE',aHL_EXT)).
fof(f180,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11,X12] : ((p(s(bool,X11)) => p(s(bool,X12))) => ((p(s(bool,X12)) => p(s(bool,X11))) => s(bool,X11) = s(bool,X12)))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.IMP_ANTISYM_AX')).
fof(f183,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f87])).
fof(f179,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))),s('type.rat.rat',X1)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X25,X26] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat',X26))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X25))),s('type.rat.rat',X26)))),
  file('thm.rat.RAT_MUL_ONE_ONE','ARITY_1_2_const.rat.rat_mul')).
fof(f178,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X21,X22] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X21)))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.rat.RAT_MUL_COMM')).
fof(f146,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.EQ_CLAUSES|split|1')).
fof(f147,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f101])).
fof(f144,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_MUL_ONE_ONE',aHL_BOOL_CASES)).
fof(f135,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.bool.TRUTH')).
fof(f4144,plain,(
  ( ! [X8] : (s('type.rat.rat',X8) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X8),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl20),
  inference(superposition,[],[f193,f2774])).
fof(f12096,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X0))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f12022,f179])).
fof(f12022,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',sK0) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f11941,f4164])).
fof(f11941,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X0))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | $spl2),
  inference(superposition,[],[f4349,f179])).
fof(f4349,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X0))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f4319,f135])).
fof(f4319,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X0))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) ) | $spl2),
  inference(superposition,[],[f188,f4243])).
fof(f4243,plain,(
  s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | $spl2),
  inference(unit_resulting_resolution,[],[f229,f147])).
fof(f2202,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,f(s(X19,sK4(X20,X21,bool,X19)))) | s(fun(X19,bool),X20) = s(fun(X19,bool),X21) | s(bool,'T_0') = s(bool,g(s(X19,sK4(X20,X21,bool,X19))))) )),
  inference(superposition,[],[f191,f144])).
fof(f12826,plain,(
  $false | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f12596,f12113])).
fof(f12596,plain,(
  ( ! [X467] : (s('type.rat.rat',sK0) != s('type.rat.rat',X467)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4163,f12113])).
fof(f4163,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | ($spl2 | $spl20)),
  inference(backward_demodulation,[],[f4162,f142])).
fof(f12475,plain,(
  ( ! [X850,X851,X849] : (s(bool,'F_0') != s(bool,f(s('type.rat.rat',X851))) | s(fun('type.rat.rat',bool),X849) = s(fun('type.rat.rat',bool),X850) | s(bool,'T_0') = s(bool,g(s('type.rat.rat',sK4(X849,X850,bool,'type.rat.rat'))))) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f2202,f12113])).
fof(f12822,plain,(
  $false | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f12357,f12113])).
fof(f12357,plain,(
  ( ! [X467] : (s('type.rat.rat',sK0) != s('type.rat.rat',X467)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4163,f12113])).
fof(f12820,plain,(
  $false | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f12819,f135])).
fof(f12819,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f12142,f12177])).
fof(f12177,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat',X0),X1)))) ) | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f12113,f1189])).
fof(f12142,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),X0))))) ) | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f188])).
fof(f12818,plain,(
  $false | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f12817,f135])).
fof(f12817,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f12144,f12177])).
fof(f12144,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),X0))))) ) | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f188])).
fof(f12815,plain,(
  $false | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f12814,f135])).
fof(f12814,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f12174,f12177])).
fof(f12174,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),X0))))) ) | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f188])).
fof(f12813,plain,(
  $false | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f12812,f135])).
fof(f12812,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f12176,f12177])).
fof(f12176,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),X0))))) ) | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f188])).
fof(f12171,plain,(
  $false | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f12022])).
fof(f12160,plain,(
  $false | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f4349])).
fof(f12139,plain,(
  $false | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f12022])).
fof(f12128,plain,(
  $false | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113,f4349])).
fof(f12114,plain,(
  $false | ($spl2 | $spl20)),
  inference(unit_resulting_resolution,[],[f4163,f12113])).
fof(f12103,plain,(
  ( ! [X30,X32] : (s('type.rat.rat',X32) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X30))) | s('type.rat.rat',sK0) = s('type.rat.rat',X30)) ) | $spl82),
  inference(cnf_transformation,[],[f12103_D])).
fof(f12103_D,plain,(
  ( ! [X30,X32] : (s('type.rat.rat',X32) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X30))) | s('type.rat.rat',sK0) = s('type.rat.rat',X30)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f12104,plain,(
  ( ! [X31] : (s('type.rat.rat',sK0) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X31))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X31)) ) | $spl84),
  inference(cnf_transformation,[],[f12104_D])).
fof(f12104_D,plain,(
  ( ! [X31] : (s('type.rat.rat',sK0) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X31))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X31)) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f12082,plain,(
  ( ! [X35,X33,X34] : (s('type.rat.rat',sK0) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X34))))) | s('type.rat.rat',X35) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X33))) | s('type.rat.rat',sK0) = s('type.rat.rat',X33) | s('type.rat.rat',sK0) = s('type.rat.rat',X34)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f12022,f8618])).
fof(f8618,plain,(
  ( ! [X39,X40] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X39))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X40))) | s('type.rat.rat',sK0) = s('type.rat.rat',X39) | s('type.rat.rat',sK0) = s('type.rat.rat',X40)) ) | ($spl2 | $spl20)),
  inference(duplicate_literal_removal,[],[f8479])).
fof(f8479,plain,(
  ( ! [X39,X40] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X39))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X40))) | s('type.rat.rat',sK0) = s('type.rat.rat',X39) | s('type.rat.rat',sK0) = s('type.rat.rat',X39) | s('type.rat.rat',sK0) = s('type.rat.rat',X40)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4185,f4901])).
fof(f4901,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X22) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | ($spl2 | $spl20)),
  inference(trivial_inequality_removal,[],[f4893])).
fof(f4893,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) != s('type.rat.rat',sK0) | s('type.rat.rat',sK0) = s('type.rat.rat',X22) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4555,f4475])).
fof(f4475,plain,(
  ( ! [X4] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X4))) | s('type.rat.rat',sK0) = s('type.rat.rat',X4)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4185,f4169])).
fof(f4169,plain,(
  ( ! [X0] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl20)),
  inference(backward_demodulation,[],[f4162,f369])).
fof(f4555,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',sK0) != s('type.rat.rat',X1) | s('type.rat.rat',sK0) = s('type.rat.rat',X0) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) ) | ($spl2 | $spl20)),
  inference(duplicate_literal_removal,[],[f4527])).
fof(f4527,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',sK0) != s('type.rat.rat',X1) | s('type.rat.rat',sK0) = s('type.rat.rat',X0) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',sK0) = s('type.rat.rat',X0)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4165,f4475])).
fof(f4165,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',X2) != s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) | s('type.rat.rat',sK0) = s('type.rat.rat',X1) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2)))) ) | ($spl2 | $spl20)),
  inference(backward_demodulation,[],[f4162,f192])).
fof(f4185,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat',X1))) | s('type.rat.rat',sK0) = s('type.rat.rat',X1)) ) | ($spl2 | $spl20)),
  inference(backward_demodulation,[],[f4162,f2310])).
fof(f12081,plain,(
  ( ! [X30,X31,X32] : (s('type.rat.rat',sK0) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X31))))) | s('type.rat.rat',X32) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X30))) | s('type.rat.rat',sK0) = s('type.rat.rat',X31) | s('type.rat.rat',sK0) = s('type.rat.rat',X30)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f12022,f8618])).
fof(f11909,plain,(
  ( ! [X12,X10] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X12)),X10,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X10),s('type.rat.rat',sK0)))) ) | $spl80),
  inference(cnf_transformation,[],[f11909_D])).
fof(f11909_D,plain,(
  ( ! [X12,X10] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X12)),X10,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X10),s('type.rat.rat',sK0)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f718,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) | $spl6),
  inference(cnf_transformation,[],[f718_D])).
fof(f718_D,plain,(
  ( ! [X11] : (p(s(bool,X11))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f11910,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X14)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X15)),X13,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X13),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f11906,f146])).
fof(f11906,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X14)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X15)),X13,'type.rat.rat')) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X13),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f11789,f246])).
fof(f11789,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X0),s('type.rat.rat',sK0)))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1)),X0,'type.rat.rat'))) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f11788,f4164])).
fof(f11788,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X1))) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X1)),X0,'type.rat.rat')) | p(s(bool,happ(s(fun('type.rat.rat',bool),X0),s('type.rat.rat',sK0))))) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f11787,f4162])).
fof(f11787,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X0),s('type.rat.rat',sK0)))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1))) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X1)),X0,'type.rat.rat'))) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f11717,f4164])).
fof(f11717,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X1)))))) | s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1))) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.rat.rat',X1)),X0,'type.rat.rat'))) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f1787,f4162])).
fof(f11908,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X12)),X10,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X10),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl20)),
  inference(subsumption_resolution,[],[f11905,f146])).
fof(f11905,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X11)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1(happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',X12)),X10,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X10),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X11)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f11789,f246])).
fof(f10522,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool))))) ) | $spl76),
  inference(cnf_transformation,[],[f10522_D])).
fof(f10522_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f10524,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl79),
  inference(cnf_transformation,[],[f10524_D])).
fof(f10524_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f10512,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK4(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2202,f246])).
fof(f10511,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK4(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK4(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2202,f246])).
fof(f10300,plain,(
  ( ! [X83,X81] : (s(bool,sK1('F_0',X81,bool)) = s(bool,sK1(sK1('F_0',X81,bool),X83,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X81),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f10300_D])).
fof(f10300_D,plain,(
  ( ! [X83,X81] : (s(bool,sK1('F_0',X81,bool)) = s(bool,sK1(sK1('F_0',X81,bool),X83,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X81),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f10299,plain,(
  ( ! [X83,X81,X82] : (p(s(bool,X82)) | s(bool,sK1('F_0',X81,bool)) = s(bool,sK1(sK1('F_0',X81,bool),X83,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X81),s(bool,'F_0')))) ) | ($spl46 | $spl66)),
  inference(subsumption_resolution,[],[f10148,f146])).
fof(f10148,plain,(
  ( ! [X83,X81,X82] : (p(s(bool,X82)) | s(bool,sK1('F_0',X81,bool)) = s(bool,sK1(sK1('F_0',X81,bool),X83,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X83),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X81),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X82)) ) | ($spl46 | $spl66)),
  inference(superposition,[],[f9631,f10004])).
fof(f9631,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,sK1(X0,X1,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) ) | $spl46),
  inference(factoring,[],[f6818])).
fof(f6818,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X12)) | s(bool,X12) = s(bool,sK1(X10,X11,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0')))) ) | $spl46),
  inference(subsumption_resolution,[],[f6756,f146])).
fof(f6756,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | p(s(bool,X12)) | s(bool,X12) = s(bool,sK1(X10,X11,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X10)) ) | $spl46),
  inference(superposition,[],[f180,f6555])).
fof(f6555,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,sK1(X0,X1,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(condensation,[],[f6554])).
fof(f6554,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,sK1(X8,X6,bool)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl46),
  inference(subsumption_resolution,[],[f6550,f144])).
fof(f6550,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') != s(bool,X7) | s(bool,X8) = s(bool,sK1(X8,X6,bool)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X6),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X7)) ) | $spl46),
  inference(superposition,[],[f6184,f246])).
fof(f6184,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(bool,bool),X13),s(bool,'F_0'))) | s(bool,X12) = s(bool,sK1(X12,X13,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl46),
  inference(superposition,[],[f6162,f144])).
fof(f10293,plain,(
  ( ! [X61,X63] : (s(bool,'F_0') = s(bool,sK1(sK1('F_0',X61,bool),X63,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X61),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X63),s(bool,'F_0')))) ) | $spl72),
  inference(cnf_transformation,[],[f10293_D])).
fof(f10293_D,plain,(
  ( ! [X61,X63] : (s(bool,'F_0') = s(bool,sK1(sK1('F_0',X61,bool),X63,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X61),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X63),s(bool,'F_0')))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f10292,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,X62)) | s(bool,'F_0') = s(bool,sK1(sK1('F_0',X61,bool),X63,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X63),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X61),s(bool,'F_0')))) ) | ($spl46 | $spl66)),
  inference(subsumption_resolution,[],[f10142,f146])).
fof(f10142,plain,(
  ( ! [X61,X62,X63] : (p(s(bool,X62)) | s(bool,'F_0') = s(bool,sK1(sK1('F_0',X61,bool),X63,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X63),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X61),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X62)) ) | ($spl46 | $spl66)),
  inference(superposition,[],[f6817,f10004])).
fof(f6817,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | s(bool,'F_0') = s(bool,sK1(X8,X9,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'F_0')))) ) | $spl46),
  inference(subsumption_resolution,[],[f6755,f146])).
fof(f6755,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | s(bool,'F_0') = s(bool,sK1(X8,X9,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X9),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) ) | $spl46),
  inference(superposition,[],[f151,f6555])).
fof(f10291,plain,(
  ( ! [X54,X52] : (s(bool,sK1('F_0',X52,bool)) != s(bool,happ(s(fun('type.rat.rat',bool),X54),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0'))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X54,'type.rat.rat'))) ) | $spl70),
  inference(cnf_transformation,[],[f10291_D])).
fof(f10291_D,plain,(
  ( ! [X54,X52] : (s(bool,sK1('F_0',X52,bool)) != s(bool,happ(s(fun('type.rat.rat',bool),X54),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0'))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X54,'type.rat.rat'))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f10290,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X53)) | s(bool,sK1('F_0',X52,bool)) != s(bool,happ(s(fun('type.rat.rat',bool),X54),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X54,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0')))) ) | ($spl14 | $spl66)),
  inference(subsumption_resolution,[],[f10139,f146])).
fof(f10139,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X53)) | s(bool,sK1('F_0',X52,bool)) != s(bool,happ(s(fun('type.rat.rat',bool),X54),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X54,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X52),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X53)) ) | ($spl14 | $spl66)),
  inference(superposition,[],[f2151,f10004])).
fof(f2151,plain,(
  ( ! [X14,X15] : (p(s(bool,X15)) | s(bool,X15) != s(bool,happ(s(fun('type.rat.rat',bool),X14),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X14,'type.rat.rat'))) ) | $spl14),
  inference(subsumption_resolution,[],[f2132,f135])).
fof(f2132,plain,(
  ( ! [X14,X15] : (~p(s(bool,'T_0')) | s(bool,X15) != s(bool,happ(s(fun('type.rat.rat',bool),X14),s('type.rat.rat',sK0))) | p(s(bool,X15)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X14,'type.rat.rat'))) ) | $spl14),
  inference(superposition,[],[f221,f1315])).
fof(f10258,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,sK1('F_0',X3,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) ) | $spl68),
  inference(cnf_transformation,[],[f10258_D])).
fof(f10258_D,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,sK1('F_0',X3,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f10257,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,sK1('F_0',X3,bool)) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0')))) ) | $spl66),
  inference(subsumption_resolution,[],[f10045,f146])).
fof(f10045,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,sK1('F_0',X3,bool)) | p(s(bool,X2)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X3),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl66),
  inference(superposition,[],[f146,f10004])).
fof(f10004,plain,(
  ( ! [X6,X8] : (s(bool,X6) = s(bool,sK1('F_0',X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl66),
  inference(cnf_transformation,[],[f10004_D])).
fof(f10004_D,plain,(
  ( ! [X6,X8] : (s(bool,X6) = s(bool,sK1('F_0',X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f10005,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,sK1('F_0',X11,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl46),
  inference(subsumption_resolution,[],[f9989,f146])).
fof(f9989,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,X9) = s(bool,sK1('F_0',X11,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl46),
  inference(superposition,[],[f9671,f246])).
fof(f9671,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) = s(bool,sK1('F_0',X1,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,'F_0')))) ) | $spl46),
  inference(condensation,[],[f9670])).
fof(f9670,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) = s(bool,sK1('F_0',X29,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0')))) ) | $spl46),
  inference(subsumption_resolution,[],[f9639,f146])).
fof(f9639,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X28) = s(bool,sK1('F_0',X29,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X29),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X27)) ) | $spl46),
  inference(superposition,[],[f6818,f144])).
fof(f10003,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) = s(bool,sK1('F_0',X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl46),
  inference(subsumption_resolution,[],[f9988,f146])).
fof(f9988,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,X6) = s(bool,sK1('F_0',X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl46),
  inference(superposition,[],[f9671,f246])).
fof(f8965,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl60),
  inference(cnf_transformation,[],[f8965_D])).
fof(f8965_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8961,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) | $spl56),
  inference(cnf_transformation,[],[f8961_D])).
fof(f8961_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f8968,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl65),
  inference(cnf_transformation,[],[f8968_D])).
fof(f8968_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f8959,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK4(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2196,f144])).
fof(f8966,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl62),
  inference(cnf_transformation,[],[f8966_D])).
fof(f8966_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f8958,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK4(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f2196,f246])).
fof(f8957,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK4(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f2196,f246])).
fof(f8960,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl54),
  inference(cnf_transformation,[],[f8960_D])).
fof(f8960_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8963,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl59),
  inference(cnf_transformation,[],[f8963_D])).
fof(f8963_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f8956,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK4(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f2196,f246])).
fof(f8955,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f2196,f246])).
fof(f7511,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,X15))) | s(bool,'T_0') = s(bool,sK1('T_0',X14,bool)) | s(bool,'F_0') = s(bool,sK1('F_0',X14,bool))) ) | $spl52),
  inference(cnf_transformation,[],[f7511_D])).
fof(f7511_D,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,X15))) | s(bool,'T_0') = s(bool,sK1('T_0',X14,bool)) | s(bool,'F_0') = s(bool,sK1('F_0',X14,bool))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7512,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool)) | s(bool,'F_0') = s(bool,sK1('F_0',X17,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) ) | $spl46),
  inference(subsumption_resolution,[],[f7492,f146])).
fof(f7492,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X19)) | s(bool,'T_0') = s(bool,sK1('T_0',X17,bool)) | s(bool,'F_0') = s(bool,sK1('F_0',X17,bool)) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X17),s(bool,X18)))) ) | $spl46),
  inference(superposition,[],[f7036,f246])).
fof(f7036,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK1('T_0',X0,bool)) | s(bool,'F_0') = s(bool,sK1('F_0',X0,bool))) ) | $spl46),
  inference(condensation,[],[f7035])).
fof(f7035,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,happ(s(fun(bool,bool),X15),s(bool,X14)))) | s(bool,'F_0') = s(bool,sK1('F_0',X15,bool)) | s(bool,'T_0') = s(bool,sK1('T_0',X15,bool)) | s(bool,'T_0') = s(bool,X16)) ) | $spl46),
  inference(subsumption_resolution,[],[f7017,f6909])).
fof(f6909,plain,(
  ( ! [X12,X13,X11] : (s(bool,X12) != s(bool,X13) | p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X13)))) | s(bool,'F_0') = s(bool,sK1('F_0',X11,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl46),
  inference(subsumption_resolution,[],[f6869,f135])).
fof(f6869,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | s(bool,X12) != s(bool,X13) | p(s(bool,happ(s(fun(bool,bool),X11),s(bool,X13)))) | s(bool,'F_0') = s(bool,sK1('F_0',X11,bool)) | s(bool,'T_0') = s(bool,X12)) ) | $spl46),
  inference(superposition,[],[f177,f6735])).
fof(f6735,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X1),s(bool,X0))) | s(bool,'F_0') = s(bool,sK1('F_0',X1,bool)) | s(bool,'T_0') = s(bool,X0)) ) | $spl46),
  inference(condensation,[],[f6734])).
fof(f6734,plain,(
  ( ! [X6,X8,X7] : (s(bool,'F_0') = s(bool,sK1('F_0',X6,bool)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,X8)) ) | $spl46),
  inference(subsumption_resolution,[],[f6726,f246])).
fof(f6726,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,X8) | s(bool,'F_0') = s(bool,sK1('F_0',X6,bool)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X6),s(bool,X7))) | s(bool,'T_0') = s(bool,X8)) ) | $spl46),
  inference(superposition,[],[f6187,f246])).
fof(f6187,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,happ(s(fun(bool,bool),X21),s(bool,X20))) | s(bool,'F_0') = s(bool,sK1('F_0',X21,bool)) | s(bool,'T_0') = s(bool,X20)) ) | $spl46),
  inference(superposition,[],[f6162,f144])).
fof(f7017,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,happ(s(fun(bool,bool),X15),s(bool,X14)))) | s(bool,'F_0') = s(bool,sK1('F_0',X15,bool)) | s(bool,'T_0') = s(bool,sK1('T_0',X15,bool)) | s(bool,X14) = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | $spl46),
  inference(superposition,[],[f6910,f246])).
fof(f6910,plain,(
  ( ! [X17,X16] : (p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X17)))) | s(bool,'F_0') = s(bool,sK1('F_0',X16,bool)) | s(bool,'T_0') = s(bool,sK1(X17,X16,bool))) ) | $spl46),
  inference(subsumption_resolution,[],[f6871,f135])).
fof(f6871,plain,(
  ( ! [X17,X16] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(bool,bool),X16),s(bool,X17)))) | s(bool,'F_0') = s(bool,sK1('F_0',X16,bool)) | s(bool,'T_0') = s(bool,sK1(X17,X16,bool))) ) | $spl46),
  inference(superposition,[],[f176,f6735])).
fof(f7510,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,sK1('T_0',X14,bool)) | s(bool,'F_0') = s(bool,sK1('F_0',X14,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,X15)))) ) | $spl46),
  inference(subsumption_resolution,[],[f7491,f146])).
fof(f7491,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,sK1('T_0',X14,bool)) | s(bool,'F_0') = s(bool,sK1('F_0',X14,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X14),s(bool,X15))) | s(bool,'T_0') = s(bool,X16)) ) | $spl46),
  inference(superposition,[],[f7036,f246])).
fof(f7038,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,X13))) | s(bool,'F_0') = s(bool,sK1('F_0',X12,bool)) | s(bool,'T_0') = s(bool,sK1(X13,X12,bool))) ) | $spl50),
  inference(cnf_transformation,[],[f7038_D])).
fof(f7038_D,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,X13))) | s(bool,'F_0') = s(bool,sK1('F_0',X12,bool)) | s(bool,'T_0') = s(bool,sK1(X13,X12,bool))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f7039,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s(bool,'F_0') = s(bool,sK1('F_0',X15,bool)) | s(bool,'T_0') = s(bool,sK1(X16,X15,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,X16)))) ) | $spl46),
  inference(subsumption_resolution,[],[f7032,f146])).
fof(f7032,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X17)) | s(bool,'F_0') = s(bool,sK1('F_0',X15,bool)) | s(bool,'T_0') = s(bool,sK1(X16,X15,bool)) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X15),s(bool,X16)))) ) | $spl46),
  inference(superposition,[],[f6910,f246])).
fof(f7037,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,sK1('F_0',X12,bool)) | s(bool,'T_0') = s(bool,sK1(X13,X12,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,X13)))) ) | $spl46),
  inference(subsumption_resolution,[],[f7031,f146])).
fof(f7031,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X14)) | s(bool,'F_0') = s(bool,sK1('F_0',X12,bool)) | s(bool,'T_0') = s(bool,sK1(X13,X12,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X12),s(bool,X13))) | s(bool,'T_0') = s(bool,X14)) ) | $spl46),
  inference(superposition,[],[f6910,f246])).
fof(f6843,plain,(
  ( ! [X6,X8] : (s(bool,'F_0') = s(bool,sK1(X6,X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl48),
  inference(cnf_transformation,[],[f6843_D])).
fof(f6843_D,plain,(
  ( ! [X6,X8] : (s(bool,'F_0') = s(bool,sK1(X6,X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f6844,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,sK1(X9,X11,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X9)) ) | $spl46),
  inference(subsumption_resolution,[],[f6833,f146])).
fof(f6833,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s(bool,'F_0') = s(bool,sK1(X9,X11,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X11),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X9)) ) | $spl46),
  inference(superposition,[],[f6817,f246])).
fof(f6842,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,sK1(X6,X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl46),
  inference(subsumption_resolution,[],[f6832,f146])).
fof(f6832,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,sK1(X6,X8,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X8),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X7)) ) | $spl46),
  inference(superposition,[],[f6817,f246])).
fof(f6161,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,X20)) ) | $spl44),
  inference(cnf_transformation,[],[f6161_D])).
fof(f6161_D,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,X20)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f6162,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,happ(s(fun(X21,bool),X22),s(X21,X23))) | s(X21,X23) = s(X21,sK1(X23,X22,X21))) ) | $spl46),
  inference(cnf_transformation,[],[f6162_D])).
fof(f6162_D,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') != s(bool,happ(s(fun(X21,bool),X22),s(X21,X23))) | s(X21,X23) = s(X21,sK1(X23,X22,X21))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f6156,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') != s(bool,happ(s(fun(X21,bool),X22),s(X21,X23))) | s(bool,'T_0') = s(bool,X20) | s(X21,X23) = s(X21,sK1(X23,X22,X21))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f6120])).
fof(f6120,plain,(
  ( ! [X23,X21,X22,X20] : (s(bool,'F_0') != s(bool,happ(s(fun(X21,bool),X22),s(X21,X23))) | s(bool,'T_0') = s(bool,X20) | s(X21,X23) = s(X21,sK1(X23,X22,X21)) | s(bool,'T_0') = s(bool,X20)) ) | $spl12),
  inference(superposition,[],[f1180,f144])).
fof(f5894,plain,(
  ( ! [X5] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X5,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X5),s('type.rat.rat',sK0)))) ) | $spl42),
  inference(cnf_transformation,[],[f5894_D])).
fof(f5894_D,plain,(
  ( ! [X5] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X5,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X5),s('type.rat.rat',sK0)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f5895,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X7,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X7),s('type.rat.rat',sK0)))) ) | $spl26),
  inference(subsumption_resolution,[],[f5891,f146])).
fof(f5891,plain,(
  ( ! [X8,X7] : (p(s(bool,X8)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X7,'type.rat.rat')) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X7),s('type.rat.rat',sK0)))) ) | $spl26),
  inference(superposition,[],[f5736,f246])).
fof(f5736,plain,(
  ( ! [X86] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X86),s('type.rat.rat',sK0)))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X86,'type.rat.rat'))) ) | $spl26),
  inference(forward_demodulation,[],[f5683,f5590])).
fof(f5683,plain,(
  ( ! [X86] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X86),s('type.rat.rat',sK0)))) | s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X86,'type.rat.rat'))) ) | $spl26),
  inference(superposition,[],[f175,f5590])).
fof(f5893,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X5,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X5),s('type.rat.rat',sK0)))) ) | $spl26),
  inference(subsumption_resolution,[],[f5890,f146])).
fof(f5890,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0)),X5,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X5),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X6)) ) | $spl26),
  inference(superposition,[],[f5736,f246])).
fof(f5614,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')),s('type.rat.rat',X37))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) | $spl40),
  inference(cnf_transformation,[],[f5614_D])).
fof(f5614_D,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')),s('type.rat.rat',X37))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f5590,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | $spl26),
  inference(cnf_transformation,[],[f5590_D])).
fof(f5590_D,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f5613,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')),s('type.rat.rat',X37))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl18 | $spl20)),
  inference(forward_demodulation,[],[f5506,f4169])).
fof(f5506,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',sK0))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')),s('type.rat.rat',X37))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl18 | $spl20)),
  inference(superposition,[],[f4517,f1942])).
fof(f4517,plain,(
  ( ! [X2,X3] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X2))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4475,f178])).
fof(f5612,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')),s('type.rat.rat',X34))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) | $spl38),
  inference(cnf_transformation,[],[f5612_D])).
fof(f5612_D,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')),s('type.rat.rat',X34))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f5611,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')),s('type.rat.rat',X34))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl16 | $spl20)),
  inference(forward_demodulation,[],[f5505,f4169])).
fof(f5505,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X34),s('type.rat.rat',sK0))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')),s('type.rat.rat',X34))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl16 | $spl20)),
  inference(superposition,[],[f4517,f1899])).
fof(f5610,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))),s('type.rat.rat',X25))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) | $spl36),
  inference(cnf_transformation,[],[f5610_D])).
fof(f5610_D,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))),s('type.rat.rat',X25))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5609,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))),s('type.rat.rat',X25))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f5502,f4169])).
fof(f5502,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat',sK0))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))),s('type.rat.rat',X25))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4517,f4517])).
fof(f5608,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))),s('type.rat.rat',X22))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | $spl34),
  inference(cnf_transformation,[],[f5608_D])).
fof(f5608_D,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))),s('type.rat.rat',X22))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f5607,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))),s('type.rat.rat',X22))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f5501,f4169])).
fof(f5501,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',sK0))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))),s('type.rat.rat',X22))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4517,f4475])).
fof(f5596,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) | $spl32),
  inference(cnf_transformation,[],[f5596_D])).
fof(f5596_D,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5595,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl18 | $spl20)),
  inference(forward_demodulation,[],[f5491,f4164])).
fof(f5491,plain,(
  ( ! [X37,X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X37))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl18 | $spl20)),
  inference(superposition,[],[f4517,f1942])).
fof(f5594,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X34),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) | $spl30),
  inference(cnf_transformation,[],[f5594_D])).
fof(f5594_D,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X34),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f5593,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X34),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl16 | $spl20)),
  inference(forward_demodulation,[],[f5490,f4164])).
fof(f5490,plain,(
  ( ! [X33,X34,X32] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X34))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X34),s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X32)),X33,'type.rat.rat')))) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X33),s('type.rat.rat',sK0)))) ) | ($spl2 | $spl16 | $spl20)),
  inference(superposition,[],[f4517,f1899])).
fof(f5592,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) | $spl28),
  inference(cnf_transformation,[],[f5592_D])).
fof(f5592_D,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f5591,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f5487,f4164])).
fof(f5487,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X25))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X23)))) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4517,f4517])).
fof(f5588,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | $spl24),
  inference(cnf_transformation,[],[f5588_D])).
fof(f5588_D,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f5587,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | ($spl2 | $spl20)),
  inference(forward_demodulation,[],[f5486,f4164])).
fof(f5486,plain,(
  ( ! [X21,X22] : (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X22))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK0),s('type.rat.rat',X21))))) | s('type.rat.rat',sK0) = s('type.rat.rat',X21)) ) | ($spl2 | $spl20)),
  inference(superposition,[],[f4517,f4475])).
fof(f4259,plain,(
  $false | ($spl2 | $spl9 | $spl20)),
  inference(subsumption_resolution,[],[f4258,f4164])).
fof(f4258,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9 | $spl20)),
  inference(forward_demodulation,[],[f4257,f4164])).
fof(f4257,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9)),
  inference(forward_demodulation,[],[f4256,f179])).
fof(f4256,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9)),
  inference(forward_demodulation,[],[f4235,f179])).
fof(f4235,plain,(
  s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9)),
  inference(unit_resulting_resolution,[],[f819,f229,f188])).
fof(f4255,plain,(
  $false | ($spl2 | $spl9 | $spl20)),
  inference(subsumption_resolution,[],[f4254,f4164])).
fof(f4254,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9 | $spl20)),
  inference(forward_demodulation,[],[f4253,f4164])).
fof(f4253,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9)),
  inference(forward_demodulation,[],[f4252,f179])).
fof(f4252,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9)),
  inference(forward_demodulation,[],[f4236,f179])).
fof(f4236,plain,(
  s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) != s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | ($spl2 | $spl9)),
  inference(unit_resulting_resolution,[],[f819,f229,f188])).
fof(f4150,plain,(
  $false | ($spl1 | $spl9 | $spl20)),
  inference(subsumption_resolution,[],[f4149,f819])).
fof(f4149,plain,(
  s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) = s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f4114,f4146])).
fof(f4146,plain,(
  s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))),s('type.rat.rat',sK0))) | ($spl1 | $spl20)),
  inference(forward_demodulation,[],[f4119,f2774])).
fof(f4119,plain,(
  s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))),s('type.rat.rat',sK0))) | ($spl1 | $spl20)),
  inference(unit_resulting_resolution,[],[f227,f2774,f193])).
fof(f4114,plain,(
  s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))),s('type.rat.rat',sK0))) | ($spl1 | $spl20)),
  inference(unit_resulting_resolution,[],[f227,f2774,f193])).
fof(f3787,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X4,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) ) | $spl22),
  inference(cnf_transformation,[],[f3787_D])).
fof(f3787_D,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X4,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f3788,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X6,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0')))) ) | $spl2),
  inference(subsumption_resolution,[],[f3780,f146])).
fof(f3780,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X6,bool)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X6),s(bool,'T_0')))) ) | $spl2),
  inference(superposition,[],[f2949,f246])).
fof(f3786,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X4,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0')))) ) | $spl2),
  inference(subsumption_resolution,[],[f3779,f146])).
fof(f3779,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,sK1('const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))),X4,bool)) | s(bool,'T_0') = s(bool,happ(s(fun(bool,bool),X4),s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X5)) ) | $spl2),
  inference(superposition,[],[f2949,f246])).
fof(f2774,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | $spl20),
  inference(cnf_transformation,[],[f2774_D])).
fof(f2774_D,plain,(
  s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2772,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | p(s(bool,X0))) ) | $spl3),
  inference(forward_demodulation,[],[f2771,f179])).
fof(f2771,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | p(s(bool,X0))) ) | $spl3),
  inference(forward_demodulation,[],[f2770,f179])).
fof(f2770,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat'))))) ) | $spl3),
  inference(subsumption_resolution,[],[f2762,f146])).
fof(f2762,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) = s('type.rat.rat',happ(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))),s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f189,f1960])).
fof(f1960,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(subsumption_resolution,[],[f1954,f135])).
fof(f1954,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f230,f246])).
fof(f1942,plain,(
  ( ! [X6,X8] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',sK0)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) | $spl18),
  inference(cnf_transformation,[],[f1942_D])).
fof(f1942_D,plain,(
  ( ! [X6,X8] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',sK0)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1943,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',sK0)),X9,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X9),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1939,f146])).
fof(f1939,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',sK0)),X9,'type.rat.rat')) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X9),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(superposition,[],[f1201,f246])).
fof(f1201,plain,(
  ( ! [X35,X36] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat',sK0)))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat'))) ) | $spl0),
  inference(forward_demodulation,[],[f1199,f228])).
fof(f1199,plain,(
  ( ! [X35,X36] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X35),s('type.rat.rat',sK0)),X36,'type.rat.rat')) | p(s(bool,happ(s(fun('type.rat.rat',bool),X36),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(backward_demodulation,[],[f228,f715])).
fof(f1941,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',sK0)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1938,f146])).
fof(f1938,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',sK0)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f1201,f246])).
fof(f1899,plain,(
  ( ! [X6,X8] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X8)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) | $spl16),
  inference(cnf_transformation,[],[f1899_D])).
fof(f1899_D,plain,(
  ( ! [X6,X8] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X8)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f1900,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X11)),X9,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X9),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1896,f146])).
fof(f1896,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X10)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X11)),X9,'type.rat.rat')) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X9),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(superposition,[],[f1200,f246])).
fof(f1200,plain,(
  ( ! [X28,X27] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X28),s('type.rat.rat',sK0)))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X27)),X28,'type.rat.rat'))) ) | $spl0),
  inference(forward_demodulation,[],[f1198,f228])).
fof(f1198,plain,(
  ( ! [X28,X27] : (s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X27)),X28,'type.rat.rat')) | p(s(bool,happ(s(fun('type.rat.rat',bool),X28),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl0),
  inference(backward_demodulation,[],[f228,f714])).
fof(f1898,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X8)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1895,f146])).
fof(f1895,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X7)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_mul_2'(s('type.rat.rat',sK0),s('type.rat.rat',X8)),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X7)) ) | $spl0),
  inference(superposition,[],[f1200,f246])).
fof(f1315,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X4),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X4,'type.rat.rat'))) ) | $spl14),
  inference(cnf_transformation,[],[f1315_D])).
fof(f1315_D,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X4),s('type.rat.rat',sK0))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X4,'type.rat.rat'))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1316,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1312,f146])).
fof(f1312,plain,(
  ( ! [X6,X7] : (p(s(bool,X7)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X6,'type.rat.rat')) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X6),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(superposition,[],[f1205,f246])).
fof(f1205,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X2),s('type.rat.rat',sK0)))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X2,'type.rat.rat'))) ) | $spl0),
  inference(forward_demodulation,[],[f1204,f228])).
fof(f1204,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun('type.rat.rat',bool),X2),s('type.rat.rat',sK0)))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X2,'type.rat.rat'))) ) | $spl0),
  inference(superposition,[],[f175,f228])).
fof(f1314,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X4,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X4),s('type.rat.rat',sK0)))) ) | $spl0),
  inference(subsumption_resolution,[],[f1311,f146])).
fof(f1311,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1('const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')),X4,'type.rat.rat')) | s(bool,'T_0') = s(bool,happ(s(fun('type.rat.rat',bool),X4),s('type.rat.rat',sK0))) | s(bool,'T_0') = s(bool,X5)) ) | $spl0),
  inference(superposition,[],[f1205,f246])).
fof(f1180,plain,(
  ( ! [X14,X12,X10,X13] : (s(bool,X10) != s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(bool,'T_0') = s(bool,X10) | s(X12,X14) = s(X12,sK1(X14,X13,X12))) ) | $spl12),
  inference(cnf_transformation,[],[f1180_D])).
fof(f1180_D,plain,(
  ( ! [X14,X12,X10,X13] : (s(bool,X10) != s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(bool,'T_0') = s(bool,X10) | s(X12,X14) = s(X12,sK1(X14,X13,X12))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f1181,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) != s(bool,happ(s(fun(X17,bool),X18),s(X17,X19))) | s(X17,X19) = s(X17,sK1(X19,X18,X17)) | s(bool,'T_0') = s(bool,X15)) ) | $spl3),
  inference(subsumption_resolution,[],[f1170,f146])).
fof(f1170,plain,(
  ( ! [X19,X17,X15,X18,X16] : (p(s(bool,X16)) | s(bool,X15) != s(bool,happ(s(fun(X17,bool),X18),s(X17,X19))) | s(X17,X19) = s(X17,sK1(X19,X18,X17)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X15)) ) | $spl3),
  inference(superposition,[],[f774,f246])).
fof(f774,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X33)) | s(bool,X33) != s(bool,happ(s(fun(X30,bool),X31),s(X30,X32))) | s(X30,X32) = s(X30,sK1(X32,X31,X30))) ) | $spl3),
  inference(subsumption_resolution,[],[f751,f135])).
fof(f751,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,'T_0')) | s(bool,X33) != s(bool,happ(s(fun(X30,bool),X31),s(X30,X32))) | p(s(bool,X33)) | s(X30,X32) = s(X30,sK1(X32,X31,X30))) ) | $spl3),
  inference(superposition,[],[f221,f721])).
fof(f721,plain,(
  ( ! [X17,X18,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X18))) | s(X16,X18) = s(X16,sK1(X18,X17,X16))) ) | $spl3),
  inference(subsumption_resolution,[],[f720,f231])).
fof(f720,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'F_0')) | s(X16,X18) = s(X16,sK1(X18,X17,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) ) | $spl3),
  inference(forward_demodulation,[],[f709,f240])).
fof(f240,plain,(
  s(bool,'F_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | $spl3),
  inference(unit_resulting_resolution,[],[f230,f151])).
fof(f709,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) | s(X16,X18) = s(X16,sK1(X18,X17,X16)) | s(bool,'T_0') = s(bool,happ(s(fun(X16,bool),X17),s(X16,X18)))) ) | $spl3),
  inference(superposition,[],[f175,f315])).
fof(f315,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(subsumption_resolution,[],[f313,f231])).
fof(f313,plain,(
  ( ! [X4] : (p(s(bool,'F_0')) | s(bool,X4) = s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | s(bool,'T_0') = s(bool,X4)) ) | $spl3),
  inference(superposition,[],[f282,f240])).
fof(f1179,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X11)) | s(bool,X10) != s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(X12,X14) = s(X12,sK1(X14,X13,X12)) | s(bool,'T_0') = s(bool,X10)) ) | $spl3),
  inference(subsumption_resolution,[],[f1169,f146])).
fof(f1169,plain,(
  ( ! [X14,X12,X10,X13,X11] : (p(s(bool,X11)) | s(bool,X10) != s(bool,happ(s(fun(X12,bool),X13),s(X12,X14))) | s(X12,X14) = s(X12,sK1(X14,X13,X12)) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X11)) ) | $spl3),
  inference(superposition,[],[f774,f246])).
fof(f928,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f833,f135])).
fof(f833,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f232,f718,f183])).
fof(f927,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f835,f135])).
fof(f835,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f233,f718,f183])).
fof(f233,plain,(
  s(bool,'T_0') != s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))) | $spl3),
  inference(unit_resulting_resolution,[],[f230,f146])).
fof(f926,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f837,f718])).
fof(f837,plain,(
  ~p(s(bool,'F_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f232,f718,f183])).
fof(f925,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f839,f718])).
fof(f839,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f233,f718,f183])).
fof(f924,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f843,f135])).
fof(f843,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f232,f718,f183])).
fof(f923,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f845,f135])).
fof(f845,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f233,f718,f183])).
fof(f922,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f847,f718])).
fof(f847,plain,(
  ~p(s(bool,'F_0')) | $spl6),
  inference(unit_resulting_resolution,[],[f232,f718,f183])).
fof(f921,plain,(
  $false | ($spl3 | $spl6)),
  inference(subsumption_resolution,[],[f849,f718])).
fof(f849,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f233,f718,f183])).
fof(f916,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f873,f135])).
fof(f873,plain,(
  ~p(s(bool,'T_0')) | $spl6),
  inference(backward_demodulation,[],[f853,f231])).
fof(f853,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl6),
  inference(unit_resulting_resolution,[],[f718,f147])).
fof(f915,plain,(
  $false | $spl6),
  inference(subsumption_resolution,[],[f914,f718])).
fof(f914,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl6),
  inference(subsumption_resolution,[],[f871,f853])).
fof(f871,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl6),
  inference(backward_demodulation,[],[f853,f150])).
fof(f854,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f232,f718,f147])).
fof(f852,plain,(
  $false | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f240,f718,f150])).
fof(f850,plain,(
  $false | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f718,f233,f718,f183])).
fof(f848,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f718,f232,f718,f183])).
fof(f846,plain,(
  $false | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f135,f233,f718,f183])).
fof(f844,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f135,f232,f718,f183])).
fof(f840,plain,(
  $false | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f718,f233,f718,f183])).
fof(f838,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f718,f232,f718,f183])).
fof(f836,plain,(
  $false | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f135,f233,f718,f183])).
fof(f834,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f135,f232,f718,f183])).
fof(f830,plain,(
  $false | ($spl3 | $spl6)),
  inference(unit_resulting_resolution,[],[f231,f240,f718,f221])).
fof(f826,plain,(
  $false | $spl6),
  inference(unit_resulting_resolution,[],[f231,f718])).
fof(f822,plain,(
  ( ! [X10,X11,X9] : (s(X9,sK2(X11,X9,X10)) != s(X9,sK3(X11,X9,X10)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X9,X10),X11)))) ) | $spl10),
  inference(cnf_transformation,[],[f822_D])).
fof(f822_D,plain,(
  ( ! [X10,X11,X9] : (s(X9,sK2(X11,X9,X10)) != s(X9,sK3(X11,X9,X10)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X9,X10),X11)))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f823,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X16)) | s(X13,sK2(X15,X13,X14)) != s(X13,sK3(X15,X13,X14)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X13,X14),X15)))) )),
  inference(subsumption_resolution,[],[f814,f146])).
fof(f814,plain,(
  ( ! [X14,X15,X13,X16] : (p(s(bool,X16)) | s(X13,sK2(X15,X13,X14)) != s(X13,sK3(X15,X13,X14)) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X13,X14),X15)))) )),
  inference(superposition,[],[f190,f246])).
fof(f821,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,X12)) | s(X9,sK2(X11,X9,X10)) != s(X9,sK3(X11,X9,X10)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X9,X10),X11)))) )),
  inference(subsumption_resolution,[],[f813,f146])).
fof(f813,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,X12)) | s(X9,sK2(X11,X9,X10)) != s(X9,sK3(X11,X9,X10)) | s(bool,'T_0') = s(bool,'const.bool.ONE_ONE_1'(s(fun(X9,X10),X11))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f190,f246])).
fof(f819,plain,(
  s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) != s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) | $spl9),
  inference(cnf_transformation,[],[f819_D])).
fof(f819_D,plain,(
  s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) != s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f817,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) != s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat'))) ) | $spl3),
  inference(subsumption_resolution,[],[f809,f146])).
fof(f809,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.rat.rat',sK2('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) != s('type.rat.rat',sK3('const.rat.rat_mul_1'(s('type.rat.rat',sK0)),'type.rat.rat','type.rat.rat')) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f190,f315])).
fof(f717,plain,(
  ( ! [X10,X8,X9] : (s(X8,X10) = s(X8,sK1(X10,X9,X8)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) ) | $spl4),
  inference(cnf_transformation,[],[f717_D])).
fof(f717_D,plain,(
  ( ! [X10,X8,X9] : (s(X8,X10) = s(X8,sK1(X10,X9,X8)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f719,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(X12,X14) = s(X12,sK1(X14,X13,X12)) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) )),
  inference(subsumption_resolution,[],[f708,f146])).
fof(f708,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,X15)) | s(X12,X14) = s(X12,sK1(X14,X13,X12)) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),X13),s(X12,X14)))) )),
  inference(superposition,[],[f175,f246])).
fof(f716,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X10) = s(X8,sK1(X10,X9,X8)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,X10)))) )),
  inference(subsumption_resolution,[],[f707,f146])).
fof(f707,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,X11)) | s(X8,X10) = s(X8,sK1(X10,X9,X8)) | s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f175,f246])).
fof(f227,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f229,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) | $spl2),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f133,plain,(
  p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) | s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f98])).
fof(f98,plain,(
  (s('type.rat.rat',sK0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0))))))) & (s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0])],[f97])).
fof(f97,plain,(
  ? [X0] : ((s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))))))) & (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0))))))))),
  inference(nnf_transformation,[],[f78])).
fof(f78,plain,(
  ? [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <~> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ~! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ~! [X0] : (~s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X0)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : (~s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X6)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : (~s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',X6)))))))),
  file('thm.rat.RAT_MUL_ONE_ONE','thm.rat.RAT_MUL_ONE_ONE')).
fof(f228,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f230,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) | $spl3),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f134,plain,(
  ~p(s(bool,'const.bool.ONE_ONE_1'(s(fun('type.rat.rat','type.rat.rat'),'const.rat.rat_mul_1'(s('type.rat.rat',sK0)))))) | s('type.rat.rat',sK0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f98])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 235
% Passive clauses: 2441
% Generated clauses: 12904
% Final active clauses: 63
% Final passive clauses: 268
% Input formulas: 33
% Initial clauses: 77
% 
% Duplicate literals: 353
% Trivial inequalities: 196
% Fw subsumption resolutions: 421
% Fw demodulations: 2491
% Bw demodulations: 172
% Condensations: 675
% Global subsumptions: 1
% 
% Simple tautologies: 86
% Equational tautologies: 79
% Forward subsumptions: 6041
% Fw demodulations to eq. taut.: 35
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 556
% Factoring: 2
% Forward superposition: 4146
% Backward superposition: 3359
% Self superposition: 114
% Equality factoring: 39
% Equality resolution: 52
% 
% Split clauses: 54
% Split components: 113
% SAT solver clauses: 2589
% SAT solver unit clauses: 31
% SAT solver binary clauses: 147
% 
% Sat splits: 54
% Sat splitting refutations: 37
% 
% Memory used [KB]: 12537
% Time elapsed: 0.901 s
% ------------------------------
% ------------------------------
% Success in time 2.209 s
