Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f21872,plain,(
  $false),
  inference(subsumption_resolution,[],[f21811,f4707])).
fof(f4707,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4706,f292])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : ! [X8,X9] : s(X11,X9) = s(X11,'const.bool.COND_3'(s(bool,'F_0'),s(X11,X8),s(X11,X9)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.bool.COND_CLAUSES|split|1')).
fof(f4706,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f4705,f438])).
fof(f438,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(resolution,[],[f264,f288])).
fof(f288,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.prim_rec.NOT_LESS_|48|')).
fof(f264,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.bool.EQ_CLAUSES|split|3')).
fof(f4705,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f4662,f249])).
fof(f249,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.arithmetic.SUB_|48||split|1')).
fof(f4662,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f299,f248])).
fof(f248,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.arithmetic.SUB|split|0')).
fof(f299,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.arithmetic.ABS_DIFF_def')).
fof(f21811,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(backward_demodulation,[],[f21810,f240])).
fof(f240,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK12)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f120])).
fof(f120,plain,(
  ? [X0] : s('type.num.num',X0) != s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ~! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.arithmetic.ABS_DIFF_ZERO|split|1')).
fof(f21810,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12)),
  inference(trivial_inequality_removal,[],[f21751])).
fof(f21751,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12) | s('type.num.num','const.num.|48|_0') = s('type.num.num',sK12)),
  inference(superposition,[],[f240,f6308])).
fof(f6308,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(forward_demodulation,[],[f6307,f291])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11] : ! [X8,X9] : s(X11,X8) = s(X11,'const.bool.COND_3'(s(bool,'T_0'),s(X11,X8),s(X11,X9)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.bool.COND_CLAUSES|split|0')).
fof(f6307,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(forward_demodulation,[],[f6306,f249])).
fof(f6306,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(forward_demodulation,[],[f6166,f248])).
fof(f6166,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(superposition,[],[f299,f4262])).
fof(f4262,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(resolution,[],[f282,f260])).
fof(f260,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.bool.EQ_CLAUSES|split|1')).
fof(f282,plain,(
  ( ! [X0] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) & (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)))),
  inference(nnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : (~s('type.num.num','const.num.|48|_0') = s('type.num.num',X6) <=> p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|1','thm.arithmetic.NOT_ZERO_LT_ZERO')).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 751
% Passive clauses: 7147
% Generated clauses: 21257
% Final active clauses: 741
% Final passive clauses: 6024
% Input formulas: 47
% Initial clauses: 174
% 
% Duplicate literals: 609
% Trivial inequalities: 92
% Fw subsumption resolutions: 2210
% Fw demodulations: 959
% Bw demodulations: 9
% 
% Simple tautologies: 263
% Equational tautologies: 57
% Forward subsumptions: 9228
% 
% Binary resolution: 7779
% Factoring: 1312
% Forward superposition: 3723
% Backward superposition: 4038
% Self superposition: 48
% Equality factoring: 30
% Equality resolution: 30
% 
% Split clauses: 571
% Split components: 1156
% SAT solver clauses: 1019
% SAT solver unit clauses: 21
% SAT solver binary clauses: 114
% 
% Sat splits: 571
% 
% TWLsolver clauses: 1013
% TWLsolver calls for satisfiability: 409
% 
% Memory used [KB]: 9594
% Time elapsed: 0.583 s
% ------------------------------
% ------------------------------
% Success in time 0.597 s
