Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (53659)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 159
% Passive clauses: 1676
% Generated clauses: 24658
% Final active clauses: 139
% Final passive clauses: 1383
% Input formulas: 22
% Initial clauses: 28
% 
% Duplicate literals: 1417
% Trivial inequalities: 340
% Fw subsumption resolutions: 3036
% Fw demodulations: 3320
% Bw demodulations: 18
% 
% Simple tautologies: 358
% Equational tautologies: 63
% Forward subsumptions: 12152
% Fw demodulations to eq. taut.: 17
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 68
% Forward superposition: 9282
% Backward superposition: 6450
% Self superposition: 272
% Equality factoring: 168
% Equality resolution: 36
% 
% Split clauses: 1330
% Split components: 2726
% SAT solver clauses: 1915
% SAT solver unit clauses: 24
% SAT solver binary clauses: 172
% 
% Sat splits: 1330
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1908
% TWLsolver calls for satisfiability: 97
% 
% Memory used [KB]: 12153
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (53671)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 452
% Generated clauses: 12303
% Final active clauses: 130
% Final passive clauses: 288
% Input formulas: 22
% Initial clauses: 28
% 
% Duplicate literals: 1704
% Trivial inequalities: 153
% Fw subsumption resolutions: 533
% Fw demodulations: 1195
% Bw demodulations: 29
% Condensations: 2358
% 
% Simple tautologies: 133
% Equational tautologies: 55
% Forward subsumptions: 4695
% Fw demodulations to eq. taut.: 5
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 22
% Forward superposition: 3134
% Backward superposition: 2527
% Self superposition: 188
% Equality factoring: 156
% Equality resolution: 22
% 
% Split clauses: 165
% Split components: 335
% SAT solver clauses: 1059
% SAT solver unit clauses: 36
% SAT solver binary clauses: 88
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 165
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1058
% TWLsolver calls for satisfiability: 183
% 
% Memory used [KB]: 8315
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% (53686)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 486
% Passive clauses: 4000
% Generated clauses: 131091
% Final active clauses: 172
% Final passive clauses: 1016
% Input formulas: 22
% Initial clauses: 28
% Discarded non-redundant clauses: 407
% 
% Duplicate literals: 11672
% Trivial inequalities: 937
% Fw subsumption resolutions: 11198
% Fw demodulations: 12470
% Bw demodulations: 2009
% Condensations: 18414
% Global subsumptions: 317
% 
% Simple tautologies: 1182
% Equational tautologies: 1847
% Forward subsumptions: 67708
% Fw demodulations to eq. taut.: 61
% Bw demodulations to eq. taut.: 51
% 
% Unit resulting resolution: 6140
% Forward superposition: 35937
% Backward superposition: 26893
% Self superposition: 969
% Equality factoring: 589
% Equality resolution: 62
% 
% Split clauses: 895
% Split components: 1795
% SAT solver clauses: 6462
% SAT solver unit clauses: 63
% SAT solver binary clauses: 755
% 
% Sat splits: 895
% Sat splitting refutations: 353
% 
% Memory used [KB]: 52451
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.rat
% Refutation not found, incomplete strategy
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation not found, incomplete strategy

% Active clauses: 14
% Passive clauses: 4
% Generated clauses: 10
% Final active clauses: 13
% Input formulas: 22
% Initial clauses: 16
% 
% Selected by SInE selection: 11
% SInE iterations: 1
% 
% Trivial inequalities: 2
% Fw demodulations: 1
% 
% Simple tautologies: 2
% Equational tautologies: 1
% Forward subsumptions: 2
% 
% Backward superposition: 2
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 4
% SAT solver binary clauses: 3
% 
% Sat splits: 2
% Sat splitting refutations: 1
% 
% Memory used [KB]: 767
% Time elapsed: 0.002 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f7984,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f96,f115,f113,f59,f166,f164,f60,f165,f163,f1089,f1200,f1207,f1425,f1426,f702,f1839,f1938,f1940,f2512,f2944,f3086,f3140,f2507,f3869,f2942,f3913,f3726,f3914,f1199,f4458,f4814,f3085,f2536,f2636,f1838,f7976,f7983])).
fof(f7983,plain,(
  $false | ($spl5 | $spl6)),
  inference(trivial_inequality_removal,[],[f7982])).
fof(f7982,plain,(
  s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) | ($spl5 | $spl6)),
  inference(forward_demodulation,[],[f163,f166])).
fof(f7976,plain,(
  $false | ($spl4 | $spl7)),
  inference(subsumption_resolution,[],[f7975,f165])).
fof(f7975,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat',sK1) | $spl4),
  inference(forward_demodulation,[],[f7958,f7736])).
fof(f7736,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0)))))) )),
  inference(superposition,[],[f7033,f64])).
fof(f64,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : s('type.rat.rat',X0) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X0)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X20] : s('type.rat.rat',X20) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.rat.rep_rat_1'(s('type.rat.rat',X20)))))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.rat.RAT')).
fof(f7033,plain,(
  ( ! [X2] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2)))))))) )),
  inference(forward_demodulation,[],[f6998,f67])).
fof(f67,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))) )),
  inference(cnf_transformation,[],[f29])).
fof(f29,plain,(
  ! [X0] : s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X0))))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X18] : s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X18))))) = s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X18)))))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.rat.RAT_AINV_CALCULATE')).
fof(f6998,plain,(
  ( ! [X2] : (s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac',X2))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.abs_rat_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X2)))))))) )),
  inference(superposition,[],[f67,f6883])).
fof(f6883,plain,(
  ( ! [X0] : (s('type.frac.frac',X0) = s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))) )),
  inference(subsumption_resolution,[],[f6882,f61])).
fof(f61,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.rat_|49||49|3_prove_ltb',aHL_TRUTH)).
fof(f6882,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.frac.frac',X0) = s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0)))))) )),
  inference(forward_demodulation,[],[f6881,f109])).
fof(f109,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) )),
  inference(unit_resulting_resolution,[],[f66,f93])).
fof(f93,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(superposition,[],[f63,f69])).
fof(f69,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.rat_|49||49|3_prove_ltb',aHL_BOOL_CASES)).
fof(f63,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f25])).
fof(f25,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.rat.rat_|49||49|3_prove_ltb',aHL_FALSITY)).
fof(f66,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) )),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X11] : p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X11))))))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.frac.FRAC_DNMPOS')).
fof(f6881,plain,(
  ( ! [X0] : (s('type.frac.frac',X0) = s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) )),
  inference(forward_demodulation,[],[f6880,f68])).
fof(f68,plain,(
  ( ! [X0] : (s('type.frac.frac',X0) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : s('type.frac.frac',X0) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X11] : s('type.frac.frac',X11) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X11))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X11)))))))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.frac.FRAC')).
fof(f6880,plain,(
  ( ! [X0] : (s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) )),
  inference(forward_demodulation,[],[f6836,f65])).
fof(f65,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : s('type.integer.int',X9) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X9)))))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.integer.INT_NEGNEG')).
fof(f6836,plain,(
  ( ! [X0] : (s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))))))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) )),
  inference(superposition,[],[f79,f3356])).
fof(f3356,plain,(
  ( ! [X0] : (s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))) )),
  inference(subsumption_resolution,[],[f3355,f61])).
fof(f3355,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))))) )),
  inference(forward_demodulation,[],[f3333,f109])).
fof(f3333,plain,(
  ( ! [X0] : (s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac',X0))) = s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))) )),
  inference(superposition,[],[f79,f68])).
fof(f79,plain,(
  ( ! [X0,X1] : (s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) | s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1)))) => s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X1))))) = s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X0),s('type.integer.int',X1))))))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X14,X15] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X15)))) => s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X14))),s('type.integer.int',X15))))) = s('type.frac.frac','const.frac.frac_ainv_1'(s('type.frac.frac','const.frac.abs_frac_1'(s('type.pair.prod'('type.integer.int','type.integer.int'),'const.pair.,_2'(s('type.integer.int',X14),s('type.integer.int',X15))))))))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.frac.FRAC_AINV_CALCULATE')).
fof(f7958,plain,(
  s('type.rat.rat',sK1) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))))) | $spl4),
  inference(superposition,[],[f7736,f164])).
fof(f1838,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f1838_D])).
fof(f1838_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2636,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | ~p(s(bool,X42)) | s(bool,X41) = s(bool,X42)) ) | $spl16),
  inference(cnf_transformation,[],[f2636_D])).
fof(f2636_D,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | ~p(s(bool,X42)) | s(bool,X41) = s(bool,X42)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2536,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f2418,f93])).
fof(f2418,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | p(s(bool,X0)) | ~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f673,f69])).
fof(f673,plain,(
  ( ! [X10,X8,X9] : (s(bool,X9) = s(bool,X10) | s(bool,X8) = s(bool,X9) | p(s(bool,X10)) | ~p(s(bool,X8))) )),
  inference(superposition,[],[f624,f93])).
fof(f624,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,X5) | s(bool,X4) = s(bool,X5) | p(s(bool,X4))) )),
  inference(superposition,[],[f61,f91])).
fof(f91,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f69,f69])).
fof(f3085,plain,(
  ( ! [X30] : (~p(s(bool,X30))) ) | $spl18),
  inference(cnf_transformation,[],[f3085_D])).
fof(f3085_D,plain,(
  ( ! [X30] : (~p(s(bool,X30))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4814,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') != s(bool,X41) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42))) ) | $spl28),
  inference(cnf_transformation,[],[f4814_D])).
fof(f4814_D,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') != s(bool,X41) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X42))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4458,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') != s(bool,X41) | ~p(s(bool,X40)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42)) ) | $spl10),
  inference(duplicate_literal_removal,[],[f4359])).
fof(f4359,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') != s(bool,X41) | ~p(s(bool,X40)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,X42) | ~p(s(bool,X40))) ) | $spl10),
  inference(superposition,[],[f1200,f2151])).
fof(f2151,plain,(
  ( ! [X14,X12,X13] : (s(bool,X12) = s(bool,X13) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X14) | ~p(s(bool,X12))) ) | $spl10),
  inference(superposition,[],[f1694,f93])).
fof(f1694,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3)) ) | $spl10),
  inference(superposition,[],[f1200,f69])).
fof(f1199,plain,(
  ( ! [X22] : (p(s(bool,X22))) ) | $spl8),
  inference(cnf_transformation,[],[f1199_D])).
fof(f1199_D,plain,(
  ( ! [X22] : (p(s(bool,X22))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3914,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) != s(bool,X42) | p(s(bool,X42))) ) | $spl26),
  inference(cnf_transformation,[],[f3914_D])).
fof(f3914_D,plain,(
  ( ! [X41,X42] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) != s(bool,X42) | p(s(bool,X42))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3726,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | p(s(bool,X42)) | s(bool,X41) != s(bool,X42) | p(s(bool,X40))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f3606])).
fof(f3606,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | p(s(bool,X42)) | s(bool,X41) != s(bool,X42) | p(s(bool,X40)) | p(s(bool,X40))) ) | $spl3),
  inference(superposition,[],[f1616,f1201])).
fof(f1201,plain,(
  ( ! [X19] : (s(bool,'F_0') = s(bool,X19) | p(s(bool,X19))) )),
  inference(subsumption_resolution,[],[f696,f61])).
fof(f696,plain,(
  ( ! [X19] : (~p(s(bool,'T_0')) | s(bool,'F_0') = s(bool,X19) | p(s(bool,X19))) )),
  inference(superposition,[],[f63,f624])).
fof(f1616,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X2) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) ) | $spl3),
  inference(superposition,[],[f1524,f624])).
fof(f1524,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X5))) ) | $spl3),
  inference(subsumption_resolution,[],[f1519,f63])).
fof(f1519,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,'F_0')) | p(s(bool,X5))) ) | $spl3),
  inference(superposition,[],[f115,f685])).
fof(f685,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | p(s(bool,X8)) | p(s(bool,X9))) )),
  inference(superposition,[],[f61,f624])).
fof(f3913,plain,(
  ( ! [X37,X38] : (s(bool,X37) != s(bool,X38) | s(bool,'T_0') != s(bool,X37) | p(s(bool,X38))) ) | $spl24),
  inference(cnf_transformation,[],[f3913_D])).
fof(f3913_D,plain,(
  ( ! [X37,X38] : (s(bool,X37) != s(bool,X38) | s(bool,'T_0') != s(bool,X37) | p(s(bool,X38))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2942,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,X38) | p(s(bool,X38)) | ~p(s(bool,X36)) | s(bool,'T_0') != s(bool,X37)) ) | ($spl3 | $spl16)),
  inference(duplicate_literal_removal,[],[f2919])).
fof(f2919,plain,(
  ( ! [X37,X38,X36] : (s(bool,X37) != s(bool,X38) | p(s(bool,X38)) | ~p(s(bool,X36)) | ~p(s(bool,X36)) | s(bool,'T_0') != s(bool,X37)) ) | ($spl3 | $spl16)),
  inference(superposition,[],[f1618,f2733])).
fof(f2733,plain,(
  ( ! [X8,X9] : (s(bool,X8) = s(bool,X9) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) ) | ($spl3 | $spl16)),
  inference(superposition,[],[f115,f2636])).
fof(f1618,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl3),
  inference(superposition,[],[f1524,f93])).
fof(f3869,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X38)) | s(bool,X37) = s(bool,X38)) ) | $spl22),
  inference(cnf_transformation,[],[f3869_D])).
fof(f3869_D,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,X37) | ~p(s(bool,X38)) | s(bool,X37) = s(bool,X38)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2507,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | s(bool,X37) = s(bool,X38) | ~p(s(bool,X38))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f2481])).
fof(f2481,plain,(
  ( ! [X37,X38,X36] : (s(bool,'T_0') != s(bool,X37) | p(s(bool,X36)) | s(bool,X37) = s(bool,X38) | p(s(bool,X36)) | ~p(s(bool,X38))) ) | $spl3),
  inference(superposition,[],[f1524,f673])).
fof(f3140,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f3121,f3085])).
fof(f3121,plain,(
  ( ! [X0] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X0))))) )),
  inference(equality_resolution,[],[f81])).
fof(f81,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) | s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0)))))) & (s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) != s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))) | p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1))))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0,X1] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X0),s('type.frac.frac',X1)))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X0))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X1))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X1))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X0))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X18,X19] : (p(s(bool,'const.rat.rat_equiv_2'(s('type.frac.frac',X18),s('type.frac.frac',X19)))) <=> s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X18))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X19))))) = s('type.integer.int','const.integer.int_mul_2'(s('type.integer.int','const.frac.frac_nmr_1'(s('type.frac.frac',X19))),s('type.integer.int','const.frac.frac_dnm_1'(s('type.frac.frac',X18))))))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.rat.rat_equiv_def')).
fof(f3086,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') != s(bool,X31)) ) | $spl20),
  inference(cnf_transformation,[],[f3086_D])).
fof(f3086_D,plain,(
  ( ! [X31] : (s(bool,'F_0') != s(bool,X31) | s(bool,'T_0') != s(bool,X31)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2944,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | ~p(s(bool,X30)) | s(bool,'T_0') != s(bool,X31)) ) | ($spl3 | $spl10 | $spl16)),
  inference(duplicate_literal_removal,[],[f2916])).
fof(f2916,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | ~p(s(bool,X30)) | ~p(s(bool,X30)) | s(bool,'T_0') != s(bool,X31)) ) | ($spl3 | $spl10 | $spl16)),
  inference(superposition,[],[f1200,f2733])).
fof(f2512,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,X42) | p(s(bool,X40)) | ~p(s(bool,X42))) )),
  inference(duplicate_literal_removal,[],[f2428])).
fof(f2428,plain,(
  ( ! [X41,X42,X40] : (s(bool,'F_0') = s(bool,X41) | s(bool,X41) = s(bool,X42) | p(s(bool,X40)) | ~p(s(bool,X42)) | p(s(bool,X40))) )),
  inference(superposition,[],[f673,f1201])).
fof(f1940,plain,(
  $false | ($spl3 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f1939,f1838])).
fof(f1939,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23)) ) | ($spl3 | $spl10 | $spl12)),
  inference(subsumption_resolution,[],[f1843,f1524])).
fof(f1843,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | ($spl10 | $spl12)),
  inference(backward_demodulation,[],[f1838,f1200])).
fof(f1938,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f1840,f61])).
fof(f1840,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(backward_demodulation,[],[f1838,f63])).
fof(f1839,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) ) | $spl14),
  inference(cnf_transformation,[],[f1839_D])).
fof(f1839_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | p(s(bool,X1))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f702,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | p(s(bool,X1))) )),
  inference(equality_factoring,[],[f624])).
fof(f1426,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f63,f1199])).
fof(f1425,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f115,f113])).
fof(f1207,plain,(
  $false | $spl2),
  inference(subsumption_resolution,[],[f1202,f61])).
fof(f1202,plain,(
  ~p(s(bool,'T_0')) | $spl2),
  inference(backward_demodulation,[],[f114,f63])).
fof(f114,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl2),
  inference(cnf_transformation,[],[f114_D])).
fof(f114_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1200,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl10),
  inference(cnf_transformation,[],[f1200_D])).
fof(f1200_D,plain,(
  ( ! [X23] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X23))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1089,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23))) ) | $spl3),
  inference(duplicate_literal_removal,[],[f1078])).
fof(f1078,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | ~p(s(bool,X23)) | p(s(bool,X22))) ) | $spl3),
  inference(superposition,[],[f820,f705])).
fof(f705,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,X16) | p(s(bool,X16))) ) | $spl3),
  inference(trivial_inequality_removal,[],[f694])).
fof(f694,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,X16) | p(s(bool,X16))) ) | $spl3),
  inference(superposition,[],[f115,f624])).
fof(f820,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl3),
  inference(superposition,[],[f728,f93])).
fof(f728,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) | $spl3),
  inference(superposition,[],[f115,f705])).
fof(f163,plain,(
  s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1))) | $spl5),
  inference(cnf_transformation,[],[f163_D])).
fof(f163_D,plain,(
  s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f165,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat',sK1) | $spl7),
  inference(cnf_transformation,[],[f165_D])).
fof(f165_D,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat',sK1) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f60,plain,(
  s('type.rat.rat',sK0) != s('type.rat.rat',sK1) | s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1)))),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1))) | s('type.rat.rat',sK0) = s('type.rat.rat',sK1)) & (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1))) | s('type.rat.rat',sK0) != s('type.rat.rat',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f51])).
fof(f51,plain,(
  ? [X0,X1] : ((s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1))) | s('type.rat.rat',X0) = s('type.rat.rat',X1)) & (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) != s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1))) | s('type.rat.rat',X0) != s('type.rat.rat',X1)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ? [X0,X1] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1))) <~> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(ennf_transformation,[],[f24])).
fof(f24,plain,(
  ~! [X0,X1] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X1))) <=> s('type.rat.rat',X0) = s('type.rat.rat',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X6))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X7))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X6))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',X7))) <=> s('type.rat.rat',X6) = s('type.rat.rat',X7))),
  file('thm.rat.rat_|49||49|3_prove_ltb','thm.rat.rat_|49||49|3_prove_ltb')).
fof(f164,plain,(
  s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1))) | $spl4),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f166,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat',sK1) | $spl6),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat',sK1) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f59,plain,(
  s('type.rat.rat',sK0) = s('type.rat.rat',sK1) | s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK0))) = s('type.rat.rat','const.rat.rat_ainv_1'(s('type.rat.rat',sK1)))),
  inference(cnf_transformation,[],[f52])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | $spl0),
  inference(cnf_transformation,[],[f113_D])).
fof(f113_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f115,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl3),
  inference(cnf_transformation,[],[f115_D])).
fof(f115_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f96,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'F_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f69])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 103
% Passive clauses: 1173
% Generated clauses: 8333
% Final active clauses: 78
% Final passive clauses: 686
% Input formulas: 22
% Initial clauses: 28
% 
% Duplicate literals: 177
% Trivial inequalities: 112
% Fw subsumption resolutions: 755
% Fw demodulations: 1488
% Bw demodulations: 267
% Global subsumptions: 2
% 
% Simple tautologies: 218
% Equational tautologies: 112
% Forward subsumptions: 4238
% Fw demodulations to eq. taut.: 50
% Bw demodulations to eq. taut.: 1
% 
% Unit resulting resolution: 558
% Forward superposition: 2318
% Backward superposition: 2010
% Self superposition: 144
% Equality factoring: 30
% Equality resolution: 16
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 1197
% SAT solver unit clauses: 78
% SAT solver binary clauses: 302
% 
% Sat splits: 12
% Sat splitting refutations: 8
% 
% Memory used [KB]: 5628
% Time elapsed: 0.269 s
% ------------------------------
% ------------------------------
% Success in time 11.988 s
