Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.transc
% (81063)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 160
% Passive clauses: 1838
% Generated clauses: 19026
% Final active clauses: 148
% Final passive clauses: 1496
% Input formulas: 39
% Initial clauses: 55
% 
% Duplicate literals: 1187
% Trivial inequalities: 99
% Fw subsumption resolutions: 2429
% Fw demodulations: 3045
% Bw demodulations: 10
% 
% Simple tautologies: 339
% Equational tautologies: 19
% Forward subsumptions: 9090
% 
% Binary resolution: 330
% Forward superposition: 7693
% Backward superposition: 3859
% Self superposition: 130
% Equality factoring: 32
% Equality resolution: 27
% 
% Split clauses: 527
% Split components: 1060
% SAT solver clauses: 1130
% SAT solver unit clauses: 6
% SAT solver binary clauses: 53
% 
% Sat splits: 527
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1129
% TWLsolver calls for satisfiability: 187
% 
% Memory used [KB]: 13048
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.transc
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.transc
% SZS output start Proof for thm.transc
fof(f4052,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4039,f120,f180,f122,f182,f220,f221,f158,f203,f169,f214,f237,f240,f118,f176,f133,f192,f238,f242,f117,f178,f366,f370,f137,f194,f233,f235,f119,f174])).
fof(f174,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0)))) | $spl1),
  inference(cnf_transformation,[],[f174_D])).
fof(f174_D,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f119,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4])],[f86])).
fof(f86,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X2),s('type.realax.real',X0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X1),s('type.realax.real',X0)))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X4),s(fun('type.realax.real','type.realax.real'),X3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X0))))))),s('type.realax.real',X0)))))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X2),s('type.realax.real',X0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X1),s('type.realax.real',X0))))) & ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X4),s(fun('type.realax.real','type.realax.real'),X3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X0))))))),s('type.realax.real',X0)))))),
  inference(ennf_transformation,[],[f41])).
fof(f41,plain,(
  ~! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X2),s('type.realax.real',X0)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X1),s('type.realax.real',X0))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X4),s(fun('type.realax.real','type.realax.real'),X3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X3),s('type.realax.real',X0))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X1))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X4),s('type.realax.real',X0))))))),s('type.realax.real',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X8),s('type.realax.real',X6)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X9))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X8))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X6))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X7))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X6))))))),s('type.realax.real',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8,X9,X10] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X8),s('type.realax.real',X6)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X9))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X8))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X6))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X7))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X6))))))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|3','thm.transc.DIFF_COMPOSITE|split|3')).
fof(f235,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0)))) | $spl85),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0)))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f233,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0)))) | ($spl1 | $spl42)),
  inference(forward_demodulation,[],[f232,f203])).
fof(f232,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',sK2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0)))) | ($spl1 | $spl42)),
  inference(forward_demodulation,[],[f174,f203])).
fof(f194,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl32),
  inference(cnf_transformation,[],[f194_D])).
fof(f194_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f137,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.transc.DIFF_COMPOSITE|split|3','thm.bool.EQ_CLAUSES|split|2')).
fof(f370,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0))) | $spl110),
  inference(cnf_transformation,[],[f370_D])).
fof(f370_D,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0))) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f366,plain,(
  s(bool,'F_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),sK4),s(fun('type.realax.real','type.realax.real'),sK3))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK0))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',sK1),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK0))))))),s('type.realax.real',sK0))) | ($spl32 | $spl85)),
  inference(resolution,[],[f235,f194])).
fof(f178,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | $spl4),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f117,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f101])).
fof(f242,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))) | $spl88),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0))) | ($spl4 | $spl28)),
  inference(resolution,[],[f192,f178])).
fof(f192,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f133,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.transc.DIFF_COMPOSITE|split|3','thm.bool.EQ_CLAUSES|split|0')).
fof(f176,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | $spl2),
  inference(cnf_transformation,[],[f176_D])).
fof(f176_D,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f118,plain,(
  p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))))),
  inference(cnf_transformation,[],[f101])).
fof(f240,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))) | $spl86),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f237,plain,(
  s(bool,'T_0') = s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0))) | ($spl2 | $spl28)),
  inference(resolution,[],[f192,f176])).
fof(f214,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) | $spl64),
  inference(cnf_transformation,[],[f214_D])).
fof(f214_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f169,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) | p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3,X4] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X2))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10,X9,X8,X7,X6] : ((p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X8),s('type.realax.real',X6)))) & p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X7),s('type.realax.real',X6))))) => p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X10),s(fun('type.realax.real','type.realax.real'),X9))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X8))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X9),s('type.realax.real',X6))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X7))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X10),s('type.realax.real',X6))))))),s('type.realax.real',X6)))))),
  file('thm.transc.DIFF_COMPOSITE|split|3','thm.lim.DIFF_MUL')).
fof(f203,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) | $spl42),
  inference(cnf_transformation,[],[f203_D])).
fof(f203_D,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f158,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X0))),s('type.realax.real',X1)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X13,X14] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) = s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X13))),s('type.realax.real',X14)))),
  file('thm.transc.DIFF_COMPOSITE|split|3','ARITY_0_2_const.realax.real_mul')).
fof(f221,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) | $spl72),
  inference(cnf_transformation,[],[f221_D])).
fof(f221_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f220,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) | ($spl42 | $spl64)),
  inference(forward_demodulation,[],[f219,f203])).
fof(f219,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),'_dst_x0x1x1_3'(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s(fun('type.realax.real','type.realax.real'),X0),s(fun('type.realax.real','type.realax.real'),X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X4))))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),happ(s(fun('type.realax.real',fun('type.realax.real','type.realax.real')),'const.realax.real_mul_0'),s('type.realax.real',X3))),s('type.realax.real',happ(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X4))))))),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X1),s('type.realax.real',X3),s('type.realax.real',X4)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),X0),s('type.realax.real',X2),s('type.realax.real',X4))))) ) | ($spl42 | $spl64)),
  inference(backward_demodulation,[],[f203,f214])).
fof(f182,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f122,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|3',aHL_FALSITY)).
fof(f180,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f120,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.transc.DIFF_COMPOSITE|split|3',aHL_TRUTH)).
fof(f4039,plain,(
  $false | ($spl6 | $spl9 | $spl72 | $spl86 | $spl88 | $spl110)),
  inference(subsumption_resolution,[],[f4038,f180])).
fof(f4038,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl9 | $spl72 | $spl86 | $spl88 | $spl110)),
  inference(forward_demodulation,[],[f4037,f242])).
fof(f4037,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl6 | $spl9 | $spl72 | $spl86 | $spl110)),
  inference(subsumption_resolution,[],[f4036,f180])).
fof(f4036,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl9 | $spl72 | $spl86 | $spl110)),
  inference(forward_demodulation,[],[f4035,f240])).
fof(f4035,plain,(
  ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl9 | $spl72 | $spl110)),
  inference(subsumption_resolution,[],[f3932,f182])).
fof(f3932,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK3),s('type.realax.real',sK1),s('type.realax.real',sK0)))) | ~p(s(bool,'const.lim.diffl_3'(s(fun('type.realax.real','type.realax.real'),sK4),s('type.realax.real',sK2),s('type.realax.real',sK0)))) | ($spl72 | $spl110)),
  inference(superposition,[],[f221,f370])).
% SZS output end Proof for thm.transc
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 74
% Passive clauses: 416
% Generated clauses: 3843
% Final active clauses: 62
% Final passive clauses: 320
% Input formulas: 39
% Initial clauses: 55
% 
% Duplicate literals: 170
% Trivial inequalities: 35
% Fw subsumption resolutions: 216
% Fw demodulations: 460
% Bw demodulations: 33
% Condensations: 145
% Global subsumptions: 4
% 
% Simple tautologies: 95
% Equational tautologies: 18
% Forward subsumptions: 1634
% 
% Binary resolution: 63
% Forward superposition: 1461
% Backward superposition: 676
% Self superposition: 70
% Equality factoring: 30
% Equality resolution: 8
% 
% Split clauses: 35
% Split components: 75
% SAT solver clauses: 708
% SAT solver unit clauses: 79
% SAT solver binary clauses: 29
% SAT solver learnt clauses: 9
% 
% Sat splits: 35
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 691
% TWLsolver calls for satisfiability: 240
% 
% Memory used [KB]: 7547
% Time elapsed: 0.199 s
% ------------------------------
% ------------------------------
% Success in time 0.816 s
