Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quote
% (102450)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 182
% Passive clauses: 4716
% Generated clauses: 23921
% Final active clauses: 163
% Final passive clauses: 4297
% Input formulas: 69
% Initial clauses: 123
% 
% Duplicate literals: 500
% Trivial inequalities: 283
% Fw subsumption resolutions: 1884
% Fw demodulations: 3839
% Bw demodulations: 16
% 
% Simple tautologies: 222
% Equational tautologies: 50
% Forward subsumptions: 11907
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 1177
% Factoring: 1
% Forward superposition: 7030
% Backward superposition: 9006
% Self superposition: 183
% Equality factoring: 66
% Equality resolution: 22
% 
% Split clauses: 282
% Split components: 565
% SAT solver clauses: 751
% SAT solver unit clauses: 7
% SAT solver binary clauses: 186
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 282
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 750
% TWLsolver calls for satisfiability: 174
% 
% Memory used [KB]: 23155
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quote
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quote
% SZS output start Proof for thm.quote
fof(f8762,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f8761,f238,f367,f365,f8423,f8424,f237,f366,f364,f8675,f282,f392,f293,f403])).
fof(f403,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))) ) | $spl56),
  inference(cnf_transformation,[],[f403_D])).
fof(f403_D,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f293,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(flattening,[],[f118])).
fof(f118,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : s(X0,happ(s(fun('type.quote.index',X0),X2),s('type.quote.index',X1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))),s(fun('type.quote.index',X0),X2),s(fun('type.quote.index',X0),X3),s(X0,X4)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9] : ! [X6,X10,X11,X12] : s(X9,happ(s(fun('type.quote.index',X9),X10),s('type.quote.index',X6))) = s(X9,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X6))),s(fun('type.quote.index',X9),X10),s(fun('type.quote.index',X9),X11),s(X9,X12)))),
  file('thm.quote.index_|49||49||split|0','thm.quote.index_case_def|split|0')).
fof(f392,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) ) | $spl34),
  inference(cnf_transformation,[],[f392_D])).
fof(f392_D,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f282,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X25] : ! [X27] : s(X25,X27) = s(X25,happ(s(fun(X25,X25),'_i_0_0'),s(X25,X27)))),
  file('thm.quote.index_|49||49||split|0','_i0_ax')).
fof(f8675,plain,(
  $false | ($spl3 | $spl34 | $spl1732)),
  inference(subsumption_resolution,[],[f8674,f367])).
fof(f8674,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',sK1) | ($spl34 | $spl1732)),
  inference(forward_demodulation,[],[f8527,f392])).
fof(f8527,plain,(
  s('type.quote.index',sK1) = s('type.quote.index',happ(s(fun('type.quote.index','type.quote.index'),'_i_0_0'),s('type.quote.index',sK0))) | ($spl34 | $spl1732)),
  inference(superposition,[],[f8424,f392])).
fof(f364,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) | $spl0),
  inference(cnf_transformation,[],[f364_D])).
fof(f364_D,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f366,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',sK1) | $spl2),
  inference(cnf_transformation,[],[f366_D])).
fof(f366_D,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',sK1) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f237,plain,(
  s('type.quote.index',sK0) = s('type.quote.index',sK1) | s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) | s('type.quote.index',sK0) = s('type.quote.index',sK1)) & (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) | s('type.quote.index',sK0) != s('type.quote.index',sK1))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f195])).
fof(f195,plain,(
  ? [X0,X1] : ((s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) | s('type.quote.index',X0) = s('type.quote.index',X1)) & (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) | s('type.quote.index',X0) != s('type.quote.index',X1)))),
  inference(nnf_transformation,[],[f174])).
fof(f174,plain,(
  ? [X0,X1] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) <~> s('type.quote.index',X0) = s('type.quote.index',X1))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X0))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X1))) <=> s('type.quote.index',X0) = s('type.quote.index',X1))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X6))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X7))) <=> s('type.quote.index',X6) = s('type.quote.index',X7))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X6))) = s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',X7))) <=> s('type.quote.index',X6) = s('type.quote.index',X7))),
  file('thm.quote.index_|49||49||split|0','thm.quote.index_|49||49||split|0')).
fof(f8424,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))) = s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK1)))) ) | $spl1732),
  inference(cnf_transformation,[],[f8424_D])).
fof(f8424_D,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))) = s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK1)))) ) <=> ~$spl1732),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1732])])).
fof(f8423,plain,(
  ( ! [X0,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK0))) = s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK1)))) ) | ($spl0 | $spl56)),
  inference(forward_demodulation,[],[f8237,f403])).
fof(f8237,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun('type.quote.index',X0),X1),s('type.quote.index',sK1))) = s(X0,'const.quote.index_CASE_4'(s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))),s(fun('type.quote.index',X0),X1),s(fun('type.quote.index',X0),X2),s(X0,X3)))) ) | ($spl0 | $spl56)),
  inference(superposition,[],[f403,f364])).
fof(f365,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) | $spl1),
  inference(cnf_transformation,[],[f365_D])).
fof(f365_D,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f367,plain,(
  s('type.quote.index',sK0) != s('type.quote.index',sK1) | $spl3),
  inference(cnf_transformation,[],[f367_D])).
fof(f367_D,plain,(
  s('type.quote.index',sK0) != s('type.quote.index',sK1) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f238,plain,(
  s('type.quote.index',sK0) != s('type.quote.index',sK1) | s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK1)))),
  inference(cnf_transformation,[],[f196])).
fof(f8761,plain,(
  $false | ($spl1 | $spl2)),
  inference(trivial_inequality_removal,[],[f8760])).
fof(f8760,plain,(
  s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) != s('type.quote.index','const.quote.Left_idx_1'(s('type.quote.index',sK0))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f365,f366])).
% SZS output end Proof for thm.quote
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 108
% Passive clauses: 868
% Generated clauses: 8421
% Final active clauses: 96
% Final passive clauses: 722
% Input formulas: 69
% Initial clauses: 123
% 
% Duplicate literals: 196
% Trivial inequalities: 49
% Fw subsumption resolutions: 513
% Fw demodulations: 1220
% Bw demodulations: 47
% Condensations: 385
% Global subsumptions: 6
% 
% Simple tautologies: 133
% Equational tautologies: 41
% Forward subsumptions: 3805
% 
% Binary resolution: 697
% Factoring: 1
% Forward superposition: 1785
% Backward superposition: 2417
% Self superposition: 59
% Equality factoring: 32
% Equality resolution: 8
% 
% Split clauses: 24
% Split components: 48
% SAT solver clauses: 1287
% SAT solver unit clauses: 144
% SAT solver binary clauses: 40
% SAT solver learnt clauses: 14
% SAT solver learnt literals: 1
% 
% Sat splits: 24
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1261
% TWLsolver calls for satisfiability: 454
% 
% Memory used [KB]: 7803
% Time elapsed: 0.296 s
% ------------------------------
% ------------------------------
% Success in time 0.903 s
