Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.rat
% (93838)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1053
% Passive clauses: 5645
% Generated clauses: 40679
% Final active clauses: 1033
% Final passive clauses: 4291
% Input formulas: 517
% Initial clauses: 879
% 
% Duplicate literals: 210
% Trivial inequalities: 50
% Fw subsumption resolutions: 1640
% Fw demodulations: 6642
% Bw demodulations: 19
% 
% Simple tautologies: 297
% Equational tautologies: 25
% Forward subsumptions: 25626
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 6096
% Factoring: 963
% Forward superposition: 14538
% Backward superposition: 9413
% Self superposition: 40
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 368
% Split components: 738
% SAT solver clauses: 515
% SAT solver unit clauses: 30
% SAT solver binary clauses: 83
% 
% Sat splits: 368
% 
% TWLsolver clauses: 515
% TWLsolver calls for satisfiability: 410
% 
% Memory used [KB]: 19189
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (93857)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 656
% Passive clauses: 2861
% Generated clauses: 30163
% Final active clauses: 638
% Final passive clauses: 2030
% Input formulas: 517
% Initial clauses: 879
% 
% Duplicate literals: 100
% Trivial inequalities: 52
% Fw subsumption resolutions: 879
% Fw demodulations: 3408
% Bw demodulations: 188
% Condensations: 102
% Global subsumptions: 38
% 
% Simple tautologies: 273
% Equational tautologies: 17
% Forward subsumptions: 18830
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 2522
% Factoring: 306
% Forward superposition: 10935
% Backward superposition: 7635
% Self superposition: 85
% Equality factoring: 2
% Equality resolution: 12
% 
% Split clauses: 6
% Split components: 12
% SAT solver clauses: 4513
% SAT solver unit clauses: 1214
% SAT solver binary clauses: 231
% SAT solver learnt clauses: 76
% SAT solver learnt literals: 16
% 
% Sat splits: 6
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 4361
% TWLsolver calls for satisfiability: 5116
% 
% Memory used [KB]: 24306
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f177138,plain,(
  $false),
  inference(subsumption_resolution,[],[f156476,f157219])).
fof(f157219,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f157218,f2406])).
fof(f2406,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))) )),
  inference(cnf_transformation,[],[f1101])).
fof(f1101,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(flattening,[],[f1100])).
fof(f1100,plain,(
  ! [X0,X1,X2] : ! [X3,X4,X5] : s(X0,happ(s(fun(X1,X0),'_dst_x1x0_2'(s(fun(X1,fun(X2,X0)),X3),s(X2,X4))),s(X1,X5))) = s(X0,happ(s(fun(X2,X0),happ(s(fun(X1,fun(X2,X0)),X3),s(X1,X5))),s(X2,X4)))),
  inference(rectify,[],[f380])).
fof(f380,axiom,(
  ! [X11,X13,X12] : ! [X14,X15,X16] : s(X11,happ(s(fun(X12,X11),happ(s(fun(X13,fun(X12,X11)),X14),s(X13,X16))),s(X12,X15))) = s(X11,happ(s(fun(X13,X11),'_dst_x1x0_2'(s(fun(X13,fun(X12,X11)),X14),s(X12,X15))),s(X13,X16)))),
  file('thm.rat.rat_of_num_def_compute|split|3','_dst_x1x0_ax')).
fof(f157218,plain,(
  ( ! [X2,X0] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,X2))))) )),
  inference(forward_demodulation,[],[f157217,f1923])).
fof(f1923,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f770])).
fof(f770,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f769])).
fof(f769,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f278])).
fof(f278,axiom,(
  ! [X24] : ! [X53,X54] : s(X24,X54) = s(X24,'const.bool.COND_3'(s(bool,'F_0'),s(X24,X53),s(X24,X54)))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.bool.COND_CLAUSES|split|1')).
fof(f157217,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),'_dst_x1x0_2'(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X2))),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2))))))) )),
  inference(forward_demodulation,[],[f156680,f2406])).
fof(f156680,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2))))),s(X0,X2))))) )),
  inference(unit_resulting_resolution,[],[f1923,f2288])).
fof(f2288,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f1522])).
fof(f1522,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) & (s(X0,X1) != s(X0,X2) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2))))))),
  inference(nnf_transformation,[],[f1076])).
fof(f1076,plain,(
  ! [X0,X1,X2] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) <=> s(X0,X1) = s(X0,X2))),
  inference(rectify,[],[f153])).
fof(f153,axiom,(
  ! [X21,X19,X20] : (p(s(bool,happ(s(fun(X21,bool),happ(s(fun(X21,fun(X21,bool)),'=_0'),s(X21,X19))),s(X21,X20)))) <=> s(X21,X19) = s(X21,X20))),
  file('thm.rat.rat_of_num_def_compute|split|3',arity_0_HL_EQ)).
fof(f156476,plain,(
  ~p(s(bool,happ(s(fun('type.rat.rat',bool),happ(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))))),
  inference(forward_demodulation,[],[f156475,f2406])).
fof(f156475,plain,(
  ~p(s(bool,happ(s(fun('type.rat.rat',bool),'_dst_x1x0_2'(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))))),
  inference(forward_demodulation,[],[f156474,f58685])).
fof(f58685,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))),s('type.rat.rat','const.rat.rat_|49|_0')))) )),
  inference(forward_demodulation,[],[f58626,f2557])).
fof(f2557,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f1693,f1681])).
fof(f1681,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f587])).
fof(f587,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f276])).
fof(f276,axiom,(
  ! [X27] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X27))))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27)))))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.numRing.num_rewrites|split|46')).
fof(f1693,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f599])).
fof(f599,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f212])).
fof(f212,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))) = s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.numeral.numeral_iisuc|split|1')).
fof(f58626,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))),s('type.rat.rat','const.rat.rat_|49|_0')))) )),
  inference(superposition,[],[f2608,f1655])).
fof(f1655,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f561])).
fof(f561,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f166])).
fof(f166,axiom,(
  ! [X6] : s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X6)))))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.numeral.numeral_suc|split|1')).
fof(f2608,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num',X0))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.rat.rat','const.rat.rat_|49|_0')))) )),
  inference(forward_demodulation,[],[f1726,f1666])).
fof(f1666,plain,(
  ( ! [X0] : (s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f572])).
fof(f572,plain,(
  ! [X0] : s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f267])).
fof(f267,axiom,(
  ! [X6] : s('type.num.num','const.numeral.iiSUC_1'(s('type.num.num',X6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.numeral.iiSUC')).
fof(f1726,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.rat.rat','const.rat.rat_|49|_0')))) )),
  inference(cnf_transformation,[],[f632])).
fof(f632,plain,(
  ! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.rat.rat_of_num_def|split|2')).
fof(f156474,plain,(
  ~p(s(bool,happ(s(fun('type.rat.rat',bool),'_dst_x1x0_2'(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))),s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))),s('type.rat.rat','const.rat.rat_|49|_0'))))))),
  inference(forward_demodulation,[],[f156268,f2406])).
fof(f156268,plain,(
  ~p(s(bool,happ(s(fun('type.rat.rat',bool),happ(s(fun('type.rat.rat',fun('type.rat.rat',bool)),'=_0'),s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))),s('type.rat.rat','const.rat.rat_|49|_0'))))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))))),
  inference(unit_resulting_resolution,[],[f2524,f2287])).
fof(f2287,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),'=_0'),s(X0,X1))),s(X0,X2)))) | s(X0,X1) = s(X0,X2)) )),
  inference(cnf_transformation,[],[f1522])).
fof(f2524,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(backward_demodulation,[],[f1620,f1578])).
fof(f1578,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(cnf_transformation,[],[f1302])).
fof(f1302,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK17))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f1149])).
fof(f1149,plain,(
  ? [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) != s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))))) = s('type.rat.rat','const.rat.rat_add_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))))),s('type.rat.rat','const.rat.rat_|49|_0')))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.rat.rat_of_num_def_compute|split|3')).
fof(f1620,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X28] : s('type.num.num',X28) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28)))),
  file('thm.rat.rat_of_num_def_compute|split|3','thm.arithmetic.NUMERAL_DEF')).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1774
% Passive clauses: 28410
% Generated clauses: 182747
% Final active clauses: 1176
% Final passive clauses: 5194
% Input formulas: 517
% Initial clauses: 879
% Discarded non-redundant clauses: 20202
% 
% Duplicate literals: 13
% Trivial inequalities: 98
% Fw subsumption resolutions: 4801
% Fw demodulations: 34845
% Bw demodulations: 10378
% Condensations: 39
% 
% Simple tautologies: 175
% Equational tautologies: 100
% Forward subsumptions: 106966
% Fw demodulations to eq. taut.: 1267
% Bw demodulations to eq. taut.: 61
% 
% Unit resulting resolution: 49515
% Factoring: 3
% Forward superposition: 37644
% Backward superposition: 34671
% Self superposition: 192
% Equality factoring: 2
% Equality resolution: 31
% 
% Split clauses: 35
% Split components: 70
% SAT solver clauses: 32361
% SAT solver unit clauses: 10625
% SAT solver binary clauses: 16987
% 
% Sat splits: 35
% Sat splitting refutations: 3879
% 
% Memory used [KB]: 86352
% Time elapsed: 8.038 s
% ------------------------------
% ------------------------------
% Success in time 9.4 s
