Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (112222)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 121
% Passive clauses: 5515
% Generated clauses: 19074
% Final active clauses: 109
% Final passive clauses: 5178
% Input formulas: 12
% Initial clauses: 21
% 
% Duplicate literals: 1066
% Trivial inequalities: 144
% Fw subsumption resolutions: 2674
% Fw demodulations: 1040
% Bw demodulations: 7
% 
% Simple tautologies: 327
% Equational tautologies: 82
% Forward subsumptions: 6788
% 
% Binary resolution: 445
% Factoring: 2
% Forward superposition: 7892
% Backward superposition: 5174
% Self superposition: 232
% Equality factoring: 181
% Equality resolution: 18
% 
% Split clauses: 669
% Split components: 1371
% SAT solver clauses: 1621
% SAT solver unit clauses: 35
% SAT solver binary clauses: 272
% 
% Sat splits: 669
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 1604
% TWLsolver calls for satisfiability: 225
% 
% Memory used [KB]: 12281
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (112237)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 71
% Passive clauses: 857
% Generated clauses: 8274
% Final active clauses: 68
% Final passive clauses: 774
% Input formulas: 12
% Initial clauses: 21
% 
% Duplicate literals: 716
% Trivial inequalities: 65
% Fw subsumption resolutions: 789
% Fw demodulations: 335
% Bw demodulations: 14
% Condensations: 1109
% 
% Simple tautologies: 118
% Equational tautologies: 49
% Forward subsumptions: 2964
% 
% Binary resolution: 287
% Factoring: 2
% Forward superposition: 2114
% Backward superposition: 1767
% Self superposition: 112
% Equality factoring: 108
% Equality resolution: 9
% 
% Split clauses: 92
% Split components: 187
% SAT solver clauses: 1320
% SAT solver unit clauses: 24
% SAT solver binary clauses: 28
% 
% Sat splits: 92
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1320
% TWLsolver calls for satisfiability: 160
% 
% Memory used [KB]: 8699
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (112258)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 664
% Passive clauses: 9652
% Generated clauses: 159337
% Final active clauses: 106
% Final passive clauses: 820
% Input formulas: 12
% Initial clauses: 21
% Discarded non-redundant clauses: 1482
% 
% Duplicate literals: 8959
% Trivial inequalities: 784
% Fw subsumption resolutions: 11576
% Fw demodulations: 28342
% Bw demodulations: 6119
% Condensations: 14242
% Global subsumptions: 1413
% 
% Simple tautologies: 1408
% Equational tautologies: 1200
% Forward subsumptions: 77997
% Fw demodulations to eq. taut.: 80
% Bw demodulations to eq. taut.: 209
% 
% Unit resulting resolution: 7115
% Factoring: 1
% Forward superposition: 36926
% Backward superposition: 34509
% Self superposition: 707
% Equality factoring: 384
% Equality resolution: 77
% 
% Split clauses: 548
% Split components: 1113
% SAT solver clauses: 13630
% SAT solver unit clauses: 34
% SAT solver binary clauses: 824
% 
% Sat splits: 548
% Sat splitting refutations: 1467
% 
% Memory used [KB]: 82642
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% (112526)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 586
% Passive clauses: 33784
% Generated clauses: 91142
% Final active clauses: 391
% Final passive clauses: 28880
% Input formulas: 12
% Initial clauses: 21
% 
% Selected by SInE selection: 12
% SInE iterations: 1
% 
% Duplicate literals: 2755
% Trivial inequalities: 1556
% Fw subsumption resolutions: 13056
% Fw demodulations: 4617
% Bw demodulations: 114
% 
% Simple tautologies: 2437
% Equational tautologies: 96
% Forward subsumptions: 31310
% Bw demodulations to eq. taut.: 14
% 
% Binary resolution: 14932
% Factoring: 9
% Forward superposition: 34972
% Backward superposition: 17656
% Self superposition: 200
% Equality factoring: 114
% Equality resolution: 94
% 
% Split clauses: 2958
% Split components: 6099
% SAT solver clauses: 3009
% SAT solver unit clauses: 16
% SAT solver binary clauses: 320
% 
% Sat splits: 2958
% Sat splitting refutations: 26
% 
% Memory used [KB]: 41321
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.arithmetic
% (112632)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 119
% Passive clauses: 2084
% Generated clauses: 14210
% Final active clauses: 72
% Final passive clauses: 1182
% Input formulas: 12
% Initial clauses: 21
% 
% Duplicate literals: 607
% Trivial inequalities: 334
% Fw subsumption resolutions: 1914
% Fw demodulations: 609
% Bw demodulations: 517
% Global subsumptions: 14
% 
% Simple tautologies: 688
% Equational tautologies: 99
% Forward subsumptions: 7658
% Bw demodulations to eq. taut.: 23
% 
% Unit resulting resolution: 522
% Forward superposition: 5303
% Backward superposition: 3227
% Self superposition: 95
% Equality factoring: 48
% Equality resolution: 14
% 
% Split clauses: 12
% Split components: 24
% SAT solver clauses: 2156
% SAT solver unit clauses: 31
% SAT solver binary clauses: 129
% 
% Sat splits: 12
% Sat splitting refutations: 44
% 
% Memory used [KB]: 8571
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.arithmetic
% (112655)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 254
% Passive clauses: 4484
% Generated clauses: 16270
% Final active clauses: 123
% Final passive clauses: 2260
% Input formulas: 12
% Initial clauses: 21
% 
% Duplicate literals: 649
% Trivial inequalities: 323
% Fw subsumption resolutions: 1015
% Bw subsumption resolutions: 339
% Fw demodulations: 793
% Bw demodulations: 864
% Condensations: 850
% Global subsumptions: 1
% 
% Simple tautologies: 524
% Equational tautologies: 54
% Forward subsumptions: 7437
% Backward subsumptions: 545
% 
% Binary resolution: 1885
% Factoring: 4
% Forward superposition: 5338
% Backward superposition: 2302
% Self superposition: 92
% Equality factoring: 78
% Equality resolution: 52
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 4651
% SAT solver unit clauses: 12
% SAT solver binary clauses: 126
% 
% Sat splits: 7
% Sat splitting refutations: 156
% 
% Memory used [KB]: 13816
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.arithmetic
% (112727)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 148
% Passive clauses: 2967
% Generated clauses: 8778
% Final active clauses: 77
% Final passive clauses: 786
% Input formulas: 12
% Initial clauses: 21
% 
% Duplicate literals: 387
% Trivial inequalities: 85
% Fw demodulations: 794
% Bw demodulations: 47
% Global subsumptions: 58
% 
% Simple tautologies: 153
% Equational tautologies: 94
% Forward subsumptions: 4037
% Backward subsumptions: 1512
% 
% Binary resolution: 960
% Unit resulting resolution: 247
% Factoring: 1
% Forward superposition: 2342
% Backward superposition: 2575
% Self superposition: 57
% Equality factoring: 54
% Equality resolution: 9
% 
% Split clauses: 147
% Split components: 299
% SAT solver clauses: 3409
% SAT solver unit clauses: 6
% SAT solver binary clauses: 139
% 
% Sat splits: 147
% Sat splitting refutations: 11
% 
% Memory used [KB]: 15095
% Time elapsed: 1.0000 s
% ------------------------------
% ------------------------------
% remaining time: 109 next slice time: 85
lrs+1_5:4_cond=on:fsr=off:fde=none:gs=on:gsem=on:lwlo=on:nm=64:nwc=1:stl=60:sos=all:spl=off_80 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f85156,plain,(
  $false),
  inference(resolution,[],[f85153,f43])).
fof(f43,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.NRC_|49|',aHL_TRUTH)).
fof(f85153,plain,(
  ~p(s(bool,'T_0'))),
  inference(duplicate_literal_removal,[],[f85152])).
fof(f85152,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'T_0'))),
  inference(forward_demodulation,[],[f85113,f20342])).
fof(f20342,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2),s(X0,X2)))) )),
  inference(equality_resolution,[],[f5489])).
fof(f5489,plain,(
  ( ! [X21,X19,X22,X20] : (s(X19,X21) != s(X19,X22) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(X19,fun(X19,bool)),X20),s('type.num.num','const.num.|48|_0'),s(X19,X21),s(X19,X22)))) )),
  inference(resolution,[],[f5369,f54])).
fof(f54,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2),s(X0,X3)))) | s(X0,X2) != s(X0,X3)) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2),s(X0,X3)))) | s(X0,X2) = s(X0,X3)) & (s(X0,X2) != s(X0,X3) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f22])).
fof(f22,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2),s(X0,X3)))) <=> s(X0,X2) = s(X0,X3))),
  inference(flattening,[],[f21])).
fof(f21,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2),s(X0,X3)))) <=> s(X0,X2) = s(X0,X3))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ! [X9,X8,X7] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X9),s('type.num.num','const.num.|48|_0'),s(X6,X8),s(X6,X7)))) <=> s(X6,X7) = s(X6,X8))),
  file('thm.arithmetic.NRC_|49|','thm.arithmetic.NRC|split|0')).
fof(f5369,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f5368])).
fof(f5368,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1) | ~p(s(bool,X1))) )),
  inference(condensation,[],[f5367])).
fof(f5367,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X10) | ~p(s(bool,X10))) )),
  inference(global_subsumption,[],[f4191,f3804])).
fof(f3804,plain,(
  ( ! [X37,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | s(bool,'T_0') = s(bool,X36)) )),
  inference(global_subsumption,[],[f3783,f2027,f167,f1248])).
fof(f1248,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))))) )),
  inference(global_subsumption,[],[f1177,f183])).
fof(f183,plain,(
  ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(resolution,[],[f79,f63])).
fof(f63,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f48])).
fof(f48,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X6] : ! [X8] : (s(X6,X8) = s(X6,X8) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.NRC_|49|','thm.bool.REFL_CLAUSE')).
fof(f79,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))))) )),
  inference(superposition,[],[f64,f61])).
fof(f61,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.arithmetic.NRC_|49|','thm.bool.IMP_ANTISYM_AX')).
fof(f64,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(forward_demodulation,[],[f42,f45])).
fof(f45,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.arithmetic.NRC_|49|','thm.arithmetic.ONE')).
fof(f42,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK0,sK2),s(sK0,sK1)))),
  inference(cnf_transformation,[],[f32])).
fof(f32,plain,(
  s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(sK0,sK2),s(sK0,sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f28])).
fof(f28,plain,(
  ? [X0,X1,X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1))) != s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X1)))),
  inference(ennf_transformation,[],[f15])).
fof(f15,plain,(
  ~! [X0,X1,X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X1)))),
  inference(flattening,[],[f14])).
fof(f14,plain,(
  ~! [X0] : ! [X1,X2,X3] : s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X3),s(X0,X2))),s(X0,X1))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X3),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X0,X2),s(X0,X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X7))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X9),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X6,X8),s(X6,X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X7))) = s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X9),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s(X6,X8),s(X6,X7)))),
  file('thm.arithmetic.NRC_|49|','thm.arithmetic.NRC_|49|')).
fof(f1177,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | ~p(s(bool,X2)) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))))) )),
  inference(superposition,[],[f127,f61])).
fof(f127,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(trivial_inequality_removal,[],[f124])).
fof(f124,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(superposition,[],[f82,f47])).
fof(f47,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.NRC_|49|',aHL_BOOL_CASES)).
fof(f82,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(superposition,[],[f64,f47])).
fof(f167,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))),
  inference(equality_resolution,[],[f77])).
fof(f77,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X0))) )),
  inference(superposition,[],[f64,f58])).
fof(f58,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f30])).
fof(f2027,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X5)) )),
  inference(global_subsumption,[],[f1567,f1697,f2008,f780,f635,f79,f1823])).
fof(f1823,plain,(
  ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(trivial_inequality_removal,[],[f1789])).
fof(f1789,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))),
  inference(resolution,[],[f99,f63])).
fof(f99,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') != s(bool,X2) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))) )),
  inference(superposition,[],[f81,f61])).
fof(f81,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(superposition,[],[f64,f47])).
fof(f635,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,X0) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) )),
  inference(resolution,[],[f83,f54])).
fof(f83,plain,(
  ( ! [X4,X5] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X5),s('type.num.num','const.num.|48|_0'),s(bool,X4),s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))))) | s(bool,X4) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) )),
  inference(superposition,[],[f64,f53])).
fof(f53,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(X0,X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f780,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))) )),
  inference(resolution,[],[f167,f61])).
fof(f2008,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X5)) | s(bool,'F_0') = s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))) )),
  inference(superposition,[],[f142,f61])).
fof(f142,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X9)) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f77,f47])).
fof(f1697,plain,(
  ( ! [X8,X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,X7) | s(bool,X8) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | ~p(s(bool,X8))) )),
  inference(resolution,[],[f1690,f61])).
fof(f1690,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | p(s(bool,X8)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(global_subsumption,[],[f1681,f1566,f807,f1415,f1229,f1022,f801])).
fof(f801,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X7),s('type.num.num','const.num.|48|_0'),s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))),s(bool,X6))))) )),
  inference(superposition,[],[f167,f53])).
fof(f1022,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))),
  inference(resolution,[],[f183,f167])).
fof(f1229,plain,(
  ( ! [X43,X44] : (s(bool,'T_0') != s(bool,X43) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(bool,fun(bool,bool)),X44),s('type.num.num','const.num.|48|_0'),s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))),s(bool,X43)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) )),
  inference(superposition,[],[f54,f127])).
fof(f1415,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) )),
  inference(factoring,[],[f1308])).
fof(f1308,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'T_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(resolution,[],[f1250,f167])).
fof(f1250,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | p(s(bool,X2))) )),
  inference(resolution,[],[f1248,f58])).
fof(f807,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(global_subsumption,[],[f806,f492,f167])).
fof(f492,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))))),
  inference(global_subsumption,[],[f458,f183,f127])).
fof(f458,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'F_0') = s(bool,X2) | s(bool,X2) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))))) )),
  inference(resolution,[],[f424,f79])).
fof(f424,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(trivial_inequality_removal,[],[f416])).
fof(f416,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X9)) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f350,f47])).
fof(f350,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) )),
  inference(resolution,[],[f317,f63])).
fof(f317,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | p(s(bool,X7)) | s(bool,X6) != s(bool,X7)) )),
  inference(global_subsumption,[],[f293,f161,f167,f183,f127])).
fof(f161,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X5)) | s(bool,X4) != s(bool,X5) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))) )),
  inference(superposition,[],[f77,f61])).
fof(f293,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | p(s(bool,'F_0')) | p(s(bool,X7)) | ~p(s(bool,X6)) | ~p(s(bool,'T_0'))) )),
  inference(superposition,[],[f173,f61])).
fof(f173,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | p(s(bool,'F_0')) | p(s(bool,X1))) )),
  inference(superposition,[],[f128,f58])).
fof(f128,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(global_subsumption,[],[f123,f82])).
fof(f123,plain,(
  s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'F_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(superposition,[],[f82,f47])).
fof(f806,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) )),
  inference(duplicate_literal_removal,[],[f776])).
fof(f776,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) )),
  inference(resolution,[],[f167,f633])).
fof(f633,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f618])).
fof(f618,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,X1) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f459,f47])).
fof(f459,plain,(
  ( ! [X4,X3] : (~p(s(bool,X4)) | s(bool,'F_0') = s(bool,X3) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) )),
  inference(resolution,[],[f424,f61])).
fof(f1566,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,X8) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f635,f47])).
fof(f1681,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | p(s(bool,X8)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(trivial_inequality_removal,[],[f1668])).
fof(f1668,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | p(s(bool,X8)) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f97,f47])).
fof(f97,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | p(s(bool,X0))) )),
  inference(superposition,[],[f81,f58])).
fof(f1567,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,X9) != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f635,f47])).
fof(f3783,plain,(
  ( ! [X37,X36] : (~p(s(bool,X36)) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X37)) | s(bool,'F_0') != s(bool,X37) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f161,f47])).
fof(f4191,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,'T_0') = s(bool,X10) | ~p(s(bool,X10)) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f4101,f61])).
fof(f4101,plain,(
  ( ! [X41,X42] : (s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,X41)) )),
  inference(resolution,[],[f3971,f63])).
fof(f3971,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(resolution,[],[f3926,f61])).
fof(f3926,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f3869])).
fof(f3869,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X1) | s(bool,X1) = s(bool,X2) | p(s(bool,X2))) )),
  inference(resolution,[],[f3804,f58])).
fof(f85113,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.|48|_0'),s(sK0,sK1),s(sK0,sK1))))),
  inference(superposition,[],[f62371,f61636])).
fof(f61636,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(global_subsumption,[],[f60725,f58780,f1823,f58936])).
fof(f58936,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))),
  inference(global_subsumption,[],[f5423,f58933,f58780,f1022])).
fof(f58933,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK5(sK1,sK2,'const.num.|48|_0',sK3,sK0))))),
  inference(condensation,[],[f58858])).
fof(f58858,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | p(s(bool,X15)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK5(sK1,sK2,'const.num.|48|_0',sK3,sK0))))) )),
  inference(resolution,[],[f773,f3804])).
fof(f773,plain,(
  p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK5(sK1,sK2,'const.num.|48|_0',sK3,sK0))))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))),
  inference(resolution,[],[f167,f55])).
fof(f55,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,sK5(X4,X3,X2,X1,X0))))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,sK5(X4,X3,X2,X1,X0))))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,sK5(X4,X3,X2,X1,X0)),s(X0,X4)))))) & (! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X5),s(X0,X4))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK5])],[f39])).
fof(f39,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4)))) | ? [X6] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X6)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X6),s(X0,X4)))))) & (! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X5),s(X0,X4))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4))))))),
  inference(rectify,[],[f38])).
fof(f38,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4)))) | ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X5),s(X0,X4)))))) & (! [X5] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X5),s(X0,X4))))) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4))))))),
  inference(nnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4)))) <=> ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X5),s(X0,X4))))))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0] : ! [X1,X2,X3,X4] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4)))) <=> ? [X5] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X5),s(X0,X4))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X9,X10,X8,X7] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s(X6,X8),s(X6,X7)))) <=> ? [X11] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X9),s(X6,X8))),s(X6,X11)))) & p(s(bool,'const.arithmetic.NRC_4'(s(fun(X6,fun(X6,bool)),X9),s('type.num.num',X10),s(X6,X11),s(X6,X7))))))),
  file('thm.arithmetic.NRC_|49|','thm.arithmetic.NRC|split|1')).
fof(f5423,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK5(sK1,sK2,'const.num.|48|_0',sK3,sK0))))),
  inference(global_subsumption,[],[f5376,f1823])).
fof(f5376,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK5(sK1,sK2,'const.num.|48|_0',sK3,sK0)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))),
  inference(resolution,[],[f95,f774])).
fof(f774,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.|48|_0'),s(sK0,sK5(sK1,sK2,'const.num.|48|_0',sK3,sK0)),s(sK0,sK1)))) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))))),
  inference(resolution,[],[f167,f56])).
fof(f56,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,sK5(X4,X3,X2,X1,X0)),s(X0,X4)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f40])).
fof(f95,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),X1),s('type.num.num','const.num.|48|_0'),s(sK0,X0),s(sK0,sK1)))) | s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,X0)))) )),
  inference(superposition,[],[f81,f53])).
fof(f58780,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(condensation,[],[f58714])).
fof(f58714,plain,(
  ( ! [X55] : (s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | s(bool,'F_0') != s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1))) | s(bool,'T_0') = s(bool,X55)) )),
  inference(resolution,[],[f2234,f17345])).
fof(f17345,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(condensation,[],[f17344])).
fof(f17344,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f17343])).
fof(f17343,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | ~p(s(bool,X1)) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(condensation,[],[f17247])).
fof(f17247,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X42) = s(bool,X43) | s(bool,'T_0') = s(bool,X43) | s(bool,'F_0') != s(bool,X44) | ~p(s(bool,X44)) | s(bool,'T_0') = s(bool,X45) | s(bool,X42) = s(bool,X45)) )),
  inference(resolution,[],[f9076,f3926])).
fof(f9076,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X10)) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12) | s(bool,'F_0') != s(bool,X11) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f7921,f61])).
fof(f7921,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f7861])).
fof(f7861,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f2494])).
fof(f2494,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) = s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'F_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f2276,f47])).
fof(f2276,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,'F_0') = s(bool,X17) | s(bool,'F_0') = s(bool,X18)) )),
  inference(global_subsumption,[],[f2241,f2104,f2080])).
fof(f2080,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(condensation,[],[f2079])).
fof(f2079,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f2029])).
fof(f2029,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | p(s(bool,X2))) )),
  inference(resolution,[],[f2027,f58])).
fof(f2104,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0)) )),
  inference(resolution,[],[f2080,f61])).
fof(f2241,plain,(
  ( ! [X17,X18] : (s(bool,X17) = s(bool,X18) | s(bool,'F_0') = s(bool,X17) | p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) | s(bool,'F_0') = s(bool,X18)) )),
  inference(resolution,[],[f2104,f1690])).
fof(f2234,plain,(
  p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))),
  inference(condensation,[],[f2233])).
fof(f2233,plain,(
  ( ! [X1] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) )),
  inference(global_subsumption,[],[f2201,f164,f2080])).
fof(f164,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) )),
  inference(superposition,[],[f77,f47])).
fof(f2201,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f164,f47])).
fof(f60725,plain,(
  ( ! [X9] : (p(s(bool,'F_0')) | s(bool,'F_0') != s(bool,X9) | p(s(bool,X9)) | s(bool,'T_0') = s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1)))) )),
  inference(backward_demodulation,[],[f60586,f164])).
fof(f60586,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s(sK0,sK2),s(sK0,sK1)))),
  inference(resolution,[],[f58936,f1823])).
fof(f62371,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,X0)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.|48|_0'),s(sK0,X0),s(sK0,sK1))))) )),
  inference(trivial_inequality_removal,[],[f61739])).
fof(f61739,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.|48|_0'),s(sK0,X0),s(sK0,sK1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,X0))))) )),
  inference(backward_demodulation,[],[f61636,f1023])).
fof(f1023,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,sK1))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.num.num','const.num.|48|_0'),s(sK0,X0),s(sK0,sK1)))) | ~p(s(bool,happ(s(fun(sK0,bool),happ(s(fun(sK0,fun(sK0,bool)),sK3),s(sK0,sK2))),s(sK0,X0))))) )),
  inference(resolution,[],[f183,f57])).
fof(f57,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s(X0,X3),s(X0,X4)))) | ~p(s(bool,'const.arithmetic.NRC_4'(s(fun(X0,fun(X0,bool)),X1),s('type.num.num',X2),s(X0,X5),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f40])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 644
% Passive clauses: 13227
% Generated clauses: 78452
% Final active clauses: 354
% Final passive clauses: 5194
% Input formulas: 12
% Initial clauses: 21
% Discarded non-redundant clauses: 15715
% 
% Duplicate literals: 3672
% Trivial inequalities: 1655
% Fw demodulations: 3582
% Bw demodulations: 1667
% Condensations: 7731
% Global subsumptions: 1110
% 
% Simple tautologies: 2577
% Equational tautologies: 84
% Forward subsumptions: 44731
% Bw demodulations to eq. taut.: 285
% 
% Binary resolution: 8904
% Factoring: 4
% Forward superposition: 38865
% Backward superposition: 10951
% Self superposition: 168
% Equality factoring: 140
% Equality resolution: 91
% 
% SAT solver clauses: 18123
% SAT solver unit clauses: 1708
% SAT solver binary clauses: 1272
% SAT solver learnt clauses: 2259
% SAT solver learnt literals: 6963
% 
% TWLsolver clauses: 14337
% TWLsolver calls for satisfiability: 370469
% 
% Memory used [KB]: 44903
% Time elapsed: 7.189 s
% ------------------------------
% ------------------------------
% Success in time 26.208 s
