Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (100183)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 2320
% Generated clauses: 20667
% Final active clauses: 123
% Final passive clauses: 2021
% Input formulas: 20
% Initial clauses: 39
% 
% Duplicate literals: 1180
% Trivial inequalities: 236
% Fw subsumption resolutions: 3218
% Fw demodulations: 2243
% Bw demodulations: 7
% 
% Simple tautologies: 437
% Equational tautologies: 65
% Forward subsumptions: 8948
% 
% Binary resolution: 405
% Forward superposition: 5806
% Backward superposition: 6856
% Self superposition: 224
% Equality factoring: 175
% Equality resolution: 26
% 
% Split clauses: 1403
% Split components: 2889
% SAT solver clauses: 2060
% SAT solver unit clauses: 29
% SAT solver binary clauses: 286
% 
% Sat splits: 1403
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2048
% TWLsolver calls for satisfiability: 164
% 
% Memory used [KB]: 9594
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (100209)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 89
% Passive clauses: 472
% Generated clauses: 10266
% Final active clauses: 81
% Final passive clauses: 354
% Input formulas: 20
% Initial clauses: 39
% 
% Duplicate literals: 1170
% Trivial inequalities: 99
% Fw subsumption resolutions: 726
% Fw demodulations: 583
% Bw demodulations: 31
% Condensations: 1632
% Global subsumptions: 4
% 
% Simple tautologies: 202
% Equational tautologies: 44
% Forward subsumptions: 3544
% 
% Binary resolution: 159
% Factoring: 1
% Forward superposition: 2416
% Backward superposition: 2680
% Self superposition: 157
% Equality factoring: 124
% Equality resolution: 14
% 
% Split clauses: 127
% Split components: 262
% SAT solver clauses: 1037
% SAT solver unit clauses: 43
% SAT solver binary clauses: 61
% SAT solver learnt clauses: 8
% 
% Sat splits: 127
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1021
% TWLsolver calls for satisfiability: 183
% 
% Memory used [KB]: 7419
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f16059,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f171,f175,f173,f263,f267,f265,f809,f804,f805,f816,f814,f806,f807,f817,f808,f813,f811,f895,f890,f891,f902,f900,f892,f893,f903,f894,f899,f897,f983,f984,f985,f986,f987,f1197,f1200,f1850,f1799,f1861,f1859,f1864,f2006,f2008,f2014,f2071,f2073,f2081,f2082,f2083,f2085,f2876,f2920,f2981,f2983,f2991,f3182,f3184,f3186,f3187,f3188,f3249,f3251,f3259,f3260,f3261,f3262,f3263,f3264,f3265,f3266,f3267,f3268,f3269,f3270,f3271,f3272,f3273,f3274,f3275,f3276,f3277,f3278,f3279,f3280,f3281,f3282,f3283,f3284,f3285,f3286,f3287,f3288,f3289,f3290,f3291,f3292,f3293,f3294,f3295,f3296,f3297,f3298,f3299,f3300,f3301,f3302,f3303,f3304,f3305,f3306,f3307,f3308,f3309,f3310,f3311,f3312,f3949,f3952,f3950,f4619,f4161,f4631,f4629,f4175,f4633,f4609,f4596,f4595,f4634,f4590,f4643,f4589,f4644,f4588,f4579,f4628,f4670,f4672,f4671,f4553,f4642,f4546,f4705,f4544,f4706,f4543,f4707,f4542,f4708,f4539,f4711,f4132,f4131,f4130,f994,f4129,f988,f4128,f993,f4733,f1199,f1198,f4735,f1196,f996,f4734,f11552,f11537,f11519,f11514,f11513,f11603,f11512,f11511,f11510,f11606,f11607,f11602,f11499,f11560,f11490,f11604,f11489,f11605,f11458,f11765,f11452,f11770,f11450,f11771,f11448,f11772,f11445,f11775,f11444,f11776,f11438,f11437,f11436,f11431,f11790,f11426,f11799,f11821,f11822,f11823,f11848,f11849,f11780,f11853,f11781,f11863,f11779,f12227,f12229,f12230,f12231,f12233,f12235,f12237,f12239,f12241,f12243,f12245,f12246,f12247,f12255,f12257,f11382,f12258,f11361,f12260,f5469,f991,f989,f12369,f12371,f12372,f12373,f12381,f12383,f11368,f12387,f4803,f4800,f1217,f1215,f13924,f13938,f14491,f14492,f14514,f14513,f15603,f15607,f15609,f15611,f4627,f15612,f992,f15613,f15950,f15952,f15976,f15978,f15980,f15982,f15984,f15988,f15995,f16037,f16038,f16039,f16040,f16041,f16042,f16043,f16044,f16045,f16046,f16047,f16048,f16049,f16050,f16051,f16052,f16053,f16054,f16056,f16058])).
fof(f16058,plain,(
  $false | ($spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987,f15975,f15951,f16055,f16057])).
fof(f16057,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl14)),
  inference(subsumption_resolution,[],[f16034,f79])).
fof(f16034,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X0))) ) | ($spl12 | $spl14)),
  inference(condensation,[],[f15999])).
fof(f15999,plain,(
  ( ! [X6,X5] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,X6)) ) | ($spl12 | $spl14)),
  inference(superposition,[],[f15949,f127])).
fof(f16055,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl14)),
  inference(subsumption_resolution,[],[f16035,f79])).
fof(f16035,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X0))) ) | ($spl12 | $spl14)),
  inference(condensation,[],[f15998])).
fof(f15998,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) ) | ($spl12 | $spl14)),
  inference(superposition,[],[f15949,f127])).
fof(f15951,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f15949,f82])).
fof(f15975,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | ($spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f114,f15949,f129])).
fof(f15987,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f15949,f79])).
fof(f15994,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0)) ) | ($spl12 | $spl14 | $spl62)),
  inference(unit_resulting_resolution,[],[f15949,f14909])).
fof(f15949,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl12 | $spl14)),
  inference(subsumption_resolution,[],[f15948,f74])).
fof(f15948,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0))) ) | ($spl12 | $spl14)),
  inference(forward_demodulation,[],[f15947,f815])).
fof(f15947,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))))) ) | $spl12),
  inference(subsumption_resolution,[],[f15906,f78])).
fof(f15906,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(superposition,[],[f288,f814])).
fof(f15892,plain,(
  ( ! [X50,X51] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),'const.pred_set.UNIV_0'),s(fun(bool,bool),X51)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X50))),s(fun(bool,bool),X51)))) | s(bool,'F_0') = s(bool,X50)) ) | $spl62),
  inference(superposition,[],[f288,f4634])).
fof(f15886,plain,(
  ( ! [X43,X42] : (p(s(bool,'const.bool.IN_2'(s(bool,X42),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),'const.pred_set.UNIV_0'),s(fun(bool,bool),X43)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X43)))) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f288,f77])).
fof(f15929,plain,(
  ( ! [X35,X33,X34] : (s(bool,'T_0') = s(bool,X35) | p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),'const.pred_set.UNIV_0'),s(fun(bool,bool),X34)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X34)))) | s(bool,X33) = s(bool,X35)) )),
  inference(forward_demodulation,[],[f15883,f122])).
fof(f15883,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,'const.bool.IN_2'(s(bool,X33),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),'const.pred_set.UNIV_0'),s(fun(bool,bool),X34)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X34)))) | s(bool,X35) = s(bool,'const.bool.IN_2'(s(bool,X35),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X33) = s(bool,X35)) )),
  inference(superposition,[],[f288,f457])).
fof(f15872,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),'const.pred_set.UNIV_0'),s(fun(bool,bool),X7)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X7)))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f288,f127])).
fof(f15871,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),'const.pred_set.UNIV_0'),s(fun(bool,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X3)))) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f288,f127])).
fof(f15870,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),'const.pred_set.UNIV_0'),s(fun(bool,bool),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f288,f77])).
fof(f14982,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,X19) | ~p(s(bool,g(s(bool,X18)))) | s(bool,X18) = s(bool,X19)) ) | $spl62),
  inference(forward_demodulation,[],[f14981,f122])).
fof(f14981,plain,(
  ( ! [X19,X18] : (~p(s(bool,g(s(bool,X18)))) | s(bool,X19) = s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X18) = s(bool,X19)) ) | $spl62),
  inference(subsumption_resolution,[],[f14939,f114])).
fof(f14939,plain,(
  ( ! [X19,X18] : (~p(s(bool,g(s(bool,X18)))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X19) = s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X18) = s(bool,X19)) ) | $spl62),
  inference(superposition,[],[f14909,f457])).
fof(f14906,plain,(
  ( ! [X52,X50,X51] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X51),s(fun(bool,bool),X52)))))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X50))),s(fun(bool,bool),X51)))) | s(bool,'F_0') = s(bool,X50)) ) | $spl62),
  inference(superposition,[],[f108,f4634])).
fof(f14905,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X48),s(fun(bool,bool),X49)))))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X47))),s(fun(bool,bool),X49)))) | s(bool,'F_0') = s(bool,X47)) ) | $spl62),
  inference(superposition,[],[f109,f4634])).
fof(f14904,plain,(
  ( ! [X45,X46,X44] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X45)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X44))),s(fun(bool,bool),X46)))) | p(s(bool,'const.bool.IN_2'(s(bool,g(s(bool,X44))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X45),s(fun(bool,bool),X46)))))) | s(bool,'F_0') = s(bool,X44)) ) | $spl62),
  inference(superposition,[],[f110,f4634])).
fof(f14903,plain,(
  ( ! [X43,X42] : (s(bool,'F_0') != s(bool,f(s(bool,sK3(X42,X43,bool,bool)))) | s(fun(bool,bool),X42) = s(fun(bool,bool),X43) | s(bool,'F_0') = s(bool,sK3(X42,X43,bool,bool))) ) | $spl62),
  inference(superposition,[],[f111,f4634])).
fof(f14918,plain,(
  ( ! [X19,X18] : (s(bool,X19) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X18)) ) | $spl62),
  inference(subsumption_resolution,[],[f14889,f76])).
fof(f14889,plain,(
  ( ! [X19,X18] : (p(s(bool,'F_0')) | s(bool,X19) = s(bool,g(s(bool,X18))) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') = s(bool,X18)) ) | $spl62),
  inference(superposition,[],[f449,f4634])).
fof(f14909,plain,(
  ( ! [X3] : (~p(s(bool,g(s(bool,X3)))) | s(bool,'F_0') = s(bool,X3)) ) | $spl62),
  inference(trivial_inequality_removal,[],[f14881])).
fof(f14881,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,'F_0') | ~p(s(bool,g(s(bool,X3)))) | s(bool,'F_0') = s(bool,X3)) ) | $spl62),
  inference(superposition,[],[f82,f4634])).
fof(f14915,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | $spl62),
  inference(subsumption_resolution,[],[f14879,f76])).
fof(f14879,plain,(
  ( ! [X1] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,g(s(bool,X1))) | s(bool,'F_0') = s(bool,X1)) ) | $spl62),
  inference(superposition,[],[f78,f4634])).
fof(f14914,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X11) = s(bool,g(s(bool,X10)))) ) | $spl62),
  inference(subsumption_resolution,[],[f14872,f114])).
fof(f14872,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X10) | s(bool,X11) = s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X11) = s(bool,g(s(bool,X10)))) ) | $spl62),
  inference(superposition,[],[f4634,f457])).
fof(f1569,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1562,f74])).
fof(f1562,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X3)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f184,f127])).
fof(f1800,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1792,f74])).
fof(f1792,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X4))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f212,f127])).
fof(f5732,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X5))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X5,X6,bool)),s(fun(bool,bool),X6))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK2(X5,X6,bool))) )),
  inference(superposition,[],[f258,f127])).
fof(f6039,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X7),s(fun(bool,bool),X6))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X5,X6,bool)),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK2(X5,X6,bool))) )),
  inference(superposition,[],[f261,f127])).
fof(f6299,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f6293,f74])).
fof(f6293,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X5),s(fun(bool,bool),X7)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f275,f127])).
fof(f6507,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X3),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f276,f77])).
fof(f6521,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),'const.pred_set.UNIV_0')))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f6513,f74])).
fof(f6513,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),'const.pred_set.UNIV_0')))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f276,f122])).
fof(f6522,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(subsumption_resolution,[],[f6515,f74])).
fof(f6515,plain,(
  ( ! [X6,X8,X7,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),X7)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X7),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X6) | s(bool,X8) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f276,f127])).
fof(f11152,plain,(
  ( ! [X420,X418,X419] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X420),s(fun(bool,bool),X419))) | s(fun(bool,bool),X418) = s(fun(bool,bool),X419) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X418,X419,bool)),s(fun(bool,bool),X418))) | s(bool,X420) = s(bool,'const.bool.IN_2'(s(bool,X420),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,sK2(X418,X419,bool))) )),
  inference(superposition,[],[f261,f457])).
fof(f11153,plain,(
  ( ! [X422,X421,X423] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,X423),s(fun(bool,bool),X421))) | s(fun(bool,bool),X421) = s(fun(bool,bool),X422) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X421,X422,bool)),s(fun(bool,bool),X422))) | s(bool,X423) = s(bool,'const.bool.IN_2'(s(bool,X423),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,sK2(X421,X422,bool))) )),
  inference(superposition,[],[f258,f457])).
fof(f1325,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f1310,f74])).
fof(f1310,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f183,f127])).
fof(f5436,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f5421,f74])).
fof(f5421,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5))))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f213,f127])).
fof(f5731,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X2))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X2,X3,bool)),s(fun(bool,bool),X3))) | s(bool,X4) = s(bool,sK2(X2,X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f258,f127])).
fof(f6038,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X3))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X2,X3,bool)),s(fun(bool,bool),X2))) | s(bool,X4) = s(bool,sK2(X2,X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f261,f127])).
fof(f11362,plain,(
  ( ! [X391,X390,X392] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X391))) | s(fun(bool,bool),X390) = s(fun(bool,bool),X391) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X390,X391,bool)),s(fun(bool,bool),X390))) | s(bool,X392) = s(bool,'const.bool.IN_2'(s(bool,X392),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X392) = s(bool,sK2(X390,X391,bool))) )),
  inference(superposition,[],[f261,f457])).
fof(f11363,plain,(
  ( ! [X393,X395,X394] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X393))) | s(fun(bool,bool),X393) = s(fun(bool,bool),X394) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X393,X394,bool)),s(fun(bool,bool),X394))) | s(bool,X395) = s(bool,'const.bool.IN_2'(s(bool,X395),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X395) = s(bool,sK2(X393,X394,bool))) )),
  inference(superposition,[],[f258,f457])).
fof(f5543,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,'T_0') = s(bool,sK2(X1,X2,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) )),
  inference(superposition,[],[f259,f77])).
fof(f5557,plain,(
  ( ! [X2,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2))) | s(fun(bool,bool),X1) = s(fun(bool,bool),X2) | s(bool,'T_0') = s(bool,sK2(X2,X1,bool)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) )),
  inference(superposition,[],[f259,f77])).
fof(f5730,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X0,X1,bool)),s(fun(bool,bool),X1))) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool))) )),
  inference(superposition,[],[f258,f77])).
fof(f6037,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,sK2(X0,X1,bool)),s(fun(bool,bool),X0))) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool))) )),
  inference(superposition,[],[f261,f77])).
fof(f14603,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X22) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X20)))) | s(bool,X19) = s(bool,X22)) )),
  inference(forward_demodulation,[],[f13011,f122])).
fof(f13011,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X20)))) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X19) = s(bool,X22)) )),
  inference(superposition,[],[f108,f457])).
fof(f14601,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),X29)))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(forward_demodulation,[],[f13013,f122])).
fof(f13013,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),X29)))) | s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f108,f457])).
fof(f14587,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X83)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X83),s(fun(bool,bool),X84)))))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(forward_demodulation,[],[f13032,f122])).
fof(f13032,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X83),s(fun(bool,bool),X84)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),'const.pred_set.UNIV_0'))),s(fun(bool,bool),X83)))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f108,f457])).
fof(f14585,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X22) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | s(bool,X19) = s(bool,X22)) )),
  inference(forward_demodulation,[],[f13282,f122])).
fof(f13282,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X19) = s(bool,X22)) )),
  inference(superposition,[],[f109,f457])).
fof(f14583,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),X30)))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(forward_demodulation,[],[f13284,f122])).
fof(f13284,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),X30)))) | s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f109,f457])).
fof(f14568,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X88)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X86),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X87),s(fun(bool,bool),X88)))))) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(forward_demodulation,[],[f13305,f122])).
fof(f13305,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'const.bool.IN_2'(s(bool,X86),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X87),s(fun(bool,bool),X88)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(bool,X86),s(fun(bool,bool),'const.pred_set.UNIV_0'))),s(fun(bool,bool),X88)))) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f109,f457])).
fof(f13602,plain,(
  ( ! [X4,X3] : (s(bool,sK2(X3,'const.pred_set.UNIV_0',bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2(X3,'const.pred_set.UNIV_0',bool)),s(fun(bool,bool),X3))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X3) | s(bool,X4) = s(bool,sK2(X3,'const.pred_set.UNIV_0',bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f107,f457])).
fof(f13619,plain,(
  ( ! [X12,X13,X11] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X11))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X12))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,X13) = s(bool,sK2(X11,X12,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f107,f127])).
fof(f13620,plain,(
  ( ! [X14,X15,X16] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X14))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X15))) | s(fun(bool,bool),X14) = s(fun(bool,bool),X15) | s(bool,X16) = s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X16) = s(bool,sK2(X14,X15,bool))) )),
  inference(superposition,[],[f107,f457])).
fof(f14567,plain,(
  ( ! [X19,X17,X18] : (s(bool,'T_0') = s(bool,sK2(X17,X18,bool)) | s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X17))) != s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X18))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,'T_0') = s(bool,X19)) )),
  inference(forward_demodulation,[],[f13621,f122])).
fof(f13621,plain,(
  ( ! [X19,X17,X18] : (s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X17))) != s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X18))) | s(fun(bool,bool),X17) = s(fun(bool,bool),X18) | s(bool,sK2(X17,X18,bool)) = s(bool,'const.bool.IN_2'(s(bool,sK2(X17,X18,bool)),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f107,f457])).
fof(f13622,plain,(
  ( ! [X21,X22,X20] : (s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X20))) != s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),X21))) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,sK2(X20,X21,bool))) )),
  inference(superposition,[],[f107,f457])).
fof(f13634,plain,(
  ( ! [X4,X3] : (s(bool,sK2('const.pred_set.UNIV_0',X3,bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2('const.pred_set.UNIV_0',X3,bool)),s(fun(bool,bool),X3))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X3) | s(bool,X4) = s(bool,sK2('const.pred_set.UNIV_0',X3,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f107,f457])).
fof(f14563,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,X22) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X20)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | s(bool,X19) = s(bool,X22)) )),
  inference(forward_demodulation,[],[f13968,f122])).
fof(f13968,plain,(
  ( ! [X21,X19,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X20)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),X21)))) | p(s(bool,'const.bool.IN_2'(s(bool,X19),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X20),s(fun(bool,bool),X21)))))) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X19) = s(bool,X22)) )),
  inference(superposition,[],[f110,f457])).
fof(f14561,plain,(
  ( ! [X30,X28,X29,X27] : (s(bool,'T_0') = s(bool,X28) | ~p(s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),X30)))) | p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(forward_demodulation,[],[f13970,f122])).
fof(f13970,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),X29)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),X30)))) | p(s(bool,'const.bool.IN_2'(s(bool,X27),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X29),s(fun(bool,bool),X30)))))) | s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f110,f457])).
fof(f14536,plain,(
  ( ! [X85,X83,X84,X82] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X83),s(fun(bool,bool),X84)))))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X84)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),X83)))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(forward_demodulation,[],[f14535,f122])).
fof(f14535,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X84)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),X83)))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),'const.pred_set.UNIV_0'))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X83),s(fun(bool,bool),X84)))))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(forward_demodulation,[],[f13989,f122])).
fof(f13989,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),X83)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),'const.pred_set.UNIV_0'))),s(fun(bool,bool),X84)))) | p(s(bool,'const.bool.IN_2'(s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),'const.pred_set.UNIV_0'))),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X83),s(fun(bool,bool),X84)))))) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f110,f457])).
fof(f149,plain,(
  s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(unit_resulting_resolution,[],[f73,f111])).
fof(f76,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INTER_UNIV|split|0',aHL_FALSITY)).
fof(f82,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.INTER_UNIV|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f83,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f52])).
fof(f150,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,X2,bool))) )),
  inference(superposition,[],[f111,f77])).
fof(f151,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK3(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f111,f77])).
fof(f183,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f108,f77])).
fof(f184,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f108,f77])).
fof(f212,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f109,f77])).
fof(f213,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f109,f77])).
fof(f258,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X3))) | s(fun(X2,bool),X3) = s(fun(X2,bool),X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X4)))) )),
  inference(superposition,[],[f107,f77])).
fof(f259,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK2(X0,X1,bool))) )),
  inference(superposition,[],[f107,f77])).
fof(f261,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X4))) | s(fun(X2,bool),X3) = s(fun(X2,bool),X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f107,f77])).
fof(f275,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X1),s(fun(bool,bool),X2)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f110,f77])).
fof(f276,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X4),s(fun(bool,bool),X5)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f110,f77])).
fof(f5542,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | s(bool,'T_0') = s(bool,sK2('const.pred_set.UNIV_0',X0,bool))) )),
  inference(superposition,[],[f259,f122])).
fof(f5544,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK2(X3,X4,bool)) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f259,f127])).
fof(f5545,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK2(X6,X8,bool)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) )),
  inference(superposition,[],[f259,f127])).
fof(f5556,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X0) | s(bool,'T_0') = s(bool,sK2(X0,'const.pred_set.UNIV_0',bool))) )),
  inference(superposition,[],[f259,f122])).
fof(f5558,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK2(X4,X3,bool)) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X3))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f259,f127])).
fof(f5559,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X8))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK2(X8,X6,bool)) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X6)))) )),
  inference(superposition,[],[f259,f127])).
fof(f11143,plain,(
  ( ! [X383,X385,X384] : (s(bool,X384) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X385))) | s(fun(bool,bool),X383) = s(fun(bool,bool),X385) | s(bool,'T_0') = s(bool,sK2(X385,X383,bool)) | s(bool,X384) = s(bool,'const.bool.IN_2'(s(bool,X384),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X383)))) )),
  inference(superposition,[],[f259,f457])).
fof(f11144,plain,(
  ( ! [X387,X388,X386] : (s(bool,X387) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X388))) | s(fun(bool,bool),X386) = s(fun(bool,bool),X388) | s(bool,'T_0') = s(bool,sK2(X386,X388,bool)) | s(bool,X387) = s(bool,'const.bool.IN_2'(s(bool,X387),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X386)))) )),
  inference(superposition,[],[f259,f457])).
fof(f11353,plain,(
  ( ! [X358,X357,X359] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X358))) | s(fun(bool,bool),X357) = s(fun(bool,bool),X358) | s(bool,'T_0') = s(bool,sK2(X358,X357,bool)) | s(bool,X359) = s(bool,'const.bool.IN_2'(s(bool,X359),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X359) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X357)))) )),
  inference(superposition,[],[f259,f457])).
fof(f11354,plain,(
  ( ! [X360,X362,X361] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X361))) | s(fun(bool,bool),X360) = s(fun(bool,bool),X361) | s(bool,'T_0') = s(bool,sK2(X360,X361,bool)) | s(bool,X362) = s(bool,'const.bool.IN_2'(s(bool,X362),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X362) = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X360)))) )),
  inference(superposition,[],[f259,f457])).
fof(f12263,plain,(
  ( ! [X166] : (s(bool,X166) = s(bool,'const.bool.IN_2'(s(bool,X166),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') = s(bool,X166)) )),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f149,f11382,f11386,f11361,f11273])).
fof(f415,plain,(
  ( ! [X78,X76,X79,X77] : (s(bool,X79) != s(bool,f(s(X76,sK3(X77,X78,bool,X76)))) | s(fun(X76,bool),X77) = s(fun(X76,bool),X78) | s(bool,'T_0') = s(bool,X79) | s(bool,'T_0') = s(bool,g(s(X76,sK3(X77,X78,bool,X76))))) )),
  inference(superposition,[],[f111,f127])).
fof(f484,plain,(
  ( ! [X70,X72,X71,X73] : (s(bool,'T_0') != s(bool,f(s(X70,sK3(X71,X72,bool,X70)))) | s(fun(X70,bool),X71) = s(fun(X70,bool),X72) | s(bool,X73) = s(bool,g(s(X70,sK3(X71,X72,bool,X70)))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f111,f127])).
fof(f1212,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(trivial_inequality_removal,[],[f1205])).
fof(f1205,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK3(X1,X2,bool,X0)))) | s(bool,'T_0') = s(bool,f(s(X0,sK3(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f151,f77])).
fof(f11150,plain,(
  ( ! [X412,X414,X413,X411] : (s(bool,X414) != s(bool,f(s(X411,sK3(X412,X413,bool,X411)))) | s(fun(X411,bool),X412) = s(fun(X411,bool),X413) | s(bool,X414) = s(bool,'const.bool.IN_2'(s(bool,X414),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,g(s(X411,sK3(X412,X413,bool,X411))))) )),
  inference(superposition,[],[f111,f457])).
fof(f11360,plain,(
  ( ! [X383,X385,X384,X386] : (s(bool,'T_0') != s(bool,f(s(X383,sK3(X384,X385,bool,X383)))) | s(fun(X383,bool),X384) = s(fun(X383,bool),X385) | s(bool,X386) = s(bool,'const.bool.IN_2'(s(bool,X386),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X386) = s(bool,g(s(X383,sK3(X384,X385,bool,X383))))) )),
  inference(superposition,[],[f111,f457])).
fof(f106,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7,X8] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X8) <=> ! [X9] : s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))),
  file('thm.pred_set.INTER_UNIV|split|0','thm.pred_set.EXTENSION')).
fof(f108,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X8)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INTER_UNIV|split|0','thm.pred_set.IN_INTER')).
fof(f109,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f110,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f71])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) )),
  inference(subsumption_resolution,[],[f280,f74])).
fof(f280,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X2))))))) )),
  inference(superposition,[],[f110,f122])).
fof(f657,plain,(
  ( ! [X30,X33,X31,X34,X32] : (p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X33)))) | s(bool,X34) = s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),'const.pred_set.INTER_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f475,f74])).
fof(f475,plain,(
  ( ! [X30,X33,X31,X34,X32] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X33)))) | s(bool,X34) = s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),'const.pred_set.INTER_2'(s(fun(X30,bool),X32),s(fun(X30,bool),X33))))) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f109,f127])).
fof(f658,plain,(
  ( ! [X39,X37,X35,X38,X36] : (p(s(bool,'const.bool.IN_2'(s(X35,X36),s(fun(X35,bool),X37)))) | s(bool,X39) = s(bool,'const.bool.IN_2'(s(X35,X36),s(fun(X35,bool),'const.pred_set.INTER_2'(s(fun(X35,bool),X37),s(fun(X35,bool),X38))))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(subsumption_resolution,[],[f476,f74])).
fof(f476,plain,(
  ( ! [X39,X37,X35,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X35,X36),s(fun(X35,bool),X37)))) | s(bool,X39) = s(bool,'const.bool.IN_2'(s(X35,X36),s(fun(X35,bool),'const.pred_set.INTER_2'(s(fun(X35,bool),X37),s(fun(X35,bool),X38))))) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f108,f127])).
fof(f659,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X40,X41),s(fun(X40,bool),X43)))) | p(s(bool,'const.bool.IN_2'(s(X40,X41),s(fun(X40,bool),'const.pred_set.INTER_2'(s(fun(X40,bool),X42),s(fun(X40,bool),X43)))))) | s(bool,X44) = s(bool,'const.bool.IN_2'(s(X40,X41),s(fun(X40,bool),X42))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f477,f74])).
fof(f477,plain,(
  ( ! [X43,X41,X44,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(X40,X41),s(fun(X40,bool),X43)))) | p(s(bool,'const.bool.IN_2'(s(X40,X41),s(fun(X40,bool),'const.pred_set.INTER_2'(s(fun(X40,bool),X42),s(fun(X40,bool),X43)))))) | s(bool,X44) = s(bool,'const.bool.IN_2'(s(X40,X41),s(fun(X40,bool),X42))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f110,f127])).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f69])).
fof(f411,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),X60)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X58),s(fun(bool,bool),X61)))) | p(s(bool,'const.bool.IN_2'(s(bool,X58),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X60),s(fun(bool,bool),X61)))))) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f110,f127])).
fof(f412,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,X63),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X64),s(fun(bool,bool),X65)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X62),s(fun(bool,bool),X65)))) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f109,f127])).
fof(f413,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X68),s(fun(bool,bool),X69)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X66),s(fun(bool,bool),X68)))) | s(bool,'T_0') = s(bool,X67) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f108,f127])).
fof(f416,plain,(
  ( ! [X80,X81,X82] : (s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),X80))) != s(bool,'const.bool.IN_2'(s(bool,X82),s(fun(bool,bool),X81))) | s(fun(bool,bool),X80) = s(fun(bool,bool),X81) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,sK2(X80,X81,bool))) )),
  inference(superposition,[],[f107,f127])).
fof(f454,plain,(
  ( ! [X19,X17,X18,X16] : (~p(s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),X17)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X18)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X17),s(fun(bool,bool),X18)))))) | s(bool,X16) = s(bool,X19) | s(bool,'T_0') = s(bool,X19)) )),
  inference(superposition,[],[f110,f127])).
fof(f480,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X54)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,X53),s(fun(bool,bool),X55)))) | p(s(bool,'const.bool.IN_2'(s(bool,X53),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X54),s(fun(bool,bool),X55)))))) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f110,f127])).
fof(f481,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X58),s(fun(bool,bool),X59)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X57),s(fun(bool,bool),X59)))) | s(bool,X57) = s(bool,X60) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f109,f127])).
fof(f482,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X62),s(fun(bool,bool),X63)))))) | p(s(bool,'const.bool.IN_2'(s(bool,X61),s(fun(bool,bool),X62)))) | s(bool,X61) = s(bool,X64) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f108,f127])).
fof(f11809,plain,(
  ( ! [X182,X183,X181] : (s(bool,'const.bool.IN_2'(s(bool,X181),s(fun(bool,bool),X182))) != s(bool,'const.bool.IN_2'(s(bool,X181),s(fun(bool,bool),X183))) | s(fun(bool,bool),X182) = s(fun(bool,bool),X183) | s(bool,'T_0') = s(bool,sK2(X182,X183,bool)) | s(bool,X181) = s(bool,'const.bool.IN_2'(s(bool,X181),s(fun(bool,bool),'const.pred_set.UNIV_0')))) )),
  inference(subsumption_resolution,[],[f11055,f114])).
fof(f11055,plain,(
  ( ! [X182,X183,X181] : (s(bool,'const.bool.IN_2'(s(bool,X181),s(fun(bool,bool),X182))) != s(bool,'const.bool.IN_2'(s(bool,X181),s(fun(bool,bool),X183))) | s(fun(bool,bool),X182) = s(fun(bool,bool),X183) | s(bool,'T_0') = s(bool,sK2(X182,X183,bool)) | s(bool,X181) = s(bool,'const.bool.IN_2'(s(bool,X181),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f259,f457])).
fof(f455,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(bool,X20),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X21),s(fun(bool,bool),X22)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X22)))) | s(bool,X20) = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f109,f127])).
fof(f456,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(bool,X24),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X25),s(fun(bool,bool),X26)))))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X25)))) | s(bool,X24) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f108,f127])).
fof(f485,plain,(
  ( ! [X76,X74,X75] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X74))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X75))) | s(fun(bool,bool),X74) = s(fun(bool,bool),X75) | s(bool,X76) = s(bool,sK2(X74,X75,bool)) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f107,f127])).
fof(f11364,plain,(
  ( ! [X397,X396,X398] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X396))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X397))) | s(fun(bool,bool),X396) = s(fun(bool,bool),X397) | s(bool,X398) = s(bool,'const.bool.IN_2'(s(bool,X398),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X398) = s(bool,sK2(X396,X397,bool))) )),
  inference(superposition,[],[f107,f457])).
fof(f10758,plain,(
  ( ! [X24,X25] : (s(bool,sK2('const.pred_set.UNIV_0',X24,bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2('const.pred_set.UNIV_0',X24,bool)),s(fun(bool,bool),X24))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X24) | s(bool,X25) = s(bool,sK2('const.pred_set.UNIV_0',X24,bool)) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f107,f457])).
fof(f10756,plain,(
  ( ! [X21,X20] : (s(bool,sK2(X20,'const.pred_set.UNIV_0',bool)) != s(bool,'const.bool.IN_2'(s(bool,sK2(X20,'const.pred_set.UNIV_0',bool)),s(fun(bool,bool),X20))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X20) | s(bool,X21) = s(bool,sK2(X20,'const.pred_set.UNIV_0',bool)) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f107,f457])).
fof(f2775,plain,(
  ( ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X6))))))) ) | $spl14),
  inference(subsumption_resolution,[],[f2741,f74])).
fof(f2741,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X6)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X6))))))) ) | $spl14),
  inference(superposition,[],[f110,f815])).
fof(f9985,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK2('const.pred_set.UNIV_0',X1,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f260,f127])).
fof(f9986,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X3))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK2('const.pred_set.UNIV_0',X3,bool))) )),
  inference(superposition,[],[f260,f127])).
fof(f11156,plain,(
  ( ! [X430,X429] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,X430),s(fun(bool,bool),X429))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X429) | s(bool,X430) = s(bool,'const.bool.IN_2'(s(bool,X430),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,sK2('const.pred_set.UNIV_0',X429,bool))) )),
  inference(superposition,[],[f260,f457])).
fof(f11366,plain,(
  ( ! [X401,X402] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X401))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X401) | s(bool,X402) = s(bool,'const.bool.IN_2'(s(bool,X402),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X402) = s(bool,sK2('const.pred_set.UNIV_0',X401,bool))) )),
  inference(superposition,[],[f260,f457])).
fof(f9583,plain,(
  ( ! [X2,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X1) | s(bool,X2) = s(bool,sK2(X1,'const.pred_set.UNIV_0',bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f257,f127])).
fof(f9584,plain,(
  ( ! [X4,X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X3))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,'T_0') = s(bool,sK2(X3,'const.pred_set.UNIV_0',bool))) )),
  inference(superposition,[],[f257,f127])).
fof(f11155,plain,(
  ( ! [X428,X427] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,X428),s(fun(bool,bool),X427))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X427) | s(bool,X428) = s(bool,'const.bool.IN_2'(s(bool,X428),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,sK2(X427,'const.pred_set.UNIV_0',bool))) )),
  inference(superposition,[],[f257,f457])).
fof(f11365,plain,(
  ( ! [X399,X400] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X399))) | s(fun(bool,bool),'const.pred_set.UNIV_0') = s(fun(bool,bool),X399) | s(bool,X400) = s(bool,'const.bool.IN_2'(s(bool,X400),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X400) = s(bool,sK2(X399,'const.pred_set.UNIV_0',bool))) )),
  inference(superposition,[],[f257,f457])).
fof(f12265,plain,(
  ( ! [X416,X415,X417] : (s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X415) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X416) = s(fun(bool,bool),X417) | s(bool,'T_0') = s(bool,sK3(X416,X417,bool,bool)) | s(bool,'T_0') = s(bool,X415)) )),
  inference(forward_demodulation,[],[f10981,f122])).
fof(f10981,plain,(
  ( ! [X416,X415,X417] : (s(bool,X415) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X416) = s(fun(bool,bool),X417) | s(bool,'T_0') = s(bool,sK3(X416,X417,bool,bool)) | s(bool,g(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(bool,g(s(bool,'F_0'))),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X415)) )),
  inference(superposition,[],[f150,f457])).
fof(f11273,plain,(
  ( ! [X165,X166,X164] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X164))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X165))) | s(fun(bool,bool),X164) = s(fun(bool,bool),X165) | s(bool,'T_0') = s(bool,sK2(X164,X165,bool)) | s(bool,X166) = s(bool,'const.bool.IN_2'(s(bool,X166),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') = s(bool,X166)) )),
  inference(superposition,[],[f259,f457])).
fof(f11386,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') != s(bool,X0)) )),
  inference(condensation,[],[f11376])).
fof(f11376,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X0) = s(bool,X1)) )),
  inference(equality_factoring,[],[f457])).
fof(f11367,plain,(
  ( ! [X405,X403,X404,X406] : (s(X405,f(s(bool,'T_0'))) != s(X405,g(s(bool,'T_0'))) | s(fun(bool,X405),X403) = s(fun(bool,X405),X404) | s(bool,X406) = s(bool,'const.bool.IN_2'(s(bool,X406),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X406) = s(bool,sK3(X403,X404,X405,bool))) )),
  inference(superposition,[],[f111,f457])).
fof(f11399,plain,(
  ( ! [X318,X320,X319] : (s(fun(X318,bool),'const.pred_set.UNIV_0') = s(fun(X318,bool),X319) | s(bool,X320) = s(bool,'const.bool.IN_2'(s(bool,X320),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X320) = s(bool,'const.bool.IN_2'(s(X318,sK2('const.pred_set.UNIV_0',X319,X318)),s(fun(X318,bool),X319)))) )),
  inference(trivial_inequality_removal,[],[f11338])).
fof(f11338,plain,(
  ( ! [X318,X320,X319] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X318,bool),'const.pred_set.UNIV_0') = s(fun(X318,bool),X319) | s(bool,X320) = s(bool,'const.bool.IN_2'(s(bool,X320),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X320) = s(bool,'const.bool.IN_2'(s(X318,sK2('const.pred_set.UNIV_0',X319,X318)),s(fun(X318,bool),X319)))) )),
  inference(superposition,[],[f260,f457])).
fof(f11400,plain,(
  ( ! [X315,X316,X317] : (s(fun(X315,bool),'const.pred_set.UNIV_0') = s(fun(X315,bool),X316) | s(bool,X317) = s(bool,'const.bool.IN_2'(s(bool,X317),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X317) = s(bool,'const.bool.IN_2'(s(X315,sK2(X316,'const.pred_set.UNIV_0',X315)),s(fun(X315,bool),X316)))) )),
  inference(trivial_inequality_removal,[],[f11337])).
fof(f11337,plain,(
  ( ! [X315,X316,X317] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X315,bool),'const.pred_set.UNIV_0') = s(fun(X315,bool),X316) | s(bool,X317) = s(bool,'const.bool.IN_2'(s(bool,X317),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X317) = s(bool,'const.bool.IN_2'(s(X315,sK2(X316,'const.pred_set.UNIV_0',X315)),s(fun(X315,bool),X316)))) )),
  inference(superposition,[],[f257,f457])).
fof(f11336,plain,(
  ( ! [X313,X312,X314,X311] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X311,sK2(X312,X313,X311)),s(fun(X311,bool),X313))) | s(fun(X311,bool),X312) = s(fun(X311,bool),X313) | s(bool,X314) = s(bool,'const.bool.IN_2'(s(bool,X314),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X314) = s(bool,'const.bool.IN_2'(s(X311,sK2(X312,X313,X311)),s(fun(X311,bool),X312)))) )),
  inference(superposition,[],[f107,f457])).
fof(f11334,plain,(
  ( ! [X304,X306,X303,X305] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X303,sK2(X304,X305,X303)),s(fun(X303,bool),X304))) | s(fun(X303,bool),X304) = s(fun(X303,bool),X305) | s(bool,X306) = s(bool,'const.bool.IN_2'(s(bool,X306),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X306) = s(bool,'const.bool.IN_2'(s(X303,sK2(X304,X305,X303)),s(fun(X303,bool),X305)))) )),
  inference(superposition,[],[f107,f457])).
fof(f11208,plain,(
  ( ! [X136,X137] : (s(bool,X136) = s(bool,'const.bool.IN_2'(s(bool,X136),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X137) = s(bool,'const.bool.IN_2'(s(bool,X137),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X136) = s(bool,X137)) )),
  inference(superposition,[],[f122,f457])).
fof(f11421,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X0) != s(bool,X1) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X1) = s(bool,X2)) )),
  inference(condensation,[],[f11420])).
fof(f11420,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,X64) != s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X64) = s(bool,X67)) )),
  inference(duplicate_literal_removal,[],[f11174])).
fof(f11174,plain,(
  ( ! [X66,X64,X67,X65] : (s(bool,X64) != s(bool,X65) | s(bool,'T_0') = s(bool,X65) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X64) = s(bool,X67)) )),
  inference(superposition,[],[f655,f457])).
fof(f11124,plain,(
  ( ! [X335,X337,X336,X338] : (s(bool,X338) != s(bool,'const.bool.IN_2'(s(X335,sK2(X336,X337,X335)),s(fun(X335,bool),X337))) | s(fun(X335,bool),X336) = s(fun(X335,bool),X337) | s(bool,X338) = s(bool,'const.bool.IN_2'(s(bool,X338),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X335,sK2(X336,X337,X335)),s(fun(X335,bool),X336)))) )),
  inference(superposition,[],[f107,f457])).
fof(f11122,plain,(
  ( ! [X327,X329,X328,X330] : (s(bool,X330) != s(bool,'const.bool.IN_2'(s(X327,sK2(X328,X329,X327)),s(fun(X327,bool),X328))) | s(fun(X327,bool),X328) = s(fun(X327,bool),X329) | s(bool,X330) = s(bool,'const.bool.IN_2'(s(bool,X330),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X327,sK2(X328,X329,X327)),s(fun(X327,bool),X329)))) )),
  inference(superposition,[],[f107,f457])).
fof(f11804,plain,(
  ( ! [X163,X161,X164,X162] : (s(X162,f(s(bool,X161))) != s(X162,g(s(bool,X161))) | s(fun(bool,X162),X163) = s(fun(bool,X162),X164) | s(bool,'T_0') = s(bool,sK3(X163,X164,X162,bool)) | s(bool,X161) = s(bool,'const.bool.IN_2'(s(bool,X161),s(fun(bool,bool),'const.pred_set.UNIV_0')))) )),
  inference(subsumption_resolution,[],[f11052,f114])).
fof(f11052,plain,(
  ( ! [X163,X161,X164,X162] : (s(X162,f(s(bool,X161))) != s(X162,g(s(bool,X161))) | s(fun(bool,X162),X163) = s(fun(bool,X162),X164) | s(bool,'T_0') = s(bool,sK3(X163,X164,X162,bool)) | s(bool,X161) = s(bool,'const.bool.IN_2'(s(bool,X161),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f150,f457])).
fof(f11454,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,X2) | s(bool,X0) = s(bool,X2) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0')))) )),
  inference(condensation,[],[f11453])).
fof(f11453,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X47) != s(bool,X48) | s(bool,'T_0') = s(bool,X46) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X47) = s(bool,'const.bool.IN_2'(s(bool,X47),s(fun(bool,bool),'const.pred_set.UNIV_0')))) )),
  inference(duplicate_literal_removal,[],[f11005])).
fof(f11005,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X47) != s(bool,X48) | s(bool,'T_0') = s(bool,X46) | s(bool,X48) = s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X47) = s(bool,'const.bool.IN_2'(s(bool,X47),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f655,f457])).
fof(f10709,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,X4) = s(bool,X5) | s(bool,X6) = s(bool,'const.bool.IN_2'(s(bool,X6),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X5) = s(bool,X6) | s(bool,X4) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f457,f127])).
fof(f10708,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X2) = s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X1) = s(bool,X2) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X0) = s(bool,X3)) )),
  inference(superposition,[],[f457,f457])).
fof(f10006,plain,(
  ( ! [X4,X5,X3] : (s(fun(X3,bool),'const.pred_set.UNIV_0') = s(fun(X3,bool),X4) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(X3,sK2('const.pred_set.UNIV_0',X4,X3)),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(trivial_inequality_removal,[],[f9996])).
fof(f9996,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X3,bool),'const.pred_set.UNIV_0') = s(fun(X3,bool),X4) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(X3,sK2('const.pred_set.UNIV_0',X4,X3)),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f260,f127])).
fof(f9604,plain,(
  ( ! [X4,X5,X3] : (s(fun(X3,bool),'const.pred_set.UNIV_0') = s(fun(X3,bool),X4) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(X3,sK2(X4,'const.pred_set.UNIV_0',X3)),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(trivial_inequality_removal,[],[f9594])).
fof(f9594,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X3,bool),'const.pred_set.UNIV_0') = s(fun(X3,bool),X4) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(X3,sK2(X4,'const.pred_set.UNIV_0',X3)),s(fun(X3,bool),X4))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f257,f127])).
fof(f8925,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X0) = s(bool,X1)) ) | $spl12),
  inference(condensation,[],[f8924])).
fof(f8924,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f8886])).
fof(f8886,plain,(
  ( ! [X24,X23,X22] : (s(bool,X23) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X22) | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,X22)) ) | $spl12),
  inference(superposition,[],[f655,f814])).
fof(f8933,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') != s(bool,X1) | s(bool,X0) = s(bool,X1)) )),
  inference(condensation,[],[f8932])).
fof(f8932,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) )),
  inference(duplicate_literal_removal,[],[f8880])).
fof(f8880,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f655,f77])).
fof(f5747,plain,(
  ( ! [X4,X2,X3] : (s(fun(X2,bool),X3) = s(fun(X2,bool),X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X3)))) )),
  inference(trivial_inequality_removal,[],[f5740])).
fof(f5740,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,'F_0') | s(fun(X2,bool),X3) = s(fun(X2,bool),X4) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X4))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X2,sK2(X3,X4,X2)),s(fun(X2,bool),X3)))) )),
  inference(superposition,[],[f258,f77])).
fof(f2557,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))),
  inference(unit_resulting_resolution,[],[f73,f107])).
fof(f815,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl14),
  inference(cnf_transformation,[],[f815_D])).
fof(f815_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f255,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))),
  inference(unit_resulting_resolution,[],[f73,f107])).
fof(f148,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(unit_resulting_resolution,[],[f73,f111])).
fof(f486,plain,(
  ( ! [X80,X78,X79,X77] : (s(X79,f(s(bool,'T_0'))) != s(X79,g(s(bool,'T_0'))) | s(fun(bool,X79),X77) = s(fun(bool,X79),X78) | s(bool,X80) = s(bool,sK3(X77,X78,X79,bool)) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f111,f127])).
fof(f479,plain,(
  ( ! [X52,X50,X51,X49] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X49,sK2(X50,X51,X49)),s(fun(X49,bool),X51))) | s(fun(X49,bool),X50) = s(fun(X49,bool),X51) | s(bool,X52) = s(bool,'const.bool.IN_2'(s(X49,sK2(X50,X51,X49)),s(fun(X49,bool),X50))) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f107,f127])).
fof(f478,plain,(
  ( ! [X47,X45,X48,X46] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X45,sK2(X46,X47,X45)),s(fun(X45,bool),X46))) | s(fun(X45,bool),X46) = s(fun(X45,bool),X47) | s(bool,X48) = s(bool,'const.bool.IN_2'(s(X45,sK2(X46,X47,X45)),s(fun(X45,bool),X47))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f107,f127])).
fof(f655,plain,(
  ( ! [X12,X10,X11] : (s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(subsumption_resolution,[],[f463,f74])).
fof(f463,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | s(bool,X10) = s(bool,X12) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f129,f127])).
fof(f457,plain,(
  ( ! [X28,X29] : (s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f122,f127])).
fof(f449,plain,(
  ( ! [X4,X3] : (p(s(bool,X3)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f74,f127])).
fof(f417,plain,(
  ( ! [X85,X83,X86,X84] : (s(X85,f(s(bool,X86))) != s(X85,g(s(bool,X86))) | s(fun(bool,X85),X83) = s(fun(bool,X85),X84) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,sK3(X83,X84,X85,bool))) )),
  inference(superposition,[],[f111,f127])).
fof(f410,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,X57) != s(bool,'const.bool.IN_2'(s(X54,sK2(X55,X56,X54)),s(fun(X54,bool),X56))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,'T_0') = s(bool,X57) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X54,sK2(X55,X56,X54)),s(fun(X54,bool),X55)))) )),
  inference(superposition,[],[f107,f127])).
fof(f409,plain,(
  ( ! [X52,X50,X53,X51] : (s(bool,X53) != s(bool,'const.bool.IN_2'(s(X50,sK2(X51,X52,X50)),s(fun(X50,bool),X51))) | s(fun(X50,bool),X51) = s(fun(X50,bool),X52) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X50,sK2(X51,X52,X50)),s(fun(X50,bool),X52)))) )),
  inference(superposition,[],[f107,f127])).
fof(f363,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f127,f127])).
fof(f260,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2('const.pred_set.UNIV_0',X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f107,f122])).
fof(f257,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X1,'const.pred_set.UNIV_0',X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),'const.pred_set.UNIV_0') = s(fun(X0,bool),X1)) )),
  inference(superposition,[],[f107,f122])).
fof(f256,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))),
  inference(unit_resulting_resolution,[],[f73,f107])).
fof(f169,plain,(
  ( ! [X6,X4,X5] : (s(bool,X6) != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(subsumption_resolution,[],[f168,f74])).
fof(f168,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,'T_0')) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f129,f122])).
fof(f129,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f82,f77])).
fof(f127,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f77,f77])).
fof(f122,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) )),
  inference(unit_resulting_resolution,[],[f101,f79])).
fof(f114,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f76,f78])).
fof(f111,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK3(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK3(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f48])).
fof(f48,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f45])).
fof(f45,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INTER_UNIV|split|0',aHL_EXT)).
fof(f101,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X9] : p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.pred_set.INTER_UNIV|split|0','thm.pred_set.IN_UNIV')).
fof(f78,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.INTER_UNIV|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f79,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f50])).
fof(f77,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INTER_UNIV|split|0',aHL_BOOL_CASES)).
fof(f74,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INTER_UNIV|split|0',aHL_TRUTH)).
fof(f73,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  s(fun(sK0,bool),sK1) != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f47])).
fof(f47,plain,(
  ? [X0,X1] : s(fun(X0,bool),X1) != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f23])).
fof(f23,plain,(
  ~! [X0,X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f22])).
fof(f22,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),X1) = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.UNIV_0'),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),X7) = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.UNIV_0'),s(fun(X6,bool),X7)))),
  file('thm.pred_set.INTER_UNIV|split|0','thm.pred_set.INTER_UNIV|split|0')).
fof(f16056,plain,(
  $false | ($spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987,f15975,f15951,f16055])).
fof(f16054,plain,(
  $false | ($spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987,f15975,f15951])).
fof(f16053,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14)),
  inference(subsumption_resolution,[],[f15960,f114])).
fof(f15960,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f2887,f15949,f129])).
fof(f2887,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ($spl0 | $spl9)),
  inference(forward_demodulation,[],[f2782,f2779])).
fof(f2779,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl9),
  inference(unit_resulting_resolution,[],[f810,f77])).
fof(f810,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl9),
  inference(cnf_transformation,[],[f810_D])).
fof(f810_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f2782,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | ($spl0 | $spl9)),
  inference(unit_resulting_resolution,[],[f1943,f810,f127])).
fof(f1943,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl0),
  inference(backward_demodulation,[],[f173,f148])).
fof(f16052,plain,(
  $false | ($spl0 | $spl5 | $spl9 | $spl12 | $spl14)),
  inference(subsumption_resolution,[],[f15961,f114])).
fof(f15961,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl5 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f14259,f15949,f129])).
fof(f14259,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl0 | $spl5 | $spl9)),
  inference(forward_demodulation,[],[f14227,f2887])).
fof(f14227,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl0 | $spl5)),
  inference(unit_resulting_resolution,[],[f264,f2175])).
fof(f2175,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(trivial_inequality_removal,[],[f2170])).
fof(f2170,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl0),
  inference(superposition,[],[f1943,f127])).
fof(f264,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl5),
  inference(cnf_transformation,[],[f264_D])).
fof(f264_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f16051,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(subsumption_resolution,[],[f15962,f114])).
fof(f15962,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl9 | $spl12 | $spl14 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(unit_resulting_resolution,[],[f14702,f15949,f129])).
fof(f14702,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl0 | $spl9 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(forward_demodulation,[],[f14659,f2887])).
fof(f14659,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ($spl0 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(unit_resulting_resolution,[],[f14605,f2175])).
fof(f14605,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | ($spl30 | $spl43 | $spl51 | $spl84)),
  inference(global_subsumption,[],[f14604,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f2557,f3952,f4047,f4059,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f9604,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11363,f11362,f11354,f11353,f11153,f11152,f11144,f11143,f6522,f6521,f6507,f6299,f6039,f6038,f6037,f5732,f5731,f5730,f5559,f5558,f5557,f5556,f5545,f5544,f5543,f5542,f5436,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603])).
fof(f4059,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl51),
  inference(trivial_inequality_removal,[],[f4054])).
fof(f4054,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,g(s(bool,'T_0'))) | s(bool,'T_0') = s(bool,X0)) ) | $spl51),
  inference(superposition,[],[f3952,f127])).
fof(f4047,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl51),
  inference(superposition,[],[f3952,f127])).
fof(f14604,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X61) = s(fun(bool,bool),X62)) ) | ($spl30 | $spl43 | $spl84)),
  inference(forward_demodulation,[],[f12830,f12523])).
fof(f12523,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl43),
  inference(unit_resulting_resolution,[],[f114,f1217,f127])).
fof(f12830,plain,(
  ( ! [X61,X62] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X61) = s(fun(bool,bool),X62)) ) | ($spl30 | $spl84)),
  inference(duplicate_literal_removal,[],[f12827])).
fof(f12827,plain,(
  ( ! [X61,X62] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X61) = s(fun(bool,bool),X62) | s(fun(bool,bool),X61) = s(fun(bool,bool),X62)) ) | ($spl30 | $spl84)),
  inference(superposition,[],[f111,f12386])).
fof(f12386,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,sK3(X16,X17,bool,bool)) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17)) ) | ($spl30 | $spl84)),
  inference(global_subsumption,[],[f4734,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f989,f1197,f1200,f148,f255,f2557,f3948,f3936,f3937,f3975,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f9604,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f149,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11363,f11362,f11354,f11353,f11153,f11152,f11144,f11143,f6522,f6521,f6507,f6299,f6039,f6038,f6037,f5732,f5731,f5730,f5559,f5558,f5557,f5556,f5545,f5544,f5543,f5542,f5436,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76])).
fof(f3975,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X66)))) | ~p(s(bool,'const.bool.IN_2'(s(bool,sK3(X64,X65,bool,bool)),s(fun(bool,bool),X67)))) | p(s(bool,'const.bool.IN_2'(s(bool,sK3(X64,X65,bool,bool)),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X66),s(fun(bool,bool),X67)))))) | s(fun(bool,bool),X64) = s(fun(bool,bool),X65)) ) | $spl30),
  inference(superposition,[],[f110,f989])).
fof(f3937,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X55),s(fun(bool,bool),X56)))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK3(X53,X54,bool,bool)),s(fun(bool,bool),X55)))) | s(fun(bool,bool),X53) = s(fun(bool,bool),X54)) ) | $spl30),
  inference(superposition,[],[f108,f989])).
fof(f3936,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X51),s(fun(bool,bool),X52)))))) | p(s(bool,'const.bool.IN_2'(s(bool,sK3(X49,X50,bool,bool)),s(fun(bool,bool),X52)))) | s(fun(bool,bool),X49) = s(fun(bool,bool),X50)) ) | $spl30),
  inference(superposition,[],[f109,f989])).
fof(f3948,plain,(
  ( ! [X21,X22,X20] : (s(bool,X22) != s(bool,sK3(X20,X21,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21)) ) | $spl30),
  inference(subsumption_resolution,[],[f3924,f74])).
fof(f3924,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | s(bool,X22) != s(bool,sK3(X20,X21,bool,bool)) | s(bool,'T_0') = s(bool,X22) | s(fun(bool,bool),X20) = s(fun(bool,bool),X21)) ) | $spl30),
  inference(superposition,[],[f129,f989])).
fof(f16050,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14 | $spl19 | $spl28)),
  inference(subsumption_resolution,[],[f15963,f114])).
fof(f15963,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl9 | $spl12 | $spl14 | $spl19 | $spl28)),
  inference(unit_resulting_resolution,[],[f15586,f15949,f129])).
fof(f15586,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | ($spl0 | $spl9 | $spl19 | $spl28)),
  inference(forward_demodulation,[],[f15394,f3611])).
fof(f3611,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | ($spl0 | $spl9 | $spl19)),
  inference(forward_demodulation,[],[f3596,f2887])).
fof(f3596,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | ($spl0 | $spl19)),
  inference(unit_resulting_resolution,[],[f896,f2175])).
fof(f896,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl19),
  inference(cnf_transformation,[],[f896_D])).
fof(f896_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f15394,plain,(
  s(bool,g(s(bool,'F_0'))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | ($spl19 | $spl28)),
  inference(unit_resulting_resolution,[],[f896,f988])).
fof(f16049,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14)),
  inference(subsumption_resolution,[],[f15971,f1943])).
fof(f15971,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ($spl0 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f2887,f15949,f129])).
fof(f16048,plain,(
  $false | ($spl0 | $spl5 | $spl9 | $spl12 | $spl14)),
  inference(subsumption_resolution,[],[f15972,f264])).
fof(f15972,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl0 | $spl5 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f14259,f15949,f129])).
fof(f16047,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(subsumption_resolution,[],[f15973,f14605])).
fof(f15973,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'T_0'))) | ($spl0 | $spl9 | $spl12 | $spl14 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(unit_resulting_resolution,[],[f14702,f15949,f129])).
fof(f16046,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14 | $spl19 | $spl28)),
  inference(subsumption_resolution,[],[f15974,f15615])).
fof(f15615,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | ($spl0 | $spl28)),
  inference(subsumption_resolution,[],[f15472,f1943])).
fof(f15472,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ($spl0 | $spl28)),
  inference(superposition,[],[f1943,f988])).
fof(f15974,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | ($spl0 | $spl9 | $spl12 | $spl14 | $spl19 | $spl28)),
  inference(unit_resulting_resolution,[],[f15586,f15949,f129])).
fof(f16045,plain,(
  $false | ($spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987,f15975])).
fof(f16044,plain,(
  $false | ($spl0 | $spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f15977,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987])).
fof(f15977,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))) ) | ($spl0 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f1943,f15949,f129])).
fof(f16043,plain,(
  $false | ($spl5 | $spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f15979,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987])).
fof(f15979,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0))))) ) | ($spl5 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f264,f15949,f129])).
fof(f16042,plain,(
  $false | ($spl12 | $spl14 | $spl30 | $spl43 | $spl51 | $spl62 | $spl84)),
  inference(global_subsumption,[],[f15981,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987])).
fof(f15981,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(bool,'T_0')))) ) | ($spl12 | $spl14 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(unit_resulting_resolution,[],[f14605,f15949,f129])).
fof(f16041,plain,(
  $false | ($spl0 | $spl12 | $spl14 | $spl28 | $spl62)),
  inference(global_subsumption,[],[f15983,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987])).
fof(f15983,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,g(s(bool,'F_0')))) ) | ($spl0 | $spl12 | $spl14 | $spl28)),
  inference(unit_resulting_resolution,[],[f15615,f15949,f129])).
fof(f16040,plain,(
  $false | ($spl9 | $spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f15985,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987])).
fof(f15985,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))) ) | ($spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f810,f15949,f129])).
fof(f16039,plain,(
  $false | ($spl12 | $spl14 | $spl19 | $spl62)),
  inference(global_subsumption,[],[f15986,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987])).
fof(f15986,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))) ) | ($spl12 | $spl14 | $spl19)),
  inference(unit_resulting_resolution,[],[f896,f15949,f129])).
fof(f16038,plain,(
  $false | ($spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994,f15987])).
fof(f16037,plain,(
  $false | ($spl12 | $spl14 | $spl62)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f815,f814,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5747,f8933,f8925,f9604,f10006,f11552,f11537,f11519,f11514,f11512,f10708,f10709,f11499,f11490,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f11382,f11386,f11361,f11273,f12265,f11365,f11155,f9584,f9583,f11366,f11156,f9986,f9985,f2775,f10756,f10758,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11360,f11150,f1212,f484,f415,f12263,f11354,f11353,f11144,f11143,f5559,f5558,f5556,f5545,f5544,f5542,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f11368,f4803,f4800,f149,f14536,f14561,f14563,f13634,f13622,f14567,f13620,f13619,f13602,f14568,f14583,f14585,f14587,f14601,f14603,f6037,f5730,f5557,f5543,f11363,f11362,f6038,f5731,f5436,f1325,f11153,f11152,f6522,f6521,f6507,f6299,f6039,f5732,f1800,f1569,f4634,f14914,f14915,f14909,f14918,f14903,f14904,f14905,f14906,f14982,f15870,f15871,f15872,f15929,f15886,f15892,f15949,f15994])).
fof(f15995,plain,(
  $false | ($spl12 | $spl14 | $spl62)),
  inference(unit_resulting_resolution,[],[f114,f15949,f14909])).
fof(f15988,plain,(
  $false | ($spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f114,f15949,f79])).
fof(f15984,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14 | $spl19 | $spl28)),
  inference(unit_resulting_resolution,[],[f15586,f15615,f15949,f129])).
fof(f15982,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14 | $spl30 | $spl43 | $spl51 | $spl84)),
  inference(unit_resulting_resolution,[],[f14702,f14605,f15949,f129])).
fof(f15980,plain,(
  $false | ($spl0 | $spl5 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f14259,f264,f15949,f129])).
fof(f15978,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f2887,f1943,f15949,f129])).
fof(f15976,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f2887,f114,f15949,f129])).
fof(f15952,plain,(
  $false | ($spl0 | $spl9 | $spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f2887,f15949,f82])).
fof(f15950,plain,(
  $false | ($spl12 | $spl14)),
  inference(unit_resulting_resolution,[],[f76,f15949])).
fof(f15613,plain,(
  ( ! [X41] : (s(bool,'F_0') = s(bool,g(s(bool,X41)))) ) | $spl134),
  inference(cnf_transformation,[],[f15613_D])).
fof(f15613_D,plain,(
  ( ! [X41] : (s(bool,'F_0') = s(bool,g(s(bool,X41)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f992,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | $spl35),
  inference(cnf_transformation,[],[f992_D])).
fof(f992_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f15612,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X41)))) ) | ($spl28 | $spl70)),
  inference(subsumption_resolution,[],[f15467,f4671])).
fof(f15467,plain,(
  ( ! [X41] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X41))) | s(bool,'T_0') = s(bool,X41)) ) | ($spl28 | $spl70)),
  inference(superposition,[],[f4671,f988])).
fof(f4627,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,X28)) ) | $spl52),
  inference(cnf_transformation,[],[f4627_D])).
fof(f4627_D,plain,(
  ( ! [X28] : (s(bool,'T_0') = s(bool,X28)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f15611,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl0 | $spl28)),
  inference(forward_demodulation,[],[f15555,f173])).
fof(f15555,plain,(
  ( ! [X39] : (s(bool,g(s(bool,'F_0'))) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl0 | $spl28)),
  inference(duplicate_literal_removal,[],[f15465])).
fof(f15465,plain,(
  ( ! [X39] : (s(bool,g(s(bool,'F_0'))) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | ($spl0 | $spl28)),
  inference(superposition,[],[f2215,f988])).
fof(f2215,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(subsumption_resolution,[],[f2199,f74])).
fof(f2199,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl0),
  inference(superposition,[],[f129,f173])).
fof(f15609,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl24 | $spl26 | $spl28)),
  inference(forward_demodulation,[],[f15557,f901])).
fof(f901,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl24),
  inference(cnf_transformation,[],[f901_D])).
fof(f901_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f15557,plain,(
  ( ! [X37] : (s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X37)) ) | ($spl26 | $spl28)),
  inference(duplicate_literal_removal,[],[f15463])).
fof(f15463,plain,(
  ( ! [X37] : (s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) ) | ($spl26 | $spl28)),
  inference(superposition,[],[f903,f988])).
fof(f15607,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl14 | $spl16 | $spl28)),
  inference(forward_demodulation,[],[f15559,f815])).
fof(f15559,plain,(
  ( ! [X35] : (s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X35)) ) | ($spl16 | $spl28)),
  inference(duplicate_literal_removal,[],[f15461])).
fof(f15461,plain,(
  ( ! [X35] : (s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X35)) ) | ($spl16 | $spl28)),
  inference(superposition,[],[f817,f988])).
fof(f15603,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl28),
  inference(forward_demodulation,[],[f15564,f122])).
fof(f15564,plain,(
  ( ! [X17,X15,X16] : (s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(X16,X17),s(fun(X16,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X15)) ) | $spl28),
  inference(duplicate_literal_removal,[],[f15452])).
fof(f15452,plain,(
  ( ! [X17,X15,X16] : (s(bool,g(s(bool,'F_0'))) != s(bool,'const.bool.IN_2'(s(X16,X17),s(fun(X16,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl28),
  inference(superposition,[],[f169,f988])).
fof(f14513,plain,(
  ( ! [X31,X32] : (s(fun(fun(bool,bool),bool),X31) = s(fun(fun(bool,bool),bool),X32) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X31,X32,bool,fun(bool,bool)))))) ) | $spl130),
  inference(cnf_transformation,[],[f14513_D])).
fof(f14513_D,plain,(
  ( ! [X31,X32] : (s(fun(fun(bool,bool),bool),X31) = s(fun(fun(bool,bool),bool),X32) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X31,X32,bool,fun(bool,bool)))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f14514,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X33)))) ) | $spl132),
  inference(cnf_transformation,[],[f14514_D])).
fof(f14514_D,plain,(
  ( ! [X33] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X33)))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f14492,plain,(
  ( ! [X35,X36,X34] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X36))) | s(fun(fun(bool,bool),bool),X34) = s(fun(fun(bool,bool),bool),X35) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X34,X35,bool,fun(bool,bool)))))) ) | $spl48),
  inference(superposition,[],[f151,f3950])).
fof(f14491,plain,(
  ( ! [X33,X31,X32] : (s(bool,'F_0') != s(bool,f(s(fun(bool,bool),X33))) | s(fun(fun(bool,bool),bool),X31) = s(fun(fun(bool,bool),bool),X32) | s(bool,'T_0') = s(bool,g(s(fun(bool,bool),sK3(X31,X32,bool,fun(bool,bool)))))) ) | $spl48),
  inference(superposition,[],[f151,f3950])).
fof(f13938,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl128),
  inference(cnf_transformation,[],[f13938_D])).
fof(f13938_D,plain,(
  ( ! [X14] : (s(bool,X14) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f13924,plain,(
  ( ! [X14,X13] : (s(bool,X14) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f13892])).
fof(f13892,plain,(
  ( ! [X14,X13] : (s(bool,X14) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X13) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X13)) ) | $spl4),
  inference(superposition,[],[f12996,f457])).
fof(f12996,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(subsumption_resolution,[],[f12969,f74])).
fof(f12969,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl4),
  inference(superposition,[],[f129,f265])).
fof(f1215,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) | $spl40),
  inference(cnf_transformation,[],[f1215_D])).
fof(f1215_D,plain,(
  ( ! [X4,X2,X3] : (s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,X4) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f1217,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl43),
  inference(cnf_transformation,[],[f1217_D])).
fof(f1217_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f4800,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f151,f127])).
fof(f4803,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK3(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f151,f127])).
fof(f12387,plain,(
  ( ! [X408,X407,X409] : (s(fun(bool,bool),X407) = s(fun(bool,bool),X408) | s(bool,X409) = s(bool,sK3(X407,X408,bool,bool)) | s(bool,X409) = s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X407,X408,bool,bool))))) ) | $spl126),
  inference(cnf_transformation,[],[f12387_D])).
fof(f12387_D,plain,(
  ( ! [X408,X407,X409] : (s(fun(bool,bool),X407) = s(fun(bool,bool),X408) | s(bool,X409) = s(bool,sK3(X407,X408,bool,bool)) | s(bool,X409) = s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X407,X408,bool,bool))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f11368,plain,(
  ( ! [X408,X407,X409] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X407) = s(fun(bool,bool),X408) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X407,X408,bool,bool)))) | s(bool,X409) = s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X409) = s(bool,sK3(X407,X408,bool,bool))) )),
  inference(superposition,[],[f151,f457])).
fof(f12383,plain,(
  $false | ($spl28 | $spl34)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f993,f5441,f5747,f8933,f9604,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f149,f11382,f11386,f11361,f11273,f12263,f12265,f12266,f5469,f11366,f11156,f9986,f9985,f11365,f11155,f9584,f9583,f10758,f10756,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11815,f11360,f11150,f1212,f484,f415,f11363,f11362,f11354,f11353,f11153,f11152,f11144,f11143,f6522,f6521,f6507,f6299,f6039,f6038,f6037,f5732,f5731,f5730,f5559,f5558,f5557,f5556,f5545,f5544,f5543,f5542,f5484,f5436,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f12267,f12370,f12367,f12380,f12382])).
fof(f12382,plain,(
  ( ! [X397,X396,X398] : (s(fun(bool,bool),X396) = s(fun(bool,bool),X397) | s(bool,X398) = s(bool,'const.bool.IN_2'(s(bool,X398),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X398) = s(bool,sK2(X396,X397,bool))) ) | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f12347,f12267])).
fof(f12347,plain,(
  ( ! [X397,X396,X398] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X397))) | s(fun(bool,bool),X396) = s(fun(bool,bool),X397) | s(bool,X398) = s(bool,'const.bool.IN_2'(s(bool,X398),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X398) = s(bool,sK2(X396,X397,bool))) ) | ($spl28 | $spl34)),
  inference(backward_demodulation,[],[f12267,f11364])).
fof(f12380,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f12334,f12267])).
fof(f12334,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl28 | $spl34)),
  inference(backward_demodulation,[],[f12267,f107])).
fof(f12367,plain,(
  ( ! [X383,X385,X384,X386] : (s(fun(X383,bool),X384) = s(fun(X383,bool),X385) | s(bool,X386) = s(bool,'const.bool.IN_2'(s(bool,X386),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X386) = s(bool,g(s(X383,sK3(X384,X385,bool,X383))))) ) | ($spl28 | $spl34)),
  inference(trivial_inequality_removal,[],[f12325])).
fof(f12325,plain,(
  ( ! [X383,X385,X384,X386] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X383,bool),X384) = s(fun(X383,bool),X385) | s(bool,X386) = s(bool,'const.bool.IN_2'(s(bool,X386),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X386) = s(bool,g(s(X383,sK3(X384,X385,bool,X383))))) ) | ($spl28 | $spl34)),
  inference(backward_demodulation,[],[f12267,f11360])).
fof(f12370,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f12269,f12267])).
fof(f12269,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl28 | $spl34)),
  inference(backward_demodulation,[],[f12267,f82])).
fof(f12267,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f988,f5484])).
fof(f5484,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(subsumption_resolution,[],[f5457,f74])).
fof(f5457,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(superposition,[],[f129,f993])).
fof(f11815,plain,(
  ( ! [X197] : (s(bool,'T_0') = s(bool,g(s(bool,X197))) | s(bool,X197) = s(bool,'const.bool.IN_2'(s(bool,X197),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl34),
  inference(subsumption_resolution,[],[f11059,f114])).
fof(f11059,plain,(
  ( ! [X197] : (s(bool,'T_0') = s(bool,g(s(bool,X197))) | s(bool,X197) = s(bool,'const.bool.IN_2'(s(bool,X197),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl34),
  inference(superposition,[],[f993,f457])).
fof(f12266,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl34),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f993,f5436,f5441,f5484,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11815,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f149,f11382,f11386,f11361,f11273,f12263,f12265])).
fof(f5441,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,g(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl34),
  inference(superposition,[],[f993,f77])).
fof(f12381,plain,(
  $false | ($spl28 | $spl34)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f993,f5441,f5747,f8933,f9604,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f149,f11382,f11386,f11361,f11273,f12263,f12265,f12266,f5469,f11366,f11156,f9986,f9985,f11365,f11155,f9584,f9583,f10758,f10756,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11815,f11360,f11150,f1212,f484,f415,f11363,f11362,f11354,f11353,f11153,f11152,f11144,f11143,f6522,f6521,f6507,f6299,f6039,f6038,f6037,f5732,f5731,f5730,f5559,f5558,f5557,f5556,f5545,f5544,f5543,f5542,f5484,f5436,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f12267,f12370,f12367,f12380])).
fof(f12373,plain,(
  $false | ($spl28 | $spl34)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f993,f5441,f5747,f8933,f9604,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f149,f11382,f11386,f11361,f11273,f12263,f12265,f12266,f5469,f11366,f11156,f9986,f9985,f11365,f11155,f9584,f9583,f10758,f10756,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11815,f11360,f11150,f1212,f484,f415,f11363,f11362,f11354,f11353,f11153,f11152,f11144,f11143,f6522,f6521,f6507,f6299,f6039,f6038,f6037,f5732,f5731,f5730,f5559,f5558,f5557,f5556,f5545,f5544,f5543,f5542,f5484,f5436,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f12267,f12370,f12367])).
fof(f12372,plain,(
  $false | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f12321,f12267])).
fof(f12321,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | ($spl28 | $spl34)),
  inference(backward_demodulation,[],[f12267,f149])).
fof(f12371,plain,(
  $false | ($spl28 | $spl34)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f993,f5441,f5747,f8933,f9604,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f11458,f11454,f11804,f11122,f11124,f11421,f11208,f11334,f11336,f11400,f11399,f11367,f149,f11382,f11386,f11361,f11273,f12263,f12265,f12266,f5469,f11366,f11156,f9986,f9985,f11365,f11155,f9584,f9583,f10758,f10756,f11364,f485,f456,f455,f11809,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f11815,f11360,f11150,f1212,f484,f415,f11363,f11362,f11354,f11353,f11153,f11152,f11144,f11143,f6522,f6521,f6507,f6299,f6039,f6038,f6037,f5732,f5731,f5730,f5559,f5558,f5557,f5556,f5545,f5544,f5543,f5542,f5484,f5436,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f12267,f12370])).
fof(f12369,plain,(
  $false | ($spl28 | $spl34)),
  inference(subsumption_resolution,[],[f12268,f74])).
fof(f12268,plain,(
  ~p(s(bool,'T_0')) | ($spl28 | $spl34)),
  inference(backward_demodulation,[],[f12267,f76])).
fof(f989,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) | $spl30),
  inference(cnf_transformation,[],[f989_D])).
fof(f989_D,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f991,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl33),
  inference(cnf_transformation,[],[f991_D])).
fof(f991_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f5469,plain,(
  ( ! [X10,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK3(X10,X11,bool,bool))) ) | $spl34),
  inference(superposition,[],[f150,f993])).
fof(f12260,plain,(
  ( ! [X389] : (s(bool,X389) = s(bool,'const.bool.IN_2'(s(bool,X389),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X389) = s(bool,g(s(bool,'F_0')))) ) | $spl124),
  inference(cnf_transformation,[],[f12260_D])).
fof(f12260_D,plain,(
  ( ! [X389] : (s(bool,X389) = s(bool,'const.bool.IN_2'(s(bool,X389),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X389) = s(bool,g(s(bool,'F_0')))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f11361,plain,(
  ( ! [X389,X387,X388] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X387) = s(fun(bool,bool),X388) | s(bool,'T_0') = s(bool,sK3(X387,X388,bool,bool)) | s(bool,X389) = s(bool,'const.bool.IN_2'(s(bool,X389),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X389) = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f150,f457])).
fof(f12258,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl122),
  inference(cnf_transformation,[],[f12258_D])).
fof(f12258_D,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f11382,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(equality_factoring,[],[f457])).
fof(f12257,plain,(
  $false | $spl52),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228,f12200,f12254,f12256])).
fof(f12256,plain,(
  ( ! [X397,X396,X398] : (s(fun(bool,bool),X396) = s(fun(bool,bool),X397) | s(bool,X398) = s(bool,'const.bool.IN_2'(s(bool,X398),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X398) = s(bool,sK2(X396,X397,bool))) ) | $spl52),
  inference(subsumption_resolution,[],[f12179,f4627])).
fof(f12179,plain,(
  ( ! [X397,X396,X398] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X397))) | s(fun(bool,bool),X396) = s(fun(bool,bool),X397) | s(bool,X398) = s(bool,'const.bool.IN_2'(s(bool,X398),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X398) = s(bool,sK2(X396,X397,bool))) ) | $spl52),
  inference(backward_demodulation,[],[f4627,f11364])).
fof(f12254,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl52),
  inference(subsumption_resolution,[],[f12166,f4627])).
fof(f12166,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | $spl52),
  inference(backward_demodulation,[],[f4627,f107])).
fof(f12200,plain,(
  ( ! [X383,X385,X384,X386] : (s(fun(X383,bool),X384) = s(fun(X383,bool),X385) | s(bool,X386) = s(bool,'const.bool.IN_2'(s(bool,X386),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X386) = s(bool,g(s(X383,sK3(X384,X385,bool,X383))))) ) | $spl52),
  inference(trivial_inequality_removal,[],[f12157])).
fof(f12157,plain,(
  ( ! [X383,X385,X384,X386] : (s(bool,'T_0') != s(bool,'T_0') | s(fun(X383,bool),X384) = s(fun(X383,bool),X385) | s(bool,X386) = s(bool,'const.bool.IN_2'(s(bool,X386),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X386) = s(bool,g(s(X383,sK3(X384,X385,bool,X383))))) ) | $spl52),
  inference(backward_demodulation,[],[f4627,f11360])).
fof(f12228,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(subsumption_resolution,[],[f11995,f4627])).
fof(f11995,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl52),
  inference(backward_demodulation,[],[f4627,f82])).
fof(f12255,plain,(
  $false | $spl52),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228,f12200,f12254])).
fof(f12247,plain,(
  $false | $spl52),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228,f12200])).
fof(f12246,plain,(
  $false | ($spl33 | $spl52)),
  inference(global_subsumption,[],[f12202,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12202,plain,(
  ( ! [X196] : (s(bool,X196) = s(bool,'const.bool.IN_2'(s(bool,X196),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl33 | $spl52)),
  inference(trivial_inequality_removal,[],[f12151])).
fof(f12151,plain,(
  ( ! [X196] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X196) = s(bool,'const.bool.IN_2'(s(bool,X196),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f11814])).
fof(f11814,plain,(
  ( ! [X196] : (s(bool,'T_0') != s(bool,f(s(bool,X196))) | s(bool,X196) = s(bool,'const.bool.IN_2'(s(bool,X196),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl33),
  inference(subsumption_resolution,[],[f11058,f114])).
fof(f11058,plain,(
  ( ! [X196] : (s(bool,'T_0') != s(bool,f(s(bool,X196))) | s(bool,X196) = s(bool,'const.bool.IN_2'(s(bool,X196),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,'F_0')) ) | $spl33),
  inference(superposition,[],[f991,f457])).
fof(f12245,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12244,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12244,plain,(
  ( ! [X27] : (p(s(bool,X27))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12145,f4627])).
fof(f12145,plain,(
  ( ! [X27] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))))))))))) | p(s(bool,X27))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f10506])).
fof(f10506,plain,(
  ( ! [X27] : (p(s(bool,X27)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X27)))))))))))))))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f10474,f78])).
fof(f10474,plain,(
  ( ! [X27] : (p(s(bool,X27)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X27))))))))))))))))) | s(bool,'T_0') = s(bool,X27)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f6283,f6687])).
fof(f6687,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,f(s(bool,f(s(bool,f(s(bool,X13))))))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f6540,f114])).
fof(f6540,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,f(s(bool,f(s(bool,f(s(bool,X13))))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X13)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5375,f77])).
fof(f5375,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X16))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5355,f76])).
fof(f5355,plain,(
  ( ! [X16] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X16))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5157,f5154])).
fof(f5154,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5130,f76])).
fof(f5130,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4953,f77])).
fof(f4953,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,X7)))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f4911,f78])).
fof(f4911,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,f(s(bool,X7))) | s(bool,'T_0') = s(bool,X7)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f83,f4830])).
fof(f4830,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4780,f77])).
fof(f4780,plain,(
  s(bool,'F_0') = s(bool,f(s(bool,'F_0'))) | ($spl0 | $spl9 | $spl33)),
  inference(forward_demodulation,[],[f4761,f2887])).
fof(f4761,plain,(
  s(bool,f(s(bool,'F_0'))) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ($spl0 | $spl33)),
  inference(unit_resulting_resolution,[],[f991,f2175])).
fof(f5157,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X16)))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5140,f78])).
fof(f5140,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4953,f4830])).
fof(f6283,plain,(
  ( ! [X50] : (p(s(bool,X50)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X50)))))))))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f6220,f78])).
fof(f6220,plain,(
  ( ! [X50] : (p(s(bool,X50)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X50))))))))))) | s(bool,'T_0') = s(bool,X50)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5201,f5876])).
fof(f5876,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,f(s(bool,f(s(bool,X13))))) | s(bool,'T_0') = s(bool,X13)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5771,f114])).
fof(f5771,plain,(
  ( ! [X13] : (s(bool,X13) = s(bool,f(s(bool,f(s(bool,X13))))) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X13)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5342,f77])).
fof(f5342,plain,(
  ( ! [X38] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5300,f76])).
fof(f5300,plain,(
  ( ! [X38] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X38))))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4953,f5154])).
fof(f5201,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X16)))))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5181,f78])).
fof(f5181,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X16))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5025,f4830])).
fof(f5025,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X16)))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5009,f78])).
fof(f5009,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X16))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4951,f4830])).
fof(f4951,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'T_0') != s(bool,f(s(bool,X4)))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f4908,f78])).
fof(f4908,plain,(
  ( ! [X4] : (p(s(bool,X4)) | s(bool,'T_0') != s(bool,f(s(bool,X4))) | s(bool,'T_0') = s(bool,X4)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f78,f4830])).
fof(f12243,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12242,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12242,plain,(
  ( ! [X70] : (p(s(bool,X70))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12116,f4627])).
fof(f12116,plain,(
  ( ! [X70] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))))))))) | p(s(bool,X70))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f8080])).
fof(f8080,plain,(
  ( ! [X70] : (p(s(bool,X70)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X70)))))))))))))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f7961,f78])).
fof(f7961,plain,(
  ( ! [X70] : (p(s(bool,X70)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X70))))))))))))))) | s(bool,'T_0') = s(bool,X70)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5703,f6687])).
fof(f5703,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X17)))))))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5679,f78])).
fof(f5679,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X17))))))))) | s(bool,'T_0') = s(bool,X17)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5201,f4830])).
fof(f12241,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12240,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12240,plain,(
  ( ! [X19] : (p(s(bool,X19))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12110,f4627])).
fof(f12110,plain,(
  ( ! [X19] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))))))) | p(s(bool,X19))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f7332])).
fof(f7332,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X19)))))))))))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f7305,f78])).
fof(f7305,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X19))))))))))))) | s(bool,'T_0') = s(bool,X19)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5703,f5876])).
fof(f12239,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12238,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12238,plain,(
  ( ! [X50] : (p(s(bool,X50))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12103,f4627])).
fof(f12103,plain,(
  ( ! [X50] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))))) | p(s(bool,X50))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f6283])).
fof(f12237,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12236,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12236,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12095,f4627])).
fof(f12095,plain,(
  ( ! [X17] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,'T_0'))))))) | p(s(bool,X17))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f5703])).
fof(f12235,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12234,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12234,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12091,f4627])).
fof(f12091,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,'T_0'))))) | p(s(bool,X16))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f5201])).
fof(f12233,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12232,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12232,plain,(
  ( ! [X16] : (p(s(bool,X16))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12089,f4627])).
fof(f12089,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | p(s(bool,X16))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f5025])).
fof(f12231,plain,(
  $false | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(global_subsumption,[],[f12226,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12226,plain,(
  ( ! [X4] : (p(s(bool,X4))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(trivial_inequality_removal,[],[f12086])).
fof(f12086,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,'T_0') | p(s(bool,X4))) ) | ($spl0 | $spl9 | $spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f4951])).
fof(f12230,plain,(
  $false | ($spl33 | $spl52)),
  inference(subsumption_resolution,[],[f12009,f4627])).
fof(f12009,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'T_0'))) | ($spl33 | $spl52)),
  inference(backward_demodulation,[],[f4627,f991])).
fof(f12229,plain,(
  $false | $spl52),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f76,f82,f83,f150,f151,f183,f184,f212,f213,f258,f259,f261,f275,f276,f1325,f1569,f1800,f1212,f484,f415,f106,f108,f109,f110,f288,f657,f658,f659,f107,f411,f412,f413,f416,f454,f480,f481,f482,f455,f456,f485,f2557,f4132,f4131,f4130,f4129,f4128,f1199,f1198,f1196,f5436,f5542,f5543,f5544,f5545,f5556,f5557,f5558,f5559,f5730,f5731,f5732,f5747,f6037,f6038,f6039,f6299,f6507,f6521,f6522,f8933,f9583,f9584,f9604,f9985,f9986,f10006,f11552,f11537,f11519,f11514,f10708,f10709,f11499,f10756,f10758,f11458,f11454,f11804,f11809,f11122,f11124,f11143,f11144,f11150,f11152,f11153,f11155,f11156,f11421,f11208,f11334,f11336,f11400,f11399,f11353,f11354,f11360,f11362,f11363,f11364,f11365,f11366,f11367,f4627,f12228])).
fof(f12227,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f11994,f74])).
fof(f11994,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(backward_demodulation,[],[f4627,f76])).
fof(f11779,plain,(
  ( ! [X86] : (s(bool,'F_0') = s(bool,f(s(bool,X86))) | s(bool,X86) = s(bool,'const.bool.IN_2'(s(bool,X86),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl112),
  inference(cnf_transformation,[],[f11779_D])).
fof(f11779_D,plain,(
  ( ! [X86] : (s(bool,'F_0') = s(bool,f(s(bool,X86))) | s(bool,X86) = s(bool,'const.bool.IN_2'(s(bool,X86),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f11863,plain,(
  ( ! [X267,X268] : (s(bool,'F_0') = s(bool,f(s(bool,X268))) | s(bool,'T_0') = s(bool,X267) | s(bool,X268) = s(bool,'const.bool.IN_2'(s(bool,X268),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f11096,f5343])).
fof(f5343,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X39))))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5301,f76])).
fof(f5301,plain,(
  ( ! [X39] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,X39))))))) | s(bool,'T_0') = s(bool,X39)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5025,f5154])).
fof(f11096,plain,(
  ( ! [X267,X268] : (s(bool,'F_0') = s(bool,f(s(bool,X268))) | s(bool,'T_0') = s(bool,X267) | s(bool,X268) = s(bool,'const.bool.IN_2'(s(bool,X268),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,X267)))))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5907,f457])).
fof(f5907,plain,(
  ( ! [X44] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X44))))))))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5839,f76])).
fof(f5839,plain,(
  ( ! [X44] : (p(s(bool,'F_0')) | s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X44))))))))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5157,f5342])).
fof(f11781,plain,(
  ( ! [X94] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X94))))) | s(bool,X94) = s(bool,'const.bool.IN_2'(s(bool,X94),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl116),
  inference(cnf_transformation,[],[f11781_D])).
fof(f11781_D,plain,(
  ( ! [X94] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X94))))) | s(bool,X94) = s(bool,'const.bool.IN_2'(s(bool,X94),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f11853,plain,(
  ( ! [X251,X252] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X252))))) | s(bool,'T_0') = s(bool,X251) | s(bool,X252) = s(bool,'const.bool.IN_2'(s(bool,X252),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f11088,f5196])).
fof(f5196,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5171,f76])).
fof(f5171,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5025,f77])).
fof(f11088,plain,(
  ( ! [X251,X252] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X252))))) | s(bool,'T_0') = s(bool,X251) | s(bool,X252) = s(bool,'const.bool.IN_2'(s(bool,X252),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X251)))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5907,f457])).
fof(f11780,plain,(
  ( ! [X90] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X90))))) | s(bool,X90) = s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl114),
  inference(cnf_transformation,[],[f11780_D])).
fof(f11780_D,plain,(
  ( ! [X90] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X90))))) | s(bool,X90) = s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f11849,plain,(
  ( ! [X246,X245] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X246))))) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'const.bool.IN_2'(s(bool,X246),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f11085,f5196])).
fof(f11085,plain,(
  ( ! [X246,X245] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X246))))) | s(bool,'T_0') = s(bool,X245) | s(bool,X246) = s(bool,'const.bool.IN_2'(s(bool,X246),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X245)))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5702,f457])).
fof(f5702,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X16))))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f5678,f76])).
fof(f5678,plain,(
  ( ! [X16] : (p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,f(s(bool,f(s(bool,X16))))))))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5201,f5154])).
fof(f11848,plain,(
  ( ! [X243,X244] : (s(bool,'F_0') = s(bool,f(s(bool,X244))) | s(bool,'T_0') = s(bool,X243) | s(bool,X244) = s(bool,'const.bool.IN_2'(s(bool,X244),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f11084,f5196])).
fof(f11084,plain,(
  ( ! [X243,X244] : (s(bool,'F_0') = s(bool,f(s(bool,X244))) | s(bool,'T_0') = s(bool,X243) | s(bool,X244) = s(bool,'const.bool.IN_2'(s(bool,X244),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,f(s(bool,X243)))))) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5375,f457])).
fof(f11823,plain,(
  ( ! [X215,X214] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X215))))) | s(bool,'T_0') = s(bool,X214) | s(bool,X215) = s(bool,'const.bool.IN_2'(s(bool,X215),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f11068,f4766])).
fof(f4766,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,f(s(bool,X0))) | s(bool,'T_0') = s(bool,X0)) ) | $spl33),
  inference(superposition,[],[f991,f77])).
fof(f11068,plain,(
  ( ! [X215,X214] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X215))))) | s(bool,'T_0') = s(bool,X214) | s(bool,X215) = s(bool,'const.bool.IN_2'(s(bool,X215),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X214)))) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5375,f457])).
fof(f11822,plain,(
  ( ! [X212,X213] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X213))))) | s(bool,'T_0') = s(bool,X212) | s(bool,X213) = s(bool,'const.bool.IN_2'(s(bool,X213),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f11067,f4766])).
fof(f11067,plain,(
  ( ! [X212,X213] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X213))))) | s(bool,'T_0') = s(bool,X212) | s(bool,X213) = s(bool,'const.bool.IN_2'(s(bool,X213),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X212)))) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5343,f457])).
fof(f11821,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,f(s(bool,X211))) | s(bool,'T_0') = s(bool,X210) | s(bool,X211) = s(bool,'const.bool.IN_2'(s(bool,X211),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f11066,f4766])).
fof(f11066,plain,(
  ( ! [X210,X211] : (s(bool,'F_0') = s(bool,f(s(bool,X211))) | s(bool,'T_0') = s(bool,X210) | s(bool,X211) = s(bool,'const.bool.IN_2'(s(bool,X211),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X210)))) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5342,f457])).
fof(f11799,plain,(
  ( ! [X125] : (s(bool,f(s(bool,X125))) = s(bool,f(s(bool,f(s(bool,X125))))) | s(bool,X125) = s(bool,'const.bool.IN_2'(s(bool,X125),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl120),
  inference(cnf_transformation,[],[f11799_D])).
fof(f11799_D,plain,(
  ( ! [X125] : (s(bool,f(s(bool,X125))) = s(bool,f(s(bool,f(s(bool,X125))))) | s(bool,X125) = s(bool,'const.bool.IN_2'(s(bool,X125),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f11426,plain,(
  ( ! [X125,X124] : (s(bool,f(s(bool,X125))) = s(bool,f(s(bool,f(s(bool,X125))))) | s(bool,'T_0') = s(bool,X124) | s(bool,X125) = s(bool,'const.bool.IN_2'(s(bool,X125),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f11041])).
fof(f11041,plain,(
  ( ! [X125,X124] : (s(bool,f(s(bool,X125))) = s(bool,f(s(bool,f(s(bool,X125))))) | s(bool,'T_0') = s(bool,X124) | s(bool,X125) = s(bool,'const.bool.IN_2'(s(bool,X125),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X124)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f8457,f457])).
fof(f8457,plain,(
  ( ! [X0] : (s(bool,f(s(bool,X0))) = s(bool,f(s(bool,f(s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl33)),
  inference(condensation,[],[f8456])).
fof(f8456,plain,(
  ( ! [X59,X58] : (s(bool,f(s(bool,X59))) = s(bool,f(s(bool,f(s(bool,X59))))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X59)) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f8232])).
fof(f8232,plain,(
  ( ! [X59,X58] : (s(bool,f(s(bool,X59))) = s(bool,f(s(bool,f(s(bool,X59))))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X59) | s(bool,'T_0') = s(bool,X58)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4830,f4934])).
fof(f4934,plain,(
  ( ! [X2,X3] : (s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f4898])).
fof(f4898,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,f(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4830,f127])).
fof(f11790,plain,(
  ( ! [X108] : (s(bool,X108) = s(bool,f(s(bool,f(s(bool,X108))))) | s(bool,X108) = s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl118),
  inference(cnf_transformation,[],[f11790_D])).
fof(f11790_D,plain,(
  ( ! [X108] : (s(bool,X108) = s(bool,f(s(bool,f(s(bool,X108))))) | s(bool,X108) = s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f11431,plain,(
  ( ! [X107,X108] : (s(bool,X108) = s(bool,f(s(bool,f(s(bool,X108))))) | s(bool,'T_0') = s(bool,X107) | s(bool,X108) = s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f11033])).
fof(f11033,plain,(
  ( ! [X107,X108] : (s(bool,X108) = s(bool,f(s(bool,f(s(bool,X108))))) | s(bool,'T_0') = s(bool,X107) | s(bool,X108) = s(bool,'const.bool.IN_2'(s(bool,X108),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X107)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5876,f457])).
fof(f11436,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X94))))) | s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,'const.bool.IN_2'(s(bool,X94),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f11026])).
fof(f11026,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') = s(bool,f(s(bool,f(s(bool,X94))))) | s(bool,'T_0') = s(bool,X93) | s(bool,X94) = s(bool,'const.bool.IN_2'(s(bool,X94),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X93)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5342,f457])).
fof(f11437,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X90))))) | s(bool,'T_0') = s(bool,X89) | s(bool,X90) = s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f11024])).
fof(f11024,plain,(
  ( ! [X90,X89] : (s(bool,'T_0') != s(bool,f(s(bool,f(s(bool,X90))))) | s(bool,'T_0') = s(bool,X89) | s(bool,X90) = s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X89)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5196,f457])).
fof(f11438,plain,(
  ( ! [X85,X86] : (s(bool,'F_0') = s(bool,f(s(bool,X86))) | s(bool,'T_0') = s(bool,X85) | s(bool,X86) = s(bool,'const.bool.IN_2'(s(bool,X86),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f11022])).
fof(f11022,plain,(
  ( ! [X85,X86] : (s(bool,'F_0') = s(bool,f(s(bool,X86))) | s(bool,'T_0') = s(bool,X85) | s(bool,X86) = s(bool,'const.bool.IN_2'(s(bool,X86),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X85)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f5154,f457])).
fof(f11776,plain,(
  ( ! [X67] : (s(bool,X67) = s(bool,f(s(bool,X67))) | s(bool,X67) = s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl110),
  inference(cnf_transformation,[],[f11776_D])).
fof(f11776_D,plain,(
  ( ! [X67] : (s(bool,X67) = s(bool,f(s(bool,X67))) | s(bool,X67) = s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f11444,plain,(
  ( ! [X66,X67] : (s(bool,X67) = s(bool,f(s(bool,X67))) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | ($spl0 | $spl9 | $spl33)),
  inference(duplicate_literal_removal,[],[f11014])).
fof(f11014,plain,(
  ( ! [X66,X67] : (s(bool,X67) = s(bool,f(s(bool,X67))) | s(bool,'T_0') = s(bool,X66) | s(bool,X67) = s(bool,'const.bool.IN_2'(s(bool,X67),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X66)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f4830,f457])).
fof(f11775,plain,(
  ( ! [X65] : (s(bool,X65) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl108),
  inference(cnf_transformation,[],[f11775_D])).
fof(f11775_D,plain,(
  ( ! [X65] : (s(bool,X65) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f11445,plain,(
  ( ! [X64,X65] : (s(bool,X65) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl0),
  inference(duplicate_literal_removal,[],[f11013])).
fof(f11013,plain,(
  ( ! [X64,X65] : (s(bool,X65) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X64)) ) | $spl0),
  inference(superposition,[],[f2215,f457])).
fof(f11772,plain,(
  ( ! [X59] : (s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X59) = s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl106),
  inference(cnf_transformation,[],[f11772_D])).
fof(f11772_D,plain,(
  ( ! [X59] : (s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X59) = s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f11448,plain,(
  ( ! [X59,X58] : (s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl26),
  inference(duplicate_literal_removal,[],[f11010])).
fof(f11010,plain,(
  ( ! [X59,X58] : (s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X58) | s(bool,X59) = s(bool,'const.bool.IN_2'(s(bool,X59),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X58)) ) | $spl26),
  inference(superposition,[],[f903,f457])).
fof(f11771,plain,(
  ( ! [X55] : (s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X55) = s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl104),
  inference(cnf_transformation,[],[f11771_D])).
fof(f11771_D,plain,(
  ( ! [X55] : (s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X55) = s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f11450,plain,(
  ( ! [X54,X55] : (s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f11008])).
fof(f11008,plain,(
  ( ! [X54,X55] : (s(bool,X55) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X54) | s(bool,X55) = s(bool,'const.bool.IN_2'(s(bool,X55),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X54)) ) | $spl16),
  inference(superposition,[],[f817,f457])).
fof(f11770,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X51) = s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f11770_D])).
fof(f11770_D,plain,(
  ( ! [X51] : (s(bool,X51) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X51) = s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f11452,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl7),
  inference(duplicate_literal_removal,[],[f11006])).
fof(f11006,plain,(
  ( ! [X50,X51] : (s(bool,X51) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X50) | s(bool,X51) = s(bool,'const.bool.IN_2'(s(bool,X51),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X50)) ) | $spl7),
  inference(superposition,[],[f687,f457])).
fof(f687,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl7),
  inference(subsumption_resolution,[],[f673,f74])).
fof(f673,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s(bool,X2) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl7),
  inference(superposition,[],[f129,f268])).
fof(f268,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl7),
  inference(unit_resulting_resolution,[],[f267,f77])).
fof(f11765,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl100),
  inference(cnf_transformation,[],[f11765_D])).
fof(f11765_D,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f11458,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0')))) )),
  inference(duplicate_literal_removal,[],[f10999])).
fof(f10999,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X25) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f169,f457])).
fof(f11605,plain,(
  ( ! [X28] : (s(bool,X28) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl96),
  inference(cnf_transformation,[],[f11605_D])).
fof(f11605_D,plain,(
  ( ! [X28] : (s(bool,X28) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f11489,plain,(
  ( ! [X64,X65] : (s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))) ) | $spl22),
  inference(duplicate_literal_removal,[],[f10745])).
fof(f10745,plain,(
  ( ! [X64,X65] : (s(bool,'T_0') = s(bool,X64) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(bool,X65),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X65) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X64)) ) | $spl22),
  inference(superposition,[],[f457,f900])).
fof(f11604,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl94),
  inference(cnf_transformation,[],[f11604_D])).
fof(f11604_D,plain,(
  ( ! [X26] : (s(bool,X26) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f11490,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.bool.IN_2'(s(bool,X62),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X62) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f10743])).
fof(f10743,plain,(
  ( ! [X61,X62] : (s(bool,'T_0') = s(bool,X61) | s(bool,X62) = s(bool,'const.bool.IN_2'(s(bool,X62),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X62) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X61)) ) | $spl12),
  inference(superposition,[],[f457,f814])).
fof(f11560,plain,(
  ( ! [X36,X34] : (s(bool,X34) = s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X34) = s(bool,X36) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(bool,X36),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl88),
  inference(cnf_transformation,[],[f11560_D])).
fof(f11560_D,plain,(
  ( ! [X36,X34] : (s(bool,X34) = s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X34) = s(bool,X36) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(bool,X36),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f11499,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,X12) | s(bool,X13) = s(bool,'const.bool.IN_2'(s(bool,X13),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X11) = s(bool,X13) | s(bool,X11) = s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNIV_0')))) )),
  inference(duplicate_literal_removal,[],[f10721])).
fof(f10721,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') = s(bool,X12) | s(bool,X13) = s(bool,'const.bool.IN_2'(s(bool,X13),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X11) = s(bool,X13) | s(bool,X11) = s(bool,'const.bool.IN_2'(s(bool,X11),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X12)) )),
  inference(superposition,[],[f457,f457])).
fof(f11602,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl90),
  inference(cnf_transformation,[],[f11602_D])).
fof(f11602_D,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f11607,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') = s(bool,X46) | s(bool,X46) = s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl0 | $spl9 | $spl33)),
  inference(subsumption_resolution,[],[f10682,f4766])).
fof(f10682,plain,(
  ( ! [X45,X46] : (s(bool,'F_0') = s(bool,X46) | s(bool,X46) = s(bool,'const.bool.IN_2'(s(bool,X46),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,f(s(bool,X45))) | s(bool,'T_0') = s(bool,X45)) ) | ($spl0 | $spl9 | $spl33)),
  inference(superposition,[],[f457,f5154])).
fof(f11606,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl98),
  inference(cnf_transformation,[],[f11606_D])).
fof(f11606_D,plain,(
  ( ! [X30] : (s(bool,X30) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f11510,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X29)) ) | $spl0),
  inference(duplicate_literal_removal,[],[f10675])).
fof(f10675,plain,(
  ( ! [X30,X29] : (s(bool,X30) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,X30) = s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) ) | $spl0),
  inference(superposition,[],[f457,f2175])).
fof(f11511,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X27)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f10674])).
fof(f10674,plain,(
  ( ! [X28,X27] : (s(bool,X28) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X28) = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) ) | $spl22),
  inference(superposition,[],[f457,f900])).
fof(f11512,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X25)) ) | $spl12),
  inference(duplicate_literal_removal,[],[f10673])).
fof(f10673,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,X26) = s(bool,'const.bool.IN_2'(s(bool,X26),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) ) | $spl12),
  inference(superposition,[],[f457,f814])).
fof(f11603,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X24) = s(bool,'const.bool.IN_2'(s(bool,X24),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) | $spl92),
  inference(cnf_transformation,[],[f11603_D])).
fof(f11603_D,plain,(
  ( ! [X24] : (s(bool,X24) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X24) = s(bool,'const.bool.IN_2'(s(bool,X24),s(fun(bool,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f11513,plain,(
  ( ! [X24,X23] : (s(bool,X24) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X24) = s(bool,'const.bool.IN_2'(s(bool,X24),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X23)) ) | $spl7),
  inference(duplicate_literal_removal,[],[f10672])).
fof(f10672,plain,(
  ( ! [X24,X23] : (s(bool,X24) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,X24) = s(bool,'const.bool.IN_2'(s(bool,X24),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X23)) ) | $spl7),
  inference(superposition,[],[f457,f646])).
fof(f646,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X15)) ) | $spl7),
  inference(subsumption_resolution,[],[f374,f114])).
fof(f374,plain,(
  ( ! [X15] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X15) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X15)) ) | $spl7),
  inference(superposition,[],[f127,f291])).
fof(f291,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl7),
  inference(unit_resulting_resolution,[],[f274,f77])).
fof(f274,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl7),
  inference(backward_demodulation,[],[f268,f149])).
fof(f11514,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,X22) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.IN_2'(s(bool,X21),s(fun(bool,bool),'const.pred_set.UNIV_0')))) )),
  inference(duplicate_literal_removal,[],[f10671])).
fof(f10671,plain,(
  ( ! [X21,X22,X20] : (s(bool,X21) = s(bool,X22) | s(bool,X22) = s(bool,'const.bool.IN_2'(s(bool,X22),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X20) | s(bool,X21) = s(bool,'const.bool.IN_2'(s(bool,X21),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f457,f457])).
fof(f11519,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(duplicate_literal_removal,[],[f10664])).
fof(f10664,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f457,f77])).
fof(f11537,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X16) = s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X14) = s(bool,X16)) )),
  inference(duplicate_literal_removal,[],[f10625])).
fof(f10625,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,'const.bool.IN_2'(s(bool,X14),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X15) | s(bool,X16) = s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X14) = s(bool,X16)) )),
  inference(superposition,[],[f457,f457])).
fof(f11552,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) = s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(bool,X36),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X34) = s(bool,X36)) )),
  inference(duplicate_literal_removal,[],[f10575])).
fof(f10575,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) = s(bool,'const.bool.IN_2'(s(bool,X34),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X35) | s(bool,'T_0') = s(bool,X35) | s(bool,X36) = s(bool,'const.bool.IN_2'(s(bool,X36),s(fun(bool,bool),'const.pred_set.UNIV_0'))) | s(bool,X34) = s(bool,X36)) )),
  inference(superposition,[],[f457,f457])).
fof(f4734,plain,(
  ( ! [X17,X16] : (s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,sK3(X16,X17,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X16,X17,bool,bool))))) ) | $spl84),
  inference(cnf_transformation,[],[f4734_D])).
fof(f4734_D,plain,(
  ( ! [X17,X16] : (s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,sK3(X16,X17,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X16,X17,bool,bool))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f996,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl39),
  inference(cnf_transformation,[],[f996_D])).
fof(f996_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f1196,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X0,X1,bool,bool)))) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool))) )),
  inference(superposition,[],[f151,f77])).
fof(f4735,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) | $spl86),
  inference(cnf_transformation,[],[f4735_D])).
fof(f4735_D,plain,(
  ( ! [X10] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(bool,'T_0') = s(bool,X10)) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1198,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,f(s(bool,X7))) | s(fun(bool,bool),X5) = s(fun(bool,bool),X6) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X5,X6,bool,bool)))) | s(bool,'T_0') = s(bool,X7) | s(bool,'T_0') = s(bool,sK3(X5,X6,bool,bool))) )),
  inference(superposition,[],[f151,f127])).
fof(f1199,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,f(s(bool,X10))) | s(fun(bool,bool),X8) = s(fun(bool,bool),X9) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X8,X9,bool,bool)))) | s(bool,'T_0') = s(bool,sK3(X8,X9,bool,bool)) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f151,f127])).
fof(f4733,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X16,X17,bool,bool)))) | s(bool,'T_0') = s(bool,sK3(X16,X17,bool,bool))) ) | $spl7),
  inference(forward_demodulation,[],[f1202,f291])).
fof(f1202,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,f(s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))))) | s(fun(bool,bool),X16) = s(fun(bool,bool),X17) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X16,X17,bool,bool)))) | s(bool,'T_0') = s(bool,sK3(X16,X17,bool,bool))) ) | $spl7),
  inference(superposition,[],[f151,f646])).
fof(f993,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl34),
  inference(cnf_transformation,[],[f993_D])).
fof(f993_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f4128,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f150,f77])).
fof(f988,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl28),
  inference(cnf_transformation,[],[f988_D])).
fof(f988_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4129,plain,(
  ( ! [X4,X2,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,sK3(X2,X3,bool,bool)) | s(bool,X4) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f150,f127])).
fof(f994,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl36),
  inference(cnf_transformation,[],[f994_D])).
fof(f994_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f4130,plain,(
  ( ! [X6,X7,X5] : (s(bool,X5) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X6) = s(fun(bool,bool),X7) | s(bool,'T_0') = s(bool,sK3(X6,X7,bool,bool)) | s(bool,'T_0') = s(bool,X5) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f150,f127])).
fof(f4131,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X9) = s(fun(bool,bool),X10) | s(bool,'T_0') = s(bool,sK3(X9,X10,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f150,f127])).
fof(f4132,plain,(
  ( ! [X12,X13,X11] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,sK3(X11,X12,bool,bool)) | s(bool,X13) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f150,f127])).
fof(f4711,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X45)))) ) | $spl82),
  inference(cnf_transformation,[],[f4711_D])).
fof(f4711_D,plain,(
  ( ! [X45] : (s(bool,X45) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X45)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f4539,plain,(
  ( ! [X45,X44] : (s(bool,X45) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,g(s(bool,X45)))) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4375])).
fof(f4375,plain,(
  ( ! [X45,X44] : (s(bool,X45) != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X44) | s(bool,'F_0') = s(bool,g(s(bool,X45))) | s(bool,'T_0') = s(bool,X44)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f2215,f4083])).
fof(f4083,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,g(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4063,f127])).
fof(f4063,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'T_0'))) | ($spl0 | $spl9 | $spl51)),
  inference(forward_demodulation,[],[f4042,f2887])).
fof(f4042,plain,(
  s(bool,g(s(bool,'T_0'))) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | ($spl0 | $spl51)),
  inference(unit_resulting_resolution,[],[f3952,f2175])).
fof(f4708,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(bool,X39)))) ) | $spl80),
  inference(cnf_transformation,[],[f4708_D])).
fof(f4708_D,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(bool,X39)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f4542,plain,(
  ( ! [X39,X38] : (s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,g(s(bool,X39)))) ) | ($spl0 | $spl9 | $spl26 | $spl51)),
  inference(duplicate_literal_removal,[],[f4372])).
fof(f4372,plain,(
  ( ! [X39,X38] : (s(bool,X39) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X38) | s(bool,'F_0') = s(bool,g(s(bool,X39))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl0 | $spl9 | $spl26 | $spl51)),
  inference(superposition,[],[f903,f4083])).
fof(f4707,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) | $spl78),
  inference(cnf_transformation,[],[f4707_D])).
fof(f4707_D,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4543,plain,(
  ( ! [X37,X36] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(bool,X37)))) ) | ($spl0 | $spl9 | $spl16 | $spl51)),
  inference(duplicate_literal_removal,[],[f4371])).
fof(f4371,plain,(
  ( ! [X37,X36] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,g(s(bool,X37))) | s(bool,'T_0') = s(bool,X36)) ) | ($spl0 | $spl9 | $spl16 | $spl51)),
  inference(superposition,[],[f817,f4083])).
fof(f4706,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X35)))) ) | $spl76),
  inference(cnf_transformation,[],[f4706_D])).
fof(f4706_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X35)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4544,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,g(s(bool,X35)))) ) | ($spl0 | $spl7 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4370])).
fof(f4370,plain,(
  ( ! [X35,X34] : (s(bool,X35) != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,X34) | s(bool,'F_0') = s(bool,g(s(bool,X35))) | s(bool,'T_0') = s(bool,X34)) ) | ($spl0 | $spl7 | $spl9 | $spl51)),
  inference(superposition,[],[f687,f4083])).
fof(f4705,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X26)))) ) | $spl74),
  inference(cnf_transformation,[],[f4705_D])).
fof(f4705_D,plain,(
  ( ! [X28,X26,X27] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,'F_0') = s(bool,g(s(bool,X26)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4546,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,g(s(bool,X26)))) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4367])).
fof(f4367,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) != s(bool,'const.bool.IN_2'(s(X27,X28),s(fun(X27,bool),'const.pred_set.UNIV_0'))) | s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,g(s(bool,X26))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f169,f4083])).
fof(f4642,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,g(s(bool,sK3(X58,X59,bool,bool)))) | s(fun(bool,bool),X58) = s(fun(bool,bool),X59)) ) | $spl64),
  inference(cnf_transformation,[],[f4642_D])).
fof(f4642_D,plain,(
  ( ! [X59,X58] : (s(bool,'F_0') = s(bool,g(s(bool,sK3(X58,X59,bool,bool)))) | s(fun(bool,bool),X58) = s(fun(bool,bool),X59)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4553,plain,(
  ( ! [X134,X132,X133] : (s(bool,'T_0') = s(bool,X134) | s(fun(bool,bool),X132) = s(fun(bool,bool),X133) | s(bool,'F_0') = s(bool,g(s(bool,sK3(X132,X133,bool,bool))))) ) | ($spl0 | $spl9 | $spl30 | $spl51)),
  inference(duplicate_literal_removal,[],[f4347])).
fof(f4347,plain,(
  ( ! [X134,X132,X133] : (s(bool,'T_0') = s(bool,X134) | s(fun(bool,bool),X132) = s(fun(bool,bool),X133) | s(bool,'F_0') = s(bool,g(s(bool,sK3(X132,X133,bool,bool)))) | s(bool,'T_0') = s(bool,X134)) ) | ($spl0 | $spl9 | $spl30 | $spl51)),
  inference(superposition,[],[f989,f4083])).
fof(f4671,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2)))) ) | $spl70),
  inference(cnf_transformation,[],[f4671_D])).
fof(f4671_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4672,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl72),
  inference(cnf_transformation,[],[f4672_D])).
fof(f4672_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4670,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2)))) ) | ($spl0 | $spl9 | $spl51)),
  inference(subsumption_resolution,[],[f4288,f78])).
fof(f4288,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') = s(bool,g(s(bool,X2))) | s(bool,'T_0') = s(bool,X3)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f78,f4083])).
fof(f4628,plain,(
  ( ! [X29,X27] : (s(bool,'F_0') = s(bool,g(s(bool,X27))) | s(bool,X27) = s(bool,X29) | s(bool,'F_0') = s(bool,g(s(bool,X29)))) ) | $spl54),
  inference(cnf_transformation,[],[f4628_D])).
fof(f4628_D,plain,(
  ( ! [X29,X27] : (s(bool,'F_0') = s(bool,g(s(bool,X27))) | s(bool,X27) = s(bool,X29) | s(bool,'F_0') = s(bool,g(s(bool,X29)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f4579,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,'F_0') = s(bool,g(s(bool,X16))) | s(bool,X14) = s(bool,X16) | s(bool,'F_0') = s(bool,g(s(bool,X14)))) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4250])).
fof(f4250,plain,(
  ( ! [X14,X15,X16] : (s(bool,'T_0') = s(bool,X15) | s(bool,'F_0') = s(bool,g(s(bool,X16))) | s(bool,X14) = s(bool,X16) | s(bool,'F_0') = s(bool,g(s(bool,X14))) | s(bool,'T_0') = s(bool,X15)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f4083])).
fof(f4588,plain,(
  ( ! [X23,X21,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,g(s(bool,X23))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,g(s(bool,X22)))) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4214])).
fof(f4214,plain,(
  ( ! [X23,X21,X22] : (s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,g(s(bool,X23))) | s(bool,'T_0') = s(bool,X21) | s(bool,'F_0') = s(bool,g(s(bool,X22))) | s(bool,'T_0') = s(bool,X21)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f4083])).
fof(f4644,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X17)))) ) | $spl68),
  inference(cnf_transformation,[],[f4644_D])).
fof(f4644_D,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X17)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4589,plain,(
  ( ! [X17,X16] : (s(bool,X17) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X17))) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4212])).
fof(f4212,plain,(
  ( ! [X17,X16] : (s(bool,X17) = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X17))) | s(bool,'T_0') = s(bool,X16) | s(bool,'T_0') = s(bool,X16)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f2175])).
fof(f4643,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) | $spl66),
  inference(cnf_transformation,[],[f4643_D])).
fof(f4643_D,plain,(
  ( ! [X15] : (s(bool,X15) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X15)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4590,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl7 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4211])).
fof(f4211,plain,(
  ( ! [X14,X15] : (s(bool,X15) = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'F_0') = s(bool,g(s(bool,X15))) | s(bool,'T_0') = s(bool,X14) | s(bool,'T_0') = s(bool,X14)) ) | ($spl0 | $spl7 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f646])).
fof(f4634,plain,(
  ( ! [X22] : (s(bool,'F_0') = s(bool,g(s(bool,X22))) | s(bool,'F_0') = s(bool,X22)) ) | $spl62),
  inference(cnf_transformation,[],[f4634_D])).
fof(f4634_D,plain,(
  ( ! [X22] : (s(bool,'F_0') = s(bool,g(s(bool,X22))) | s(bool,'F_0') = s(bool,X22)) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4595,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4206])).
fof(f4206,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,X1) | s(bool,'F_0') = s(bool,g(s(bool,X1))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f77])).
fof(f4596,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,g(s(bool,sK3(X58,X59,bool,bool)))) | s(fun(bool,bool),X58) = s(fun(bool,bool),X59)) ) | ($spl0 | $spl9 | $spl30 | $spl51)),
  inference(duplicate_literal_removal,[],[f4205])).
fof(f4205,plain,(
  ( ! [X59,X60,X58] : (s(bool,'T_0') = s(bool,X60) | s(bool,'F_0') = s(bool,g(s(bool,sK3(X58,X59,bool,bool)))) | s(bool,'T_0') = s(bool,X60) | s(fun(bool,bool),X58) = s(fun(bool,bool),X59)) ) | ($spl0 | $spl9 | $spl30 | $spl51)),
  inference(superposition,[],[f4083,f989])).
fof(f4609,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,g(s(bool,X24))) | s(bool,'F_0') = s(bool,g(s(bool,X26))) | s(bool,X24) = s(bool,X26)) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4186])).
fof(f4186,plain,(
  ( ! [X26,X24,X25] : (s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,g(s(bool,X24))) | s(bool,'T_0') = s(bool,X25) | s(bool,'F_0') = s(bool,g(s(bool,X26))) | s(bool,X24) = s(bool,X26)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f4083])).
fof(f4633,plain,(
  ( ! [X23,X21] : (s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,g(s(bool,X21))) | s(bool,'T_0') = s(bool,X23)) ) | $spl60),
  inference(cnf_transformation,[],[f4633_D])).
fof(f4633_D,plain,(
  ( ! [X23,X21] : (s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,g(s(bool,X21))) | s(bool,'T_0') = s(bool,X23)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f4175,plain,(
  ( ! [X23,X21,X22] : (s(bool,'F_0') = s(bool,X22) | s(bool,X21) = s(bool,X23) | s(bool,'T_0') = s(bool,X23) | s(bool,'F_0') = s(bool,g(s(bool,X22))) | s(bool,'T_0') = s(bool,g(s(bool,X21)))) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f4083])).
fof(f4629,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X49)) ) | $spl56),
  inference(cnf_transformation,[],[f4629_D])).
fof(f4629_D,plain,(
  ( ! [X50,X48,X49] : (s(bool,X49) = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X50) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X49)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4631,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | $spl58),
  inference(cnf_transformation,[],[f4631_D])).
fof(f4631_D,plain,(
  s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4161,plain,(
  ( ! [X50,X48,X49] : (s(bool,'F_0') = s(bool,g(s(bool,'F_0'))) | s(bool,X49) = s(bool,g(s(bool,X48))) | s(bool,'T_0') = s(bool,X49) | s(bool,X48) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f4083])).
fof(f4619,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,g(s(bool,X29))) | s(bool,X27) = s(bool,X29)) ) | ($spl0 | $spl9 | $spl51)),
  inference(duplicate_literal_removal,[],[f4149])).
fof(f4149,plain,(
  ( ! [X28,X29,X27] : (s(bool,'F_0') = s(bool,g(s(bool,X27))) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'F_0') = s(bool,g(s(bool,X29))) | s(bool,X27) = s(bool,X29)) ) | ($spl0 | $spl9 | $spl51)),
  inference(superposition,[],[f4083,f4083])).
fof(f3950,plain,(
  ( ! [X45,X44] : (s(fun(bool,bool),X44) = s(fun(bool,bool),X45)) ) | $spl48),
  inference(cnf_transformation,[],[f3950_D])).
fof(f3950_D,plain,(
  ( ! [X45,X44] : (s(fun(bool,bool),X44) = s(fun(bool,bool),X45)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3952,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | $spl51),
  inference(cnf_transformation,[],[f3952_D])).
fof(f3952_D,plain,(
  s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f3949,plain,(
  ( ! [X45,X44] : (s(bool,'T_0') != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X44) = s(fun(bool,bool),X45)) ) | ($spl30 | $spl43)),
  inference(forward_demodulation,[],[f3939,f1225])).
fof(f1225,plain,(
  s(bool,'T_0') = s(bool,f(s(bool,'T_0'))) | $spl43),
  inference(unit_resulting_resolution,[],[f114,f1217,f127])).
fof(f3939,plain,(
  ( ! [X45,X44] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X44) = s(fun(bool,bool),X45)) ) | $spl30),
  inference(duplicate_literal_removal,[],[f3934])).
fof(f3934,plain,(
  ( ! [X45,X44] : (s(bool,f(s(bool,'T_0'))) != s(bool,g(s(bool,'T_0'))) | s(fun(bool,bool),X44) = s(fun(bool,bool),X45) | s(fun(bool,bool),X44) = s(fun(bool,bool),X45)) ) | $spl30),
  inference(superposition,[],[f111,f989])).
fof(f3312,plain,(
  $false | ($spl19 | $spl21 | $spl25)),
  inference(global_subsumption,[],[f896,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f902])).
fof(f1742,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl25),
  inference(unit_resulting_resolution,[],[f76,f902,f449])).
fof(f931,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl25),
  inference(trivial_inequality_removal,[],[f925])).
fof(f925,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl25),
  inference(superposition,[],[f902,f127])).
fof(f3311,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f902])).
fof(f3310,plain,(
  $false | ($spl21 | $spl25 | $spl44)),
  inference(global_subsumption,[],[f1859,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76])).
fof(f3309,plain,(
  $false | ($spl21 | $spl25 | $spl46)),
  inference(global_subsumption,[],[f1860,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76])).
fof(f1860,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl46),
  inference(cnf_transformation,[],[f1860_D])).
fof(f1860_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f3308,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76])).
fof(f3307,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82])).
fof(f3306,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83])).
fof(f3305,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150])).
fof(f3304,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151])).
fof(f3303,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184,f183])).
fof(f3302,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212,f184])).
fof(f3301,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213,f212])).
fof(f3300,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258,f213])).
fof(f3299,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259,f258])).
fof(f3298,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261,f259])).
fof(f3297,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275,f261])).
fof(f3296,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276,f275])).
fof(f3295,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899,f276])).
fof(f3294,plain,(
  $false | ($spl7 | $spl21 | $spl25)),
  inference(global_subsumption,[],[f291,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899])).
fof(f3293,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f899])).
fof(f3292,plain,(
  $false | ($spl21 | $spl25 | $spl43)),
  inference(global_subsumption,[],[f1226,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325])).
fof(f1226,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,X0))) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1)) ) | $spl43),
  inference(superposition,[],[f1217,f127])).
fof(f3291,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325])).
fof(f3290,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569])).
fof(f3289,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800])).
fof(f3288,plain,(
  $false | ($spl9 | $spl21 | $spl25)),
  inference(global_subsumption,[],[f2779,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212])).
fof(f3287,plain,(
  $false | ($spl0 | $spl9 | $spl21 | $spl25)),
  inference(global_subsumption,[],[f2887,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212])).
fof(f3286,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212])).
fof(f3285,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484])).
fof(f3284,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415])).
fof(f3283,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106])).
fof(f3282,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108])).
fof(f3281,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109])).
fof(f3280,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110])).
fof(f3279,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288])).
fof(f3278,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657])).
fof(f3277,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658])).
fof(f3276,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659])).
fof(f3275,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107])).
fof(f3274,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411])).
fof(f3273,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412])).
fof(f3272,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416,f413])).
fof(f3271,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454,f416])).
fof(f3270,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480,f454])).
fof(f3269,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481,f480])).
fof(f3268,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f454,f416,f480,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482,f481])).
fof(f3267,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f481,f454,f416,f480,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455,f482])).
fof(f3266,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f482,f481,f454,f416,f480,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456,f455])).
fof(f3265,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f482,f481,f454,f416,f480,f455,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485,f456])).
fof(f3264,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557,f485])).
fof(f3263,plain,(
  $false | ($spl14 | $spl21 | $spl25)),
  inference(global_subsumption,[],[f2775,f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f485,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557])).
fof(f3262,plain,(
  $false | ($spl21 | $spl25)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f902,f931,f983,f984,f985,f986,f987,f1197,f1200,f1799,f148,f255,f1742,f485,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f899,f2557])).
fof(f3261,plain,(
  $false | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f3238,f74])).
fof(f3238,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl47)),
  inference(backward_demodulation,[],[f3190,f1861])).
fof(f3190,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f3189,f74])).
fof(f3189,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,X0)) ) | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f3175,f903])).
fof(f3175,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl47),
  inference(superposition,[],[f1861,f127])).
fof(f3260,plain,(
  $false | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f3237,f74])).
fof(f3237,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl47)),
  inference(backward_demodulation,[],[f3190,f3167])).
fof(f3167,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),sK1),s(fun(sK0,bool),X0))))))) ) | $spl47),
  inference(unit_resulting_resolution,[],[f1861,f108])).
fof(f3259,plain,(
  $false | ($spl26 | $spl47)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f485,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f1861,f2557,f903,f3173,f3167,f3190,f3250,f3258])).
fof(f3258,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f3223,f3190])).
fof(f3223,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl26 | $spl47)),
  inference(backward_demodulation,[],[f3190,f107])).
fof(f3250,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f3192,f3190])).
fof(f3192,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl26 | $spl47)),
  inference(backward_demodulation,[],[f3190,f82])).
fof(f3173,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f1861,f78])).
fof(f3251,plain,(
  $false | ($spl26 | $spl47)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f485,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f1861,f2557,f903,f3173,f3167,f3190,f3250])).
fof(f3249,plain,(
  $false | ($spl26 | $spl47)),
  inference(subsumption_resolution,[],[f3191,f74])).
fof(f3191,plain,(
  ~p(s(bool,'T_0')) | ($spl26 | $spl47)),
  inference(backward_demodulation,[],[f3190,f76])).
fof(f3188,plain,(
  $false | ($spl21 | $spl47)),
  inference(subsumption_resolution,[],[f3168,f899])).
fof(f3168,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f1861,f83])).
fof(f3187,plain,(
  $false | ($spl21 | $spl47)),
  inference(subsumption_resolution,[],[f3169,f899])).
fof(f3169,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl47),
  inference(unit_resulting_resolution,[],[f114,f1861,f449])).
fof(f3186,plain,(
  $false | ($spl0 | $spl9 | $spl21 | $spl47)),
  inference(subsumption_resolution,[],[f3185,f899])).
fof(f3185,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl9 | $spl47)),
  inference(forward_demodulation,[],[f3170,f2887])).
fof(f3170,plain,(
  s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl0 | $spl47)),
  inference(unit_resulting_resolution,[],[f1943,f1861,f449])).
fof(f3184,plain,(
  $false | ($spl7 | $spl21 | $spl47)),
  inference(subsumption_resolution,[],[f3183,f899])).
fof(f3183,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl7 | $spl47)),
  inference(forward_demodulation,[],[f3171,f291])).
fof(f3171,plain,(
  s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl7 | $spl47)),
  inference(unit_resulting_resolution,[],[f274,f1861,f449])).
fof(f3182,plain,(
  $false | ($spl9 | $spl21 | $spl47)),
  inference(subsumption_resolution,[],[f3181,f899])).
fof(f3181,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | ($spl9 | $spl47)),
  inference(forward_demodulation,[],[f3172,f2779])).
fof(f3172,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | ($spl9 | $spl47)),
  inference(unit_resulting_resolution,[],[f810,f1861,f449])).
fof(f2991,plain,(
  $false | ($spl18 | $spl22)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f897,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2879,f2880,f2874,f2885,f2921,f2982,f2990])).
fof(f2990,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2953,f2921])).
fof(f2953,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl18 | $spl22)),
  inference(backward_demodulation,[],[f2921,f107])).
fof(f2982,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2923,f2921])).
fof(f2923,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl18 | $spl22)),
  inference(backward_demodulation,[],[f2921,f82])).
fof(f2921,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f900,f2880])).
fof(f2885,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl18),
  inference(subsumption_resolution,[],[f2841,f73])).
fof(f2841,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))) | $spl18),
  inference(superposition,[],[f107,f897])).
fof(f2874,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl18),
  inference(subsumption_resolution,[],[f2842,f74])).
fof(f2842,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl18),
  inference(superposition,[],[f109,f897])).
fof(f2880,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl18),
  inference(subsumption_resolution,[],[f2852,f74])).
fof(f2852,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl18),
  inference(superposition,[],[f129,f897])).
fof(f2879,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl18),
  inference(subsumption_resolution,[],[f2844,f74])).
fof(f2844,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl18),
  inference(superposition,[],[f110,f897])).
fof(f2983,plain,(
  $false | ($spl18 | $spl22)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f890,f891,f892,f893,f894,f897,f983,f984,f985,f986,f987,f1197,f1200,f148,f255,f485,f456,f455,f482,f481,f480,f454,f416,f413,f412,f411,f107,f659,f658,f657,f288,f110,f109,f108,f106,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f2879,f2880,f2874,f2885,f2921,f2982])).
fof(f2981,plain,(
  $false | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f2922,f74])).
fof(f2922,plain,(
  ~p(s(bool,'T_0')) | ($spl18 | $spl22)),
  inference(backward_demodulation,[],[f2921,f76])).
fof(f2920,plain,(
  $false | ($spl18 | $spl47)),
  inference(subsumption_resolution,[],[f1861,f2874])).
fof(f2876,plain,(
  $false | ($spl18 | $spl25)),
  inference(subsumption_resolution,[],[f2875,f76])).
fof(f2875,plain,(
  p(s(bool,'F_0')) | ($spl18 | $spl25)),
  inference(forward_demodulation,[],[f2874,f1742])).
fof(f2085,plain,(
  $false | ($spl9 | $spl15 | $spl16)),
  inference(subsumption_resolution,[],[f810,f2017])).
fof(f2017,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f816,f843,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f1740,f148,f485,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f106,f288,f110,f109,f108,f657,f658,f659,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f255,f817])).
fof(f1740,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl15),
  inference(unit_resulting_resolution,[],[f76,f816,f449])).
fof(f843,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(trivial_inequality_removal,[],[f837])).
fof(f837,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X0)) ) | $spl15),
  inference(superposition,[],[f816,f127])).
fof(f2083,plain,(
  $false | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f816,f2017])).
fof(f2082,plain,(
  $false | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2060,f2017])).
fof(f2060,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f2017,f255])).
fof(f2081,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f816,f843,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f1740,f148,f485,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f106,f288,f110,f109,f108,f657,f658,f659,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f255,f817,f2017,f2072,f2080])).
fof(f2080,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2048,f2017])).
fof(f2048,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f2017,f107])).
fof(f2072,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2019,f2017])).
fof(f2019,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f2017,f82])).
fof(f2073,plain,(
  $false | ($spl15 | $spl16)),
  inference(global_subsumption,[],[f73,f74,f77,f79,f78,f101,f111,f114,f122,f127,f129,f149,f169,f171,f256,f257,f260,f263,f363,f409,f410,f417,f449,f457,f655,f478,f479,f486,f804,f805,f806,f807,f808,f816,f843,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f1740,f148,f485,f482,f481,f454,f416,f480,f456,f455,f413,f412,f411,f107,f106,f288,f110,f109,f108,f657,f658,f659,f415,f484,f1212,f1800,f1569,f1325,f276,f275,f261,f259,f258,f213,f212,f184,f183,f151,f150,f83,f82,f76,f255,f817,f2017,f2072])).
fof(f2071,plain,(
  $false | ($spl15 | $spl16)),
  inference(subsumption_resolution,[],[f2018,f74])).
fof(f2018,plain,(
  ~p(s(bool,'T_0')) | ($spl15 | $spl16)),
  inference(backward_demodulation,[],[f2017,f76])).
fof(f2014,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f73,f74,f76,f77,f79,f78,f83,f82,f101,f107,f106,f110,f109,f108,f111,f114,f122,f127,f129,f149,f150,f151,f169,f171,f183,f184,f212,f213,f256,f255,f257,f258,f259,f260,f261,f263,f275,f276,f288,f363,f409,f410,f411,f412,f413,f415,f416,f417,f449,f454,f455,f456,f457,f655,f657,f658,f659,f478,f479,f480,f481,f482,f484,f485,f486,f804,f805,f806,f807,f808,f811,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f1212,f1325,f1569,f1800,f1853,f1854,f148,f1848,f1944,f1946,f2007,f2013])).
fof(f2013,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f1977,f1946])).
fof(f1977,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f1946,f107])).
fof(f2007,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f1948,f1946])).
fof(f1948,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f1946,f82])).
fof(f1946,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f814,f1854])).
fof(f1944,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl8),
  inference(subsumption_resolution,[],[f1814,f73])).
fof(f1814,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(fun(sK0,bool),sK1) = s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))) | $spl8),
  inference(superposition,[],[f107,f811])).
fof(f1848,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl8),
  inference(subsumption_resolution,[],[f1815,f74])).
fof(f1815,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1)))) | $spl8),
  inference(superposition,[],[f109,f811])).
fof(f1854,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl8),
  inference(subsumption_resolution,[],[f1825,f74])).
fof(f1825,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl8),
  inference(superposition,[],[f129,f811])).
fof(f1853,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f1817,f74])).
fof(f1817,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),X0)))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))),s(fun(sK0,bool),X0))))))) ) | $spl8),
  inference(superposition,[],[f110,f811])).
fof(f2008,plain,(
  $false | ($spl8 | $spl12)),
  inference(global_subsumption,[],[f73,f74,f76,f77,f79,f78,f83,f82,f101,f107,f106,f110,f109,f108,f111,f114,f122,f127,f129,f149,f150,f151,f169,f171,f183,f184,f212,f213,f256,f255,f257,f258,f259,f260,f261,f263,f275,f276,f288,f363,f409,f410,f411,f412,f413,f415,f416,f417,f449,f454,f455,f456,f457,f655,f657,f658,f659,f478,f479,f480,f481,f482,f484,f485,f486,f804,f805,f806,f807,f808,f811,f890,f891,f892,f893,f894,f983,f984,f985,f986,f987,f1197,f1200,f1212,f1325,f1569,f1800,f1853,f1854,f148,f1848,f1944,f1946,f2007])).
fof(f2006,plain,(
  $false | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f1947,f74])).
fof(f1947,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl12)),
  inference(backward_demodulation,[],[f1946,f76])).
fof(f1864,plain,(
  $false | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f1863,f76])).
fof(f1863,plain,(
  p(s(bool,'F_0')) | ($spl8 | $spl15)),
  inference(backward_demodulation,[],[f1740,f1848])).
fof(f1859,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),X29)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X28),s(fun(bool,bool),X29))))) | s(bool,'T_0') = s(bool,X30)) ) | $spl44),
  inference(cnf_transformation,[],[f1859_D])).
fof(f1859_D,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),X29)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X28),s(fun(bool,bool),X29))))) | s(bool,'T_0') = s(bool,X30)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1861,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | $spl47),
  inference(cnf_transformation,[],[f1861_D])).
fof(f1861_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f1799,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1)))) | p(s(bool,'const.bool.IN_2'(s(bool,X30),s(fun(bool,bool),X29)))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INTER_2'(s(fun(bool,bool),X28),s(fun(bool,bool),X29)))))) ) | $spl25),
  inference(superposition,[],[f212,f931])).
fof(f1850,plain,(
  $false | ($spl3 | $spl8 | $spl15)),
  inference(subsumption_resolution,[],[f1849,f76])).
fof(f1849,plain,(
  p(s(bool,'F_0')) | ($spl3 | $spl8 | $spl15)),
  inference(forward_demodulation,[],[f1848,f844])).
fof(f844,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl15)),
  inference(forward_demodulation,[],[f833,f189])).
fof(f189,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f182,f77])).
fof(f182,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl3),
  inference(backward_demodulation,[],[f176,f148])).
fof(f176,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f175,f77])).
fof(f833,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | ($spl3 | $spl15)),
  inference(unit_resulting_resolution,[],[f816,f645])).
fof(f645,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(subsumption_resolution,[],[f373,f114])).
fof(f373,plain,(
  ( ! [X14] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X14) = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,X14)) ) | $spl3),
  inference(superposition,[],[f127,f189])).
fof(f1200,plain,(
  ( ! [X12,X13,X11] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X11) = s(fun(bool,bool),X12) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X11,X12,bool,bool)))) | s(bool,X13) = s(bool,sK3(X11,X12,bool,bool)) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f151,f127])).
fof(f1197,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X2) = s(fun(bool,bool),X3) | s(bool,'T_0') = s(bool,g(s(bool,sK3(X2,X3,bool,bool)))) | s(bool,X4) = s(bool,sK3(X2,X3,bool,bool)) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f151,f127])).
fof(f987,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK3(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f150,f77])).
fof(f986,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK3(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f150,f127])).
fof(f985,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK3(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f150,f127])).
fof(f984,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK3(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f150,f127])).
fof(f983,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f150,f127])).
fof(f897,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl18),
  inference(cnf_transformation,[],[f897_D])).
fof(f897_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f899,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl21),
  inference(cnf_transformation,[],[f899_D])).
fof(f899_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f894,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f256,f77])).
fof(f903,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl26),
  inference(cnf_transformation,[],[f903_D])).
fof(f903_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f893,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f256,f127])).
fof(f892,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f256,f127])).
fof(f900,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl22),
  inference(cnf_transformation,[],[f900_D])).
fof(f900_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f902,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | $spl25),
  inference(cnf_transformation,[],[f902_D])).
fof(f902_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f891,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f256,f127])).
fof(f890,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f256,f127])).
fof(f895,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl3),
  inference(forward_demodulation,[],[f889,f189])).
fof(f889,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl3),
  inference(superposition,[],[f256,f645])).
fof(f811,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl8),
  inference(cnf_transformation,[],[f811_D])).
fof(f811_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f813,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl11),
  inference(cnf_transformation,[],[f813_D])).
fof(f813_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f808,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f255,f77])).
fof(f817,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) | $spl16),
  inference(cnf_transformation,[],[f817_D])).
fof(f817_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f807,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f255,f127])).
fof(f806,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f255,f127])).
fof(f814,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f814_D])).
fof(f814_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f816,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | $spl15),
  inference(cnf_transformation,[],[f816_D])).
fof(f816_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f805,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f255,f127])).
fof(f804,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f255,f127])).
fof(f809,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl3),
  inference(forward_demodulation,[],[f803,f189])).
fof(f803,plain,(
  s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),sK1))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1))))) | $spl3),
  inference(superposition,[],[f255,f645])).
fof(f265,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl4),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f267,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f267_D])).
fof(f267_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f263,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3(sK1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),bool,sK0))))),
  inference(superposition,[],[f149,f77])).
fof(f173,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl0),
  inference(cnf_transformation,[],[f173_D])).
fof(f173_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f175,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | $spl3),
  inference(cnf_transformation,[],[f175_D])).
fof(f175_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f171,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK3('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.UNIV_0'),s(fun(sK0,bool),sK1)),sK1,bool,sK0))))),
  inference(superposition,[],[f148,f77])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 186
% Passive clauses: 916
% Generated clauses: 16587
% Final active clauses: 58
% Final passive clauses: 166
% Input formulas: 20
% Initial clauses: 39
% 
% Duplicate literals: 914
% Trivial inequalities: 156
% Fw subsumption resolutions: 1332
% Fw demodulations: 3172
% Bw demodulations: 529
% Condensations: 954
% Global subsumptions: 96
% 
% Simple tautologies: 222
% Equational tautologies: 214
% Forward subsumptions: 8391
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 16
% 
% Unit resulting resolution: 1244
% Forward superposition: 3877
% Backward superposition: 3197
% Self superposition: 120
% Equality factoring: 74
% Equality resolution: 25
% 
% Split clauses: 93
% Split components: 197
% SAT solver clauses: 1327
% SAT solver unit clauses: 30
% SAT solver binary clauses: 249
% 
% Sat splits: 93
% Sat splitting refutations: 127
% 
% Memory used [KB]: 6396
% Time elapsed: 0.484 s
% ------------------------------
% ------------------------------
% Success in time 1.803 s
