Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (117591)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 117
% Passive clauses: 2114
% Generated clauses: 19085
% Final active clauses: 104
% Final passive clauses: 1766
% Input formulas: 24
% Initial clauses: 48
% 
% Duplicate literals: 887
% Trivial inequalities: 368
% Fw subsumption resolutions: 3265
% Fw demodulations: 2269
% Bw demodulations: 11
% 
% Simple tautologies: 432
% Equational tautologies: 63
% Forward subsumptions: 8088
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 394
% Factoring: 5
% Forward superposition: 5227
% Backward superposition: 6078
% Self superposition: 236
% Equality factoring: 177
% Equality resolution: 27
% 
% Split clauses: 1019
% Split components: 2087
% SAT solver clauses: 1614
% SAT solver unit clauses: 49
% SAT solver binary clauses: 354
% 
% Sat splits: 1019
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 1603
% TWLsolver calls for satisfiability: 127
% 
% Memory used [KB]: 8315
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (117607)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 92
% Passive clauses: 642
% Generated clauses: 11771
% Final active clauses: 77
% Final passive clauses: 482
% Input formulas: 24
% Initial clauses: 48
% 
% Duplicate literals: 889
% Trivial inequalities: 145
% Fw subsumption resolutions: 1273
% Fw demodulations: 945
% Bw demodulations: 74
% Condensations: 1318
% Global subsumptions: 9
% 
% Simple tautologies: 230
% Equational tautologies: 46
% Forward subsumptions: 5238
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 257
% Factoring: 3
% Forward superposition: 2980
% Backward superposition: 2928
% Self superposition: 148
% Equality factoring: 106
% Equality resolution: 17
% 
% Split clauses: 49
% Split components: 107
% SAT solver clauses: 1123
% SAT solver unit clauses: 63
% SAT solver binary clauses: 49
% SAT solver learnt clauses: 18
% 
% Sat splits: 49
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1087
% TWLsolver calls for satisfiability: 248
% 
% Memory used [KB]: 8315
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f662,plain,(
  $false),
  inference(subsumption_resolution,[],[f661,f135])).
fof(f135,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f90,f87])).
fof(f87,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.CROSS_EMPTY|split|0',aHL_TRUTH)).
fof(f90,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f28])).
fof(f28,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.CROSS_EMPTY|split|0','thm.bool.NOT_CLAUSES|split|1')).
fof(f661,plain,(
  p(s(bool,'F_0'))),
  inference(forward_demodulation,[],[f660,f143])).
fof(f143,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0')))) )),
  inference(unit_resulting_resolution,[],[f123,f101])).
fof(f101,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.CROSS_EMPTY|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f123,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : ! [X12] : ~p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.CROSS_EMPTY|split|0','thm.pred_set.NOT_IN_EMPTY')).
fof(f660,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'type.pair.prod'(sK0,sK1))))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))),
  inference(subsumption_resolution,[],[f654,f87])).
fof(f654,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK1,'const.pair.SND_1'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'type.pair.prod'(sK0,sK1))))),s(fun(sK1,bool),'const.pred_set.EMPTY_0'))))),
  inference(superposition,[],[f132,f553])).
fof(f553,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f135,f545,f177])).
fof(f177,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f174,f135])).
fof(f174,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f126,f94])).
fof(f94,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.CROSS_EMPTY|split|0',aHL_BOOL_CASES)).
fof(f126,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.CROSS_EMPTY|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f545,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(forward_demodulation,[],[f526,f143])).
fof(f526,plain,(
  s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK0,sK1),sK3('const.pred_set.EMPTY_0','const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')),'type.pair.prod'(sK0,sK1))),s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))))),
  inference(unit_resulting_resolution,[],[f86,f125])).
fof(f125,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f48])).
fof(f48,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : ! [X13,X11] : (s(fun(X6,bool),X11) = s(fun(X6,bool),X13) <=> ! [X12] : s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X11))) = s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X13))))),
  file('thm.pred_set.CROSS_EMPTY|split|0','thm.pred_set.EXTENSION')).
fof(f86,plain,(
  s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK0,sK1),bool),'const.pred_set.CROSS_2'(s(fun(sK0,bool),sK2),s(fun(sK1,bool),'const.pred_set.EMPTY_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f55])).
fof(f55,plain,(
  ? [X0,X1,X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(ennf_transformation,[],[f27])).
fof(f27,plain,(
  ~! [X0,X1,X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f26])).
fof(f26,plain,(
  ~! [X0,X1] : ! [X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8] : s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8] : s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),X8),s(fun(X7,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.CROSS_EMPTY|split|0','thm.pred_set.CROSS_EMPTY|split|0')).
fof(f132,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))) )),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))))))))),
  inference(flattening,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))))))))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X8,X14,X12] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X6,X7),X12),s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),X8),s(fun(X7,bool),X14)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X7),X12))),s(fun(X6,bool),X8)))) & p(s(bool,'const.bool.IN_2'(s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),X12))),s(fun(X7,bool),X14))))))),
  file('thm.pred_set.CROSS_EMPTY|split|0','thm.pred_set.IN_CROSS')).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 26
% Passive clauses: 58
% Generated clauses: 564
% Final active clauses: 22
% Final passive clauses: 30
% Input formulas: 24
% Initial clauses: 48
% 
% Duplicate literals: 44
% Trivial inequalities: 11
% Fw subsumption resolutions: 35
% Fw demodulations: 23
% Bw demodulations: 6
% Condensations: 59
% 
% Simple tautologies: 40
% Equational tautologies: 31
% Forward subsumptions: 257
% 
% Unit resulting resolution: 68
% Factoring: 1
% Forward superposition: 79
% Backward superposition: 138
% Self superposition: 23
% Equality factoring: 26
% Equality resolution: 3
% 
% SAT solver clauses: 60
% SAT solver unit clauses: 17
% SAT solver binary clauses: 15
% 
% Memory used [KB]: 1023
% Time elapsed: 0.017 s
% ------------------------------
% ------------------------------
% Success in time 1.341 s
