Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (97027)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 139
% Passive clauses: 2871
% Generated clauses: 21868
% Final active clauses: 120
% Final passive clauses: 2449
% Input formulas: 18
% Initial clauses: 39
% 
% Duplicate literals: 948
% Trivial inequalities: 387
% Fw subsumption resolutions: 2977
% Fw demodulations: 3017
% Bw demodulations: 16
% 
% Simple tautologies: 456
% Equational tautologies: 67
% Forward subsumptions: 8069
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 772
% Factoring: 5
% Forward superposition: 6171
% Backward superposition: 6991
% Self superposition: 209
% Equality factoring: 159
% Equality resolution: 31
% 
% Split clauses: 773
% Split components: 1587
% SAT solver clauses: 1366
% SAT solver unit clauses: 27
% SAT solver binary clauses: 278
% 
% Sat splits: 773
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1356
% TWLsolver calls for satisfiability: 211
% 
% Memory used [KB]: 10746
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (97055)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 145
% Passive clauses: 456
% Generated clauses: 19761
% Final active clauses: 97
% Final passive clauses: 167
% Input formulas: 18
% Initial clauses: 39
% 
% Duplicate literals: 1063
% Trivial inequalities: 87
% Fw subsumption resolutions: 1946
% Fw demodulations: 2163
% Bw demodulations: 136
% Condensations: 1759
% Global subsumptions: 4
% 
% Simple tautologies: 255
% Equational tautologies: 51
% Forward subsumptions: 10029
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 234
% Factoring: 4
% Forward superposition: 4854
% Backward superposition: 5943
% Self superposition: 219
% Equality factoring: 108
% Equality resolution: 23
% 
% Split clauses: 115
% Split components: 231
% SAT solver clauses: 1132
% SAT solver unit clauses: 43
% SAT solver binary clauses: 47
% SAT solver learnt clauses: 14
% SAT solver learnt literals: 11
% 
% Sat splits: 115
% Sat splitting refutations: 56
% 
% TWLsolver clauses: 1110
% TWLsolver calls for satisfiability: 273
% 
% Memory used [KB]: 13304
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f4821,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f68,f110,f108,f67,f109,f107,f272,f622,f626,f624,f627,f631,f726,f763,f1050,f629,f818,f868,f1068,f1066,f1139,f1140,f1141,f1142,f1145,f1148,f1149,f1153,f1154,f1250,f1251,f1252,f1253,f1260,f1254,f1262,f1259,f1486,f1487,f1493,f1491,f1488,f1489,f1496,f1490,f1498,f1495,f2152,f2155,f2157,f2159,f2161,f2164,f2167,f2170,f2173,f2176,f2254,f2256,f2259,f2264,f2275,f3688,f3689,f3800,f3802,f3804,f3805,f3806,f3808,f3810,f3812,f3813,f3814,f3816,f3818,f3820,f3821,f3822,f3904,f3906,f3909,f3912,f3749,f3746,f1257,f1255,f4786,f4788,f4790,f4792,f4794,f4796,f4798,f4800,f4802,f4804,f4806,f4808,f4810,f4812,f4814,f4816,f4818,f4820])).
fof(f4820,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4819,f111])).
fof(f111,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f71,f73])).
fof(f73,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f23])).
fof(f23,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.ABSORPTION','thm.bool.EQ_CLAUSES|split|0')).
fof(f71,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f22])).
fof(f22,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.ABSORPTION',aHL_FALSITY)).
fof(f4819,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f4748,f4425])).
fof(f4425,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK2)))))) ) | $spl20),
  inference(unit_resulting_resolution,[],[f4355,f74])).
fof(f74,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f47])).
fof(f4355,plain,(
  ( ! [X8] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X8),s(fun(sK0,bool),sK2))))))) ) | $spl20),
  inference(subsumption_resolution,[],[f4320,f69])).
fof(f69,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.ABSORPTION',aHL_TRUTH)).
fof(f4320,plain,(
  ( ! [X8] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X8),s(fun(sK0,bool),sK2))))))) ) | $spl20),
  inference(superposition,[],[f100,f1256])).
fof(f1256,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl20),
  inference(cnf_transformation,[],[f1256_D])).
fof(f1256_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f100,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f36])).
fof(f36,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X7,X12,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X12),s(fun(X6,bool),X8)))))) <=> (s(X6,X7) = s(X6,X12) | p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.ABSORPTION','thm.pred_set.IN_INSERT')).
fof(f4748,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f1258,f72])).
fof(f72,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.ABSORPTION',aHL_BOOL_CASES)).
fof(f1258,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(cnf_transformation,[],[f1258_D])).
fof(f1258_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f4818,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4817,f69])).
fof(f4817,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f4749,f4425])).
fof(f4749,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f1258,f74])).
fof(f4816,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4815,f111])).
fof(f4815,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f4750,f4425])).
fof(f4750,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f111,f1258,f125])).
fof(f125,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f72,f72])).
fof(f4814,plain,(
  $false | ($spl3 | $spl8 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4813,f1075])).
fof(f1075,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8)),
  inference(backward_demodulation,[],[f629,f374])).
fof(f374,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f110,f105])).
fof(f105,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.ABSORPTION',aHL_EXT)).
fof(f4813,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f4751,f4425])).
fof(f4751,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f1075,f1258,f125])).
fof(f4812,plain,(
  $false | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4811,f642])).
fof(f642,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7)),
  inference(backward_demodulation,[],[f636,f373])).
fof(f373,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl3),
  inference(unit_resulting_resolution,[],[f110,f105])).
fof(f636,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl7),
  inference(unit_resulting_resolution,[],[f71,f626,f161])).
fof(f161,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f156,f71])).
fof(f156,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f101,f72])).
fof(f101,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.ABSORPTION','thm.bool.IMP_ANTISYM_AX')).
fof(f4811,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f4752,f4425])).
fof(f4752,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f642,f1258,f125])).
fof(f4810,plain,(
  $false | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(subsumption_resolution,[],[f4809,f1499])).
fof(f1499,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl32)),
  inference(backward_demodulation,[],[f1495,f572])).
fof(f572,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f110,f97])).
fof(f97,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f62])).
fof(f62,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X8,X11] : (s(fun(X6,bool),X8) = s(fun(X6,bool),X11) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X11))))),
  file('thm.pred_set.ABSORPTION','thm.pred_set.EXTENSION')).
fof(f4809,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f4753,f4425])).
fof(f4753,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1499,f1258,f125])).
fof(f4808,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4807,f111])).
fof(f4807,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f4756,f4425])).
fof(f4756,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f111,f1258,f125])).
fof(f4806,plain,(
  $false | ($spl3 | $spl8 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4805,f1075])).
fof(f4805,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f4757,f4425])).
fof(f4757,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f1075,f1258,f125])).
fof(f4804,plain,(
  $false | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4803,f642])).
fof(f4803,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f4758,f4425])).
fof(f4758,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f642,f1258,f125])).
fof(f4802,plain,(
  $false | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(subsumption_resolution,[],[f4801,f1499])).
fof(f4801,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f4759,f4425])).
fof(f4759,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1499,f1258,f125])).
fof(f4800,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4799,f111])).
fof(f4799,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f4762,f4425])).
fof(f4762,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl23),
  inference(unit_resulting_resolution,[],[f71,f1258,f161])).
fof(f4798,plain,(
  $false | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4797,f642])).
fof(f4797,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f4764,f4425])).
fof(f4764,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f1258,f933])).
fof(f933,plain,(
  ( ! [X38] : (s(bool,X38) = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl3 | $spl7)),
  inference(trivial_inequality_removal,[],[f819])).
fof(f819,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X38) = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X38)) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f642,f125])).
fof(f4796,plain,(
  $false | ($spl3 | $spl8 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4795,f1075])).
fof(f4795,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f4765,f4425])).
fof(f4765,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl8 | $spl23)),
  inference(unit_resulting_resolution,[],[f1258,f1178])).
fof(f1178,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl8)),
  inference(trivial_inequality_removal,[],[f1172])).
fof(f1172,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl8)),
  inference(superposition,[],[f1075,f125])).
fof(f4794,plain,(
  $false | ($spl20 | $spl23 | $spl24)),
  inference(subsumption_resolution,[],[f4793,f1256])).
fof(f4793,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl20 | $spl23 | $spl24)),
  inference(forward_demodulation,[],[f4766,f4425])).
fof(f4766,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl23 | $spl24)),
  inference(unit_resulting_resolution,[],[f1258,f1260])).
fof(f4792,plain,(
  $false | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(subsumption_resolution,[],[f4791,f1499])).
fof(f4791,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl20 | $spl23 | $spl32)),
  inference(forward_demodulation,[],[f4768,f4425])).
fof(f4768,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl23 | $spl32)),
  inference(unit_resulting_resolution,[],[f1258,f1527])).
fof(f1527,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl32)),
  inference(trivial_inequality_removal,[],[f1521])).
fof(f1521,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | ($spl3 | $spl32)),
  inference(superposition,[],[f1499,f125])).
fof(f4790,plain,(
  $false | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4789,f636])).
fof(f4789,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f4769,f4425])).
fof(f4769,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl7 | $spl23)),
  inference(unit_resulting_resolution,[],[f1258,f1908])).
fof(f1908,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl3 | $spl7)),
  inference(subsumption_resolution,[],[f1846,f110])).
fof(f1846,plain,(
  ( ! [X25] : (s(bool,X25) != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X25)) ) | ($spl3 | $spl7)),
  inference(superposition,[],[f105,f933])).
fof(f4788,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4787,f69])).
fof(f4787,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f4771,f4425])).
fof(f4771,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f69,f1258,f104])).
fof(f104,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f43])).
fof(f4786,plain,(
  $false | ($spl20 | $spl23)),
  inference(subsumption_resolution,[],[f4785,f69])).
fof(f4785,plain,(
  ~p(s(bool,'T_0')) | ($spl20 | $spl23)),
  inference(forward_demodulation,[],[f4774,f4425])).
fof(f4774,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) | $spl23),
  inference(unit_resulting_resolution,[],[f69,f1258,f104])).
fof(f1255,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f1255_D])).
fof(f1255_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1257,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(cnf_transformation,[],[f1257_D])).
fof(f1257_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f3746,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f3393,f125])).
fof(f3393,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f110,f97])).
fof(f3749,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X3) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X3)) ) | $spl3),
  inference(superposition,[],[f3393,f125])).
fof(f3912,plain,(
  $false | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f3886,f69])).
fof(f3886,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f3823,f3787])).
fof(f3787,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl21),
  inference(unit_resulting_resolution,[],[f69,f1257,f104])).
fof(f3823,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f3797,f1260])).
fof(f3797,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(trivial_inequality_removal,[],[f3789])).
fof(f3789,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X0)) ) | $spl21),
  inference(superposition,[],[f1257,f125])).
fof(f3909,plain,(
  $false | ($spl21 | $spl24)),
  inference(global_subsumption,[],[f68,f67,f69,f72,f74,f73,f104,f101,f105,f111,f125,f161,f273,f689,f731,f732,f745,f784,f1064,f1065,f823,f824,f832,f1601,f1603,f1606,f1621,f1623,f1626,f1641,f831,f744,f827,f800,f799,f797,f796,f742,f741,f740,f739,f738,f99,f96,f100,f98,f1069,f1070,f97,f743,f830,f1969,f1628,f1625,f1608,f1605,f575,f574,f573,f485,f484,f479,f478,f376,f375,f321,f320,f249,f247,f243,f241,f78,f77,f71,f825,f828,f798,f826,f829,f1257,f1260,f3787,f3823,f3905,f3908])).
fof(f3908,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f3856,f3823])).
fof(f3856,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f3823,f97])).
fof(f3905,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f3825,f3823])).
fof(f3825,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f3823,f77])).
fof(f829,plain,(
  ( ! [X87,X85,X86,X84] : (p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X85),s(fun(bool,bool),X86)))))) | s(bool,X84) != s(bool,X85) | s(bool,X84) = s(bool,X87) | s(bool,'T_0') = s(bool,X87)) )),
  inference(superposition,[],[f99,f125])).
fof(f826,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X73),s(fun(bool,bool),X74)))))) | s(bool,X72) = s(bool,X73) | p(s(bool,'const.bool.IN_2'(s(bool,X72),s(fun(bool,bool),X74)))) | s(bool,X72) = s(bool,X75) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f98,f125])).
fof(f798,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'const.bool.IN_2'(s(bool,X45),s(fun(bool,bool),X46)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X47),s(fun(bool,bool),X46)))))) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f100,f125])).
fof(f828,plain,(
  ( ! [X80,X83,X81,X82] : (p(s(bool,'const.bool.IN_2'(s(bool,X81),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X82)))))) | s(bool,X80) != s(bool,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X83)) )),
  inference(superposition,[],[f99,f125])).
fof(f825,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.bool.IN_2'(s(bool,X69),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X70)))))) | s(bool,X68) = s(bool,X69) | p(s(bool,'const.bool.IN_2'(s(bool,X69),s(fun(bool,bool),X70)))) | s(bool,X68) = s(bool,X71) | s(bool,'T_0') = s(bool,X71)) )),
  inference(superposition,[],[f98,f125])).
fof(f77,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.ABSORPTION','thm.bool.EQ_CLAUSES|split|3')).
fof(f78,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f49])).
fof(f241,plain,(
  ( ! [X4,X2,X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(bool,X2),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X3),s(fun(bool,bool),X4)))))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f99,f72])).
fof(f243,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X9),s(fun(bool,bool),X10)))))) | s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f99,f72])).
fof(f247,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X2),s(fun(bool,bool),X4)))))) | s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f99,f72])).
fof(f249,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,'const.bool.IN_2'(s(bool,X9),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X10)))))) | s(bool,X8) != s(bool,X9) | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f99,f72])).
fof(f320,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),X1)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X2),s(fun(bool,bool),X1)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f100,f72])).
fof(f321,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X4)))) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X5),s(fun(bool,bool),X4)))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f100,f72])).
fof(f375,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK4(X0,X1,X2,bool))) )),
  inference(superposition,[],[f105,f72])).
fof(f376,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK4(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK4(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f105,f72])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2)))))) | s(bool,'F_0') = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f98,f72])).
fof(f479,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X4),s(fun(bool,bool),X5)))))) | s(bool,X3) = s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X3),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f98,f72])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2)))))) | s(bool,'F_0') = s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),X2)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f98,f72])).
fof(f485,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X5)))))) | s(bool,X3) = s(bool,X4) | p(s(bool,'const.bool.IN_2'(s(bool,X4),s(fun(bool,bool),X5)))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f98,f72])).
fof(f573,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X2)))) )),
  inference(superposition,[],[f97,f72])).
fof(f574,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X0))) != s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),X1))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK3(X0,X1,bool))) )),
  inference(superposition,[],[f97,f72])).
fof(f575,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,sK3(X1,X2,X0)),s(fun(X0,bool),X1)))) )),
  inference(superposition,[],[f97,f72])).
fof(f1605,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X16),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X17),s(fun(bool,bool),X18))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f273,f72])).
fof(f1608,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X28),s(fun(bool,bool),X29))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f273,f72])).
fof(f1625,plain,(
  ( ! [X17,X18,X16] : (s(bool,'F_0') != s(bool,X17) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X17),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X16),s(fun(bool,bool),X18))))) | s(bool,'T_0') = s(bool,X16)) )),
  inference(superposition,[],[f273,f72])).
fof(f1628,plain,(
  ( ! [X28,X29,X27] : (s(bool,X27) != s(bool,X28) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X28),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X29))))) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f273,f72])).
fof(f1969,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X1))))))) )),
  inference(condensation,[],[f1968])).
fof(f1968,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X26)))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) )),
  inference(trivial_inequality_removal,[],[f1949])).
fof(f1949,plain,(
  ( ! [X26,X24,X25] : (s(bool,X24) != s(bool,X24) | p(s(bool,'const.bool.IN_2'(s(bool,X25),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'F_0'),s(fun(bool,bool),X26)))))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f241,f72])).
fof(f830,plain,(
  ( ! [X94,X92,X93,X91] : (s(bool,'T_0') != s(bool,f(s(X91,sK4(X92,X93,bool,X91)))) | s(fun(X91,bool),X92) = s(fun(X91,bool),X93) | s(bool,X94) = s(bool,g(s(X91,sK4(X92,X93,bool,X91)))) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f105,f125])).
fof(f743,plain,(
  ( ! [X103,X101,X102,X100] : (s(bool,X103) != s(bool,f(s(X100,sK4(X101,X102,bool,X100)))) | s(fun(X100,bool),X101) = s(fun(X100,bool),X102) | s(bool,'T_0') = s(bool,X103) | s(bool,'T_0') = s(bool,g(s(X100,sK4(X101,X102,bool,X100))))) )),
  inference(superposition,[],[f105,f125])).
fof(f1070,plain,(
  ( ! [X52,X50,X53,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(X49,X50),s(fun(X49,bool),'const.pred_set.INSERT_2'(s(X49,X52),s(fun(X49,bool),X51)))))) | s(bool,X53) = s(bool,'const.bool.IN_2'(s(X49,X50),s(fun(X49,bool),X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f822,f69])).
fof(f822,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X49,X50),s(fun(X49,bool),'const.pred_set.INSERT_2'(s(X49,X52),s(fun(X49,bool),X51)))))) | s(bool,X53) = s(bool,'const.bool.IN_2'(s(X49,X50),s(fun(X49,bool),X51))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f100,f125])).
fof(f1069,plain,(
  ( ! [X39,X43,X41,X42,X40] : (s(X39,X40) = s(X39,X41) | p(s(bool,'const.bool.IN_2'(s(X39,X40),s(fun(X39,bool),X42)))) | s(bool,X43) = s(bool,'const.bool.IN_2'(s(X39,X40),s(fun(X39,bool),'const.pred_set.INSERT_2'(s(X39,X41),s(fun(X39,bool),X42))))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f820,f69])).
fof(f820,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'T_0')) | s(X39,X40) = s(X39,X41) | p(s(bool,'const.bool.IN_2'(s(X39,X40),s(fun(X39,bool),X42)))) | s(bool,X43) = s(bool,'const.bool.IN_2'(s(X39,X40),s(fun(X39,bool),'const.pred_set.INSERT_2'(s(X39,X41),s(fun(X39,bool),X42))))) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f98,f125])).
fof(f98,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f65])).
fof(f96,plain,(
  ( ! [X4,X2,X0,X1] : (s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f99,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) != s(X0,X2)) )),
  inference(cnf_transformation,[],[f65])).
fof(f738,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'const.bool.IN_2'(s(bool,X78),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X77),s(fun(bool,bool),X79)))))) | s(bool,X76) = s(bool,X78) | p(s(bool,'const.bool.IN_2'(s(bool,X78),s(fun(bool,bool),X79)))) | s(bool,'T_0') = s(bool,X77) | s(bool,'T_0') = s(bool,X76)) )),
  inference(superposition,[],[f98,f125])).
fof(f739,plain,(
  ( ! [X80,X83,X81,X82] : (~p(s(bool,'const.bool.IN_2'(s(bool,X81),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X82),s(fun(bool,bool),X83)))))) | s(bool,X80) = s(bool,X82) | p(s(bool,'const.bool.IN_2'(s(bool,X80),s(fun(bool,bool),X83)))) | s(bool,'T_0') = s(bool,X81) | s(bool,'T_0') = s(bool,X80)) )),
  inference(superposition,[],[f98,f125])).
fof(f740,plain,(
  ( ! [X87,X85,X86,X84] : (~p(s(bool,'const.bool.IN_2'(s(bool,X85),s(fun(bool,bool),X86)))) | p(s(bool,'const.bool.IN_2'(s(bool,X84),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X87),s(fun(bool,bool),X86)))))) | s(bool,'T_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X84)) )),
  inference(superposition,[],[f100,f125])).
fof(f741,plain,(
  ( ! [X90,X88,X91,X89] : (p(s(bool,'const.bool.IN_2'(s(bool,X90),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X89),s(fun(bool,bool),X91)))))) | s(bool,X88) != s(bool,X90) | s(bool,'T_0') = s(bool,X89) | s(bool,'T_0') = s(bool,X88)) )),
  inference(superposition,[],[f99,f125])).
fof(f742,plain,(
  ( ! [X94,X92,X95,X93] : (p(s(bool,'const.bool.IN_2'(s(bool,X93),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X94),s(fun(bool,bool),X95)))))) | s(bool,X92) != s(bool,X94) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f99,f125])).
fof(f796,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.bool.IN_2'(s(bool,X38),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X37),s(fun(bool,bool),X39)))))) | s(bool,'T_0') = s(bool,X38) | p(s(bool,'const.bool.IN_2'(s(bool,X38),s(fun(bool,bool),X39)))) | s(bool,X37) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f98,f125])).
fof(f797,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.bool.IN_2'(s(bool,X41),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X42),s(fun(bool,bool),X43)))))) | s(bool,'T_0') = s(bool,X42) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X43)))) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f98,f125])).
fof(f799,plain,(
  ( ! [X52,X50,X51,X49] : (p(s(bool,'const.bool.IN_2'(s(bool,X50),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X49),s(fun(bool,bool),X51)))))) | s(bool,'T_0') != s(bool,X50) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f99,f125])).
fof(f800,plain,(
  ( ! [X54,X56,X55,X53] : (p(s(bool,'const.bool.IN_2'(s(bool,X53),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X54),s(fun(bool,bool),X55)))))) | s(bool,'T_0') != s(bool,X54) | s(bool,X53) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f99,f125])).
fof(f827,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X77)))) | p(s(bool,'const.bool.IN_2'(s(bool,X76),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X78),s(fun(bool,bool),X77)))))) | s(bool,X76) = s(bool,X79) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f100,f125])).
fof(f744,plain,(
  ( ! [X105,X106,X104] : (s(bool,'const.bool.IN_2'(s(bool,X106),s(fun(bool,bool),X104))) != s(bool,'const.bool.IN_2'(s(bool,X106),s(fun(bool,bool),X105))) | s(fun(bool,bool),X104) = s(fun(bool,bool),X105) | s(bool,'T_0') = s(bool,X106) | s(bool,'T_0') = s(bool,sK3(X104,X105,bool))) )),
  inference(superposition,[],[f97,f125])).
fof(f831,plain,(
  ( ! [X97,X95,X96] : (s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X95))) != s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),X96))) | s(fun(bool,bool),X95) = s(fun(bool,bool),X96) | s(bool,X97) = s(bool,sK3(X95,X96,bool)) | s(bool,'T_0') = s(bool,X97)) )),
  inference(superposition,[],[f97,f125])).
fof(f1641,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))) )),
  inference(equality_resolution,[],[f273])).
fof(f1626,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X20),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,'T_0'),s(fun(bool,bool),X21))))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f273,f125])).
fof(f1623,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X10),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X8),s(fun(bool,bool),X11))))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f273,f125])).
fof(f1621,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X1),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X0),s(fun(bool,bool),X2))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f273,f125])).
fof(f1606,plain,(
  ( ! [X21,X19,X22,X20] : (s(bool,X19) != s(bool,X20) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X20),s(fun(bool,bool),X21))))) | s(bool,X19) = s(bool,X22) | s(bool,'T_0') = s(bool,X22)) )),
  inference(superposition,[],[f273,f125])).
fof(f1603,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,X9) != s(bool,X10) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X8),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X10),s(fun(bool,bool),X11))))) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) )),
  inference(superposition,[],[f273,f125])).
fof(f1601,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),'const.pred_set.INSERT_2'(s(bool,X1),s(fun(bool,bool),X2))))) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f273,f125])).
fof(f832,plain,(
  ( ! [X101,X99,X100,X98] : (s(X100,f(s(bool,'T_0'))) != s(X100,g(s(bool,'T_0'))) | s(fun(bool,X100),X98) = s(fun(bool,X100),X99) | s(bool,X101) = s(bool,sK4(X98,X99,X100,bool)) | s(bool,'T_0') = s(bool,X101)) )),
  inference(superposition,[],[f105,f125])).
fof(f824,plain,(
  ( ! [X61,X59,X60,X58] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X58,sK3(X59,X60,X58)),s(fun(X58,bool),X60))) | s(fun(X58,bool),X59) = s(fun(X58,bool),X60) | s(bool,X61) = s(bool,'const.bool.IN_2'(s(X58,sK3(X59,X60,X58)),s(fun(X58,bool),X59))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f97,f125])).
fof(f823,plain,(
  ( ! [X57,X54,X56,X55] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X54,sK3(X55,X56,X54)),s(fun(X54,bool),X55))) | s(fun(X54,bool),X55) = s(fun(X54,bool),X56) | s(bool,X57) = s(bool,'const.bool.IN_2'(s(X54,sK3(X55,X56,X54)),s(fun(X54,bool),X56))) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f97,f125])).
fof(f1065,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f808,f69])).
fof(f808,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f104,f125])).
fof(f1064,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f804,f69])).
fof(f804,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f77,f125])).
fof(f784,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f73,f125])).
fof(f745,plain,(
  ( ! [X109,X107,X110,X108] : (s(X109,f(s(bool,X110))) != s(X109,g(s(bool,X110))) | s(fun(bool,X109),X107) = s(fun(bool,X109),X108) | s(bool,'T_0') = s(bool,X110) | s(bool,'T_0') = s(bool,sK4(X107,X108,X109,bool))) )),
  inference(superposition,[],[f105,f125])).
fof(f732,plain,(
  ( ! [X61,X64,X62,X63] : (s(bool,X64) != s(bool,'const.bool.IN_2'(s(X61,sK3(X62,X63,X61)),s(fun(X61,bool),X63))) | s(fun(X61,bool),X62) = s(fun(X61,bool),X63) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X61,sK3(X62,X63,X61)),s(fun(X61,bool),X62)))) )),
  inference(superposition,[],[f97,f125])).
fof(f731,plain,(
  ( ! [X59,X57,X60,X58] : (s(bool,X60) != s(bool,'const.bool.IN_2'(s(X57,sK3(X58,X59,X57)),s(fun(X57,bool),X58))) | s(fun(X57,bool),X58) = s(fun(X57,bool),X59) | s(bool,'T_0') = s(bool,X60) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X57,sK3(X58,X59,X57)),s(fun(X57,bool),X59)))) )),
  inference(superposition,[],[f97,f125])).
fof(f689,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f125,f125])).
fof(f273,plain,(
  ( ! [X6,X4,X7,X5] : (s(X4,X5) != s(X4,X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.INSERT_2'(s(X4,X6),s(fun(X4,bool),X7)))))) )),
  inference(subsumption_resolution,[],[f255,f71])).
fof(f255,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,'F_0')) | s(X4,X5) != s(X4,X6) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X4,X5),s(fun(X4,bool),'const.pred_set.INSERT_2'(s(X4,X6),s(fun(X4,bool),X7)))))) )),
  inference(superposition,[],[f99,f72])).
fof(f3906,plain,(
  $false | ($spl21 | $spl24)),
  inference(global_subsumption,[],[f68,f67,f69,f72,f74,f73,f104,f101,f105,f111,f125,f161,f273,f689,f731,f732,f745,f784,f1064,f1065,f823,f824,f832,f1601,f1603,f1606,f1621,f1623,f1626,f1641,f831,f744,f827,f800,f799,f797,f796,f742,f741,f740,f739,f738,f99,f96,f100,f98,f1069,f1070,f97,f743,f830,f1969,f1628,f1625,f1608,f1605,f575,f574,f573,f485,f484,f479,f478,f376,f375,f321,f320,f249,f247,f243,f241,f78,f77,f71,f825,f828,f798,f826,f829,f1257,f1260,f3787,f3823,f3905])).
fof(f3904,plain,(
  $false | ($spl21 | $spl24)),
  inference(subsumption_resolution,[],[f3824,f69])).
fof(f3824,plain,(
  ~p(s(bool,'T_0')) | ($spl21 | $spl24)),
  inference(backward_demodulation,[],[f3823,f71])).
fof(f3822,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3759,f1262])).
fof(f3759,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f1257,f72])).
fof(f3821,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3761,f1262])).
fof(f3761,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f111,f1257,f125])).
fof(f3820,plain,(
  $false | ($spl3 | $spl8 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3819,f1262])).
fof(f3819,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21)),
  inference(forward_demodulation,[],[f3762,f1165])).
fof(f1165,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8)),
  inference(unit_resulting_resolution,[],[f1075,f72])).
fof(f3762,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1075,f1257,f125])).
fof(f3818,plain,(
  $false | ($spl3 | $spl7 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3817,f1262])).
fof(f3817,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl7 | $spl21)),
  inference(forward_demodulation,[],[f3763,f648])).
fof(f648,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7)),
  inference(unit_resulting_resolution,[],[f642,f72])).
fof(f3763,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl7 | $spl21)),
  inference(unit_resulting_resolution,[],[f642,f1257,f125])).
fof(f3816,plain,(
  $false | ($spl3 | $spl21 | $spl27 | $spl32)),
  inference(subsumption_resolution,[],[f3815,f1262])).
fof(f3815,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f3764,f1514])).
fof(f1514,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl32)),
  inference(unit_resulting_resolution,[],[f1499,f72])).
fof(f3764,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(unit_resulting_resolution,[],[f1499,f1257,f125])).
fof(f3814,plain,(
  $false | ($spl3 | $spl21 | $spl23)),
  inference(subsumption_resolution,[],[f3767,f1258])).
fof(f3767,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f3393,f1257,f125])).
fof(f3813,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3768,f1262])).
fof(f3768,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f111,f1257,f125])).
fof(f3812,plain,(
  $false | ($spl3 | $spl8 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3811,f1262])).
fof(f3811,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21)),
  inference(forward_demodulation,[],[f3769,f1165])).
fof(f3769,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1075,f1257,f125])).
fof(f3810,plain,(
  $false | ($spl3 | $spl7 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3809,f1262])).
fof(f3809,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl7 | $spl21)),
  inference(forward_demodulation,[],[f3770,f648])).
fof(f3770,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl7 | $spl21)),
  inference(unit_resulting_resolution,[],[f642,f1257,f125])).
fof(f3808,plain,(
  $false | ($spl3 | $spl21 | $spl27 | $spl32)),
  inference(subsumption_resolution,[],[f3807,f1262])).
fof(f3807,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f3771,f1514])).
fof(f3771,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(unit_resulting_resolution,[],[f1499,f1257,f125])).
fof(f3806,plain,(
  $false | ($spl3 | $spl21 | $spl23)),
  inference(subsumption_resolution,[],[f3774,f1258])).
fof(f3774,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | ($spl3 | $spl21)),
  inference(unit_resulting_resolution,[],[f3393,f1257,f125])).
fof(f3805,plain,(
  $false | ($spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3775,f1262])).
fof(f3775,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl21),
  inference(unit_resulting_resolution,[],[f71,f1257,f161])).
fof(f3804,plain,(
  $false | ($spl3 | $spl7 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3803,f1262])).
fof(f3803,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl7 | $spl21)),
  inference(forward_demodulation,[],[f3778,f648])).
fof(f3778,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl7 | $spl21)),
  inference(unit_resulting_resolution,[],[f1257,f933])).
fof(f3802,plain,(
  $false | ($spl3 | $spl8 | $spl21 | $spl27)),
  inference(subsumption_resolution,[],[f3801,f1262])).
fof(f3801,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21)),
  inference(forward_demodulation,[],[f3779,f1165])).
fof(f3779,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl8 | $spl21)),
  inference(unit_resulting_resolution,[],[f1257,f1178])).
fof(f3800,plain,(
  $false | ($spl3 | $spl21 | $spl27 | $spl32)),
  inference(subsumption_resolution,[],[f3799,f1262])).
fof(f3799,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(forward_demodulation,[],[f3781,f1514])).
fof(f3781,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl21 | $spl32)),
  inference(unit_resulting_resolution,[],[f1257,f1527])).
fof(f3689,plain,(
  $false | ($spl0 | $spl3 | $spl22)),
  inference(subsumption_resolution,[],[f3681,f279])).
fof(f279,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl0),
  inference(unit_resulting_resolution,[],[f107,f74])).
fof(f3681,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | ($spl3 | $spl22)),
  inference(backward_demodulation,[],[f3676,f1263])).
fof(f1263,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl22)),
  inference(backward_demodulation,[],[f1259,f571])).
fof(f571,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(unit_resulting_resolution,[],[f110,f97])).
fof(f3676,plain,(
  s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)) | ($spl3 | $spl22)),
  inference(subsumption_resolution,[],[f3675,f71])).
fof(f3675,plain,(
  p(s(bool,'F_0')) | s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)) | ($spl3 | $spl22)),
  inference(forward_demodulation,[],[f3674,f1365])).
fof(f1365,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl22)),
  inference(unit_resulting_resolution,[],[f1263,f72])).
fof(f3674,plain,(
  s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl22),
  inference(subsumption_resolution,[],[f3629,f69])).
fof(f3629,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl22),
  inference(superposition,[],[f98,f1259])).
fof(f3688,plain,(
  $false | ($spl0 | $spl3 | $spl22)),
  inference(subsumption_resolution,[],[f3687,f111])).
fof(f3687,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3 | $spl22)),
  inference(forward_demodulation,[],[f3679,f279])).
fof(f3679,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | ($spl3 | $spl22)),
  inference(backward_demodulation,[],[f3676,f1365])).
fof(f2275,plain,(
  $false | ($spl8 | $spl14)),
  inference(global_subsumption,[],[f68,f67,f69,f71,f72,f74,f73,f78,f77,f97,f96,f100,f99,f98,f104,f101,f105,f111,f125,f161,f241,f247,f273,f249,f243,f320,f321,f375,f376,f478,f479,f484,f485,f573,f574,f575,f689,f731,f732,f738,f739,f740,f741,f742,f743,f744,f745,f784,f796,f797,f798,f799,f800,f1064,f1065,f1069,f1070,f823,f824,f825,f826,f827,f828,f829,f830,f831,f832,f629,f1066,f1601,f1603,f1605,f1606,f1608,f1621,f1623,f1625,f1626,f1628,f1641,f1969,f2143,f2255,f2258,f2263,f2274])).
fof(f2274,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2123,f2143])).
fof(f2123,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,X0) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | $spl14),
  inference(equality_factoring,[],[f1066])).
fof(f2263,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2262,f2143])).
fof(f2262,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | p(s(bool,X1))) ) | $spl14),
  inference(subsumption_resolution,[],[f2056,f73])).
fof(f2056,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | p(s(bool,X1)) | s(bool,'T_0') = s(bool,X1)) ) | $spl14),
  inference(superposition,[],[f73,f1066])).
fof(f2258,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2209,f2143])).
fof(f2209,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) ) | ($spl8 | $spl14)),
  inference(backward_demodulation,[],[f2143,f97])).
fof(f2255,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2178,f2143])).
fof(f2178,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | ($spl8 | $spl14)),
  inference(backward_demodulation,[],[f2143,f77])).
fof(f2143,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2)) ) | ($spl8 | $spl14)),
  inference(duplicate_literal_removal,[],[f2048])).
fof(f2048,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,X2) | s(bool,'T_0') = s(bool,X2)) ) | ($spl8 | $spl14)),
  inference(superposition,[],[f1066,f629])).
fof(f2264,plain,(
  $false | ($spl8 | $spl14)),
  inference(global_subsumption,[],[f68,f67,f69,f71,f72,f74,f73,f78,f77,f97,f96,f100,f99,f98,f104,f101,f105,f111,f125,f161,f241,f247,f273,f249,f243,f320,f321,f375,f376,f478,f479,f484,f485,f573,f574,f575,f689,f731,f732,f738,f739,f740,f741,f742,f743,f744,f745,f784,f796,f797,f798,f799,f800,f1064,f1065,f1069,f1070,f823,f824,f825,f826,f827,f828,f829,f830,f831,f832,f629,f1066,f1601,f1603,f1605,f1606,f1608,f1621,f1623,f1625,f1626,f1628,f1641,f1969,f2143,f2255,f2258,f2263])).
fof(f2259,plain,(
  $false | ($spl8 | $spl14)),
  inference(global_subsumption,[],[f68,f67,f69,f71,f72,f74,f73,f78,f77,f97,f96,f100,f99,f98,f104,f101,f105,f111,f125,f161,f241,f247,f273,f249,f243,f320,f321,f375,f376,f478,f479,f484,f485,f573,f574,f575,f689,f731,f732,f738,f739,f740,f741,f742,f743,f744,f745,f784,f796,f797,f798,f799,f800,f1064,f1065,f1069,f1070,f823,f824,f825,f826,f827,f828,f829,f830,f831,f832,f629,f1066,f1601,f1603,f1605,f1606,f1608,f1621,f1623,f1625,f1626,f1628,f1641,f1969,f2143,f2255,f2258])).
fof(f2256,plain,(
  $false | ($spl8 | $spl14)),
  inference(global_subsumption,[],[f68,f67,f69,f71,f72,f74,f73,f78,f77,f97,f96,f100,f99,f98,f104,f101,f105,f111,f125,f161,f241,f247,f273,f249,f243,f320,f321,f375,f376,f478,f479,f484,f485,f573,f574,f575,f689,f731,f732,f738,f739,f740,f741,f742,f743,f744,f745,f784,f796,f797,f798,f799,f800,f1064,f1065,f1069,f1070,f823,f824,f825,f826,f827,f828,f829,f830,f831,f832,f629,f1066,f1601,f1603,f1605,f1606,f1608,f1621,f1623,f1625,f1626,f1628,f1641,f1969,f2143,f2255])).
fof(f2254,plain,(
  $false | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2177,f69])).
fof(f2177,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl14)),
  inference(backward_demodulation,[],[f2143,f71])).
fof(f2176,plain,(
  $false | ($spl3 | $spl8 | $spl14 | $spl32)),
  inference(subsumption_resolution,[],[f2175,f111])).
fof(f2175,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl8 | $spl14 | $spl32)),
  inference(forward_demodulation,[],[f2174,f629])).
fof(f2174,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl14 | $spl32)),
  inference(subsumption_resolution,[],[f2043,f1499])).
fof(f2043,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl14 | $spl32)),
  inference(superposition,[],[f1066,f1514])).
fof(f2173,plain,(
  $false | ($spl3 | $spl8 | $spl14 | $spl22)),
  inference(subsumption_resolution,[],[f2172,f111])).
fof(f2172,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl8 | $spl14 | $spl22)),
  inference(forward_demodulation,[],[f2171,f629])).
fof(f2171,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl14 | $spl22)),
  inference(subsumption_resolution,[],[f2042,f1263])).
fof(f2042,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl14 | $spl22)),
  inference(superposition,[],[f1066,f1365])).
fof(f2170,plain,(
  $false | ($spl3 | $spl7 | $spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2169,f111])).
fof(f2169,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl7 | $spl8 | $spl14)),
  inference(forward_demodulation,[],[f2168,f629])).
fof(f2168,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl7 | $spl14)),
  inference(subsumption_resolution,[],[f2039,f642])).
fof(f2039,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7 | $spl14)),
  inference(superposition,[],[f1066,f648])).
fof(f2167,plain,(
  $false | ($spl3 | $spl7 | $spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2166,f642])).
fof(f2166,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7 | $spl8 | $spl14)),
  inference(condensation,[],[f2165])).
fof(f2165,plain,(
  ( ! [X9] : (s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl8 | $spl14)),
  inference(forward_demodulation,[],[f2144,f629])).
fof(f2144,plain,(
  ( ! [X9] : (s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl14)),
  inference(duplicate_literal_removal,[],[f2037])).
fof(f2037,plain,(
  ( ! [X9] : (s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | ($spl3 | $spl7 | $spl14)),
  inference(superposition,[],[f1066,f933])).
fof(f2164,plain,(
  $false | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2163,f111])).
fof(f2163,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl14)),
  inference(condensation,[],[f2162])).
fof(f2162,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X0)) ) | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f2149,f629])).
fof(f2149,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f2032])).
fof(f2032,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(superposition,[],[f1066,f72])).
fof(f2161,plain,(
  $false | ($spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2160,f111])).
fof(f2160,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl8 | $spl14)),
  inference(forward_demodulation,[],[f2027,f629])).
fof(f2027,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl14),
  inference(unit_resulting_resolution,[],[f111,f1066])).
fof(f2159,plain,(
  $false | ($spl3 | $spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2158,f1075])).
fof(f2158,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8 | $spl14)),
  inference(forward_demodulation,[],[f2028,f629])).
fof(f2028,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl8 | $spl14)),
  inference(unit_resulting_resolution,[],[f1075,f1066])).
fof(f2157,plain,(
  $false | ($spl3 | $spl7 | $spl8 | $spl14)),
  inference(subsumption_resolution,[],[f2156,f642])).
fof(f2156,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | ($spl3 | $spl7 | $spl8 | $spl14)),
  inference(forward_demodulation,[],[f2029,f629])).
fof(f2029,plain,(
  s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl7 | $spl14)),
  inference(unit_resulting_resolution,[],[f642,f1066])).
fof(f2155,plain,(
  $false | ($spl3 | $spl8 | $spl14 | $spl22)),
  inference(subsumption_resolution,[],[f2154,f111])).
fof(f2154,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl8 | $spl14 | $spl22)),
  inference(forward_demodulation,[],[f2153,f629])).
fof(f2153,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl14 | $spl22)),
  inference(forward_demodulation,[],[f2030,f1365])).
fof(f2030,plain,(
  s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl14 | $spl22)),
  inference(unit_resulting_resolution,[],[f1263,f1066])).
fof(f2152,plain,(
  $false | ($spl3 | $spl8 | $spl14 | $spl32)),
  inference(subsumption_resolution,[],[f2151,f111])).
fof(f2151,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl3 | $spl8 | $spl14 | $spl32)),
  inference(forward_demodulation,[],[f2150,f629])).
fof(f2150,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | ($spl3 | $spl14 | $spl32)),
  inference(forward_demodulation,[],[f2031,f1514])).
fof(f2031,plain,(
  s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | ($spl3 | $spl14 | $spl32)),
  inference(unit_resulting_resolution,[],[f1499,f1066])).
fof(f1495,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl32),
  inference(cnf_transformation,[],[f1495_D])).
fof(f1495_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1498,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl37),
  inference(cnf_transformation,[],[f1498_D])).
fof(f1498_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f1490,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(superposition,[],[f572,f72])).
fof(f1496,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl34),
  inference(cnf_transformation,[],[f1496_D])).
fof(f1496_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f1489,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) ) | $spl3),
  inference(superposition,[],[f572,f125])).
fof(f1488,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f572,f125])).
fof(f1491,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl28),
  inference(cnf_transformation,[],[f1491_D])).
fof(f1491_D,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f1493,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | $spl31),
  inference(cnf_transformation,[],[f1493_D])).
fof(f1493_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1487,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f572,f125])).
fof(f1486,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f572,f125])).
fof(f1259,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl22),
  inference(cnf_transformation,[],[f1259_D])).
fof(f1259_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f1262,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl27),
  inference(cnf_transformation,[],[f1262_D])).
fof(f1262_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1254,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | $spl3),
  inference(superposition,[],[f571,f72])).
fof(f1260,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) | $spl24),
  inference(cnf_transformation,[],[f1260_D])).
fof(f1260_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1253,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))))) ) | $spl3),
  inference(superposition,[],[f571,f125])).
fof(f1252,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X2)) ) | $spl3),
  inference(superposition,[],[f571,f125])).
fof(f1251,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X1) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X1)) ) | $spl3),
  inference(superposition,[],[f571,f125])).
fof(f1250,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))) | s(bool,'T_0') = s(bool,X0)) ) | $spl3),
  inference(superposition,[],[f571,f125])).
fof(f1154,plain,(
  $false | ($spl3 | $spl8 | $spl11)),
  inference(subsumption_resolution,[],[f1075,f1138])).
fof(f1138,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(trivial_inequality_removal,[],[f1137])).
fof(f1137,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(superposition,[],[f631,f72])).
fof(f1153,plain,(
  $false | ($spl3 | $spl8 | $spl11)),
  inference(subsumption_resolution,[],[f1152,f110])).
fof(f1152,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | ($spl8 | $spl11)),
  inference(subsumption_resolution,[],[f1094,f1138])).
fof(f1094,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl8),
  inference(superposition,[],[f105,f629])).
fof(f1149,plain,(
  $false | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f1138,f1068])).
fof(f1148,plain,(
  $false | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f1147,f1068])).
fof(f1147,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(condensation,[],[f1146])).
fof(f1146,plain,(
  ( ! [X3] : (s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | $spl11),
  inference(subsumption_resolution,[],[f1136,f72])).
fof(f1136,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | s(bool,'T_0') = s(bool,X3) | s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | $spl11),
  inference(superposition,[],[f631,f125])).
fof(f1145,plain,(
  $false | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f1144,f1068])).
fof(f1144,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(condensation,[],[f1143])).
fof(f1143,plain,(
  ( ! [X2] : (s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(subsumption_resolution,[],[f1135,f72])).
fof(f1135,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X2)) ) | $spl11),
  inference(superposition,[],[f631,f125])).
fof(f1142,plain,(
  $false | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f1125,f1068])).
fof(f1125,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f631,f72])).
fof(f1141,plain,(
  $false | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f1127,f1068])).
fof(f1127,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f111,f631,f125])).
fof(f1140,plain,(
  $false | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f1130,f1068])).
fof(f1130,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f111,f631,f125])).
fof(f1139,plain,(
  $false | ($spl11 | $spl17)),
  inference(subsumption_resolution,[],[f1131,f1068])).
fof(f1131,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(unit_resulting_resolution,[],[f71,f631,f161])).
fof(f1066,plain,(
  ( ! [X37] : (s(bool,X37) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X37)) ) | $spl14),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  ( ! [X37] : (s(bool,X37) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f1068,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl17),
  inference(cnf_transformation,[],[f1068_D])).
fof(f1068_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f868,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,X37) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X37)) ) | $spl3),
  inference(superposition,[],[f374,f125])).
fof(f818,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,X37) = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X37)) ) | $spl3),
  inference(superposition,[],[f374,f125])).
fof(f629,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl8),
  inference(cnf_transformation,[],[f629_D])).
fof(f629_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1050,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X40)) ) | $spl12),
  inference(cnf_transformation,[],[f1050_D])).
fof(f1050_D,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X40)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f763,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X40)) ) | $spl3),
  inference(superposition,[],[f374,f125])).
fof(f726,plain,(
  ( ! [X40] : (s(bool,X40) != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,X40) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0))))) ) | $spl3),
  inference(superposition,[],[f374,f125])).
fof(f631,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f631_D])).
fof(f631_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f627,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),bool,sK0)))) | $spl3),
  inference(superposition,[],[f374,f72])).
fof(f624,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl4),
  inference(cnf_transformation,[],[f624_D])).
fof(f624_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f626,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl7),
  inference(cnf_transformation,[],[f626_D])).
fof(f626_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f622,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)),sK2,bool,sK0)))) | $spl3),
  inference(superposition,[],[f373,f72])).
fof(f272,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f271,f71])).
fof(f271,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f270,f119])).
fof(f119,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl1),
  inference(unit_resulting_resolution,[],[f108,f78])).
fof(f270,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl2),
  inference(equality_resolution,[],[f253])).
fof(f253,plain,(
  ( ! [X0] : (s(sK0,sK1) != s(sK0,X0) | p(s(bool,'const.bool.IN_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))) ) | $spl2),
  inference(superposition,[],[f99,f109])).
fof(f107,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl0),
  inference(cnf_transformation,[],[f107_D])).
fof(f107_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f109,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(cnf_transformation,[],[f109_D])).
fof(f109_D,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f67,plain,(
  s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f46])).
fof(f46,plain,(
  (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) & (~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f45])).
fof(f45,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) <~> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))),
  inference(ennf_transformation,[],[f21])).
fof(f21,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) <=> s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) <=> s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),X8))))),
  file('thm.pred_set.ABSORPTION','thm.pred_set.ABSORPTION')).
fof(f108,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl1),
  inference(cnf_transformation,[],[f108_D])).
fof(f108_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f110,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl3),
  inference(cnf_transformation,[],[f110_D])).
fof(f110_D,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f68,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f46])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 96
% Passive clauses: 303
% Generated clauses: 4858
% Final active clauses: 41
% Final passive clauses: 77
% Input formulas: 18
% Initial clauses: 39
% 
% Duplicate literals: 207
% Trivial inequalities: 52
% Fw subsumption resolutions: 291
% Fw demodulations: 1054
% Bw demodulations: 161
% Condensations: 254
% Global subsumptions: 7
% 
% Simple tautologies: 68
% Equational tautologies: 93
% Forward subsumptions: 2445
% 
% Unit resulting resolution: 493
% Factoring: 1
% Forward superposition: 1041
% Backward superposition: 1006
% Self superposition: 38
% Equality factoring: 38
% Equality resolution: 11
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 416
% SAT solver unit clauses: 6
% SAT solver binary clauses: 114
% 
% Sat splits: 20
% Sat splitting refutations: 64
% 
% Memory used [KB]: 3198
% Time elapsed: 0.128 s
% ------------------------------
% ------------------------------
% Success in time 1.444 s
