Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% (55927)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 179
% Passive clauses: 3033
% Generated clauses: 26141
% Final active clauses: 157
% Final passive clauses: 2577
% Input formulas: 30
% Initial clauses: 69
% 
% Duplicate literals: 831
% Trivial inequalities: 687
% Fw subsumption resolutions: 3437
% Fw demodulations: 4533
% Bw demodulations: 20
% 
% Simple tautologies: 540
% Equational tautologies: 40
% Forward subsumptions: 12424
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 428
% Factoring: 2
% Forward superposition: 8396
% Backward superposition: 7358
% Self superposition: 154
% Equality factoring: 73
% Equality resolution: 49
% 
% Split clauses: 329
% Split components: 686
% SAT solver clauses: 774
% SAT solver unit clauses: 20
% SAT solver binary clauses: 143
% 
% Sat splits: 329
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 772
% TWLsolver calls for satisfiability: 180
% 
% Memory used [KB]: 13560
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.arithmetic
% (55957)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 135
% Passive clauses: 774
% Generated clauses: 14297
% Final active clauses: 129
% Final passive clauses: 623
% Input formulas: 30
% Initial clauses: 69
% 
% Duplicate literals: 576
% Trivial inequalities: 201
% Fw subsumption resolutions: 1400
% Fw demodulations: 1792
% Bw demodulations: 15
% Condensations: 990
% Global subsumptions: 5
% 
% Simple tautologies: 266
% Equational tautologies: 30
% Forward subsumptions: 6865
% Bw demodulations to eq. taut.: 5
% 
% Binary resolution: 378
% Factoring: 1
% Forward superposition: 3965
% Backward superposition: 3900
% Self superposition: 108
% Equality factoring: 58
% Equality resolution: 29
% 
% Split clauses: 93
% Split components: 199
% SAT solver clauses: 1406
% SAT solver unit clauses: 74
% SAT solver binary clauses: 33
% SAT solver learnt clauses: 10
% 
% Sat splits: 93
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1386
% TWLsolver calls for satisfiability: 428
% 
% Memory used [KB]: 9978
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.arithmetic
% (55994)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 421
% Passive clauses: 6159
% Generated clauses: 427779
% Final active clauses: 105
% Final passive clauses: 802
% Input formulas: 30
% Initial clauses: 69
% Discarded non-redundant clauses: 1551
% 
% Duplicate literals: 3331
% Trivial inequalities: 891
% Fw subsumption resolutions: 6199
% Fw demodulations: 250700
% Bw demodulations: 3899
% Condensations: 6730
% Global subsumptions: 1527
% 
% Simple tautologies: 975
% Equational tautologies: 537
% Forward subsumptions: 102994
% Fw demodulations to eq. taut.: 44617
% Bw demodulations to eq. taut.: 150
% 
% Unit resulting resolution: 102566
% Factoring: 2
% Forward superposition: 19442
% Backward superposition: 18220
% Self superposition: 417
% Equality factoring: 191
% Equality resolution: 74
% 
% Split clauses: 268
% Split components: 554
% SAT solver clauses: 9765
% SAT solver unit clauses: 272
% SAT solver binary clauses: 1277
% 
% Sat splits: 268
% Sat splitting refutations: 1541
% 
% Memory used [KB]: 115264
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.arithmetic
% (56256)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 697
% Passive clauses: 26484
% Generated clauses: 143920
% Final active clauses: 455
% Final passive clauses: 20702
% Input formulas: 30
% Initial clauses: 67
% 
% Selected by SInE selection: 28
% SInE iterations: 1
% 
% Duplicate literals: 3192
% Trivial inequalities: 2919
% Fw subsumption resolutions: 17549
% Fw demodulations: 27191
% Bw demodulations: 123
% 
% Simple tautologies: 2900
% Equational tautologies: 70
% Forward subsumptions: 62146
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 6342
% Factoring: 4
% Forward superposition: 48949
% Backward superposition: 35335
% Self superposition: 254
% Equality factoring: 166
% Equality resolution: 225
% 
% Split clauses: 3387
% Split components: 7044
% SAT solver clauses: 3486
% SAT solver unit clauses: 16
% SAT solver binary clauses: 714
% 
% Sat splits: 3387
% Sat splitting refutations: 89
% 
% Memory used [KB]: 56928
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.arithmetic
% (56379)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 66
% Passive clauses: 1357
% Generated clauses: 14737
% Final active clauses: 54
% Final passive clauses: 857
% Input formulas: 30
% Initial clauses: 69
% 
% Duplicate literals: 262
% Trivial inequalities: 103
% Fw subsumption resolutions: 773
% Fw demodulations: 4281
% Bw demodulations: 400
% Global subsumptions: 2
% 
% Simple tautologies: 311
% Equational tautologies: 57
% Forward subsumptions: 6485
% Fw demodulations to eq. taut.: 990
% Bw demodulations to eq. taut.: 10
% 
% Unit resulting resolution: 2279
% Forward superposition: 3916
% Backward superposition: 2199
% Self superposition: 161
% Equality factoring: 18
% Equality resolution: 9
% 
% Split clauses: 8
% Split components: 17
% SAT solver clauses: 1485
% SAT solver unit clauses: 131
% SAT solver binary clauses: 91
% 
% Sat splits: 8
% Sat splitting refutations: 116
% 
% Memory used [KB]: 8187
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f5539,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f117,f260,f258,f532,f631,f629,f1120,f1609,f2266,f2267,f2281,f2282,f2767,f3459,f3533,f3531,f3219,f3567,f3565,f5474,f5490,f5508,f5509,f5511,f5512,f5513,f5514,f5515,f5516,f5520,f5528,f5530,f5533,f5535])).
fof(f5535,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1608,f1628,f1680,f2137,f2139,f2141,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183,f2180,f1543,f2146,f2106,f1526,f1524,f2143,f2102,f2100,f2097,f2094,f5527,f1504,f5529,f1577,f1500,f5532,f1498,f5534])).
fof(f5534,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f2164,f4254])).
fof(f1498,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',X11))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',X11))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) ) | $spl6),
  inference(superposition,[],[f163,f631])).
fof(f163,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X1,X0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ((~p(s(bool,sK3(X1,X0))) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | p(s(bool,sK3(X1,X0)))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X1,X0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f108])).
fof(f108,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ? [X2] : ((~p(s(bool,X2)) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | p(s(bool,X2))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,X2),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ? [X2] : (((~p(s(bool,X2)) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | p(s(bool,X2)))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,X2),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ? [X2] : ((p(s(bool,X2)) <=> s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,X2),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))),
  inference(ennf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) => ? [X2] : ((p(s(bool,X2)) <=> s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.bool.COND_3'(s(bool,X2),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0)))))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) => ? [X9] : ((p(s(bool,X9)) <=> s('type.num.num',X6) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',X6)))))) & s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6))) = s('type.num.num','const.bool.COND_3'(s(bool,X9),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',X6)))))))))),
  file('thm.arithmetic.SUC_MOD','thm.arithmetic.arithmetic_395_prove_ltb')).
fof(f5532,plain,(
  ( ! [X0] : (s('type.num.num',sK1) != s('type.num.num',X0)) ) | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f2166,f4256])).
fof(f1500,plain,(
  ( ! [X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))) | s('type.num.num',X13) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))) ) | $spl6),
  inference(superposition,[],[f168,f631])).
fof(f168,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X22,X6] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))) <=> s('type.num.num',X6) = s('type.num.num',X22))),
  file('thm.arithmetic.SUC_MOD','thm.prim_rec.INV_SUC_EQ')).
fof(f1577,plain,(
  s(bool,'T_0') = s(bool,sK3(sK1,sK0)) | $spl6),
  inference(resolution,[],[f1538,f127])).
fof(f127,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.EQ_CLAUSES|split|1')).
fof(f5529,plain,(
  ( ! [X2] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1)))))) ) | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f2168,f4258])).
fof(f1504,plain,(
  ( ! [X17] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))) | s('type.num.num',X17) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))) ) | $spl6),
  inference(superposition,[],[f169,f631])).
fof(f169,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f110])).
fof(f5527,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f2170,f4260])).
fof(f2094,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) | ~p(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f161,f1629])).
fof(f161,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | ~p(s(bool,sK3(X1,X0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f2097,plain,(
  ( ! [X4] : (s('type.num.num',X4) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) | p(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),X4))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f162,f1629])).
fof(f162,plain,(
  ( ! [X0,X1] : (p(s(bool,sK3(X1,X0))) | s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f109])).
fof(f2100,plain,(
  ( ! [X7] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),X7)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f163,f1629])).
fof(f2102,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f168,f1629])).
fof(f2143,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f2142,f1629])).
fof(f2142,plain,(
  ( ! [X0] : (~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))))) ) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f2093,f149])).
fof(f149,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f50])).
fof(f50,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))),
  file('thm.arithmetic.SUC_MOD','thm.num.NOT_SUC')).
fof(f2093,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f161,f1629])).
fof(f1524,plain,(
  ( ! [X88,X87,X85,X89,X86] : (s('type.num.num','const.bool.COND_3'(s(bool,X85),s('type.num.num',X86),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,X87),s('type.num.num',X88),s('type.num.num',X89))) | s('type.num.num',X89) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | s('type.num.num',X86) != s('type.num.num',X88) | s(bool,X85) != s(bool,X87)) ) | $spl6),
  inference(superposition,[],[f160,f631])).
fof(f160,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : (s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6)) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) != s(bool,X2) | (p(s(bool,X2)) & s(X0,X3) != s(X0,X4)) | (~p(s(bool,X2)) & s(X0,X5) != s(X0,X6))) | s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2,X3,X4,X5,X6] : ((s(bool,X1) = s(bool,X2) & (p(s(bool,X2)) => s(X0,X3) = s(X0,X4)) & (~p(s(bool,X2)) => s(X0,X5) = s(X0,X6))) => s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X13] : ! [X19,X20,X14,X17,X15,X18] : ((s(bool,X19) = s(bool,X20) & (p(s(bool,X20)) => s(X13,X14) = s(X13,X17)) & (~p(s(bool,X20)) => s(X13,X15) = s(X13,X18))) => s(X13,'const.bool.COND_3'(s(bool,X19),s(X13,X14),s(X13,X15))) = s(X13,'const.bool.COND_3'(s(bool,X20),s(X13,X17),s(X13,X18))))),
  file('thm.arithmetic.SUC_MOD','thm.bool.COND_CONG')).
fof(f1526,plain,(
  ( ! [X99,X97,X95,X98,X96] : (s('type.num.num','const.bool.COND_3'(s(bool,X95),s('type.num.num',sK0),s('type.num.num',X96))) = s('type.num.num','const.bool.COND_3'(s(bool,X97),s('type.num.num',X98),s('type.num.num',X99))) | s('type.num.num',X96) != s('type.num.num',X99) | s('type.num.num',X98) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | s(bool,X95) != s(bool,X97)) ) | $spl6),
  inference(superposition,[],[f160,f631])).
fof(f2106,plain,(
  ( ! [X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f169,f1629])).
fof(f2146,plain,(
  ( ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X6,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))))))) ) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f2099,f1629])).
fof(f2099,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X6,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f163,f1629])).
fof(f1543,plain,(
  ( ! [X4] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',sK0))))) | ~p(s(bool,sK3(X4,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1542,f119])).
fof(f119,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.arithmetic.SUC_MOD',aHL_TRUTH)).
fof(f1542,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',sK0))))) | ~p(s(bool,sK3(X4,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(forward_demodulation,[],[f1541,f205])).
fof(f1541,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',sK0))))) | ~p(s(bool,sK3(X4,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(forward_demodulation,[],[f1491,f631])).
fof(f1491,plain,(
  ( ! [X4] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',sK0))))) | ~p(s(bool,sK3(X4,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))))) ) | $spl6),
  inference(superposition,[],[f161,f631])).
fof(f2180,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) | ~p(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl6),
  inference(superposition,[],[f161,f1608])).
fof(f2183,plain,(
  ( ! [X4] : (s('type.num.num',X4) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) | p(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)),X4))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) ) | $spl6),
  inference(superposition,[],[f162,f1608])).
fof(f2229,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | ~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) ) | $spl6),
  inference(forward_demodulation,[],[f2228,f1608])).
fof(f2228,plain,(
  ( ! [X0] : (~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f2179,f149])).
fof(f2179,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) ) | $spl6),
  inference(superposition,[],[f161,f1608])).
fof(f1549,plain,(
  ( ! [X7] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK0))))) | p(s(bool,sK3(X7,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1548,f119])).
fof(f1548,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK0))))) | p(s(bool,sK3(X7,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(forward_demodulation,[],[f1547,f205])).
fof(f1547,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK0))))) | p(s(bool,sK3(X7,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(forward_demodulation,[],[f1494,f631])).
fof(f1494,plain,(
  ( ! [X7] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK0))))) | p(s(bool,sK3(X7,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))))) ) | $spl6),
  inference(superposition,[],[f162,f631])).
fof(f2186,plain,(
  ( ! [X7] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)),X7)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) ) | $spl6),
  inference(superposition,[],[f163,f1608])).
fof(f2188,plain,(
  ( ! [X9] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))) ) | $spl6),
  inference(superposition,[],[f168,f1608])).
fof(f1555,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f1554,f119])).
fof(f1554,plain,(
  ( ! [X10] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(forward_demodulation,[],[f1553,f205])).
fof(f1553,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK0)))))))) ) | $spl6),
  inference(forward_demodulation,[],[f1497,f631])).
fof(f1497,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))))) ) | $spl6),
  inference(superposition,[],[f163,f631])).
fof(f1471,plain,(
  ( ! [X0] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl6),
  inference(superposition,[],[f631,f168])).
fof(f2232,plain,(
  ( ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X6,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))))))) ) | $spl6),
  inference(forward_demodulation,[],[f2185,f1608])).
fof(f2185,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X6,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) ) | $spl6),
  inference(superposition,[],[f163,f1608])).
fof(f1473,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',X0))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl6),
  inference(superposition,[],[f631,f168])).
fof(f2212,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s('type.num.num','const.bool.COND_3'(s(bool,X81),s('type.num.num',X82),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,X83),s('type.num.num',X84),s('type.num.num',X85))) | s('type.num.num',X85) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num',X82) != s('type.num.num',X84) | s(bool,X81) != s(bool,X83)) ) | $spl6),
  inference(superposition,[],[f160,f1608])).
fof(f2078,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1629,f168])).
fof(f1477,plain,(
  ( ! [X0] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))) ) | $spl6),
  inference(superposition,[],[f631,f169])).
fof(f5488,plain,(
  ( ! [X169,X168] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X168),s('type.num.num',X169))) | s('type.num.num',X168) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))) ) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f1629,f5193])).
fof(f1629,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f258,f1608])).
fof(f1483,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))) ) | $spl6),
  inference(superposition,[],[f168,f631])).
fof(f2080,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num',sK2) != s('type.num.num',X0)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1629,f169])).
fof(f2192,plain,(
  ( ! [X13] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))) | s('type.num.num',X13) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))) ) | $spl6),
  inference(superposition,[],[f169,f1608])).
fof(f2082,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2)))))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1629,f168])).
fof(f2126,plain,(
  ( ! [X85,X83,X81,X84,X82] : (s('type.num.num','const.bool.COND_3'(s(bool,X81),s('type.num.num',X82),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,X83),s('type.num.num',X84),s('type.num.num',X85))) | s('type.num.num',X85) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | s('type.num.num',X82) != s('type.num.num',X84) | s(bool,X81) != s(bool,X83)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f160,f1629])).
fof(f2084,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f1629,f168])).
fof(f2214,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s('type.num.num','const.bool.COND_3'(s(bool,X91),s('type.num.num','const.num.|48|_0'),s('type.num.num',X92))) = s('type.num.num','const.bool.COND_3'(s(bool,X93),s('type.num.num',X94),s('type.num.num',X95))) | s('type.num.num',X92) != s('type.num.num',X95) | s('type.num.num',X94) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s(bool,X91) != s(bool,X93)) ) | $spl6),
  inference(superposition,[],[f160,f1608])).
fof(f1492,plain,(
  ( ! [X5] : (s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',X5))))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) ) | $spl6),
  inference(superposition,[],[f161,f631])).
fof(f5487,plain,(
  p(s(bool,sK3(sK1,sK0))) | $spl6),
  inference(condensation,[],[f1991])).
fof(f1991,plain,(
  ( ! [X1] : (p(s(bool,sK3(sK1,sK0))) | p(s(bool,X1))) ) | $spl6),
  inference(subsumption_resolution,[],[f1884,f126])).
fof(f126,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f91])).
fof(f1884,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,X1) | p(s(bool,sK3(sK1,sK0))) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f1577,f164])).
fof(f164,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.IMP_ANTISYM_AX')).
fof(f2128,plain,(
  ( ! [X94,X92,X95,X93,X91] : (s('type.num.num','const.bool.COND_3'(s(bool,X91),s('type.num.num','const.num.|48|_0'),s('type.num.num',X92))) = s('type.num.num','const.bool.COND_3'(s(bool,X93),s('type.num.num',X94),s('type.num.num',X95))) | s('type.num.num',X92) != s('type.num.num',X95) | s('type.num.num',X94) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | s(bool,X91) != s(bool,X93)) ) | ($spl0 | $spl6)),
  inference(superposition,[],[f160,f1629])).
fof(f1495,plain,(
  ( ! [X8] : (s('type.num.num',X8) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',X8))))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))),X8))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8))))) ) | $spl6),
  inference(superposition,[],[f162,f631])).
fof(f3121,plain,(
  ( ! [X1] : (s('type.num.num',sK0) != s('type.num.num',X1) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2879,f168])).
fof(f5429,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X78),s(bool,X79))) | s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X80),s(bool,X81),s(bool,X82))) | s(bool,X79) != s(bool,X82) | s(bool,X78) != s(bool,X81) | s(bool,X77) != s(bool,X80)) )),
  inference(forward_demodulation,[],[f5396,f122])).
fof(f122,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))) )),
  inference(cnf_transformation,[],[f34])).
fof(f34,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X21] : s(bool,'T_0') = s(bool,'const.bool.BOUNDED_1'(s(bool,X21)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.BOUNDED_THM')).
fof(f5396,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (s(bool,'T_0') != s(bool,'const.bool.COND_3'(s(bool,X80),s(bool,X81),s(bool,X82))) | s(bool,'const.bool.BOUNDED_1'(s(bool,X83))) = s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X78),s(bool,X79))) | s(bool,X79) != s(bool,X82) | s(bool,X78) != s(bool,X81) | s(bool,X77) != s(bool,X80)) )),
  inference(superposition,[],[f5355,f160])).
fof(f5424,plain,(
  ( ! [X47,X46] : (s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X46),s(bool,X47))) | s(bool,'T_0') != s(bool,X47)) )),
  inference(forward_demodulation,[],[f5391,f122])).
fof(f5391,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') != s(bool,X47) | s(bool,'const.bool.BOUNDED_1'(s(bool,X48))) = s(bool,'const.bool.COND_3'(s(bool,'F_0'),s(bool,X46),s(bool,X47)))) )),
  inference(superposition,[],[f5355,f154])).
fof(f154,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f58])).
fof(f58,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X13] : ! [X10,X11] : s(X13,X11) = s(X13,'const.bool.COND_3'(s(bool,'F_0'),s(X13,X10),s(X13,X11)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.COND_CLAUSES|split|1')).
fof(f5406,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,'const.bool.BOUNDED_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(trivial_inequality_removal,[],[f5368])).
fof(f5368,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X0) = s(bool,'const.bool.BOUNDED_1'(s(bool,X1))) | s(bool,'F_0') = s(bool,X0)) )),
  inference(superposition,[],[f5355,f124])).
fof(f124,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.arithmetic.SUC_MOD',aHL_BOOL_CASES)).
fof(f5360,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,X7) | s(bool,X7) = s(bool,'const.bool.BOUNDED_1'(s(bool,X8))) | s(bool,'F_0') = s(bool,X6)) )),
  inference(superposition,[],[f5355,f124])).
fof(f5346,plain,(
  ( ! [X80,X78,X83,X81,X79,X77,X82] : (~p(s(bool,'const.bool.COND_3'(s(bool,X80),s(bool,X81),s(bool,X82)))) | s(bool,'T_0') != s(bool,X83) | s(bool,X83) = s(bool,'const.bool.COND_3'(s(bool,X77),s(bool,X78),s(bool,X79))) | s(bool,X79) != s(bool,X82) | s(bool,X78) != s(bool,X81) | s(bool,X77) != s(bool,X80)) )),
  inference(superposition,[],[f5242,f160])).
fof(f5345,plain,(
  ( ! [X70,X76,X74,X72,X71,X75,X73] : (~p(s(bool,'const.bool.COND_3'(s(bool,X73),s(bool,X74),s(bool,X75)))) | s(bool,'T_0') != s(bool,X76) | s(bool,X76) = s(bool,'const.bool.COND_3'(s(bool,X70),s(bool,X71),s(bool,X72))) | s(bool,X72) != s(bool,X75) | p(s(bool,X73)) | s(bool,X70) != s(bool,X73)) )),
  inference(superposition,[],[f5242,f159])).
fof(f159,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | s(X0,X5) != s(X0,X6) | p(s(bool,X2)) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f79])).
fof(f5344,plain,(
  ( ! [X68,X66,X64,X69,X67,X65,X63] : (~p(s(bool,'const.bool.COND_3'(s(bool,X66),s(bool,X67),s(bool,X68)))) | s(bool,'T_0') != s(bool,X69) | s(bool,X69) = s(bool,'const.bool.COND_3'(s(bool,X63),s(bool,X64),s(bool,X65))) | s(bool,X65) != s(bool,X68) | p(s(bool,X63)) | s(bool,X63) != s(bool,X66)) )),
  inference(superposition,[],[f5242,f159])).
fof(f5343,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X56] : (~p(s(bool,'const.bool.COND_3'(s(bool,X59),s(bool,X60),s(bool,X61)))) | s(bool,'T_0') != s(bool,X62) | s(bool,X62) = s(bool,'const.bool.COND_3'(s(bool,X56),s(bool,X57),s(bool,X58))) | ~p(s(bool,X59)) | s(bool,X57) != s(bool,X60) | s(bool,X56) != s(bool,X59)) )),
  inference(superposition,[],[f5242,f158])).
fof(f158,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X3),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X2),s(X0,X4),s(X0,X6))) | ~p(s(bool,X2)) | s(X0,X3) != s(X0,X4) | s(bool,X1) != s(bool,X2)) )),
  inference(cnf_transformation,[],[f79])).
fof(f5342,plain,(
  ( ! [X54,X52,X50,X55,X53,X51,X49] : (~p(s(bool,'const.bool.COND_3'(s(bool,X52),s(bool,X53),s(bool,X54)))) | s(bool,'T_0') != s(bool,X55) | s(bool,X55) = s(bool,'const.bool.COND_3'(s(bool,X49),s(bool,X50),s(bool,X51))) | ~p(s(bool,X49)) | s(bool,X50) != s(bool,X53) | s(bool,X49) != s(bool,X52)) )),
  inference(superposition,[],[f5242,f158])).
fof(f5355,plain,(
  ( ! [X37,X38] : (s(bool,'T_0') != s(bool,X38) | s(bool,X38) = s(bool,'const.bool.BOUNDED_1'(s(bool,X37)))) )),
  inference(subsumption_resolution,[],[f5335,f119])).
fof(f5335,plain,(
  ( ! [X37,X38] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X38) | s(bool,X38) = s(bool,'const.bool.BOUNDED_1'(s(bool,X37)))) )),
  inference(superposition,[],[f5242,f122])).
fof(f5354,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,X36) | s(bool,X36) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(subsumption_resolution,[],[f5334,f119])).
fof(f5334,plain,(
  ( ! [X36] : (~p(s(bool,'T_0')) | s(bool,'T_0') != s(bool,X36) | s(bool,X36) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f5242,f205])).
fof(f5317,plain,(
  ( ! [X59,X60] : (s(bool,'T_0') != s(bool,X59) | s(bool,X59) = s(bool,sK3(X60,sK0)) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X60),s('type.num.num',sK0)))))) )),
  inference(resolution,[],[f5242,f203])).
fof(f5316,plain,(
  ( ! [X57,X58,X56] : (s(bool,'T_0') != s(bool,X56) | s(bool,X56) = s(bool,sK3(X57,X58)) | s('type.num.num',X58) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X57),s('type.num.num',X58))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X58))))) )),
  inference(resolution,[],[f5242,f162])).
fof(f5300,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') != s(bool,X26) | s(bool,X26) = s(bool,X27) | s(bool,X27) != s(bool,X28) | ~p(s(bool,X28))) )),
  inference(resolution,[],[f5242,f191])).
fof(f191,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(condensation,[],[f190])).
fof(f190,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f184])).
fof(f184,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X14,X17,X15,X18] : ((s(bool,X14) = s(bool,X17) & (p(s(bool,X17)) => s(bool,X15) = s(bool,X18))) => ((p(s(bool,X14)) => p(s(bool,X15))) <=> (p(s(bool,X17)) => p(s(bool,X18)))))),
  file('thm.arithmetic.SUC_MOD','thm.bool.IMP_CONG')).
fof(f5298,plain,(
  ( ! [X21,X22,X20] : (s(bool,'T_0') != s(bool,X20) | s(bool,X20) = s(bool,X21) | s(bool,X21) = s(bool,X22) | p(s(bool,X22))) )),
  inference(resolution,[],[f5242,f164])).
fof(f5297,plain,(
  ( ! [X14,X12,X19,X17,X15,X13,X18,X16] : (s(bool,'T_0') != s(bool,X12) | s(bool,X12) = s(bool,X13) | s(X14,'const.bool.COND_3'(s(bool,X13),s(X14,X18),s(X14,X19))) = s(X14,'const.bool.COND_3'(s(bool,X15),s(X14,X16),s(X14,X17))) | s(X14,X17) != s(X14,X19) | s(bool,X13) != s(bool,X15)) )),
  inference(resolution,[],[f5242,f159])).
fof(f5294,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) )),
  inference(resolution,[],[f5242,f131])).
fof(f131,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.EQ_CLAUSES|split|2')).
fof(f5292,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') != s(bool,X3)) )),
  inference(resolution,[],[f5242,f126])).
fof(f5233,plain,(
  ( ! [X163,X161,X162,X160] : (s('type.num.num',X162) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X161))) | s('type.num.num',X162) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X160))),s('type.num.num',X163))) | s('type.num.num',X160) != s('type.num.num',X161)) )),
  inference(superposition,[],[f5098,f169])).
fof(f5230,plain,(
  ( ! [X154,X155,X153] : (s('type.num.num',sK0) != s('type.num.num',X154) | s('type.num.num',X154) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X153),s('type.num.num',sK0))))),s('type.num.num',X155))) | s(bool,'F_0') = s(bool,sK3(X153,sK0))) )),
  inference(superposition,[],[f5098,f571])).
fof(f5227,plain,(
  ( ! [X144,X142,X143,X141] : (s('type.num.num',X142) != s('type.num.num',X143) | s('type.num.num',X143) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X141),s('type.num.num',X144))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X141))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X142)))) )),
  inference(superposition,[],[f5098,f168])).
fof(f5226,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (s(fun(X135,X136),X138) != s(fun(X135,X136),X139) | s(fun(X135,X136),X139) = s(fun(X135,X136),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X135,X136),X137),s(fun(X135,X136),X140))) | s(X136,f(s(X135,sK4(X138,X137,X136,X135)))) != s(X136,g(s(X135,sK4(X138,X137,X136,X135))))) )),
  inference(superposition,[],[f5098,f170])).
fof(f170,plain,(
  ( ! [X2,X0,X3,X1] : (s(fun(X0,X1),X2) = s(fun(X0,X1),X3) | s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f83])).
fof(f83,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.arithmetic.SUC_MOD',aHL_EXT)).
fof(f5225,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (s(fun(X129,X130),X132) != s(fun(X129,X130),X133) | s(fun(X129,X130),X133) = s(fun(X129,X130),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X129,X130),X131),s(fun(X129,X130),X134))) | s(X130,f(s(X129,sK4(X131,X132,X130,X129)))) != s(X130,g(s(X129,sK4(X131,X132,X130,X129))))) )),
  inference(superposition,[],[f5098,f170])).
fof(f5218,plain,(
  ( ! [X114,X115,X113] : (s(bool,X113) != s(bool,X114) | s(bool,X114) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'F_0'),s(bool,X115))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5215,plain,(
  ( ! [X105,X106,X104] : (s(bool,X104) != s(bool,X105) | s(bool,X105) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,'T_0'),s(bool,X106))) | s(bool,'F_0') = s(bool,X104)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5204,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X68) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X67),s(bool,X69))) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5203,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,X65) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X64),s(bool,X66))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5201,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X53,X51,X49] : (s(X46,X53) != s(X46,'const.bool.COND_3'(s(bool,X50),s(X46,X51),s(X46,X52))) | s(X46,X53) = s(X46,'const.bool.COND_3'(s(bool,'T_0'),s(X46,'const.bool.COND_3'(s(bool,X47),s(X46,X48),s(X46,X49))),s(X46,X54))) | s(X46,X49) != s(X46,X52) | s(X46,X48) != s(X46,X51) | s(bool,X47) != s(bool,X50)) )),
  inference(superposition,[],[f5098,f160])).
fof(f5261,plain,(
  ( ! [X171,X172,X170] : (s('type.num.num',X172) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X170,X171)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X170),s('type.num.num',X171))))))) | s('type.num.num',X172) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X170))),s('type.num.num',X171))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X171))))) )),
  inference(forward_demodulation,[],[f5194,f153])).
fof(f153,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X1) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X13] : ! [X10,X11] : s(X13,X10) = s(X13,'const.bool.COND_3'(s(bool,'T_0'),s(X13,X10),s(X13,X11)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.COND_CLAUSES|split|0')).
fof(f5194,plain,(
  ( ! [X173,X171,X172,X170] : (s('type.num.num',X172) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X170))),s('type.num.num',X171))) | s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X172),s('type.num.num',X173))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X170,X171)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X170),s('type.num.num',X171))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X171))))) )),
  inference(superposition,[],[f5098,f163])).
fof(f5193,plain,(
  ( ! [X169,X168] : (s('type.num.num',X168) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X168),s('type.num.num',X169)))) ) | $spl0),
  inference(superposition,[],[f5098,f258])).
fof(f5191,plain,(
  ( ! [X163,X161,X162,X160] : (s('type.num.num',X162) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X161))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X160))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X162),s('type.num.num',X163))) | s('type.num.num',X160) != s('type.num.num',X161)) )),
  inference(superposition,[],[f5098,f169])).
fof(f5189,plain,(
  ( ! [X156,X157] : (s('type.num.num',X156) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X156),s('type.num.num',X157)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f5098,f2873])).
fof(f5188,plain,(
  ( ! [X154,X155,X153] : (s('type.num.num',sK0) != s('type.num.num',X154) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X153),s('type.num.num',sK0))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X154),s('type.num.num',X155))) | s(bool,'F_0') = s(bool,sK3(X153,sK0))) )),
  inference(superposition,[],[f5098,f571])).
fof(f5260,plain,(
  ( ! [X151,X149,X150] : (s('type.num.num',X151) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X149),s('type.num.num',X150))))) | s('type.num.num',X150) != s('type.num.num',X151) | ~p(s(bool,sK3(X149,X150))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X150))))) )),
  inference(forward_demodulation,[],[f5187,f153])).
fof(f5187,plain,(
  ( ! [X152,X151,X149,X150] : (s('type.num.num',X150) != s('type.num.num',X151) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X149),s('type.num.num',X150))))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X151),s('type.num.num',X152))) | ~p(s(bool,sK3(X149,X150))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X150))))) )),
  inference(superposition,[],[f5098,f161])).
fof(f5185,plain,(
  ( ! [X144,X142,X143,X141] : (s('type.num.num',X142) != s('type.num.num',X143) | s('type.num.num',X141) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X143),s('type.num.num',X144))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X141))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X142)))) )),
  inference(superposition,[],[f5098,f168])).
fof(f5184,plain,(
  ( ! [X140,X138,X136,X135,X139,X137] : (s(fun(X135,X136),X138) != s(fun(X135,X136),X139) | s(fun(X135,X136),X137) = s(fun(X135,X136),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X135,X136),X139),s(fun(X135,X136),X140))) | s(X136,f(s(X135,sK4(X138,X137,X136,X135)))) != s(X136,g(s(X135,sK4(X138,X137,X136,X135))))) )),
  inference(superposition,[],[f5098,f170])).
fof(f5183,plain,(
  ( ! [X134,X132,X130,X133,X131,X129] : (s(fun(X129,X130),X132) != s(fun(X129,X130),X133) | s(fun(X129,X130),X131) = s(fun(X129,X130),'const.bool.COND_3'(s(bool,'T_0'),s(fun(X129,X130),X133),s(fun(X129,X130),X134))) | s(X130,f(s(X129,sK4(X131,X132,X130,X129)))) != s(X130,g(s(X129,sK4(X131,X132,X130,X129))))) )),
  inference(superposition,[],[f5098,f170])).
fof(f5176,plain,(
  ( ! [X114,X115,X113] : (s(bool,X113) != s(bool,X114) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X114),s(bool,X115))) | s(bool,'T_0') = s(bool,X113)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5252,plain,(
  ( ! [X107,X108] : (s(bool,'F_0') = s(bool,X108) | s(bool,X107) != s(bool,X108) | p(s(bool,X107))) )),
  inference(forward_demodulation,[],[f5174,f153])).
fof(f5174,plain,(
  ( ! [X109,X107,X108] : (s(bool,X107) != s(bool,X108) | s(bool,'F_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X108),s(bool,X109))) | p(s(bool,X107))) )),
  inference(superposition,[],[f5098,f133])).
fof(f133,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.EQ_CLAUSES|split|3')).
fof(f5173,plain,(
  ( ! [X105,X106,X104] : (s(bool,X104) != s(bool,X105) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X105),s(bool,X106))) | s(bool,'F_0') = s(bool,X104)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5250,plain,(
  ( ! [X99,X98] : (s(bool,'T_0') = s(bool,X99) | s(bool,X98) != s(bool,X99) | ~p(s(bool,X98))) )),
  inference(forward_demodulation,[],[f5171,f153])).
fof(f5171,plain,(
  ( ! [X99,X100,X98] : (s(bool,X98) != s(bool,X99) | s(bool,'T_0') = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X99),s(bool,X100))) | ~p(s(bool,X98))) )),
  inference(superposition,[],[f5098,f129])).
fof(f129,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.arithmetic.SUC_MOD','thm.bool.EQ_CLAUSES|split|0')).
fof(f5248,plain,(
  ( ! [X92,X90,X91] : (s(bool,X90) = s(bool,X92) | s(bool,X91) != s(bool,X92) | ~p(s(bool,X90)) | ~p(s(bool,X91))) )),
  inference(forward_demodulation,[],[f5169,f153])).
fof(f5169,plain,(
  ( ! [X92,X90,X93,X91] : (s(bool,X91) != s(bool,X92) | s(bool,X90) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X92),s(bool,X93))) | ~p(s(bool,X90)) | ~p(s(bool,X91))) )),
  inference(superposition,[],[f5098,f167])).
fof(f167,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f82])).
fof(f5246,plain,(
  ( ! [X83,X84,X82] : (s(bool,X82) = s(bool,X84) | s(bool,X83) != s(bool,X84) | p(s(bool,X83)) | p(s(bool,X82))) )),
  inference(forward_demodulation,[],[f5167,f153])).
fof(f5167,plain,(
  ( ! [X85,X83,X84,X82] : (s(bool,X83) != s(bool,X84) | s(bool,X82) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X84),s(bool,X85))) | p(s(bool,X83)) | p(s(bool,X82))) )),
  inference(superposition,[],[f5098,f164])).
fof(f5244,plain,(
  ( ! [X76,X77] : (s(bool,X76) = s(bool,X77) | s(bool,'F_0') != s(bool,X77) | p(s(bool,X76))) )),
  inference(forward_demodulation,[],[f5165,f153])).
fof(f5165,plain,(
  ( ! [X78,X76,X77] : (s(bool,'F_0') != s(bool,X77) | s(bool,X76) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X77),s(bool,X78))) | p(s(bool,X76))) )),
  inference(superposition,[],[f5098,f131])).
fof(f5242,plain,(
  ( ! [X70,X71] : (~p(s(bool,X70)) | s(bool,'T_0') != s(bool,X71) | s(bool,X70) = s(bool,X71)) )),
  inference(forward_demodulation,[],[f5163,f153])).
fof(f5163,plain,(
  ( ! [X70,X72,X71] : (s(bool,'T_0') != s(bool,X71) | s(bool,X70) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X71),s(bool,X72))) | ~p(s(bool,X70))) )),
  inference(superposition,[],[f5098,f127])).
fof(f5162,plain,(
  ( ! [X68,X69,X67] : (s(bool,'F_0') != s(bool,X68) | s(bool,X67) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X68),s(bool,X69))) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5161,plain,(
  ( ! [X66,X64,X65] : (s(bool,'T_0') != s(bool,X65) | s(bool,X64) = s(bool,'const.bool.COND_3'(s(bool,'T_0'),s(bool,X65),s(bool,X66))) | s(bool,'F_0') = s(bool,X64)) )),
  inference(superposition,[],[f5098,f124])).
fof(f5159,plain,(
  ( ! [X47,X54,X52,X50,X48,X46,X53,X51,X49] : (s(X46,X53) != s(X46,'const.bool.COND_3'(s(bool,X50),s(X46,X51),s(X46,X52))) | s(X46,'const.bool.COND_3'(s(bool,'T_0'),s(X46,X53),s(X46,X54))) = s(X46,'const.bool.COND_3'(s(bool,X47),s(X46,X48),s(X46,X49))) | s(X46,X49) != s(X46,X52) | s(X46,X48) != s(X46,X51) | s(bool,X47) != s(bool,X50)) )),
  inference(superposition,[],[f5098,f160])).
fof(f5241,plain,(
  ( ! [X39,X37,X43,X41,X38,X44,X42,X40] : (s(X37,X44) = s(X37,'const.bool.COND_3'(s(bool,X38),s(X37,X39),s(X37,X40))) | s(X37,X44) != s(X37,'const.bool.COND_3'(s(bool,X41),s(X37,X42),s(X37,X43))) | s(X37,X40) != s(X37,X43) | p(s(bool,X41)) | s(bool,X38) != s(bool,X41)) )),
  inference(forward_demodulation,[],[f5158,f153])).
fof(f5158,plain,(
  ( ! [X39,X37,X45,X43,X41,X38,X44,X42,X40] : (s(X37,X44) != s(X37,'const.bool.COND_3'(s(bool,X41),s(X37,X42),s(X37,X43))) | s(X37,'const.bool.COND_3'(s(bool,'T_0'),s(X37,X44),s(X37,X45))) = s(X37,'const.bool.COND_3'(s(bool,X38),s(X37,X39),s(X37,X40))) | s(X37,X40) != s(X37,X43) | p(s(bool,X41)) | s(bool,X38) != s(bool,X41)) )),
  inference(superposition,[],[f5098,f159])).
fof(f5240,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X34,X32] : (s(X28,X35) = s(X28,'const.bool.COND_3'(s(bool,X29),s(X28,X30),s(X28,X31))) | s(X28,X35) != s(X28,'const.bool.COND_3'(s(bool,X32),s(X28,X33),s(X28,X34))) | s(X28,X31) != s(X28,X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) )),
  inference(forward_demodulation,[],[f5157,f153])).
fof(f5157,plain,(
  ( ! [X30,X28,X35,X33,X31,X29,X36,X34,X32] : (s(X28,X35) != s(X28,'const.bool.COND_3'(s(bool,X32),s(X28,X33),s(X28,X34))) | s(X28,'const.bool.COND_3'(s(bool,'T_0'),s(X28,X35),s(X28,X36))) = s(X28,'const.bool.COND_3'(s(bool,X29),s(X28,X30),s(X28,X31))) | s(X28,X31) != s(X28,X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) )),
  inference(superposition,[],[f5098,f159])).
fof(f5239,plain,(
  ( ! [X26,X24,X23,X21,X19,X25,X22,X20] : (s(X19,X26) = s(X19,'const.bool.COND_3'(s(bool,X20),s(X19,X21),s(X19,X22))) | s(X19,X26) != s(X19,'const.bool.COND_3'(s(bool,X23),s(X19,X24),s(X19,X25))) | ~p(s(bool,X23)) | s(X19,X21) != s(X19,X24) | s(bool,X20) != s(bool,X23)) )),
  inference(forward_demodulation,[],[f5156,f153])).
fof(f5156,plain,(
  ( ! [X26,X24,X23,X21,X19,X27,X25,X22,X20] : (s(X19,X26) != s(X19,'const.bool.COND_3'(s(bool,X23),s(X19,X24),s(X19,X25))) | s(X19,'const.bool.COND_3'(s(bool,'T_0'),s(X19,X26),s(X19,X27))) = s(X19,'const.bool.COND_3'(s(bool,X20),s(X19,X21),s(X19,X22))) | ~p(s(bool,X23)) | s(X19,X21) != s(X19,X24) | s(bool,X20) != s(bool,X23)) )),
  inference(superposition,[],[f5098,f158])).
fof(f5238,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X16] : (s(X10,X17) = s(X10,'const.bool.COND_3'(s(bool,X11),s(X10,X12),s(X10,X13))) | s(X10,X17) != s(X10,'const.bool.COND_3'(s(bool,X14),s(X10,X15),s(X10,X16))) | ~p(s(bool,X11)) | s(X10,X12) != s(X10,X15) | s(bool,X11) != s(bool,X14)) )),
  inference(forward_demodulation,[],[f5155,f153])).
fof(f5155,plain,(
  ( ! [X14,X12,X10,X17,X15,X13,X11,X18,X16] : (s(X10,X17) != s(X10,'const.bool.COND_3'(s(bool,X14),s(X10,X15),s(X10,X16))) | s(X10,'const.bool.COND_3'(s(bool,'T_0'),s(X10,X17),s(X10,X18))) = s(X10,'const.bool.COND_3'(s(bool,X11),s(X10,X12),s(X10,X13))) | ~p(s(bool,X11)) | s(X10,X12) != s(X10,X15) | s(bool,X11) != s(bool,X14)) )),
  inference(superposition,[],[f5098,f158])).
fof(f5154,plain,(
  ( ! [X6,X8,X7,X5,X9] : (s(X5,X7) != s(X5,X8) | s(X5,'const.bool.COND_3'(s(bool,'T_0'),s(X5,X8),s(X5,X9))) = s(X5,'const.bool.COND_3'(s(bool,'F_0'),s(X5,X6),s(X5,X7)))) )),
  inference(superposition,[],[f5098,f154])).
fof(f5153,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(X0,X1) != s(X0,X3) | s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X1),s(X0,X2))) = s(X0,'const.bool.COND_3'(s(bool,'T_0'),s(X0,X3),s(X0,X4)))) )),
  inference(superposition,[],[f5098,f153])).
fof(f5138,plain,(
  ( ! [X167,X165,X163,X161,X169,X166,X164,X162,X170,X168] : (~p(s(bool,'const.bool.COND_3'(s(bool,X164),s(bool,X165),s(bool,X166)))) | s(X167,X168) = s(X167,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X161),s(bool,X162),s(bool,X163))),s(X167,X169),s(X167,X170))) | s(X167,X168) != s(X167,X169) | s(bool,X163) != s(bool,X166) | s(bool,X162) != s(bool,X165) | s(bool,X161) != s(bool,X164)) )),
  inference(superposition,[],[f412,f160])).
fof(f5137,plain,(
  ( ! [X158,X156,X154,X152,X151,X159,X157,X155,X153,X160] : (~p(s(bool,'const.bool.COND_3'(s(bool,X154),s(bool,X155),s(bool,X156)))) | s(X157,X158) = s(X157,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X151),s(bool,X152),s(bool,X153))),s(X157,X159),s(X157,X160))) | s(X157,X158) != s(X157,X159) | s(bool,X153) != s(bool,X156) | p(s(bool,X154)) | s(bool,X151) != s(bool,X154)) )),
  inference(superposition,[],[f412,f159])).
fof(f5136,plain,(
  ( ! [X146,X144,X142,X149,X147,X145,X143,X141,X150,X148] : (~p(s(bool,'const.bool.COND_3'(s(bool,X144),s(bool,X145),s(bool,X146)))) | s(X147,X148) = s(X147,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X141),s(bool,X142),s(bool,X143))),s(X147,X149),s(X147,X150))) | s(X147,X148) != s(X147,X149) | s(bool,X143) != s(bool,X146) | p(s(bool,X141)) | s(bool,X141) != s(bool,X144)) )),
  inference(superposition,[],[f412,f159])).
fof(f5135,plain,(
  ( ! [X134,X132,X140,X138,X136,X135,X133,X131,X139,X137] : (~p(s(bool,'const.bool.COND_3'(s(bool,X134),s(bool,X135),s(bool,X136)))) | s(X137,X138) = s(X137,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X131),s(bool,X132),s(bool,X133))),s(X137,X139),s(X137,X140))) | s(X137,X138) != s(X137,X139) | ~p(s(bool,X134)) | s(bool,X132) != s(bool,X135) | s(bool,X131) != s(bool,X134)) )),
  inference(superposition,[],[f412,f158])).
fof(f5134,plain,(
  ( ! [X127,X125,X123,X121,X130,X128,X126,X124,X122,X129] : (~p(s(bool,'const.bool.COND_3'(s(bool,X124),s(bool,X125),s(bool,X126)))) | s(X127,X128) = s(X127,'const.bool.COND_3'(s(bool,'const.bool.COND_3'(s(bool,X121),s(bool,X122),s(bool,X123))),s(X127,X129),s(X127,X130))) | s(X127,X128) != s(X127,X129) | ~p(s(bool,X121)) | s(bool,X122) != s(bool,X125) | s(bool,X121) != s(bool,X124)) )),
  inference(superposition,[],[f412,f158])).
fof(f5109,plain,(
  ( ! [X140,X138,X141,X139,X137] : (s(X137,X138) = s(X137,'const.bool.COND_3'(s(bool,sK3(X139,sK0)),s(X137,X140),s(X137,X141))) | s(X137,X138) != s(X137,X140) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X139),s('type.num.num',sK0)))))) )),
  inference(resolution,[],[f412,f203])).
fof(f5108,plain,(
  ( ! [X134,X132,X136,X135,X133,X131] : (s(X131,X132) = s(X131,'const.bool.COND_3'(s(bool,sK3(X133,X134)),s(X131,X135),s(X131,X136))) | s(X131,X132) != s(X131,X135) | s('type.num.num',X134) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X133),s('type.num.num',X134))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X134))))) )),
  inference(resolution,[],[f412,f162])).
fof(f5098,plain,(
  ( ! [X92,X90,X91,X89] : (s(X89,X90) != s(X89,X91) | s(X89,X90) = s(X89,'const.bool.COND_3'(s(bool,'T_0'),s(X89,X91),s(X89,X92)))) )),
  inference(resolution,[],[f412,f119])).
fof(f5092,plain,(
  ( ! [X57,X54,X58,X56,X55,X53] : (s(X53,X54) = s(X53,'const.bool.COND_3'(s(bool,X55),s(X53,X56),s(X53,X57))) | s(X53,X54) != s(X53,X56) | s(bool,X55) != s(bool,X58) | ~p(s(bool,X58))) )),
  inference(resolution,[],[f412,f191])).
fof(f5090,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s(X41,X42) = s(X41,'const.bool.COND_3'(s(bool,X43),s(X41,X44),s(X41,X45))) | s(X41,X42) != s(X41,X44) | s(bool,X43) = s(bool,X46) | p(s(bool,X46))) )),
  inference(resolution,[],[f412,f164])).
fof(f5089,plain,(
  ( ! [X30,X39,X37,X35,X33,X31,X38,X36,X34,X32,X40] : (s(X30,X31) = s(X30,'const.bool.COND_3'(s(bool,X32),s(X30,X33),s(X30,X34))) | s(X30,X31) != s(X30,X33) | s(X35,'const.bool.COND_3'(s(bool,X32),s(X35,X39),s(X35,X40))) = s(X35,'const.bool.COND_3'(s(bool,X36),s(X35,X37),s(X35,X38))) | s(X35,X38) != s(X35,X40) | s(bool,X32) != s(bool,X36)) )),
  inference(resolution,[],[f412,f159])).
fof(f5086,plain,(
  ( ! [X19,X17,X15,X18,X16] : (s(X15,X16) = s(X15,'const.bool.COND_3'(s(bool,X17),s(X15,X18),s(X15,X19))) | s(X15,X16) != s(X15,X18) | s(bool,'F_0') = s(bool,X17)) )),
  inference(resolution,[],[f412,f131])).
fof(f4716,plain,(
  ( ! [X8,X7] : (s('type.num.num',X8) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) )),
  inference(superposition,[],[f279,f168])).
fof(f4714,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X3))) | s('type.num.num',X3) != s('type.num.num',X4)) )),
  inference(superposition,[],[f279,f169])).
fof(f4712,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) | ($spl0 | $spl5)),
  inference(superposition,[],[f279,f2873])).
fof(f4711,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3(X2,sK0))) )),
  inference(superposition,[],[f279,f571])).
fof(f4710,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))))) | ~p(s(bool,sK3(X0,X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(superposition,[],[f279,f161])).
fof(f4708,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X44),s('type.num.num',X45),s('type.num.num',X46))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X41),s('type.num.num',X42),s('type.num.num',X43))))) | s('type.num.num',X43) != s('type.num.num',X46) | s('type.num.num',X42) != s('type.num.num',X45) | s(bool,X41) != s(bool,X44)) )),
  inference(superposition,[],[f279,f160])).
fof(f4707,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X38),s('type.num.num',X39),s('type.num.num',X40))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X35),s('type.num.num',X36),s('type.num.num',X37))))) | s('type.num.num',X37) != s('type.num.num',X40) | p(s(bool,X38)) | s(bool,X35) != s(bool,X38)) )),
  inference(superposition,[],[f279,f159])).
fof(f4706,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X32),s('type.num.num',X33),s('type.num.num',X34))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X29),s('type.num.num',X30),s('type.num.num',X31))))) | s('type.num.num',X31) != s('type.num.num',X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) )),
  inference(superposition,[],[f279,f159])).
fof(f4705,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X23),s('type.num.num',X24),s('type.num.num',X25))))) | ~p(s(bool,X26)) | s('type.num.num',X24) != s('type.num.num',X27) | s(bool,X23) != s(bool,X26)) )),
  inference(superposition,[],[f279,f158])).
fof(f4704,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X17),s('type.num.num',X18),s('type.num.num',X19))))) | ~p(s(bool,X17)) | s('type.num.num',X18) != s('type.num.num',X21) | s(bool,X17) != s(bool,X20)) )),
  inference(superposition,[],[f279,f158])).
fof(f4701,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK3(X11,X12)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',X12))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))))) )),
  inference(superposition,[],[f279,f163])).
fof(f4698,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | s('type.num.num',X7) != s('type.num.num',X8)) )),
  inference(superposition,[],[f279,f169])).
fof(f4719,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0))))))) | s(bool,'F_0') = s(bool,sK3(X6,sK0))) )),
  inference(trivial_inequality_removal,[],[f4695])).
fof(f4695,plain,(
  ( ! [X6] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0))))))) | s(bool,'F_0') = s(bool,sK3(X6,sK0))) )),
  inference(superposition,[],[f279,f571])).
fof(f4720,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5))))))) | ~p(s(bool,sK3(X4,X5)))) )),
  inference(subsumption_resolution,[],[f4694,f222])).
fof(f4694,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5))))))) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) )),
  inference(superposition,[],[f279,f161])).
fof(f4690,plain,(
  ( ! [X4,X5] : (s('type.num.num',X4) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) )),
  inference(superposition,[],[f279,f168])).
fof(f4688,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) | s('type.num.num',sK0) != s('type.num.num',X0)) )),
  inference(superposition,[],[f279,f169])).
fof(f4686,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f279,f168])).
fof(f4675,plain,(
  ( ! [X8,X7] : (s('type.num.num',X8) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))) )),
  inference(superposition,[],[f277,f168])).
fof(f4673,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',sK0))) | s('type.num.num',X3) != s('type.num.num',X4)) )),
  inference(superposition,[],[f277,f169])).
fof(f4671,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))),s('type.num.num',sK0))) | ($spl0 | $spl5)),
  inference(superposition,[],[f277,f2873])).
fof(f4670,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X2,sK0))) )),
  inference(superposition,[],[f277,f571])).
fof(f4669,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK0))) | ~p(s(bool,sK3(X0,X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(superposition,[],[f277,f161])).
fof(f4667,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X44),s('type.num.num',X45),s('type.num.num',X46))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X41),s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num',sK0))) | s('type.num.num',X43) != s('type.num.num',X46) | s('type.num.num',X42) != s('type.num.num',X45) | s(bool,X41) != s(bool,X44)) )),
  inference(superposition,[],[f277,f160])).
fof(f4666,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X38),s('type.num.num',X39),s('type.num.num',X40))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X35),s('type.num.num',X36),s('type.num.num',X37))),s('type.num.num',sK0))) | s('type.num.num',X37) != s('type.num.num',X40) | p(s(bool,X38)) | s(bool,X35) != s(bool,X38)) )),
  inference(superposition,[],[f277,f159])).
fof(f4665,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X32),s('type.num.num',X33),s('type.num.num',X34))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X29),s('type.num.num',X30),s('type.num.num',X31))),s('type.num.num',sK0))) | s('type.num.num',X31) != s('type.num.num',X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) )),
  inference(superposition,[],[f277,f159])).
fof(f4664,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X23),s('type.num.num',X24),s('type.num.num',X25))),s('type.num.num',sK0))) | ~p(s(bool,X26)) | s('type.num.num',X24) != s('type.num.num',X27) | s(bool,X23) != s(bool,X26)) )),
  inference(superposition,[],[f277,f158])).
fof(f4663,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X17),s('type.num.num',X18),s('type.num.num',X19))),s('type.num.num',sK0))) | ~p(s(bool,X17)) | s('type.num.num',X18) != s('type.num.num',X21) | s(bool,X17) != s(bool,X20)) )),
  inference(superposition,[],[f277,f158])).
fof(f4660,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,sK3(X11,X12)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',X12))))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))))) )),
  inference(superposition,[],[f277,f163])).
fof(f4657,plain,(
  ( ! [X8,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',sK0))) | s('type.num.num',X7) != s('type.num.num',X8)) )),
  inference(superposition,[],[f277,f169])).
fof(f4653,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5))))),s('type.num.num',sK0))) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) )),
  inference(superposition,[],[f277,f161])).
fof(f4649,plain,(
  ( ! [X4,X5] : (s('type.num.num',X4) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f277,f168])).
fof(f4647,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK0))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(superposition,[],[f277,f169])).
fof(f4645,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f277,f168])).
fof(f4281,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X44),s('type.num.num',X45),s('type.num.num',X46))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X41),s('type.num.num',X42),s('type.num.num',X43))))),s('type.num.num',sK0))) | s('type.num.num',X43) != s('type.num.num',X46) | s('type.num.num',X42) != s('type.num.num',X45) | s(bool,X41) != s(bool,X44)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f160])).
fof(f4280,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X38),s('type.num.num',X39),s('type.num.num',X40))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X35),s('type.num.num',X36),s('type.num.num',X37))))),s('type.num.num',sK0))) | s('type.num.num',X37) != s('type.num.num',X40) | p(s(bool,X38)) | s(bool,X35) != s(bool,X38)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f159])).
fof(f4279,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X32),s('type.num.num',X33),s('type.num.num',X34))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X29),s('type.num.num',X30),s('type.num.num',X31))))),s('type.num.num',sK0))) | s('type.num.num',X31) != s('type.num.num',X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f159])).
fof(f4278,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X23),s('type.num.num',X24),s('type.num.num',X25))))),s('type.num.num',sK0))) | ~p(s(bool,X26)) | s('type.num.num',X24) != s('type.num.num',X27) | s(bool,X23) != s(bool,X26)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f158])).
fof(f4277,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X17),s('type.num.num',X18),s('type.num.num',X19))))),s('type.num.num',sK0))) | ~p(s(bool,X17)) | s('type.num.num',X18) != s('type.num.num',X21) | s(bool,X17) != s(bool,X20)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f158])).
fof(f4274,plain,(
  ( ! [X12,X11] : (s('type.num.num',sK1) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,sK3(X11,X12)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',X12))))))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f163])).
fof(f4271,plain,(
  ( ! [X8,X7] : (s('type.num.num',sK1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))),s('type.num.num',sK0))) | s('type.num.num',X7) != s('type.num.num',X8)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f169])).
fof(f4269,plain,(
  ( ! [X6] : (s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0))))))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X6,sK0))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f571])).
fof(f4268,plain,(
  ( ! [X4,X5] : (s('type.num.num',sK1) != s('type.num.num',X5) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5))))))),s('type.num.num',sK0))) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f161])).
fof(f4266,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK1) != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f168])).
fof(f4264,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4256,f168])).
fof(f4262,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4146,f168])).
fof(f4260,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4146,f168])).
fof(f4258,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4146,f168])).
fof(f4256,plain,(
  ( ! [X0] : (s('type.num.num',sK1) != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4146,f169])).
fof(f4254,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f4146,f168])).
fof(f4179,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),sK0)) | ($spl0 | $spl5)),
  inference(superposition,[],[f601,f2873])).
fof(f4178,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),sK0)) | ($spl0 | $spl5)),
  inference(superposition,[],[f599,f2873])).
fof(f4177,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),sK0)) | ($spl0 | $spl5)),
  inference(superposition,[],[f571,f2873])).
fof(f4222,plain,(
  ( ! [X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))))))) ) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f4164,f2873])).
fof(f4164,plain,(
  ( ! [X10] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f163,f2873])).
fof(f4162,plain,(
  ( ! [X8] : (s('type.num.num',X8) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X8))))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),X8))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f162,f2873])).
fof(f4220,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | p(s(bool,sK3(X7,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))))) ) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f4161,f2873])).
fof(f4161,plain,(
  ( ! [X7] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | p(s(bool,sK3(X7,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f162,f2873])).
fof(f4159,plain,(
  ( ! [X5] : (s('type.num.num',X5) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X5))))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f161,f2873])).
fof(f4218,plain,(
  ( ! [X4] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | ~p(s(bool,sK3(X4,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))))) ) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f4158,f2873])).
fof(f4158,plain,(
  ( ! [X4] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | ~p(s(bool,sK3(X4,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f161,f2873])).
fof(f4153,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))) | ($spl0 | $spl5)),
  inference(superposition,[],[f222,f2873])).
fof(f4152,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))),s('type.num.num',sK0)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f220,f2873])).
fof(f4148,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f168,f2873])).
fof(f4146,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl5)),
  inference(superposition,[],[f149,f2873])).
fof(f4135,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2873,f169])).
fof(f4131,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2873,f168])).
fof(f4129,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2873,f168])).
fof(f4127,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2873,f168])).
fof(f4124,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK0) != s('type.num.num',X7) | s(bool,'F_0') != s(bool,sK3(X6,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0)))))))) )),
  inference(superposition,[],[f601,f168])).
fof(f4122,plain,(
  ( ! [X2,X3] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s(bool,'F_0') != s(bool,sK3(X2,sK0)) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f601,f169])).
fof(f4118,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s(bool,'F_0') != s(bool,sK3(X0,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0)))))) )),
  inference(superposition,[],[f601,f168])).
fof(f4115,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s(bool,'F_0') != s(bool,sK3(X1,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f601,f168])).
fof(f4113,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X44),s('type.num.num',X45),s('type.num.num',X46))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,X41),s('type.num.num',X42),s('type.num.num',X43)),sK0)) | s('type.num.num',X43) != s('type.num.num',X46) | s('type.num.num',X42) != s('type.num.num',X45) | s(bool,X41) != s(bool,X44)) )),
  inference(superposition,[],[f601,f160])).
fof(f4112,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X38),s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,X35),s('type.num.num',X36),s('type.num.num',X37)),sK0)) | s('type.num.num',X37) != s('type.num.num',X40) | p(s(bool,X38)) | s(bool,X35) != s(bool,X38)) )),
  inference(superposition,[],[f601,f159])).
fof(f4111,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X32),s('type.num.num',X33),s('type.num.num',X34))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,X29),s('type.num.num',X30),s('type.num.num',X31)),sK0)) | s('type.num.num',X31) != s('type.num.num',X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) )),
  inference(superposition,[],[f601,f159])).
fof(f4110,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,X23),s('type.num.num',X24),s('type.num.num',X25)),sK0)) | ~p(s(bool,X26)) | s('type.num.num',X24) != s('type.num.num',X27) | s(bool,X23) != s(bool,X26)) )),
  inference(superposition,[],[f601,f158])).
fof(f4109,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,X17),s('type.num.num',X18),s('type.num.num',X19)),sK0)) | ~p(s(bool,X17)) | s('type.num.num',X18) != s('type.num.num',X21) | s(bool,X17) != s(bool,X20)) )),
  inference(superposition,[],[f601,f158])).
fof(f4108,plain,(
  ( ! [X15,X16] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X16),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X15),s('type.num.num',X16)),sK0))) )),
  inference(superposition,[],[f601,f154])).
fof(f4107,plain,(
  ( ! [X14,X13] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X13),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X13),s('type.num.num',X14)),sK0))) )),
  inference(superposition,[],[f601,f153])).
fof(f4106,plain,(
  ( ! [X12,X11] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.bool.COND_3'(s(bool,sK3(X11,X12)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',X12)))))),sK0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))))) )),
  inference(superposition,[],[f601,f163])).
fof(f4105,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),sK0)) | $spl0),
  inference(superposition,[],[f601,f258])).
fof(f4103,plain,(
  ( ! [X8,X7] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num',X7)),sK0)) | s('type.num.num',X7) != s('type.num.num',X8)) )),
  inference(superposition,[],[f601,f169])).
fof(f4102,plain,(
  ( ! [X6] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0)))),sK0)) | s(bool,'F_0') = s(bool,sK3(X6,sK0))) )),
  inference(superposition,[],[f601,f571])).
fof(f4101,plain,(
  ( ! [X4,X5] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5)))),sK0)) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) )),
  inference(superposition,[],[f601,f161])).
fof(f4099,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3(X0,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f601,f168])).
fof(f4095,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK0) != s('type.num.num',X7) | s(bool,'T_0') = s(bool,sK3(X6,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0)))))))) )),
  inference(superposition,[],[f599,f168])).
fof(f4093,plain,(
  ( ! [X2,X3] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s(bool,'T_0') = s(bool,sK3(X2,sK0)) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f599,f169])).
fof(f4089,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s(bool,'T_0') = s(bool,sK3(X0,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0)))))) )),
  inference(superposition,[],[f599,f168])).
fof(f4086,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s(bool,'T_0') = s(bool,sK3(X1,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f599,f168])).
fof(f4084,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X44),s('type.num.num',X45),s('type.num.num',X46))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X41),s('type.num.num',X42),s('type.num.num',X43)),sK0)) | s('type.num.num',X43) != s('type.num.num',X46) | s('type.num.num',X42) != s('type.num.num',X45) | s(bool,X41) != s(bool,X44)) )),
  inference(superposition,[],[f599,f160])).
fof(f4083,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X38),s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X35),s('type.num.num',X36),s('type.num.num',X37)),sK0)) | s('type.num.num',X37) != s('type.num.num',X40) | p(s(bool,X38)) | s(bool,X35) != s(bool,X38)) )),
  inference(superposition,[],[f599,f159])).
fof(f4082,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X32),s('type.num.num',X33),s('type.num.num',X34))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X29),s('type.num.num',X30),s('type.num.num',X31)),sK0)) | s('type.num.num',X31) != s('type.num.num',X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) )),
  inference(superposition,[],[f599,f159])).
fof(f4081,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X23),s('type.num.num',X24),s('type.num.num',X25)),sK0)) | ~p(s(bool,X26)) | s('type.num.num',X24) != s('type.num.num',X27) | s(bool,X23) != s(bool,X26)) )),
  inference(superposition,[],[f599,f158])).
fof(f4080,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X17),s('type.num.num',X18),s('type.num.num',X19)),sK0)) | ~p(s(bool,X17)) | s('type.num.num',X18) != s('type.num.num',X21) | s(bool,X17) != s(bool,X20)) )),
  inference(superposition,[],[f599,f158])).
fof(f4079,plain,(
  ( ! [X15,X16] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X16),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X15),s('type.num.num',X16)),sK0))) )),
  inference(superposition,[],[f599,f154])).
fof(f4078,plain,(
  ( ! [X14,X13] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X13),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X13),s('type.num.num',X14)),sK0))) )),
  inference(superposition,[],[f599,f153])).
fof(f4077,plain,(
  ( ! [X12,X11] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.bool.COND_3'(s(bool,sK3(X11,X12)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',X12)))))),sK0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))))) )),
  inference(superposition,[],[f599,f163])).
fof(f4076,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),sK0)) | $spl0),
  inference(superposition,[],[f599,f258])).
fof(f4074,plain,(
  ( ! [X8,X7] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',X7)),sK0)) | s('type.num.num',X7) != s('type.num.num',X8)) )),
  inference(superposition,[],[f599,f169])).
fof(f4073,plain,(
  ( ! [X6] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0)))),sK0)) | s(bool,'F_0') = s(bool,sK3(X6,sK0))) )),
  inference(superposition,[],[f599,f571])).
fof(f4072,plain,(
  ( ! [X4,X5] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5)))),sK0)) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) )),
  inference(superposition,[],[f599,f161])).
fof(f4070,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3(X0,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f599,f168])).
fof(f3515,plain,(
  ( ! [X177,X175,X178,X176,X174,X179] : (s('type.num.num',sK0) != s('type.num.num',X175) | s('type.num.num','const.bool.COND_3'(s(bool,X178),s('type.num.num',X179),s('type.num.num',X175))) = s('type.num.num','const.bool.COND_3'(s(bool,X176),s('type.num.num',X177),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X174),s('type.num.num',sK0))))))) | s('type.num.num',X177) != s('type.num.num',X179) | s(bool,X176) != s(bool,X178) | s(bool,'F_0') = s(bool,sK3(X174,sK0))) )),
  inference(superposition,[],[f160,f571])).
fof(f3513,plain,(
  ( ! [X167,X165,X163,X166,X164,X162] : (s('type.num.num',sK0) != s('type.num.num',X163) | s('type.num.num','const.bool.COND_3'(s(bool,X166),s('type.num.num',X163),s('type.num.num',X167))) = s('type.num.num','const.bool.COND_3'(s(bool,X164),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X162),s('type.num.num',sK0))))),s('type.num.num',X165))) | s('type.num.num',X165) != s('type.num.num',X167) | s(bool,X164) != s(bool,X166) | s(bool,'F_0') = s(bool,sK3(X162,sK0))) )),
  inference(superposition,[],[f160,f571])).
fof(f3511,plain,(
  ( ! [X154,X152,X151,X155,X153,X150] : (s('type.num.num','const.bool.COND_3'(s(bool,X151),s('type.num.num',sK0),s('type.num.num',X152))) = s('type.num.num','const.bool.COND_3'(s(bool,X153),s('type.num.num',X154),s('type.num.num',X155))) | s('type.num.num',X152) != s('type.num.num',X155) | s('type.num.num',X154) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X150),s('type.num.num',sK0))))) | s(bool,X151) != s(bool,X153) | s(bool,'F_0') = s(bool,sK3(X150,sK0))) )),
  inference(superposition,[],[f160,f571])).
fof(f3509,plain,(
  ( ! [X142,X140,X138,X143,X141,X139] : (s('type.num.num','const.bool.COND_3'(s(bool,X139),s('type.num.num',X140),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,X141),s('type.num.num',X142),s('type.num.num',X143))) | s('type.num.num',X143) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X138),s('type.num.num',sK0))))) | s('type.num.num',X140) != s('type.num.num',X142) | s(bool,X139) != s(bool,X141) | s(bool,'F_0') = s(bool,sK3(X138,sK0))) )),
  inference(superposition,[],[f160,f571])).
fof(f3508,plain,(
  ( ! [X134,X132,X136,X135,X133,X137] : (s('type.num.num',sK0) != s('type.num.num',X133) | s('type.num.num','const.bool.COND_3'(s(bool,X134),s('type.num.num',X135),s('type.num.num',X133))) = s('type.num.num','const.bool.COND_3'(s(bool,X136),s('type.num.num',X137),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X132),s('type.num.num',sK0))))))) | p(s(bool,X136)) | s(bool,X134) != s(bool,X136) | s(bool,'F_0') = s(bool,sK3(X132,sK0))) )),
  inference(superposition,[],[f159,f571])).
fof(f3507,plain,(
  ( ! [X127,X130,X128,X126,X131,X129] : (s('type.num.num',sK0) != s('type.num.num',X127) | s('type.num.num','const.bool.COND_3'(s(bool,X130),s('type.num.num',X131),s('type.num.num',X127))) = s('type.num.num','const.bool.COND_3'(s(bool,X128),s('type.num.num',X129),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X126),s('type.num.num',sK0))))))) | p(s(bool,X130)) | s(bool,X128) != s(bool,X130) | s(bool,'F_0') = s(bool,sK3(X126,sK0))) )),
  inference(superposition,[],[f159,f571])).
fof(f3504,plain,(
  ( ! [X111,X109,X112,X110,X108,X113] : (s('type.num.num','const.bool.COND_3'(s(bool,X112),s('type.num.num',X113),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,X109),s('type.num.num',X110),s('type.num.num',X111))) | s('type.num.num',X111) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X108),s('type.num.num',sK0))))) | p(s(bool,X112)) | s(bool,X109) != s(bool,X112) | s(bool,'F_0') = s(bool,sK3(X108,sK0))) )),
  inference(superposition,[],[f159,f571])).
fof(f3503,plain,(
  ( ! [X103,X107,X105,X102,X106,X104] : (s('type.num.num','const.bool.COND_3'(s(bool,X103),s('type.num.num',X104),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,X105),s('type.num.num',X106),s('type.num.num',X107))) | s('type.num.num',X107) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X102),s('type.num.num',sK0))))) | p(s(bool,X105)) | s(bool,X103) != s(bool,X105) | s(bool,'F_0') = s(bool,sK3(X102,sK0))) )),
  inference(superposition,[],[f159,f571])).
fof(f3502,plain,(
  ( ! [X101,X99,X97,X100,X98,X96] : (s('type.num.num',sK0) != s('type.num.num',X97) | s('type.num.num','const.bool.COND_3'(s(bool,X98),s('type.num.num',X97),s('type.num.num',X99))) = s('type.num.num','const.bool.COND_3'(s(bool,X100),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X96),s('type.num.num',sK0))))),s('type.num.num',X101))) | ~p(s(bool,X100)) | s(bool,X98) != s(bool,X100) | s(bool,'F_0') = s(bool,sK3(X96,sK0))) )),
  inference(superposition,[],[f158,f571])).
fof(f3501,plain,(
  ( ! [X94,X92,X90,X95,X93,X91] : (s('type.num.num',sK0) != s('type.num.num',X91) | s('type.num.num','const.bool.COND_3'(s(bool,X94),s('type.num.num',X91),s('type.num.num',X95))) = s('type.num.num','const.bool.COND_3'(s(bool,X92),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X90),s('type.num.num',sK0))))),s('type.num.num',X93))) | ~p(s(bool,X94)) | s(bool,X92) != s(bool,X94) | s(bool,'F_0') = s(bool,sK3(X90,sK0))) )),
  inference(superposition,[],[f158,f571])).
fof(f3500,plain,(
  ( ! [X88,X87,X85,X89,X86,X84] : (s('type.num.num','const.bool.COND_3'(s(bool,X88),s('type.num.num',sK0),s('type.num.num',X89))) = s('type.num.num','const.bool.COND_3'(s(bool,X85),s('type.num.num',X86),s('type.num.num',X87))) | ~p(s(bool,X88)) | s('type.num.num',X86) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X84),s('type.num.num',sK0))))) | s(bool,X85) != s(bool,X88) | s(bool,'F_0') = s(bool,sK3(X84,sK0))) )),
  inference(superposition,[],[f158,f571])).
fof(f3499,plain,(
  ( ! [X80,X78,X83,X81,X79,X82] : (s('type.num.num','const.bool.COND_3'(s(bool,X79),s('type.num.num',sK0),s('type.num.num',X80))) = s('type.num.num','const.bool.COND_3'(s(bool,X81),s('type.num.num',X82),s('type.num.num',X83))) | ~p(s(bool,X81)) | s('type.num.num',X82) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X78),s('type.num.num',sK0))))) | s(bool,X79) != s(bool,X81) | s(bool,'F_0') = s(bool,sK3(X78,sK0))) )),
  inference(superposition,[],[f158,f571])).
fof(f3492,plain,(
  ( ! [X57] : (s('type.num.num',sK0) != s('type.num.num',sK1) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X57),s('type.num.num',sK0))))))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X57,sK0))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f571])).
fof(f3491,plain,(
  ( ! [X56] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X56),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3(X56,sK0))) ) | $spl5),
  inference(superposition,[],[f3104,f571])).
fof(f3517,plain,(
  ( ! [X54] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X54),s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,sK3(X54,sK0))) )),
  inference(trivial_inequality_removal,[],[f3489])).
fof(f3489,plain,(
  ( ! [X54] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X54),s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,sK3(X54,sK0))) )),
  inference(superposition,[],[f222,f571])).
fof(f3488,plain,(
  ( ! [X53] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X53),s('type.num.num',sK0))))),s('type.num.num',sK0)))) | s(bool,'F_0') = s(bool,sK3(X53,sK0))) )),
  inference(superposition,[],[f220,f571])).
fof(f3486,plain,(
  ( ! [X50,X49] : (s('type.num.num',sK0) != s('type.num.num',X50) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X50))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X49),s('type.num.num',sK0))))))) | s(bool,'F_0') = s(bool,sK3(X49,sK0))) )),
  inference(superposition,[],[f169,f571])).
fof(f3484,plain,(
  ( ! [X45,X46] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X46))) | s('type.num.num',X46) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X45),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3(X45,sK0))) )),
  inference(superposition,[],[f169,f571])).
fof(f3480,plain,(
  ( ! [X37,X38] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X38))) | s('type.num.num',X38) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X37),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3(X37,sK0))) )),
  inference(superposition,[],[f168,f571])).
fof(f3530,plain,(
  ( ! [X35,X36] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0))))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X36,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X36),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0))))))))))) | s(bool,'F_0') = s(bool,sK3(X35,sK0))) )),
  inference(subsumption_resolution,[],[f3529,f119])).
fof(f3529,plain,(
  ( ! [X35,X36] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0))))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X36,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X36),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0))))))))))) | s(bool,'F_0') = s(bool,sK3(X35,sK0))) )),
  inference(forward_demodulation,[],[f3479,f205])).
fof(f3479,plain,(
  ( ! [X35,X36] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0))))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X36,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X36),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X35),s('type.num.num',sK0))))))))))) | s(bool,'F_0') = s(bool,sK3(X35,sK0))) )),
  inference(superposition,[],[f163,f571])).
fof(f3478,plain,(
  ( ! [X33,X34] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))),s('type.num.num',X34))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X33),s('type.num.num',sK0)))),X34)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',X34))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X34)))) | s(bool,'F_0') = s(bool,sK3(X33,sK0))) )),
  inference(superposition,[],[f163,f571])).
fof(f3477,plain,(
  ( ! [X31,X32] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X32))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X32,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',sK0)))))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X32),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X31),s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,sK3(X31,sK0))) )),
  inference(superposition,[],[f163,f571])).
fof(f3528,plain,(
  ( ! [X30,X29] : (p(s(bool,sK3(X30,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))))))) | s(bool,'F_0') = s(bool,sK3(X29,sK0))) )),
  inference(subsumption_resolution,[],[f3527,f119])).
fof(f3527,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | p(s(bool,sK3(X30,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))))))) | s(bool,'F_0') = s(bool,sK3(X29,sK0))) )),
  inference(forward_demodulation,[],[f3476,f205])).
fof(f3476,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | p(s(bool,sK3(X30,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X30),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X29),s('type.num.num',sK0))))))))) | s(bool,'F_0') = s(bool,sK3(X29,sK0))) )),
  inference(superposition,[],[f162,f571])).
fof(f3475,plain,(
  ( ! [X28,X27] : (s('type.num.num',X28) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',X28))))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X27),s('type.num.num',sK0)))),X28))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X28)))) | s(bool,'F_0') = s(bool,sK3(X27,sK0))) )),
  inference(superposition,[],[f162,f571])).
fof(f3474,plain,(
  ( ! [X26,X25] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X26),s('type.num.num',sK0))))) | p(s(bool,sK3(X26,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X25),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X25),s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,sK3(X25,sK0))) )),
  inference(superposition,[],[f162,f571])).
fof(f3526,plain,(
  ( ! [X24,X23] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))))))) | ~p(s(bool,sK3(X24,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))))) | s(bool,'F_0') = s(bool,sK3(X23,sK0))) )),
  inference(subsumption_resolution,[],[f3525,f119])).
fof(f3525,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))))))) | ~p(s(bool,sK3(X24,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))))) | s(bool,'F_0') = s(bool,sK3(X23,sK0))) )),
  inference(forward_demodulation,[],[f3473,f205])).
fof(f3473,plain,(
  ( ! [X24,X23] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X24),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))))))) | ~p(s(bool,sK3(X24,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X23),s('type.num.num',sK0))))))) | s(bool,'F_0') = s(bool,sK3(X23,sK0))) )),
  inference(superposition,[],[f161,f571])).
fof(f3472,plain,(
  ( ! [X21,X22] : (s('type.num.num',X22) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',X22))))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X21),s('type.num.num',sK0)))),X22))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X22)))) | s(bool,'F_0') = s(bool,sK3(X21,sK0))) )),
  inference(superposition,[],[f161,f571])).
fof(f3471,plain,(
  ( ! [X19,X20] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X20),s('type.num.num',sK0))))) | ~p(s(bool,sK3(X20,'const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X19),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X19),s('type.num.num',sK0)))))))) | s(bool,'F_0') = s(bool,sK3(X19,sK0))) )),
  inference(superposition,[],[f161,f571])).
fof(f3518,plain,(
  ( ! [X16] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X16),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X16,sK0))) ) | $spl5),
  inference(trivial_inequality_removal,[],[f3468])).
fof(f3468,plain,(
  ( ! [X16] : (s('type.num.num',sK0) != s('type.num.num',sK0) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X16),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X16,sK0))) ) | $spl5),
  inference(superposition,[],[f3104,f571])).
fof(f3466,plain,(
  ( ! [X14] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK0)))))) | s(bool,'F_0') = s(bool,sK3(X14,sK0))) )),
  inference(superposition,[],[f222,f571])).
fof(f3465,plain,(
  ( ! [X13] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X13),s('type.num.num',sK0))),s('type.num.num',sK0)))) | s(bool,'F_0') = s(bool,sK3(X13,sK0))) )),
  inference(superposition,[],[f220,f571])).
fof(f3461,plain,(
  ( ! [X6,X5] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) | s('type.num.num',X6) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X5,sK0))) )),
  inference(superposition,[],[f168,f571])).
fof(f3524,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X2,sK0))) )),
  inference(subsumption_resolution,[],[f3523,f119])).
fof(f3523,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3(X2,sK0))) )),
  inference(forward_demodulation,[],[f3456,f205])).
fof(f3456,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s(bool,'F_0') = s(bool,sK3(X2,sK0))) )),
  inference(superposition,[],[f163,f571])).
fof(f3453,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK0) = s('type.num.num',X7) | s(bool,'F_0') = s(bool,sK3(X6,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0)))))))) )),
  inference(superposition,[],[f571,f168])).
fof(f3451,plain,(
  ( ! [X2,X3] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s(bool,'F_0') = s(bool,sK3(X2,sK0)) | s('type.num.num',X3) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f571,f169])).
fof(f3449,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s(bool,'F_0') = s(bool,sK3(X0,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0)))))) )),
  inference(superposition,[],[f571,f168])).
fof(f3448,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)) | $spl0),
  inference(superposition,[],[f571,f258])).
fof(f3446,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))))) | s(bool,'F_0') = s(bool,sK3(X1,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f571,f168])).
fof(f3444,plain,(
  ( ! [X45,X43,X41,X46,X44,X42] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X44),s('type.num.num',X45),s('type.num.num',X46))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X41),s('type.num.num',X42),s('type.num.num',X43)),sK0)) | s('type.num.num',X43) != s('type.num.num',X46) | s('type.num.num',X42) != s('type.num.num',X45) | s(bool,X41) != s(bool,X44)) )),
  inference(superposition,[],[f571,f160])).
fof(f3443,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X38),s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X35),s('type.num.num',X36),s('type.num.num',X37)),sK0)) | s('type.num.num',X37) != s('type.num.num',X40) | p(s(bool,X38)) | s(bool,X35) != s(bool,X38)) )),
  inference(superposition,[],[f571,f159])).
fof(f3442,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X32),s('type.num.num',X33),s('type.num.num',X34))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X29),s('type.num.num',X30),s('type.num.num',X31)),sK0)) | s('type.num.num',X31) != s('type.num.num',X34) | p(s(bool,X29)) | s(bool,X29) != s(bool,X32)) )),
  inference(superposition,[],[f571,f159])).
fof(f3441,plain,(
  ( ! [X28,X26,X24,X23,X27,X25] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X26),s('type.num.num',X27),s('type.num.num',X28))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X23),s('type.num.num',X24),s('type.num.num',X25)),sK0)) | ~p(s(bool,X26)) | s('type.num.num',X24) != s('type.num.num',X27) | s(bool,X23) != s(bool,X26)) )),
  inference(superposition,[],[f571,f158])).
fof(f3440,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.bool.COND_3'(s(bool,X20),s('type.num.num',X21),s('type.num.num',X22))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,X17),s('type.num.num',X18),s('type.num.num',X19)),sK0)) | ~p(s(bool,X17)) | s('type.num.num',X18) != s('type.num.num',X21) | s(bool,X17) != s(bool,X20)) )),
  inference(superposition,[],[f571,f158])).
fof(f3439,plain,(
  ( ! [X15,X16] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X16),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X15),s('type.num.num',X16)),sK0))) )),
  inference(superposition,[],[f571,f154])).
fof(f3438,plain,(
  ( ! [X14,X13] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X13),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X13),s('type.num.num',X14)),sK0))) )),
  inference(superposition,[],[f571,f153])).
fof(f3437,plain,(
  ( ! [X12,X11] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.bool.COND_3'(s(bool,sK3(X11,X12)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X11),s('type.num.num',X12)))))),sK0)) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12))))) )),
  inference(superposition,[],[f571,f163])).
fof(f3436,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),sK0)) | $spl0),
  inference(superposition,[],[f571,f258])).
fof(f3434,plain,(
  ( ! [X8,X7] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',X7)),sK0)) | s('type.num.num',X7) != s('type.num.num',X8)) )),
  inference(superposition,[],[f571,f169])).
fof(f3433,plain,(
  ( ! [X6] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK0),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num',sK0)))),sK0)) | s(bool,'F_0') = s(bool,sK3(X6,sK0))) )),
  inference(superposition,[],[f571,f571])).
fof(f3432,plain,(
  ( ! [X4,X5] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X5),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5)))),sK0)) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) )),
  inference(superposition,[],[f571,f161])).
fof(f3430,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3(X0,sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f571,f168])).
fof(f3407,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X43),s('type.num.num',X44),s('type.num.num',X45))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X40),s('type.num.num',X41),s('type.num.num',X42))))),s('type.num.num',sK0))) | s('type.num.num',X42) != s('type.num.num',X45) | s('type.num.num',X41) != s('type.num.num',X44) | s(bool,X40) != s(bool,X43)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f160])).
fof(f3406,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X37),s('type.num.num',X38),s('type.num.num',X39))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X34),s('type.num.num',X35),s('type.num.num',X36))))),s('type.num.num',sK0))) | s('type.num.num',X36) != s('type.num.num',X39) | p(s(bool,X37)) | s(bool,X34) != s(bool,X37)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f159])).
fof(f3405,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X31),s('type.num.num',X32),s('type.num.num',X33))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X28),s('type.num.num',X29),s('type.num.num',X30))))),s('type.num.num',sK0))) | s('type.num.num',X30) != s('type.num.num',X33) | p(s(bool,X28)) | s(bool,X28) != s(bool,X31)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f159])).
fof(f3404,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X25),s('type.num.num',X26),s('type.num.num',X27))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))))),s('type.num.num',sK0))) | ~p(s(bool,X25)) | s('type.num.num',X23) != s('type.num.num',X26) | s(bool,X22) != s(bool,X25)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f158])).
fof(f3403,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s('type.num.num',sK1) != s('type.num.num','const.bool.COND_3'(s(bool,X19),s('type.num.num',X20),s('type.num.num',X21))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))))),s('type.num.num',sK0))) | ~p(s(bool,X16)) | s('type.num.num',X17) != s('type.num.num',X20) | s(bool,X16) != s(bool,X19)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f158])).
fof(f3400,plain,(
  ( ! [X10,X11] : (s('type.num.num',sK1) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',X11))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))))))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f163])).
fof(f3397,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))),s('type.num.num',sK0))) | s('type.num.num',X6) != s('type.num.num',X7)) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f169])).
fof(f3396,plain,(
  ( ! [X4,X5] : (s('type.num.num',sK1) != s('type.num.num',X5) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5))))))),s('type.num.num',sK0))) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f161])).
fof(f3394,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK1) != s('type.num.num',X1) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f168])).
fof(f3392,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f3114,f168])).
fof(f3389,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK0) != s('type.num.num',X7) | s('type.num.num',X6) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) ) | $spl5),
  inference(superposition,[],[f3104,f168])).
fof(f3387,plain,(
  ( ! [X2,X3] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | s('type.num.num',X2) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | s('type.num.num',X2) != s('type.num.num',X3)) ) | $spl5),
  inference(superposition,[],[f3104,f169])).
fof(f3386,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))) | ~p(s(bool,sK3(X0,X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl5),
  inference(superposition,[],[f3104,f161])).
fof(f3384,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X43),s('type.num.num',X44),s('type.num.num',X45))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,X40),s('type.num.num',X41),s('type.num.num',X42))) | s('type.num.num',X42) != s('type.num.num',X45) | s('type.num.num',X41) != s('type.num.num',X44) | s(bool,X40) != s(bool,X43)) ) | $spl5),
  inference(superposition,[],[f3104,f160])).
fof(f3383,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X37),s('type.num.num',X38),s('type.num.num',X39))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,X34),s('type.num.num',X35),s('type.num.num',X36))) | s('type.num.num',X36) != s('type.num.num',X39) | p(s(bool,X37)) | s(bool,X34) != s(bool,X37)) ) | $spl5),
  inference(superposition,[],[f3104,f159])).
fof(f3382,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X31),s('type.num.num',X32),s('type.num.num',X33))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,X28),s('type.num.num',X29),s('type.num.num',X30))) | s('type.num.num',X30) != s('type.num.num',X33) | p(s(bool,X28)) | s(bool,X28) != s(bool,X31)) ) | $spl5),
  inference(superposition,[],[f3104,f159])).
fof(f3381,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X25),s('type.num.num',X26),s('type.num.num',X27))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))) | ~p(s(bool,X25)) | s('type.num.num',X23) != s('type.num.num',X26) | s(bool,X22) != s(bool,X25)) ) | $spl5),
  inference(superposition,[],[f3104,f158])).
fof(f3380,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X19),s('type.num.num',X20),s('type.num.num',X21))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))) | ~p(s(bool,X16)) | s('type.num.num',X17) != s('type.num.num',X20) | s(bool,X16) != s(bool,X19)) ) | $spl5),
  inference(superposition,[],[f3104,f158])).
fof(f3379,plain,(
  ( ! [X14,X15] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num',X14),s('type.num.num',X15)))) ) | $spl5),
  inference(superposition,[],[f3104,f154])).
fof(f3378,plain,(
  ( ! [X12,X13] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X12))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num',X12),s('type.num.num',X13)))) ) | $spl5),
  inference(superposition,[],[f3104,f153])).
fof(f3377,plain,(
  ( ! [X10,X11] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',X11))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) ) | $spl5),
  inference(superposition,[],[f3104,f163])).
fof(f3374,plain,(
  ( ! [X6,X7] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | s('type.num.num',X6) != s('type.num.num',X7)) ) | $spl5),
  inference(superposition,[],[f3104,f169])).
fof(f3373,plain,(
  ( ! [X4,X5] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5))))) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) ) | $spl5),
  inference(superposition,[],[f3104,f161])).
fof(f3371,plain,(
  ( ! [X0,X1] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) ) | $spl5),
  inference(superposition,[],[f3104,f168])).
fof(f3342,plain,(
  ( ! [X6,X7] : (s('type.num.num',X7) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f222,f168])).
fof(f3340,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X2)))) | s('type.num.num',X2) != s('type.num.num',X3)) )),
  inference(superposition,[],[f222,f169])).
fof(f3339,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1)))))) | ~p(s(bool,sK3(X0,X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(superposition,[],[f222,f161])).
fof(f3337,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X43),s('type.num.num',X44),s('type.num.num',X45))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X40),s('type.num.num',X41),s('type.num.num',X42)))))) | s('type.num.num',X42) != s('type.num.num',X45) | s('type.num.num',X41) != s('type.num.num',X44) | s(bool,X40) != s(bool,X43)) )),
  inference(superposition,[],[f222,f160])).
fof(f3336,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X37),s('type.num.num',X38),s('type.num.num',X39))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X34),s('type.num.num',X35),s('type.num.num',X36)))))) | s('type.num.num',X36) != s('type.num.num',X39) | p(s(bool,X37)) | s(bool,X34) != s(bool,X37)) )),
  inference(superposition,[],[f222,f159])).
fof(f3335,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X31),s('type.num.num',X32),s('type.num.num',X33))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X28),s('type.num.num',X29),s('type.num.num',X30)))))) | s('type.num.num',X30) != s('type.num.num',X33) | p(s(bool,X28)) | s(bool,X28) != s(bool,X31)) )),
  inference(superposition,[],[f222,f159])).
fof(f3334,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X25),s('type.num.num',X26),s('type.num.num',X27))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24)))))) | ~p(s(bool,X25)) | s('type.num.num',X23) != s('type.num.num',X26) | s(bool,X22) != s(bool,X25)) )),
  inference(superposition,[],[f222,f158])).
fof(f3333,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X19),s('type.num.num',X20),s('type.num.num',X21))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18)))))) | ~p(s(bool,X16)) | s('type.num.num',X17) != s('type.num.num',X20) | s(bool,X16) != s(bool,X19)) )),
  inference(superposition,[],[f222,f158])).
fof(f3330,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',X11))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11)))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) )),
  inference(superposition,[],[f222,f163])).
fof(f3327,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) | s('type.num.num',X6) != s('type.num.num',X7)) )),
  inference(superposition,[],[f222,f169])).
fof(f3345,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5)))))))) | ~p(s(bool,sK3(X4,X5)))) )),
  inference(subsumption_resolution,[],[f3326,f222])).
fof(f3326,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5)))))))) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) )),
  inference(superposition,[],[f222,f161])).
fof(f3322,plain,(
  ( ! [X4,X5] : (s('type.num.num',X4) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0)))))) )),
  inference(superposition,[],[f222,f168])).
fof(f3320,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | s('type.num.num',sK0) != s('type.num.num',X0)) )),
  inference(superposition,[],[f222,f169])).
fof(f3318,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f222,f168])).
fof(f3311,plain,(
  ( ! [X6,X7] : (s('type.num.num',X7) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))) )),
  inference(superposition,[],[f220,f168])).
fof(f3309,plain,(
  ( ! [X2,X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',sK0)))) | s('type.num.num',X2) != s('type.num.num',X3)) )),
  inference(superposition,[],[f220,f169])).
fof(f3308,plain,(
  ( ! [X0,X1] : (s('type.num.num',X1) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',sK0)))) | ~p(s(bool,sK3(X0,X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) )),
  inference(superposition,[],[f220,f161])).
fof(f3306,plain,(
  ( ! [X45,X43,X41,X44,X42,X40] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X43),s('type.num.num',X44),s('type.num.num',X45))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X40),s('type.num.num',X41),s('type.num.num',X42))),s('type.num.num',sK0)))) | s('type.num.num',X42) != s('type.num.num',X45) | s('type.num.num',X41) != s('type.num.num',X44) | s(bool,X40) != s(bool,X43)) )),
  inference(superposition,[],[f220,f160])).
fof(f3305,plain,(
  ( ! [X39,X37,X35,X38,X36,X34] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X37),s('type.num.num',X38),s('type.num.num',X39))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X34),s('type.num.num',X35),s('type.num.num',X36))),s('type.num.num',sK0)))) | s('type.num.num',X36) != s('type.num.num',X39) | p(s(bool,X37)) | s(bool,X34) != s(bool,X37)) )),
  inference(superposition,[],[f220,f159])).
fof(f3304,plain,(
  ( ! [X30,X28,X33,X31,X29,X32] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X31),s('type.num.num',X32),s('type.num.num',X33))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X28),s('type.num.num',X29),s('type.num.num',X30))),s('type.num.num',sK0)))) | s('type.num.num',X30) != s('type.num.num',X33) | p(s(bool,X28)) | s(bool,X28) != s(bool,X31)) )),
  inference(superposition,[],[f220,f159])).
fof(f3303,plain,(
  ( ! [X26,X24,X23,X27,X25,X22] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X25),s('type.num.num',X26),s('type.num.num',X27))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X22),s('type.num.num',X23),s('type.num.num',X24))),s('type.num.num',sK0)))) | ~p(s(bool,X25)) | s('type.num.num',X23) != s('type.num.num',X26) | s(bool,X22) != s(bool,X25)) )),
  inference(superposition,[],[f220,f158])).
fof(f3302,plain,(
  ( ! [X21,X19,X17,X20,X18,X16] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.bool.COND_3'(s(bool,X19),s('type.num.num',X20),s('type.num.num',X21))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,X16),s('type.num.num',X17),s('type.num.num',X18))),s('type.num.num',sK0)))) | ~p(s(bool,X16)) | s('type.num.num',X17) != s('type.num.num',X20) | s(bool,X16) != s(bool,X19)) )),
  inference(superposition,[],[f220,f158])).
fof(f3299,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',X11))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.bool.COND_3'(s(bool,sK3(X10,X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X10),s('type.num.num',X11))))))),s('type.num.num',sK0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) )),
  inference(superposition,[],[f220,f163])).
fof(f3296,plain,(
  ( ! [X6,X7] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num',sK0)))) | s('type.num.num',X6) != s('type.num.num',X7)) )),
  inference(superposition,[],[f220,f169])).
fof(f3295,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',X5))))),s('type.num.num',sK0)))) | ~p(s(bool,sK3(X4,X5))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X5))))) )),
  inference(superposition,[],[f220,f161])).
fof(f3291,plain,(
  ( ! [X4,X5] : (s('type.num.num',X4) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f220,f168])).
fof(f3289,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK0)))) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X0)) )),
  inference(superposition,[],[f220,f169])).
fof(f3287,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(superposition,[],[f220,f168])).
fof(f3182,plain,(
  ( ! [X7] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),X7)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X7))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))))) ) | $spl0),
  inference(superposition,[],[f163,f258])).
fof(f3228,plain,(
  ( ! [X6] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X6,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))))))) ) | $spl0),
  inference(forward_demodulation,[],[f3181,f258])).
fof(f3181,plain,(
  ( ! [X6] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X6,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))))) ) | $spl0),
  inference(superposition,[],[f163,f258])).
fof(f3179,plain,(
  ( ! [X4] : (s('type.num.num',X4) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X4))))) | p(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),X4))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4))))) ) | $spl0),
  inference(superposition,[],[f162,f258])).
fof(f3226,plain,(
  ( ! [X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | p(s(bool,sK3(X3,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))))) ) | $spl0),
  inference(forward_demodulation,[],[f3178,f258])).
fof(f3178,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | p(s(bool,sK3(X3,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))))) ) | $spl0),
  inference(superposition,[],[f162,f258])).
fof(f3176,plain,(
  ( ! [X1] : (s('type.num.num',X1) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X1))))) | ~p(s(bool,sK3('const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)),X1))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))))) ) | $spl0),
  inference(superposition,[],[f161,f258])).
fof(f3224,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | ~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))))) ) | $spl0),
  inference(forward_demodulation,[],[f3175,f258])).
fof(f3175,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | ~p(s(bool,sK3(X0,'const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))))) ) | $spl0),
  inference(superposition,[],[f161,f258])).
fof(f3223,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | $spl0),
  inference(subsumption_resolution,[],[f3222,f119])).
fof(f3222,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | $spl0),
  inference(forward_demodulation,[],[f3172,f205])).
fof(f3172,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl0),
  inference(superposition,[],[f163,f258])).
fof(f3221,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | $spl0),
  inference(subsumption_resolution,[],[f3220,f119])).
fof(f3220,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | $spl0),
  inference(forward_demodulation,[],[f3171,f205])).
fof(f3171,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl0),
  inference(superposition,[],[f162,f258])).
fof(f3166,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl0),
  inference(superposition,[],[f258,f168])).
fof(f3164,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2)))))) ) | $spl0),
  inference(superposition,[],[f258,f168])).
fof(f3162,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num',sK2) != s('type.num.num',X0)) ) | $spl0),
  inference(superposition,[],[f258,f169])).
fof(f3160,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl0),
  inference(superposition,[],[f258,f168])).
fof(f3118,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2879,f168])).
fof(f3116,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1)))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2879,f168])).
fof(f3114,plain,(
  ( ! [X0] : (s('type.num.num',sK1) != s('type.num.num',X0) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2879,f169])).
fof(f3112,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f2879,f168])).
fof(f3104,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))) ) | $spl5),
  inference(superposition,[],[f818,f169])).
fof(f3101,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))) ) | $spl5),
  inference(superposition,[],[f818,f168])).
fof(f3099,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',X0))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl5),
  inference(superposition,[],[f818,f168])).
fof(f3097,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl5),
  inference(superposition,[],[f818,f168])).
fof(f636,plain,(
  s(bool,'T_0') != s(bool,sK3(sK2,sK0)) | $spl5),
  inference(resolution,[],[f629,f126])).
fof(f818,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) | $spl5),
  inference(subsumption_resolution,[],[f711,f121])).
fof(f121,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f33])).
fof(f33,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.arithmetic.SUC_MOD',aHL_FALSITY)).
fof(f711,plain,(
  p(s(bool,'F_0')) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) | $spl5),
  inference(superposition,[],[f203,f638])).
fof(f2879,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f2878,f119])).
fof(f2878,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f2877,f205])).
fof(f2877,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f820,f2873])).
fof(f820,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl5),
  inference(subsumption_resolution,[],[f714,f121])).
fof(f714,plain,(
  p(s(bool,'F_0')) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl5),
  inference(superposition,[],[f162,f638])).
fof(f2873,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f2872,f119])).
fof(f2872,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f2871,f205])).
fof(f2871,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f822,f258])).
fof(f822,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl5),
  inference(forward_demodulation,[],[f715,f154])).
fof(f715,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl5),
  inference(superposition,[],[f163,f638])).
fof(f2861,plain,(
  s(bool,'F_0') = s(bool,sK3(sK2,sK0)) | $spl5),
  inference(trivial_inequality_removal,[],[f684])).
fof(f684,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,sK3(sK2,sK0)) | $spl5),
  inference(superposition,[],[f636,f124])).
fof(f2862,plain,(
  ~p(s(bool,sK3(sK2,sK0))) | $spl5),
  inference(trivial_inequality_removal,[],[f686])).
fof(f686,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ~p(s(bool,sK3(sK2,sK0))) | $spl5),
  inference(superposition,[],[f636,f127])).
fof(f577,plain,(
  ( ! [X19,X18] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X18),s('type.num.num',sK0))))) | s(bool,X19) != s(bool,sK3(X18,sK0)) | ~p(s(bool,X19))) )),
  inference(resolution,[],[f202,f191])).
fof(f575,plain,(
  ( ! [X14,X15] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X14),s('type.num.num',sK0))))) | s(bool,X15) = s(bool,sK3(X14,sK0)) | p(s(bool,X15))) )),
  inference(resolution,[],[f202,f164])).
fof(f574,plain,(
  ( ! [X12,X10,X8,X7,X13,X11,X9] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',sK0))))) | s(X8,'const.bool.COND_3'(s(bool,X9),s(X8,X10),s(X8,X11))) = s(X8,'const.bool.COND_3'(s(bool,sK3(X7,sK0)),s(X8,X12),s(X8,X13))) | s(X8,X11) != s(X8,X13) | s(bool,X9) != s(bool,sK3(X7,sK0))) )),
  inference(resolution,[],[f202,f159])).
fof(f571,plain,(
  ( ! [X4] : (s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3(X4,sK0))) )),
  inference(resolution,[],[f202,f131])).
fof(f569,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK3(X2,sK0)) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0)))))) )),
  inference(resolution,[],[f202,f126])).
fof(f202,plain,(
  ( ! [X0] : (~p(s(bool,sK3(X0,sK0))) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X0),s('type.num.num',sK0)))))) )),
  inference(resolution,[],[f116,f161])).
fof(f608,plain,(
  ( ! [X19,X20] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X19),s('type.num.num',sK0))))) | s(bool,X20) != s(bool,sK3(X19,sK0)) | p(s(bool,X20))) )),
  inference(resolution,[],[f203,f191])).
fof(f606,plain,(
  ( ! [X15,X16] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X15),s('type.num.num',sK0))))) | s(bool,X16) = s(bool,sK3(X15,sK0)) | ~p(s(bool,X16))) )),
  inference(resolution,[],[f203,f167])).
fof(f605,plain,(
  ( ! [X14,X12,X10,X8,X13,X11,X9] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',sK0))))) | s(X9,'const.bool.COND_3'(s(bool,X10),s(X9,X11),s(X9,X12))) = s(X9,'const.bool.COND_3'(s(bool,sK3(X8,sK0)),s(X9,X13),s(X9,X14))) | s(X9,X11) != s(X9,X13) | s(bool,X10) != s(bool,sK3(X8,sK0))) )),
  inference(resolution,[],[f203,f158])).
fof(f601,plain,(
  ( ! [X4] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X4),s('type.num.num',sK0))))) | s(bool,'F_0') != s(bool,sK3(X4,sK0))) )),
  inference(resolution,[],[f203,f130])).
fof(f130,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f93])).
fof(f599,plain,(
  ( ! [X2] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,sK3(X2,sK0))) )),
  inference(resolution,[],[f203,f127])).
fof(f203,plain,(
  ( ! [X1] : (p(s(bool,sK3(X1,sK0))) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',sK0)))))) )),
  inference(resolution,[],[f116,f162])).
fof(f204,plain,(
  ( ! [X2] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(X2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0)))))))) )),
  inference(resolution,[],[f116,f163])).
fof(f2227,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(subsumption_resolution,[],[f2226,f119])).
fof(f2226,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl6),
  inference(forward_demodulation,[],[f2176,f205])).
fof(f2176,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl6),
  inference(superposition,[],[f163,f1608])).
fof(f2225,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0))) | $spl6),
  inference(subsumption_resolution,[],[f2224,f119])).
fof(f2224,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0))) | $spl6),
  inference(forward_demodulation,[],[f2175,f205])).
fof(f2175,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl6),
  inference(superposition,[],[f162,f1608])).
fof(f2223,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0))) | $spl6),
  inference(subsumption_resolution,[],[f2222,f119])).
fof(f2222,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0))) | $spl6),
  inference(forward_demodulation,[],[f2174,f205])).
fof(f2174,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl6),
  inference(superposition,[],[f161,f1608])).
fof(f2170,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',X0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl6),
  inference(superposition,[],[f1608,f168])).
fof(f2168,plain,(
  ( ! [X2] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1)))))) ) | $spl6),
  inference(superposition,[],[f1608,f168])).
fof(f2166,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num',sK1) != s('type.num.num',X0)) ) | $spl6),
  inference(superposition,[],[f1608,f169])).
fof(f2164,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl6),
  inference(superposition,[],[f1608,f168])).
fof(f2141,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f2140,f119])).
fof(f2140,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f2090,f205])).
fof(f2090,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl6)),
  inference(superposition,[],[f163,f1629])).
fof(f2139,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f2138,f119])).
fof(f2138,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f2089,f205])).
fof(f2089,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl6)),
  inference(superposition,[],[f162,f1629])).
fof(f2137,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f2136,f119])).
fof(f2136,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f2088,f205])).
fof(f2088,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl6)),
  inference(superposition,[],[f161,f1629])).
fof(f1680,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | $spl6),
  inference(superposition,[],[f1487,f168])).
fof(f1628,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | $spl6),
  inference(forward_demodulation,[],[f118,f1608])).
fof(f1608,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | $spl6),
  inference(forward_demodulation,[],[f1606,f153])).
fof(f1606,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl6),
  inference(backward_demodulation,[],[f1577,f1540])).
fof(f1487,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | $spl6),
  inference(superposition,[],[f150,f631])).
fof(f150,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6] : ~s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))),
  file('thm.arithmetic.SUC_MOD','thm.arithmetic.SUC_NOT')).
fof(f1540,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(sK1,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl6),
  inference(subsumption_resolution,[],[f1539,f119])).
fof(f1539,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(sK1,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | $spl6),
  inference(forward_demodulation,[],[f1482,f205])).
fof(f1482,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(sK1,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl6),
  inference(superposition,[],[f163,f631])).
fof(f1538,plain,(
  p(s(bool,sK3(sK1,sK0))) | $spl6),
  inference(subsumption_resolution,[],[f1537,f119])).
fof(f1537,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,sK3(sK1,sK0))) | $spl6),
  inference(forward_demodulation,[],[f1532,f205])).
fof(f1532,plain,(
  p(s(bool,sK3(sK1,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl6),
  inference(trivial_inequality_removal,[],[f1481])).
fof(f1481,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK0) | p(s(bool,sK3(sK1,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl6),
  inference(superposition,[],[f162,f631])).
fof(f850,plain,(
  ( ! [X52,X48,X53,X51,X49] : (s(X49,X51) = s(X49,'const.bool.COND_3'(s(bool,X48),s(X49,X52),s(X49,X53))) | s(X49,X51) != s(X49,X53) | p(s(bool,X48))) ) | $spl5),
  inference(forward_demodulation,[],[f849,f154])).
fof(f849,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(X49,'const.bool.COND_3'(s(bool,'F_0'),s(X49,X50),s(X49,X51))) = s(X49,'const.bool.COND_3'(s(bool,X48),s(X49,X52),s(X49,X53))) | s(X49,X51) != s(X49,X53) | p(s(bool,X48))) ) | $spl5),
  inference(forward_demodulation,[],[f848,f638])).
fof(f848,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(X49,'const.bool.COND_3'(s(bool,X48),s(X49,X52),s(X49,X53))) = s(X49,'const.bool.COND_3'(s(bool,sK3(sK2,sK0)),s(X49,X50),s(X49,X51))) | s(X49,X51) != s(X49,X53) | p(s(bool,X48))) ) | $spl5),
  inference(subsumption_resolution,[],[f744,f131])).
fof(f744,plain,(
  ( ! [X52,X50,X48,X53,X51,X49] : (s(bool,'F_0') != s(bool,X48) | s(X49,'const.bool.COND_3'(s(bool,X48),s(X49,X52),s(X49,X53))) = s(X49,'const.bool.COND_3'(s(bool,sK3(sK2,sK0)),s(X49,X50),s(X49,X51))) | s(X49,X51) != s(X49,X53) | p(s(bool,X48))) ) | $spl5),
  inference(superposition,[],[f159,f638])).
fof(f844,plain,(
  ( ! [X30,X35,X33,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(X30,X35) = s(X30,'const.bool.COND_3'(s(bool,X31),s(X30,X32),s(X30,X33))) | s(X30,X33) != s(X30,X35)) ) | $spl5),
  inference(forward_demodulation,[],[f843,f638])).
fof(f843,plain,(
  ( ! [X30,X35,X33,X31,X32] : (s(X30,X35) = s(X30,'const.bool.COND_3'(s(bool,X31),s(X30,X32),s(X30,X33))) | s(X30,X33) != s(X30,X35) | s(bool,X31) != s(bool,sK3(sK2,sK0))) ) | $spl5),
  inference(subsumption_resolution,[],[f842,f121])).
fof(f842,plain,(
  ( ! [X30,X35,X33,X31,X32] : (p(s(bool,'F_0')) | s(X30,X35) = s(X30,'const.bool.COND_3'(s(bool,X31),s(X30,X32),s(X30,X33))) | s(X30,X33) != s(X30,X35) | s(bool,X31) != s(bool,sK3(sK2,sK0))) ) | $spl5),
  inference(forward_demodulation,[],[f841,f638])).
fof(f841,plain,(
  ( ! [X30,X35,X33,X31,X32] : (s(X30,X35) = s(X30,'const.bool.COND_3'(s(bool,X31),s(X30,X32),s(X30,X33))) | s(X30,X33) != s(X30,X35) | p(s(bool,sK3(sK2,sK0))) | s(bool,X31) != s(bool,sK3(sK2,sK0))) ) | $spl5),
  inference(forward_demodulation,[],[f741,f154])).
fof(f741,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(X30,'const.bool.COND_3'(s(bool,'F_0'),s(X30,X34),s(X30,X35))) = s(X30,'const.bool.COND_3'(s(bool,X31),s(X30,X32),s(X30,X33))) | s(X30,X33) != s(X30,X35) | p(s(bool,sK3(sK2,sK0))) | s(bool,X31) != s(bool,sK3(sK2,sK0))) ) | $spl5),
  inference(superposition,[],[f159,f638])).
fof(f669,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'F_0') != s(bool,X1) | s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) | s(X0,X3) != s(X0,X5)) ) | $spl5),
  inference(forward_demodulation,[],[f668,f638])).
fof(f668,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X4),s(X0,X5))) = s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) | s(X0,X3) != s(X0,X5) | s(bool,X1) != s(bool,sK3(sK2,sK0))) ) | $spl5),
  inference(forward_demodulation,[],[f641,f638])).
fof(f641,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,'const.bool.COND_3'(s(bool,X1),s(X0,X2),s(X0,X3))) = s(X0,'const.bool.COND_3'(s(bool,sK3(sK2,sK0)),s(X0,X4),s(X0,X5))) | s(X0,X3) != s(X0,X5) | s(bool,X1) != s(bool,sK3(sK2,sK0))) ) | $spl5),
  inference(resolution,[],[f629,f159])).
fof(f638,plain,(
  s(bool,'F_0') = s(bool,sK3(sK2,sK0)) | $spl5),
  inference(resolution,[],[f629,f131])).
fof(f412,plain,(
  ( ! [X26,X23,X21,X25,X22] : (~p(s(bool,X21)) | s(X22,X23) = s(X22,'const.bool.COND_3'(s(bool,X21),s(X22,X25),s(X22,X26))) | s(X22,X23) != s(X22,X25)) )),
  inference(forward_demodulation,[],[f411,f153])).
fof(f411,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(X22,'const.bool.COND_3'(s(bool,'T_0'),s(X22,X23),s(X22,X24))) = s(X22,'const.bool.COND_3'(s(bool,X21),s(X22,X25),s(X22,X26))) | ~p(s(bool,X21)) | s(X22,X23) != s(X22,X25)) )),
  inference(forward_demodulation,[],[f410,f205])).
fof(f410,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(X22,'const.bool.COND_3'(s(bool,X21),s(X22,X25),s(X22,X26))) = s(X22,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))),s(X22,X23),s(X22,X24))) | ~p(s(bool,X21)) | s(X22,X23) != s(X22,X25)) )),
  inference(subsumption_resolution,[],[f314,f127])).
fof(f314,plain,(
  ( ! [X26,X24,X23,X21,X25,X22] : (s(bool,'T_0') != s(bool,X21) | s(X22,'const.bool.COND_3'(s(bool,X21),s(X22,X25),s(X22,X26))) = s(X22,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))),s(X22,X23),s(X22,X24))) | ~p(s(bool,X21)) | s(X22,X23) != s(X22,X25)) )),
  inference(superposition,[],[f158,f205])).
fof(f406,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(bool,'T_0') != s(bool,X4) | s(X3,X7) = s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) | s(X3,X5) != s(X3,X7)) )),
  inference(forward_demodulation,[],[f405,f205])).
fof(f405,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X3,X7) = s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) | s(X3,X5) != s(X3,X7) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(subsumption_resolution,[],[f404,f119])).
fof(f404,plain,(
  ( ! [X6,X4,X7,X5,X3] : (~p(s(bool,'T_0')) | s(X3,X7) = s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) | s(X3,X5) != s(X3,X7) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(forward_demodulation,[],[f403,f205])).
fof(f403,plain,(
  ( ! [X6,X4,X7,X5,X3] : (s(X3,X7) = s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s(X3,X5) != s(X3,X7) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(forward_demodulation,[],[f311,f153])).
fof(f311,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(X3,'const.bool.COND_3'(s(bool,'T_0'),s(X3,X7),s(X3,X8))) = s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | s(X3,X5) != s(X3,X7) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f158,f205])).
fof(f279,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(superposition,[],[f205,f168])).
fof(f277,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK0)))) )),
  inference(superposition,[],[f205,f168])).
fof(f222,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(superposition,[],[f116,f168])).
fof(f220,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',sK0))))) )),
  inference(superposition,[],[f116,f168])).
fof(f242,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(bool,'T_0') != s(bool,X4) | s(X3,'const.bool.COND_3'(s(bool,'T_0'),s(X3,X7),s(X3,X8))) = s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) | s(X3,X5) != s(X3,X7)) )),
  inference(forward_demodulation,[],[f241,f205])).
fof(f241,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(X3,'const.bool.COND_3'(s(bool,'T_0'),s(X3,X7),s(X3,X8))) = s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) | s(X3,X5) != s(X3,X7) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(forward_demodulation,[],[f211,f205])).
fof(f211,plain,(
  ( ! [X6,X4,X8,X7,X5,X3] : (s(X3,'const.bool.COND_3'(s(bool,X4),s(X3,X5),s(X3,X6))) = s(X3,'const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))),s(X3,X7),s(X3,X8))) | s(X3,X5) != s(X3,X7) | s(bool,X4) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) )),
  inference(resolution,[],[f116,f158])).
fof(f239,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(forward_demodulation,[],[f207,f205])).
fof(f207,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))),
  inference(resolution,[],[f116,f130])).
fof(f205,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))),
  inference(resolution,[],[f116,f127])).
fof(f116,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) & (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))) & (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f89])).
fof(f89,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0)))) & (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0)))))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & ((s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0)))) & (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) & (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X0))) <~> s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) => (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X1),s('type.num.num',X0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X2),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) => (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X6))) <=> s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)))) => (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X7))),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X6))) <=> s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X7),s('type.num.num',X6))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',X8),s('type.num.num',X6)))))),
  file('thm.arithmetic.SUC_MOD','thm.arithmetic.SUC_MOD')).
fof(f118,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f90])).
fof(f5533,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1498,f1608,f1628,f1680,f2137,f2139,f2141,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183,f2180,f1543,f2146,f2106,f1526,f1524,f2143,f2102,f2100,f2097,f2094,f5527,f1504,f5529,f1577,f1500,f5532])).
fof(f5530,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1498,f1500,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183,f2180,f1543,f2146,f2106,f1526,f1524,f2143,f2102,f2100,f2097,f2094,f5527,f1504,f5529])).
fof(f5528,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(global_subsumption,[],[f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1498,f1500,f1504,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183,f2180,f1543,f2146,f2106,f1526,f1524,f2143,f2102,f2100,f2097,f2094,f5527])).
fof(f5520,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f1608,f4146])).
fof(f5516,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5504,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5504,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(backward_demodulation,[],[f5492,f4638])).
fof(f4638,plain,(
  s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4637,f2873])).
fof(f4637,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4614,f3945])).
fof(f3945,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl13)),
  inference(backward_demodulation,[],[f3942,f3696])).
fof(f3696,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f3695,f119])).
fof(f3695,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) | ($spl0 | $spl13)),
  inference(forward_demodulation,[],[f3694,f205])).
fof(f3694,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl13)),
  inference(forward_demodulation,[],[f3693,f258])).
fof(f3693,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl13)),
  inference(forward_demodulation,[],[f3586,f154])).
fof(f3586,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl13)),
  inference(superposition,[],[f163,f3569])).
fof(f3569,plain,(
  s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)) | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f3448,f3568])).
fof(f3568,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f3221,f3565])).
fof(f3942,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f3941,f119])).
fof(f3941,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) | ($spl0 | $spl13)),
  inference(forward_demodulation,[],[f3940,f205])).
fof(f3940,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl13)),
  inference(forward_demodulation,[],[f3939,f3696])).
fof(f3939,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl13)),
  inference(forward_demodulation,[],[f3836,f154])).
fof(f3836,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl13)),
  inference(superposition,[],[f163,f3818])).
fof(f3818,plain,(
  s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0)) | ($spl0 | $spl13)),
  inference(trivial_inequality_removal,[],[f3811])).
fof(f3811,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK0) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0)) | ($spl0 | $spl13)),
  inference(superposition,[],[f3568,f571])).
fof(f4614,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(superposition,[],[f4026,f2873])).
fof(f4026,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num',X0)),sK0))) ) | ($spl0 | $spl13)),
  inference(superposition,[],[f3946,f168])).
fof(f3946,plain,(
  s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0)) | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f3841,f121])).
fof(f3841,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK1)),sK0)) | ($spl0 | $spl13)),
  inference(superposition,[],[f126,f3818])).
fof(f5492,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl6 | $spl13)),
  inference(forward_demodulation,[],[f5491,f154])).
fof(f5491,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl0 | $spl6 | $spl13)),
  inference(forward_demodulation,[],[f2227,f3818])).
fof(f5515,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5503,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5503,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(backward_demodulation,[],[f5492,f4597])).
fof(f4597,plain,(
  s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4596,f2873])).
fof(f4596,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4573,f3945])).
fof(f4573,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(superposition,[],[f3779,f2873])).
fof(f3779,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num',X0)),sK0))) ) | ($spl0 | $spl13)),
  inference(superposition,[],[f3699,f168])).
fof(f3699,plain,(
  s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)) | ($spl0 | $spl13)),
  inference(subsumption_resolution,[],[f3591,f121])).
fof(f3591,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') != s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0)) | ($spl0 | $spl13)),
  inference(superposition,[],[f126,f3569])).
fof(f5514,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5502,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5502,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(backward_demodulation,[],[f5492,f4556])).
fof(f4556,plain,(
  s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4555,f2873])).
fof(f4555,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4532,f3945])).
fof(f4532,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(superposition,[],[f3822,f2873])).
fof(f3822,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',X0)),sK0))) ) | ($spl0 | $spl13)),
  inference(superposition,[],[f3818,f168])).
fof(f5513,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5501,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5501,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(backward_demodulation,[],[f5492,f4515])).
fof(f4515,plain,(
  s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0)))),sK0)) | s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4514,f2873])).
fof(f4514,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4491,f3945])).
fof(f4491,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))),sK0)) | ($spl0 | $spl5 | $spl13)),
  inference(superposition,[],[f3572,f2873])).
fof(f3572,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s(bool,'F_0') = s(bool,sK3('const.num.SUC_1'(s('type.num.num',X0)),sK0))) ) | ($spl0 | $spl13)),
  inference(superposition,[],[f3569,f168])).
fof(f5512,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5499,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5499,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X11))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X11))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) ) | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(backward_demodulation,[],[f5492,f4223])).
fof(f4223,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))),s('type.num.num',X11))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X11))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) ) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4165,f3945])).
fof(f4165,plain,(
  ( ! [X11] : (s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))),s('type.num.num',X11))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3('const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))),X11)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',X11))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X11))))) ) | ($spl0 | $spl5)),
  inference(superposition,[],[f163,f2873])).
fof(f5511,plain,(
  $false | ($spl0 | $spl3 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5510,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5510,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl3 | $spl5 | $spl6 | $spl13)),
  inference(backward_demodulation,[],[f5498,f5458])).
fof(f5458,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl3 | $spl5)),
  inference(forward_demodulation,[],[f5435,f2873])).
fof(f5435,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))) | ($spl0 | $spl3 | $spl5)),
  inference(superposition,[],[f2255,f2873])).
fof(f2255,plain,(
  ( ! [X1] : (s('type.num.num',X1) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))) ) | $spl3),
  inference(superposition,[],[f259,f168])).
fof(f259,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | $spl3),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f5498,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | ($spl0 | $spl6 | $spl13)),
  inference(backward_demodulation,[],[f5492,f3945])).
fof(f5509,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5493,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5493,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f4679,f5492])).
fof(f4679,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4678,f2873])).
fof(f4678,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4655,f3945])).
fof(f4655,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))),s('type.num.num',sK0))) | ($spl0 | $spl5)),
  inference(superposition,[],[f277,f2873])).
fof(f5508,plain,(
  $false | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(global_subsumption,[],[f5494,f118,f117,f116,f205,f239,f242,f220,f222,f277,f279,f406,f412,f629,f638,f669,f844,f850,f631,f1538,f1540,f1487,f1543,f1498,f1500,f1504,f1524,f1526,f1577,f1608,f1628,f1680,f2137,f2139,f2141,f2143,f2094,f2097,f2146,f2100,f2102,f2106,f2164,f2166,f2168,f2170,f2223,f2225,f2227,f2180,f204,f203,f599,f601,f605,f606,f608,f202,f569,f571,f574,f575,f577,f258,f2862,f2861,f2873,f2879,f818,f636,f3097,f3099,f3101,f3104,f3112,f3114,f3116,f3118,f3160,f3162,f3164,f3166,f3219,f3221,f3223,f3224,f3176,f3226,f3179,f3228,f3182,f3287,f3289,f3291,f3295,f3296,f3299,f3302,f3303,f3304,f3305,f3306,f3308,f3309,f3311,f3318,f3320,f3322,f3345,f3327,f3330,f3333,f3334,f3335,f3336,f3337,f3339,f3340,f3342,f3371,f3373,f3374,f3377,f3378,f3379,f3380,f3381,f3382,f3383,f3384,f3386,f3387,f3389,f3392,f3394,f3396,f3397,f3400,f3403,f3404,f3405,f3406,f3407,f3430,f3432,f3433,f3434,f3436,f3437,f3438,f3439,f3440,f3441,f3442,f3443,f3444,f3446,f3448,f3449,f3451,f3453,f3524,f3459,f3461,f3465,f3466,f3518,f3471,f3472,f3526,f3474,f3475,f3528,f3477,f3478,f3530,f3480,f3484,f3486,f3488,f3517,f3491,f3492,f3499,f3500,f3501,f3502,f3503,f3504,f3507,f3508,f3509,f3511,f3513,f3515,f4070,f4072,f4073,f4074,f4076,f4077,f4078,f4079,f4080,f4081,f4082,f4083,f4084,f4086,f4089,f4093,f4095,f4099,f4101,f4102,f4103,f4105,f4106,f4107,f4108,f4109,f4110,f4111,f4112,f4113,f4115,f4118,f4122,f4124,f4127,f4129,f4131,f4135,f4146,f4148,f4152,f4153,f4218,f4159,f4220,f4162,f4222,f4177,f4178,f4179,f4254,f4256,f4258,f4260,f4262,f4264,f4266,f4268,f4269,f4271,f4274,f4277,f4278,f4279,f4280,f4281,f4645,f4647,f4649,f4653,f4657,f4660,f4663,f4664,f4665,f4666,f4667,f4669,f4670,f4671,f4673,f4675,f4686,f4688,f4690,f4720,f4719,f4698,f4701,f4704,f4705,f4706,f4707,f4708,f4710,f4711,f4712,f4714,f4716,f5086,f5089,f5090,f5092,f5098,f5108,f5109,f5134,f5135,f5136,f5137,f5138,f5153,f5154,f5238,f5239,f5240,f5241,f5159,f5161,f5162,f5242,f5244,f5246,f5248,f5250,f5173,f5252,f5176,f5183,f5184,f5185,f5260,f5188,f5189,f5191,f5193,f5261,f5201,f5203,f5204,f5215,f5218,f5225,f5226,f5227,f5230,f5233,f5292,f5294,f5297,f5298,f5300,f5316,f5317,f5354,f5355,f5342,f5343,f5344,f5345,f5346,f5360,f5406,f5424,f5429,f3121,f1495,f2128,f5487,f1492,f2214,f2084,f2126,f2082,f2192,f2080,f1483,f1629,f5488,f1477,f2078,f2212,f1473,f2232,f1471,f1555,f2188,f2186,f1549,f2229,f2183])).
fof(f5494,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0)))) | ($spl0 | $spl5 | $spl6 | $spl13)),
  inference(subsumption_resolution,[],[f4234,f5492])).
fof(f4234,plain,(
  p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))),s('type.num.num',sK0)))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4233,f2873])).
fof(f4233,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))))),s('type.num.num',sK0))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))),s('type.num.num',sK0)))) | ($spl0 | $spl5 | $spl13)),
  inference(forward_demodulation,[],[f4175,f3945])).
fof(f4175,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))),s('type.num.num',sK0)))) | ($spl0 | $spl5)),
  inference(superposition,[],[f220,f2873])).
fof(f5490,plain,(
  $false | ($spl0 | $spl5 | $spl6)),
  inference(subsumption_resolution,[],[f5489,f4146])).
fof(f5489,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl6)),
  inference(backward_demodulation,[],[f1629,f258])).
fof(f5474,plain,(
  $false | ($spl0 | $spl3 | $spl5 | $spl7)),
  inference(subsumption_resolution,[],[f5473,f2864])).
fof(f2864,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | $spl7),
  inference(subsumption_resolution,[],[f2863,f119])).
fof(f2863,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | $spl7),
  inference(forward_demodulation,[],[f2763,f205])).
fof(f2763,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl7),
  inference(forward_demodulation,[],[f2660,f154])).
fof(f2660,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl7),
  inference(superposition,[],[f163,f2588])).
fof(f2588,plain,(
  s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | $spl7),
  inference(resolution,[],[f2583,f131])).
fof(f2583,plain,(
  ~p(s(bool,sK3(sK1,sK0))) | $spl7),
  inference(subsumption_resolution,[],[f2582,f119])).
fof(f2582,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK3(sK1,sK0))) | $spl7),
  inference(forward_demodulation,[],[f2581,f205])).
fof(f2581,plain,(
  ~p(s(bool,sK3(sK1,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl7),
  inference(trivial_inequality_removal,[],[f2576])).
fof(f2576,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK0) | ~p(s(bool,sK3(sK1,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl7),
  inference(superposition,[],[f630,f161])).
fof(f630,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | $spl7),
  inference(cnf_transformation,[],[f630_D])).
fof(f630_D,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f5473,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl3 | $spl5)),
  inference(forward_demodulation,[],[f5455,f2873])).
fof(f5455,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) | $spl3),
  inference(equality_resolution,[],[f2255])).
fof(f3565,plain,(
  ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | $spl13),
  inference(cnf_transformation,[],[f3565_D])).
fof(f3565_D,plain,(
  ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f3567,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | $spl14),
  inference(cnf_transformation,[],[f3567_D])).
fof(f3567_D,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f3219,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | $spl0),
  inference(subsumption_resolution,[],[f3218,f119])).
fof(f3218,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | $spl0),
  inference(forward_demodulation,[],[f3170,f205])).
fof(f3170,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))))) | ~p(s(bool,sK3('const.num.SUC_1'(s('type.num.num',sK2)),sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl0),
  inference(superposition,[],[f161,f258])).
fof(f3531,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,sK3(X3,sK0))) ) | $spl8),
  inference(cnf_transformation,[],[f3531_D])).
fof(f3531_D,plain,(
  ( ! [X3] : (s(bool,'F_0') = s(bool,sK3(X3,sK0))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f3533,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | $spl11),
  inference(cnf_transformation,[],[f3533_D])).
fof(f3533_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f3459,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',sK0) | s(bool,'F_0') = s(bool,sK3(X3,sK0))) )),
  inference(superposition,[],[f149,f571])).
fof(f2767,plain,(
  $false | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f2766,f189])).
fof(f189,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f151])).
fof(f151,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X13] : ! [X14] : (s(X13,X14) = s(X13,X14) <=> p(s(bool,'T_0')))),
  file('thm.arithmetic.SUC_MOD','thm.bool.REFL_CLAUSE')).
fof(f2766,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl4 | $spl7)),
  inference(forward_demodulation,[],[f2765,f205])).
fof(f2765,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl4 | $spl7)),
  inference(subsumption_resolution,[],[f2764,f149])).
fof(f2764,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl0 | $spl4 | $spl7)),
  inference(forward_demodulation,[],[f2763,f2287])).
fof(f2287,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f258,f2286])).
fof(f2286,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | $spl4),
  inference(subsumption_resolution,[],[f2285,f119])).
fof(f2285,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | $spl4),
  inference(forward_demodulation,[],[f1796,f205])).
fof(f1796,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl4),
  inference(forward_demodulation,[],[f1693,f153])).
fof(f1693,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl4),
  inference(superposition,[],[f163,f1632])).
fof(f1632,plain,(
  s(bool,'T_0') = s(bool,sK3(sK2,sK0)) | $spl4),
  inference(resolution,[],[f628,f127])).
fof(f628,plain,(
  p(s(bool,sK3(sK2,sK0))) | $spl4),
  inference(cnf_transformation,[],[f628_D])).
fof(f628_D,plain,(
  p(s(bool,sK3(sK2,sK0))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f2282,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1577,f2259])).
fof(f2259,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,sK3(X2,sK0))) ) | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f569,f2257])).
fof(f2257,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))) ) | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f2256,f1483])).
fof(f2256,plain,(
  ( ! [X0] : (s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) | s('type.num.num',X0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0)))) ) | ($spl3 | $spl4)),
  inference(forward_demodulation,[],[f2254,f1630])).
fof(f1630,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))))) | $spl4),
  inference(resolution,[],[f628,f202])).
fof(f2254,plain,(
  ( ! [X0] : (s('type.num.num',X0) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) | s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0)))))) ) | $spl3),
  inference(superposition,[],[f259,f168])).
fof(f2281,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1632,f2259])).
fof(f2267,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f1630,f2257])).
fof(f2266,plain,(
  $false | ($spl3 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f631,f2257])).
fof(f1609,plain,(
  $false | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f1608,f1240])).
fof(f1240,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f1238,f261])).
fof(f261,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) != s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | $spl2),
  inference(subsumption_resolution,[],[f118,f260])).
fof(f1238,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f1236,f153])).
fof(f1236,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'T_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f1205,f1201])).
fof(f1201,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(sK2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0))) | ($spl2 | $spl6)),
  inference(backward_demodulation,[],[f631,f536])).
fof(f536,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(sK2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | $spl2),
  inference(subsumption_resolution,[],[f535,f119])).
fof(f535,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(sK2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | $spl2),
  inference(forward_demodulation,[],[f485,f205])).
fof(f485,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,sK3(sK2,sK0)),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl2),
  inference(superposition,[],[f163,f260])).
fof(f1205,plain,(
  s(bool,'T_0') = s(bool,sK3(sK2,sK0)) | ($spl2 | $spl6)),
  inference(resolution,[],[f1200,f127])).
fof(f1200,plain,(
  p(s(bool,sK3(sK2,sK0))) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f534,f631])).
fof(f534,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | p(s(bool,sK3(sK2,sK0))) | $spl2),
  inference(subsumption_resolution,[],[f533,f119])).
fof(f533,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | p(s(bool,sK3(sK2,sK0))) | $spl2),
  inference(forward_demodulation,[],[f484,f205])).
fof(f484,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | p(s(bool,sK3(sK2,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl2),
  inference(superposition,[],[f162,f260])).
fof(f1120,plain,(
  $false | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1119,f189])).
fof(f1119,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f1118,f205])).
fof(f1118,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1117,f261])).
fof(f1117,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f1116,f825])).
fof(f825,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f824,f119])).
fof(f824,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f823,f205])).
fof(f823,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f822,f260])).
fof(f1116,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f1013,f154])).
fof(f1013,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f163,f929])).
fof(f929,plain,(
  s(bool,'F_0') = s(bool,sK3(sK1,sK0)) | ($spl2 | $spl5)),
  inference(resolution,[],[f921,f131])).
fof(f921,plain,(
  ~p(s(bool,sK3(sK1,sK0))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f920,f119])).
fof(f920,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,sK3(sK1,sK0))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f919,f205])).
fof(f919,plain,(
  ~p(s(bool,sK3(sK1,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl2 | $spl5)),
  inference(trivial_inequality_removal,[],[f914])).
fof(f914,plain,(
  s('type.num.num',sK0) != s('type.num.num',sK0) | ~p(s(bool,sK3(sK1,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | ($spl2 | $spl5)),
  inference(superposition,[],[f632,f161])).
fof(f632,plain,(
  s('type.num.num',sK0) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f534,f629])).
fof(f629,plain,(
  ~p(s(bool,sK3(sK2,sK0))) | $spl5),
  inference(cnf_transformation,[],[f629_D])).
fof(f629_D,plain,(
  ~p(s(bool,sK3(sK2,sK0))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f631,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | $spl6),
  inference(cnf_transformation,[],[f631_D])).
fof(f631_D,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f532,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | ~p(s(bool,sK3(sK2,sK0))) | $spl2),
  inference(subsumption_resolution,[],[f531,f119])).
fof(f531,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | ~p(s(bool,sK3(sK2,sK0))) | $spl2),
  inference(forward_demodulation,[],[f483,f205])).
fof(f483,plain,(
  s('type.num.num',sK0) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))))) | ~p(s(bool,sK3(sK2,sK0))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',sK0)))) | $spl2),
  inference(superposition,[],[f161,f260])).
fof(f258,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) | $spl0),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f260,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | $spl2),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f117,plain,(
  s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK1),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num',sK2),s('type.num.num',sK0))) | s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK1))),s('type.num.num',sK0))) = s('type.num.num','const.arithmetic.MOD_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',sK2))),s('type.num.num',sK0)))),
  inference(cnf_transformation,[],[f90])).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 133
% Passive clauses: 916
% Generated clauses: 5491
% Final active clauses: 100
% Final passive clauses: 597
% Input formulas: 30
% Initial clauses: 69
% 
% Duplicate literals: 35
% Trivial inequalities: 84
% Fw subsumption resolutions: 396
% Bw subsumption resolutions: 22
% Fw demodulations: 1372
% Bw demodulations: 54
% Condensations: 103
% Global subsumptions: 12
% 
% Simple tautologies: 113
% Equational tautologies: 1
% Forward subsumptions: 2476
% Backward subsumptions: 22
% 
% Binary resolution: 239
% Forward superposition: 1467
% Backward superposition: 1527
% Self superposition: 1
% Equality resolution: 26
% 
% Split clauses: 4
% Split components: 8
% SAT solver clauses: 956
% SAT solver unit clauses: 5
% SAT solver binary clauses: 84
% 
% Sat splits: 4
% Sat splitting refutations: 22
% 
% Memory used [KB]: 4989
% Time elapsed: 0.260 s
% ------------------------------
% ------------------------------
% Success in time 16.638 s
