Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.util_prob
% (84826)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1060
% Passive clauses: 9359
% Generated clauses: 25031
% Final active clauses: 1049
% Final passive clauses: 7970
% Input formulas: 53
% Initial clauses: 218
% 
% Duplicate literals: 359
% Trivial inequalities: 81
% Fw subsumption resolutions: 2949
% Fw demodulations: 1639
% Bw demodulations: 10
% 
% Simple tautologies: 270
% Equational tautologies: 36
% Forward subsumptions: 9600
% 
% Binary resolution: 9861
% Factoring: 673
% Forward superposition: 2956
% Backward superposition: 5942
% Self superposition: 27
% Equality factoring: 26
% Equality resolution: 13
% 
% Split clauses: 678
% Split components: 1367
% SAT solver clauses: 1004
% SAT solver unit clauses: 24
% SAT solver binary clauses: 223
% 
% Sat splits: 678
% 
% TWLsolver clauses: 1004
% TWLsolver calls for satisfiability: 471
% 
% Memory used [KB]: 13432
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.util_prob
% (84837)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 578
% Passive clauses: 4042
% Generated clauses: 17889
% Final active clauses: 564
% Final passive clauses: 3439
% Input formulas: 53
% Initial clauses: 218
% 
% Duplicate literals: 254
% Trivial inequalities: 119
% Fw subsumption resolutions: 1227
% Fw demodulations: 1136
% Bw demodulations: 37
% Condensations: 509
% Global subsumptions: 17
% 
% Simple tautologies: 321
% Equational tautologies: 18
% Forward subsumptions: 5658
% 
% Binary resolution: 4999
% Factoring: 164
% Forward superposition: 2788
% Backward superposition: 2290
% Self superposition: 29
% Equality factoring: 28
% Equality resolution: 14
% 
% Split clauses: 92
% Split components: 189
% SAT solver clauses: 4810
% SAT solver unit clauses: 221
% SAT solver binary clauses: 93
% SAT solver learnt clauses: 34
% SAT solver learnt literals: 10
% 
% Sat splits: 92
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 4742
% TWLsolver calls for satisfiability: 2042
% 
% Memory used [KB]: 19573
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.util_prob
% (84850)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 3066
% Passive clauses: 12713
% Generated clauses: 148752
% Final active clauses: 2786
% Final passive clauses: 6158
% Input formulas: 53
% Initial clauses: 218
% Discarded non-redundant clauses: 4242
% 
% Duplicate literals: 6680
% Trivial inequalities: 184
% Fw subsumption resolutions: 6724
% Fw demodulations: 20973
% Bw demodulations: 840
% Condensations: 11907
% Global subsumptions: 43
% 
% Simple tautologies: 225
% Equational tautologies: 1119
% Forward subsumptions: 87160
% Fw demodulations to eq. taut.: 3
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 30208
% Factoring: 1
% Forward superposition: 31693
% Backward superposition: 38626
% Self superposition: 200
% Equality factoring: 158
% Equality resolution: 6
% 
% Split clauses: 320
% Split components: 649
% SAT solver clauses: 14095
% SAT solver unit clauses: 1617
% SAT solver binary clauses: 3802
% 
% Sat splits: 320
% Sat splitting refutations: 696
% 
% Memory used [KB]: 71896
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.util_prob
% (85102)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2281
% Passive clauses: 56550
% Generated clauses: 192075
% Final active clauses: 2205
% Final passive clauses: 52166
% Input formulas: 53
% Initial clauses: 208
% 
% Selected by SInE selection: 49
% SInE iterations: 1
% 
% Duplicate literals: 5840
% Trivial inequalities: 639
% Fw subsumption resolutions: 16212
% Fw demodulations: 3258
% Bw demodulations: 63
% 
% Simple tautologies: 5377
% Equational tautologies: 1482
% Forward subsumptions: 95655
% 
% Binary resolution: 147625
% Factoring: 309
% Forward superposition: 13792
% Backward superposition: 3057
% Self superposition: 11
% Equality factoring: 2
% Equality resolution: 38
% 
% Split clauses: 8992
% Split components: 20739
% SAT solver clauses: 9603
% SAT solver unit clauses: 43
% SAT solver binary clauses: 4015
% 
% Sat splits: 8992
% Sat splitting refutations: 45
% 
% Memory used [KB]: 49636
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.util_prob
% (85199)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 580
% Passive clauses: 3987
% Generated clauses: 15210
% Final active clauses: 567
% Final passive clauses: 3286
% Input formulas: 53
% Initial clauses: 218
% 
% Duplicate literals: 313
% Trivial inequalities: 56
% Fw subsumption resolutions: 697
% Fw demodulations: 1364
% Bw demodulations: 76
% Global subsumptions: 3
% 
% Simple tautologies: 353
% Equational tautologies: 50
% Forward subsumptions: 8026
% 
% Unit resulting resolution: 4037
% Factoring: 2
% Forward superposition: 3570
% Backward superposition: 4729
% Self superposition: 65
% Equality factoring: 6
% Equality resolution: 6
% 
% Split clauses: 17
% Split components: 35
% SAT solver clauses: 4078
% SAT solver unit clauses: 448
% SAT solver binary clauses: 734
% 
% Sat splits: 17
% Sat splitting refutations: 69
% 
% Memory used [KB]: 14200
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.util_prob
% (85219)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1359
% Passive clauses: 7054
% Generated clauses: 23799
% Final active clauses: 640
% Final passive clauses: 3500
% Input formulas: 53
% Initial clauses: 218
% 
% Duplicate literals: 1053
% Trivial inequalities: 66
% Fw subsumption resolutions: 845
% Bw subsumption resolutions: 1551
% Fw demodulations: 1087
% Bw demodulations: 57
% Condensations: 1651
% Global subsumptions: 110
% 
% Simple tautologies: 1569
% Equational tautologies: 36
% Forward subsumptions: 10159
% Backward subsumptions: 1408
% 
% Binary resolution: 7825
% Factoring: 59
% Forward superposition: 3275
% Backward superposition: 3251
% Self superposition: 75
% Equality factoring: 4
% Equality resolution: 12
% 
% Split clauses: 20
% Split components: 40
% SAT solver clauses: 7395
% SAT solver unit clauses: 425
% SAT solver binary clauses: 1366
% 
% Sat splits: 20
% Sat splitting refutations: 208
% 
% Memory used [KB]: 14328
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.util_prob
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.util_prob
% SZS output start Proof for thm.util_prob
fof(f4952,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f295,f520,f294,f522,f994,f993,f992,f990,f989,f988,f984,f983,f982,f979,f978,f977,f975,f974,f973,f971,f970,f969,f967,f966,f965,f945,f944,f943,f941,f940,f939,f937,f936,f935,f933,f932,f931,f929,f928,f927,f925,f924,f923,f921,f920,f919,f917,f916,f915,f908,f907,f906,f903,f902,f901,f898,f897,f896,f893,f892,f891,f882,f881,f880,f879,f878,f877,f874,f873,f872,f870,f869,f868,f865,f864,f863,f861,f860,f859,f857,f856,f855,f853,f852,f851,f842,f841,f840,f839,f838,f837,f1036,f834,f833,f832,f830,f829,f828,f814,f813,f812,f810,f809,f808,f806,f805,f804,f1077,f798,f797,f796,f794,f793,f792,f1081,f789,f788,f787,f786,f785,f1079,f784,f1083,f776,f775,f774,f772,f771,f770,f768,f767,f766,f1085,f747,f746,f745,f744,f743,f742,f1089,f739,f738,f737,f736,f735,f1087,f734,f1091,f733,f732,f731,f730,f1099,f729,f728,f727,f1095,f726,f1103,f725,f1093,f724,f1101,f723,f1097,f722,f1105,f719,f718,f717,f715,f714,f713,f1109,f711,f710,f709,f1107,f707,f706,f705,f1111,f691,f690,f689,f687,f686,f685,f683,f682,f681,f1113,f657,f656,f655,f653,f652,f651,f649,f648,f647,f1115,f634,f633,f632,f631,f1123,f630,f629,f1119,f628,f627,f1127,f626,f1117,f625,f1121,f624,f1125,f623,f1129,f615,f614,f613,f611,f610,f609,f607,f606,f605,f1131,f597,f596,f595,f593,f592,f591,f589,f588,f587,f1133,f581,f580,f579,f577,f576,f575,f573,f572,f571,f1135,f570,f569,f568,f567,f565,f564,f563,f561,f560,f559,f557,f556,f555,f1145,f553,f552,f551,f1143,f550,f549,f1141,f548,f1137,f547,f1139,f546,f545,f544,f543,f542,f541,f1149,f540,f1147,f539,f1151,f537,f535,f528,f527,f526,f525,f1034,f1002,f1003,f1153,f1173,f1172,f1171,f1180,f1170,f1182,f1163,f1162,f1161,f1184,f1160,f1186,f1492,f1490,f1483,f1472,f1462,f1461,f1460,f1459,f1458,f1456,f1454,f1447,f1441,f1439,f1437,f1436,f1435,f1562,f1434,f1560,f1432,f1431,f1430,f1564,f1418,f1566,f1416,f1414,f1413,f1412,f1570,f1411,f1568,f1409,f1408,f1407,f1572,f1403,f1574,f1380,f1576,f1377,f1375,f1374,f1373,f1580,f1372,f1578,f1370,f1369,f1368,f1582,f1365,f1363,f1362,f1361,f1355,f1353,f1352,f1351,f1350,f1348,f1347,f1346,f1345,f1343,f1342,f1341,f1590,f1340,f1338,f1337,f1336,f1586,f1330,f1328,f1327,f1326,f1325,f1323,f1322,f1321,f1320,f1318,f1317,f1316,f1594,f1315,f1584,f1313,f1312,f1311,f1592,f1305,f1303,f1302,f1301,f1300,f1298,f1297,f1296,f1295,f1588,f1293,f1292,f1291,f1596,f1289,f1287,f1286,f1285,f1600,f1284,f1598,f1282,f1281,f1280,f1602,f1277,f1604,f1276,f1267,f1265,f1264,f1263,f1610,f1262,f1608,f1260,f1259,f1258,f1612,f1256,f1254,f1253,f1252,f1616,f1251,f1614,f1249,f1248,f1247,f1618,f1246,f1244,f1243,f1242,f1622,f1241,f1620,f1239,f1238,f1237,f1606,f1236,f1234,f1233,f1232,f1231,f1229,f1228,f1227,f1226,f1225,f1224,f1632,f1223,f1630,f1222,f1220,f1219,f1218,f1628,f1217,f1624,f1215,f1214,f1213,f1626,f1206,f1205,f1204,f1201,f1200,f1199,f1193,f1192,f1531,f1191,f1190,f1188,f1187,f1494,f1638,f1495,f1636,f1496,f1497,f1639,f1634,f1701,f1704,f1702,f1705,f1787,f1786,f1785,f1790,f1799,f1791,f1797,f1792,f1795,f1961,f1967,f1969,f1985,f1984,f1983,f1987,f1997,f1988,f1995,f1989,f1993,f2025,f2031,f2033,f2042,f2048,f2050,f2087,f2089,f2094,f2096,f2101,f2103,f2108,f2112,f2117,f2121,f2151,f2160,f2146,f2163,f2185,f2184,f2183,f2306,f2170,f2169,f2168,f2316,f2279,f2328,f2280,f2330,f2479,f2485,f2488,f2507,f2509,f2515,f2517,f2523,f2525,f2946,f2958,f2945,f2960,f2944,f2962,f2943,f2964,f3070,f3241,f3094,f3095,f3250,f3248,f3097,f3249,f3098,f3251,f3099,f3100,f3101,f3253,f3102,f3252,f3103,f3254,f3104,f3255,f3107,f3256,f3110,f3257,f3111,f3258,f3112,f3259,f3118,f3260,f3125,f3261,f3126,f3262,f3127,f3263,f3132,f3264,f3133,f3265,f3135,f3266,f3137,f3267,f3138,f3268,f3139,f3269,f3277,f3848,f3850,f3885,f3884,f3888,f3883,f3882,f3881,f3890,f3898,f3897,f3901,f4199,f4205,f4207,f4273,f4279,f4284,f4295,f4301,f4306,f4329,f4333,f4339,f4343,f4349,f4353,f4729,f4722,f4721,f4720,f4719,f4711,f4698,f4690,f4689,f4688,f4687,f4684,f4683,f4680,f4679,f4678,f4677,f4669,f4663,f4662,f4657,f4656,f4655,f4654,f4810,f4652,f4808,f4651,f4650,f4649,f4813,f4635,f4818,f4631,f4630,f4629,f4628,f4824,f4625,f4822,f4624,f4623,f4622,f4828,f4618,f4832,f4595,f4840,f4568,f4848,f4558,f4557,f4556,f4555,f4855,f4553,f4853,f4552,f4551,f4550,f4858,f4547,f4546,f4545,f4544,f4543,f4542,f4541,f4540,f4539,f4538,f4537,f4536,f4535,f4534,f4533,f4532,f4866,f4530,f4529,f4528,f4527,f4862,f4526,f4525,f4524,f4523,f4522,f4521,f4520,f4519,f4518,f4517,f4516,f4515,f4871,f4513,f4860,f4512,f4511,f4510,f4869,f4509,f4508,f4507,f4506,f4505,f4504,f4503,f4502,f4501,f4864,f4500,f4499,f4498,f4874,f4494,f4493,f4492,f4491,f4880,f4488,f4878,f4487,f4486,f4485,f4884,f4475,f4474,f4473,f4472,f4892,f4469,f4890,f4468,f4467,f4466,f4895,f4462,f4898,f4461,f4900,f4460,f4902,f4459,f4455,f4906,f4451,f4908,f4444,f4443,f4442,f4441,f4914,f4438,f4912,f4437,f4436,f4435,f4918,f4431,f4430,f4429,f4428,f4924,f4425,f4922,f4424,f4423,f4422,f4928,f4421,f4930,f4420,f4419,f4418,f4417,f4934,f4416,f4932,f4415,f4414,f4413,f4904,f4412,f4411,f4410,f4409,f4408,f4407,f4406,f4405,f4404,f4403,f4402,f4944,f4401,f4942,f4400,f4399,f4398,f4397,f4940,f4396,f4936,f4395,f4394,f4393,f4938,f4387,f4386,f4385,f4383,f4382,f4381,f4375,f4374,f4769,f4373,f4372,f4947,f4948,f3243,f4949,f4734,f4735,f4951,f4946])).
fof(f4946,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl396),
  inference(cnf_transformation,[],[f4946_D])).
fof(f4946_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f4951,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(global_subsumption,[],[f295,f294,f1636,f4729,f4722,f4711,f4698,f4690,f4689,f4688,f4687,f4684,f4680,f4669,f4663,f4658,f4657,f4656,f4655,f4654,f4652,f4651,f4650,f4649,f4648,f4647,f4640,f4637,f4635,f4634,f4631,f4630,f4629,f4628,f4625,f4624,f4623,f4622,f4599,f4596,f4595,f4594,f4593,f4586,f4583,f4582,f4575,f4572,f4571,f4568,f4558,f4557,f4556,f4555,f4553,f4552,f4551,f4550,f4549,f4548,f4547,f4546,f4545,f4544,f4543,f4542,f4541,f4540,f4539,f4538,f4537,f4536,f4535,f4534,f4533,f4532,f4530,f4529,f4528,f4527,f4526,f4525,f4524,f4523,f4522,f4521,f4520,f4519,f4518,f4517,f4516,f4515,f4513,f4512,f4511,f4510,f4509,f4508,f4507,f4506,f4505,f4504,f4503,f4502,f4501,f4500,f4499,f4498,f4497,f4494,f4493,f4492,f4491,f4488,f4487,f4486,f4485,f4458,f4457,f4454,f4453,f4450,f4449,f4448,f4447,f4444,f4443,f4442,f4441,f4438,f4437,f4436,f4435,f4434,f4431,f4430,f4429,f4428,f4425,f4424,f4423,f4422,f4420,f4419,f4418,f4417,f4416,f4415,f4414,f4413,f4412,f4411,f4410,f4409,f4408,f4407,f4406,f4405,f4404,f4403,f4402,f4401,f4400,f4399,f4398,f4397,f4396,f4395,f4394,f4393,f4392,f4391,f4387,f4386,f4385,f4383,f4382,f4381,f4375,f4374,f4373,f4372,f4733,f4734,f4735,f4740,f4745,f4746,f4747,f4750,f4757,f4760,f4761,f4763,f4764,f4765,f4766])).
fof(f4766,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(superposition,[],[f1636,f303])).
fof(f303,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE',aHL_BOOL_CASES)).
fof(f4765,plain,(
  ( ! [X43,X42] : (~p(s(bool,X42)) | sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X42,X43)) ) | $spl172),
  inference(resolution,[],[f1636,f500])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f291])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f162,f181,f180])).
fof(f180,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f181,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.sat.sat_27_prove_ltb')).
fof(f4764,plain,(
  ( ! [X41,X40] : (p(s(bool,X40)) | sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X41,X40)) ) | $spl172),
  inference(resolution,[],[f1636,f499])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f291])).
fof(f4763,plain,(
  ( ! [X39,X38] : (p(s(bool,X38)) | sP11(X39,X38,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | p(s(bool,X39))) ) | $spl172),
  inference(resolution,[],[f1636,f498])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP11(X2,X0,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f4761,plain,(
  ( ! [X35,X34] : (~p(s(bool,X34)) | sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X34,X35)) ) | $spl172),
  inference(resolution,[],[f1636,f483])).
fof(f483,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f132,f178,f177])).
fof(f177,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f178,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f132,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f50])).
fof(f50,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.sat.sat_25_prove_ltb')).
fof(f4760,plain,(
  ( ! [X33,X32] : (~p(s(bool,X32)) | sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X33,X32)) ) | $spl172),
  inference(resolution,[],[f1636,f482])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f283])).
fof(f4757,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X26,X27)) ) | $spl172),
  inference(resolution,[],[f1636,f466])).
fof(f466,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f130,f175,f174])).
fof(f174,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f175,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f130,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f49])).
fof(f49,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) & p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.sat.sat_23_prove_ltb')).
fof(f4750,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X13)) ) | $spl172),
  inference(resolution,[],[f1636,f409])).
fof(f409,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1] : ((((p(s(bool,X1)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)))) | sP1(X1,X0)) & (~sP1(X1,X0) | ((~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X0)) | p(s(bool,X1))))))),
  inference(nnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> sP1(X1,X0))),
  inference(definition_folding,[],[f117,f167])).
fof(f167,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f117,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) <=> ~p(s(bool,X0))) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X21,X22] : ((p(s(bool,X22)) <=> ~p(s(bool,X21))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.sat.sat_29_prove_ltb')).
fof(f4747,plain,(
  ( ! [X10] : (p(s(bool,X10)) | sP1(X10,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(resolution,[],[f1636,f406])).
fof(f406,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f243])).
fof(f243,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f242])).
fof(f242,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f167])).
fof(f4746,plain,(
  ( ! [X9] : (p(s(bool,X9)) | sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X9)) ) | $spl172),
  inference(resolution,[],[f1636,f403])).
fof(f403,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f243])).
fof(f4745,plain,(
  ( ! [X8] : (~p(s(bool,X8)) | ~sP1(X8,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(resolution,[],[f1636,f402])).
fof(f402,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f243])).
fof(f4740,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(resolution,[],[f1636,f381])).
fof(f381,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.IMP_ANTISYM_AX')).
fof(f4733,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK12,X2),s(fun(sK12,bool),sK13)))) | s(sK12,X2) = s(sK12,happ(s(fun('type.num.num',sK12),sK14),s('type.num.num',sK31(X2,sK13,'const.pred_set.UNIV_0',sK14,sK12,'type.num.num'))))) ) | $spl172),
  inference(resolution,[],[f1636,f419])).
fof(f419,plain,(
  ( ! [X4,X2,X0,X8,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK31(X8,X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f254,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK31(X8,X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,sK31(X8,X4,X3,X2,X1,X0)))))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK32(X4,X3,X2,X1,X0)),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,sK32(X4,X3,X2,X1,X0)))),s(fun(X1,bool),X4))))) | (p(s(bool,'const.bool.IN_2'(s(X1,sK33(X4,X3,X2,X1,X0)),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,sK33(X4,X3,X2,X1,X0)) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK31,sK32,sK33])],[f253])).
fof(f253,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X10] : (~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) & ! [X8] : (~p(s(bool,'const.bool.IN_2'(s(X1,X8),s(fun(X1,bool),X4)))) | ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X0,X9),s(fun(X0,bool),X3)))) & s(X1,X8) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X9))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f252])).
fof(f252,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & (? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f251])).
fof(f251,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6))))))) & ((? [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) & ~p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) | s(X1,X5) != s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (~p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) | ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(ennf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (! [X7] : (p(s(bool,'const.bool.IN_2'(s(X0,X7),s(fun(X0,bool),X3)))) => p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X7))),s(fun(X1,bool),X4))))) & ! [X5] : (p(s(bool,'const.bool.IN_2'(s(X1,X5),s(fun(X1,bool),X4)))) => ? [X6] : (p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X3)))) & s(X1,X5) = s(X1,happ(s(fun(X0,X1),X2),s(X0,X6)))))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6,X16] : ! [X8,X7,X11] : (p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X16),X8),s(fun(X6,bool),X7),s(fun(X16,bool),X11)))) <=> (! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.IN_2'(s(X16,happ(s(fun(X6,X16),X8),s(X6,X12))),s(fun(X16,bool),X11))))) & ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X16,X12),s(fun(X16,bool),X11)))) => ? [X13] : (p(s(bool,'const.bool.IN_2'(s(X6,X13),s(fun(X6,bool),X7)))) & s(X16,X12) = s(X16,happ(s(fun(X6,X16),X8),s(X6,X13)))))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.pred_set.SURJ_DEF')).
fof(f4391,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f392])).
fof(f392,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,'F_0')) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0')))) & (((~p(s(bool,X1)) | p(s(bool,'F_0'))) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(definition_folding,[],[f155,f165])).
fof(f165,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f155,plain,(
  ! [X0,X1] : ((p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (((p(s(bool,X1)) | p(s(bool,X0))) | p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) & ~p(s(bool,'F_0'))) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(ennf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ((~(p(s(bool,X1)) | p(s(bool,X0))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X1)) => p(s(bool,'F_0'))) => (~p(s(bool,X0)) => p(s(bool,'F_0')))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X16,X6] : ((~(p(s(bool,X6)) | p(s(bool,X16))) => p(s(bool,'F_0'))) <=> ((p(s(bool,X6)) => p(s(bool,'F_0'))) => (~p(s(bool,X16)) => p(s(bool,'F_0')))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.sat.sat_9_prove_ltb')).
fof(f302,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f59])).
fof(f59,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE',aHL_FALSITY)).
fof(f4392,plain,(
  ( ! [X0] : (sP0('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f393])).
fof(f393,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f238])).
fof(f4434,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f438])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f262])).
fof(f262,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f261])).
fof(f261,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f260])).
fof(f260,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f4447,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f442])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f264])).
fof(f264,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f263])).
fof(f263,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f4448,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f446])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f4449,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f447])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f4450,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f267])).
fof(f4453,plain,(
  ( ! [X0] : (sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'const.combin.I_1'(s(bool,X0)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f354,f1636,f450])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | sP5(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f128,f172,f171,f170,f169])).
fof(f172,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.sat.sat_2|49|_prove_ltb')).
fof(f354,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X12] : s(X6,X12) = s(X6,'const.combin.I_1'(s(X6,X12)))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.combin.I_THM')).
fof(f4454,plain,(
  ( ! [X0] : (sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,X0)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f356,f1636,f450])).
fof(f356,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6] : ! [X24] : s(X6,'const.combin.I_1'(s(X6,X24))) = s(X6,happ(s(fun(X6,X6),'const.combin.I_0'),s(X6,X24)))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','ARITY_0_1_const.combin.I')).
fof(f4457,plain,(
  ( ! [X0] : (sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,X0)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f354,f1636,f450])).
fof(f4458,plain,(
  ( ! [X0] : (sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),'const.combin.I_1'(s(bool,X0)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f356,f1636,f450])).
fof(f4497,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f174])).
fof(f4548,plain,(
  ( ! [X0] : (~sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f467])).
fof(f467,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f275])).
fof(f4549,plain,(
  ( ! [X0] : (~sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0',X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f4571,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f478])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f280])).
fof(f280,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f279])).
fof(f279,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f177])).
fof(f4572,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f479])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f281])).
fof(f4575,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f482])).
fof(f296,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE',aHL_TRUTH)).
fof(f4582,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'T_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f482])).
fof(f4583,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f482])).
fof(f4586,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f483])).
fof(f4593,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f483])).
fof(f4594,plain,(
  ( ! [X0] : (sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f483])).
fof(f4596,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f283])).
fof(f4599,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f486])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f283])).
fof(f4634,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X1)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f496])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f289])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f288])).
fof(f288,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f287])).
fof(f287,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f180])).
fof(f4637,plain,(
  ( ! [X0] : (sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f499])).
fof(f4640,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f500])).
fof(f4647,plain,(
  ( ! [X0] : (sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f500])).
fof(f4648,plain,(
  ( ! [X0] : (sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f500])).
fof(f4658,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f503])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f291])).
fof(f4735,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(resolution,[],[f1636,f315])).
fof(f315,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.EQ_CLAUSES|split|1')).
fof(f4734,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(resolution,[],[f1636,f313])).
fof(f313,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.EQ_CLAUSES|split|0')).
fof(f4949,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,happ(s(fun('type.num.num',sK12),sK14),s('type.num.num',X0))),s(fun(sK12,bool),sK13))))) ) | $spl398),
  inference(cnf_transformation,[],[f4949_D])).
fof(f4949_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,happ(s(fun('type.num.num',sK12),sK14),s('type.num.num',X0))),s(fun(sK12,bool),sK13))))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f3243,plain,(
  ~p(s(bool,'T_0')) | $spl239),
  inference(cnf_transformation,[],[f3243_D])).
fof(f3243_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl239),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl239])])).
fof(f4948,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(sK12,happ(s(fun('type.num.num',sK12),sK14),s('type.num.num',X0))),s(fun(sK12,bool),sK13))))) ) | ($spl2 | $spl170 | $spl172)),
  inference(forward_demodulation,[],[f4731,f1641])).
fof(f1641,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(backward_demodulation,[],[f1634,f1207])).
fof(f1207,plain,(
  ( ! [X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))) = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f381])).
fof(f355,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))) )),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0] : ! [X1] : p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0'))))),
  inference(rectify,[],[f40])).
fof(f40,axiom,(
  ! [X6] : ! [X12] : p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.UNIV_0'))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.pred_set.IN_UNIV')).
fof(f4731,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s('type.num.num',X0),s(fun('type.num.num',bool),'const.pred_set.UNIV_0')))) | p(s(bool,'const.bool.IN_2'(s(sK12,happ(s(fun('type.num.num',sK12),sK14),s('type.num.num',X0))),s(fun(sK12,bool),sK13))))) ) | $spl172),
  inference(resolution,[],[f1636,f417])).
fof(f417,plain,(
  ( ! [X4,X2,X0,X10,X3,X1] : (~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X10),s(fun(X0,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,happ(s(fun(X0,X1),X2),s(X0,X10))),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f254])).
fof(f4947,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl172)),
  inference(forward_demodulation,[],[f4370,f2544])).
fof(f2544,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,sK12)),s(fun(sK12,bool),sK13)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f523,f381])).
fof(f523,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK12,sK17(X0,sK13,sK12)),s(fun(sK12,bool),sK13))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f360])).
fof(f360,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),sK15(X1,X0)),s('type.num.num',sK16(X6,X1,X0)))))) & (! [X2] : (p(s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,sK17(X2,X1,X0)) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16,sK17])],[f214])).
fof(f214,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ? [X5] : ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | ? [X7] : s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),X5),s('type.num.num',X7))))) & (! [X2] : ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,X3) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) | ? [X2] : ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4))))) & (! [X2] : ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & ! [X4] : s(X0,X3) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(flattening,[],[f97])).
fof(f97,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) => ? [X4] : s(X0,X3) = s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4)))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X7)))) <=> ? [X8] : ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X7)))) => ? [X23] : s(X6,X12) = s(X6,happ(s(fun('type.num.num',X6),X8),s('type.num.num',X23)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.util_prob.countable_def')).
fof(f4370,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK12,sK17(sK14,sK13,sK12)),s(fun(sK12,bool),sK13)))) | ($spl1 | $spl172)),
  inference(unit_resulting_resolution,[],[f524,f1636,f419])).
fof(f524,plain,(
  ( ! [X0,X1] : (s(sK12,sK17(X0,sK13,sK12)) != s(sK12,happ(s(fun('type.num.num',sK12),X0),s('type.num.num',X1)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f361])).
fof(f361,plain,(
  ( ! [X4,X2,X0,X1] : (s(X0,sK17(X2,X1,X0)) != s(X0,happ(s(fun('type.num.num',X0),X2),s('type.num.num',X4))) | p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f4372,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f313])).
fof(f4373,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f315])).
fof(f4769,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl303),
  inference(cnf_transformation,[],[f4769_D])).
fof(f4769_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f4374,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f316])).
fof(f316,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.EQ_CLAUSES|split|2')).
fof(f4375,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f318])).
fof(f318,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.EQ_CLAUSES|split|3')).
fof(f4381,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f381])).
fof(f514,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f357])).
fof(f357,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f95])).
fof(f95,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X12] : (s(X6,X12) = s(X6,X12) <=> p(s(bool,'T_0')))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.REFL_CLAUSE')).
fof(f4382,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f381])).
fof(f297,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.TRUTH')).
fof(f4383,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f381])).
fof(f4385,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f381])).
fof(f4386,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f381])).
fof(f4387,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f381])).
fof(f4938,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl389),
  inference(cnf_transformation,[],[f4938_D])).
fof(f4938_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl389),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl389])])).
fof(f4393,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f402])).
fof(f4394,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f402])).
fof(f4395,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f402])).
fof(f4936,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl387),
  inference(cnf_transformation,[],[f4936_D])).
fof(f4936_D,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl387),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl387])])).
fof(f4396,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f402])).
fof(f4940,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl391),
  inference(cnf_transformation,[],[f4940_D])).
fof(f4940_D,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl391),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl391])])).
fof(f4397,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f402])).
fof(f4398,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f402])).
fof(f4399,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f402])).
fof(f4400,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f402])).
fof(f4942,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl392),
  inference(cnf_transformation,[],[f4942_D])).
fof(f4942_D,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f4401,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f403])).
fof(f4944,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl394),
  inference(cnf_transformation,[],[f4944_D])).
fof(f4944_D,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f4402,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f406])).
fof(f4403,plain,(
  sP1('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f407])).
fof(f407,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP1(X1,X0) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f244])).
fof(f4404,plain,(
  sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f408])).
fof(f408,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X1,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f4405,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f409])).
fof(f4406,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f409])).
fof(f4407,plain,(
  ~sP1('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f409])).
fof(f4408,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f409])).
fof(f4409,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f409])).
fof(f4410,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f409])).
fof(f4411,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f409])).
fof(f4412,plain,(
  ~sP1('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f409])).
fof(f4904,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl363),
  inference(cnf_transformation,[],[f4904_D])).
fof(f4904_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl363),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl363])])).
fof(f4413,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f259])).
fof(f259,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f257])).
fof(f257,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f172])).
fof(f4414,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f433])).
fof(f4415,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f433])).
fof(f4932,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl383),
  inference(cnf_transformation,[],[f4932_D])).
fof(f4932_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl383),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl383])])).
fof(f4416,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f433])).
fof(f4934,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl385),
  inference(cnf_transformation,[],[f4934_D])).
fof(f4934_D,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl385),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl385])])).
fof(f4417,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f433])).
fof(f4418,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f433])).
fof(f4419,plain,(
  ~sP5('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f433])).
fof(f4420,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f433])).
fof(f4930,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl380),
  inference(cnf_transformation,[],[f4930_D])).
fof(f4930_D,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f4421,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f1269,f1268,f1257,f1636,f434])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | ~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f1257,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f438])).
fof(f1268,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f442])).
fof(f1269,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f446])).
fof(f4928,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl379),
  inference(cnf_transformation,[],[f4928_D])).
fof(f4928_D,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl379),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl379])])).
fof(f4422,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f262])).
fof(f4423,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f437])).
fof(f4424,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f437])).
fof(f4922,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl375),
  inference(cnf_transformation,[],[f4922_D])).
fof(f4922_D,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl375),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl375])])).
fof(f4425,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f437])).
fof(f4924,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl377),
  inference(cnf_transformation,[],[f4924_D])).
fof(f4924_D,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl377),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl377])])).
fof(f4428,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f437])).
fof(f4429,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f437])).
fof(f4430,plain,(
  ~sP4('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f437])).
fof(f4431,plain,(
  ~sP4('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f437])).
fof(f4918,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl373),
  inference(cnf_transformation,[],[f4918_D])).
fof(f4918_D,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f4435,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f441])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f265])).
fof(f4436,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f441])).
fof(f4437,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f441])).
fof(f4912,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl369),
  inference(cnf_transformation,[],[f4912_D])).
fof(f4912_D,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl369),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl369])])).
fof(f4438,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f441])).
fof(f4914,plain,(
  ~sP3('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl371),
  inference(cnf_transformation,[],[f4914_D])).
fof(f4914_D,plain,(
  ~sP3('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl371),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl371])])).
fof(f4441,plain,(
  ~sP3('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f441])).
fof(f4442,plain,(
  ~sP3('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f441])).
fof(f4443,plain,(
  ~sP3('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f441])).
fof(f4444,plain,(
  ~sP3('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f441])).
fof(f4908,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl366),
  inference(cnf_transformation,[],[f4908_D])).
fof(f4908_D,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f4451,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl84 | $spl172)),
  inference(unit_resulting_resolution,[],[f1153,f1636,f450])).
fof(f4906,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl364),
  inference(cnf_transformation,[],[f4906_D])).
fof(f4906_D,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f4455,plain,(
  sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | ($spl84 | $spl172)),
  inference(unit_resulting_resolution,[],[f1153,f1636,f450])).
fof(f4459,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | ($spl7 | $spl172)),
  inference(unit_resulting_resolution,[],[f1036,f1636,f451])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f268])).
fof(f4902,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl361),
  inference(cnf_transformation,[],[f4902_D])).
fof(f4902_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl361),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl361])])).
fof(f4460,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl5 | $spl172)),
  inference(unit_resulting_resolution,[],[f1034,f1636,f451])).
fof(f4900,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl359),
  inference(cnf_transformation,[],[f4900_D])).
fof(f4900_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl359),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl359])])).
fof(f4461,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl7 | $spl172)),
  inference(unit_resulting_resolution,[],[f1036,f1636,f451])).
fof(f4898,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl357),
  inference(cnf_transformation,[],[f4898_D])).
fof(f4898_D,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl357),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl357])])).
fof(f4462,plain,(
  ~sP5('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | ($spl5 | $spl172)),
  inference(unit_resulting_resolution,[],[f1034,f1636,f451])).
fof(f4895,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl354),
  inference(cnf_transformation,[],[f4895_D])).
fof(f4895_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f4466,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f514,f1290,f1636,f456])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (~sP6(X1,X2,X0) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f270])).
fof(f270,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f269])).
fof(f269,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f175])).
fof(f1290,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f461])).
fof(f4467,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f297,f1290,f1636,f456])).
fof(f4468,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f296,f1290,f1636,f456])).
fof(f4890,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl350),
  inference(cnf_transformation,[],[f4890_D])).
fof(f4890_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f4469,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f1636,f1290,f1636,f456])).
fof(f4892,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl352),
  inference(cnf_transformation,[],[f4892_D])).
fof(f4892_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f4472,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f514,f1290,f1636,f456])).
fof(f4473,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f297,f1290,f1636,f456])).
fof(f4474,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f296,f1290,f1636,f456])).
fof(f4475,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl172)),
  inference(unit_resulting_resolution,[],[f1636,f1290,f1636,f456])).
fof(f4884,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl349),
  inference(cnf_transformation,[],[f4884_D])).
fof(f4884_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl349),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl349])])).
fof(f4485,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f460])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f273])).
fof(f4486,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f460])).
fof(f4487,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f460])).
fof(f4878,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl345),
  inference(cnf_transformation,[],[f4878_D])).
fof(f4878_D,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f4488,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f460])).
fof(f4880,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl347),
  inference(cnf_transformation,[],[f4880_D])).
fof(f4880_D,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl347),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl347])])).
fof(f4491,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f460])).
fof(f4492,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f460])).
fof(f4493,plain,(
  ~sP6('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f460])).
fof(f4494,plain,(
  ~sP6('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f460])).
fof(f4874,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl342),
  inference(cnf_transformation,[],[f4874_D])).
fof(f4874_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f4498,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f514,f1636,f466])).
fof(f4499,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f514,f1636,f466])).
fof(f4500,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f514,f1636,f466])).
fof(f4864,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl334),
  inference(cnf_transformation,[],[f4864_D])).
fof(f4864_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f4501,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f514,f1636,f466])).
fof(f4502,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f297,f1636,f466])).
fof(f4503,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f297,f1636,f466])).
fof(f4504,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f297,f1636,f466])).
fof(f4505,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f297,f1636,f466])).
fof(f4506,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f296,f1636,f466])).
fof(f4507,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f296,f1636,f466])).
fof(f4508,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f296,f1636,f466])).
fof(f4509,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f296,f1636,f466])).
fof(f4869,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl338),
  inference(cnf_transformation,[],[f4869_D])).
fof(f4869_D,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f4510,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f1636,f466])).
fof(f4511,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f1636,f466])).
fof(f4512,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f1636,f466])).
fof(f4860,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl330),
  inference(cnf_transformation,[],[f4860_D])).
fof(f4860_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f4513,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f1636,f466])).
fof(f4871,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl340),
  inference(cnf_transformation,[],[f4871_D])).
fof(f4871_D,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f4515,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f514,f1636,f466])).
fof(f4516,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f514,f1636,f466])).
fof(f4517,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f514,f1636,f466])).
fof(f4518,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f514,f1636,f466])).
fof(f4519,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f297,f1636,f466])).
fof(f4520,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f297,f1636,f466])).
fof(f4521,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f297,f1636,f466])).
fof(f4522,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f297,f1636,f466])).
fof(f4523,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f296,f1636,f466])).
fof(f4524,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f296,f1636,f466])).
fof(f4525,plain,(
  sP7('T_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f296,f1636,f466])).
fof(f4526,plain,(
  sP7('T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f296,f1636,f466])).
fof(f4862,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl332),
  inference(cnf_transformation,[],[f4862_D])).
fof(f4862_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f4527,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f1636,f466])).
fof(f4528,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f1636,f466])).
fof(f4529,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f1636,f466])).
fof(f4530,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f1636,f466])).
fof(f4866,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl336),
  inference(cnf_transformation,[],[f4866_D])).
fof(f4866_D,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f4532,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f514,f1636,f466])).
fof(f4533,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f514,f1636,f466])).
fof(f4534,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f514,f1636,f466])).
fof(f4535,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f514,f1636,f466])).
fof(f4536,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f297,f1636,f466])).
fof(f4537,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f297,f1636,f466])).
fof(f4538,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f297,f1636,f466])).
fof(f4539,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f297,f1636,f466])).
fof(f4540,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f514,f296,f1636,f466])).
fof(f4541,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f297,f296,f1636,f466])).
fof(f4542,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f296,f296,f1636,f466])).
fof(f4543,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f296,f1636,f466])).
fof(f4544,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f1636,f1636,f466])).
fof(f4545,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f1636,f1636,f466])).
fof(f4546,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f1636,f1636,f466])).
fof(f4547,plain,(
  sP7('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f1636,f1636,f466])).
fof(f4858,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl329),
  inference(cnf_transformation,[],[f4858_D])).
fof(f4858_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl329),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl329])])).
fof(f4550,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f514,f1636,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f275])).
fof(f4551,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f297,f1636,f469])).
fof(f4552,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f296,f1636,f469])).
fof(f4853,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl325),
  inference(cnf_transformation,[],[f4853_D])).
fof(f4853_D,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f4553,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f1636,f469])).
fof(f4855,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl327),
  inference(cnf_transformation,[],[f4855_D])).
fof(f4855_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl327),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl327])])).
fof(f4555,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f514,f1636,f469])).
fof(f4556,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f297,f1636,f469])).
fof(f4557,plain,(
  ~sP7('F_0','T_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f296,f1636,f469])).
fof(f4558,plain,(
  ~sP7('F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f1636,f469])).
fof(f4848,plain,(
  ~sP8('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl323),
  inference(cnf_transformation,[],[f4848_D])).
fof(f4848_D,plain,(
  ~sP8('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl323),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl323])])).
fof(f4568,plain,(
  ~sP8('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f1636,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f281])).
fof(f4840,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl321),
  inference(cnf_transformation,[],[f4840_D])).
fof(f4840_D,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl321),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl321])])).
fof(f4595,plain,(
  ~sP9('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f1636,f484])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f283])).
fof(f4832,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl318),
  inference(cnf_transformation,[],[f4832_D])).
fof(f4832_D,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f4618,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl1 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f783,f1636,f493])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | sP11(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f285])).
fof(f285,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f181])).
fof(f783,plain,(
  ( ! [X0,X1] : (sP10('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f497])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f4828,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl317),
  inference(cnf_transformation,[],[f4828_D])).
fof(f4828_D,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl317),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl317])])).
fof(f4622,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f514,f1636,f494])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f289])).
fof(f4623,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f297,f1636,f494])).
fof(f4624,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f296,f1636,f494])).
fof(f4822,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl313),
  inference(cnf_transformation,[],[f4822_D])).
fof(f4822_D,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f4625,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f1636,f494])).
fof(f4824,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl315),
  inference(cnf_transformation,[],[f4824_D])).
fof(f4824_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl315),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl315])])).
fof(f4628,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f1636,f494])).
fof(f4629,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f297,f302,f1636,f494])).
fof(f4630,plain,(
  ~sP10('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f296,f302,f1636,f494])).
fof(f4631,plain,(
  ~sP10('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f1636,f494])).
fof(f4818,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl310),
  inference(cnf_transformation,[],[f4818_D])).
fof(f4818_D,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f4635,plain,(
  sP11('F_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f1636,f498])).
fof(f4813,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl309),
  inference(cnf_transformation,[],[f4813_D])).
fof(f4813_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl309),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl309])])).
fof(f4649,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f514,f1636,f501])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (~sP11(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f291])).
fof(f4650,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f297,f1636,f501])).
fof(f4651,plain,(
  ~sP11('T_0','F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f296,f1636,f501])).
fof(f4808,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl305),
  inference(cnf_transformation,[],[f4808_D])).
fof(f4808_D,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl305),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl305])])).
fof(f4652,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f1636,f501])).
fof(f4810,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl307),
  inference(cnf_transformation,[],[f4810_D])).
fof(f4810_D,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl307),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl307])])).
fof(f4654,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f514,f1636,f501])).
fof(f4655,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f297,f1636,f501])).
fof(f4656,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl172),
  inference(unit_resulting_resolution,[],[f302,f296,f1636,f501])).
fof(f4657,plain,(
  ~sP11('const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f1636,f1636,f501])).
fof(f4662,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl84 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f302,f1153,f1636,f505])).
fof(f505,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f163])).
fof(f163,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X12,X18,X13,X19] : ((s(bool,X12) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X13) = s(bool,X19))) => ((p(s(bool,X12)) => p(s(bool,X13))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.bool.IMP_CONG')).
fof(f4663,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f354,f1636,f505])).
fof(f4669,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f354,f1636,f507])).
fof(f507,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f4677,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f514,f1641,f1636,f507])).
fof(f4678,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f297,f1641,f1636,f507])).
fof(f4679,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f296,f1641,f1636,f507])).
fof(f4680,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f354,f1636,f507])).
fof(f4683,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl84 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f302,f1153,f1636,f509])).
fof(f509,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f4684,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f354,f1636,f509])).
fof(f4687,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f514,f1636,f510])).
fof(f510,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f4688,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f297,f1636,f510])).
fof(f4689,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f296,f1636,f510])).
fof(f4690,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f1636,f1636,f510])).
fof(f4698,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f1636,f1636,f510])).
fof(f4711,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f1636,f302,f354,f1636,f511])).
fof(f511,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X2)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f4719,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f514,f1641,f1636,f511])).
fof(f4720,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f297,f1641,f1636,f511])).
fof(f4721,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ($spl2 | $spl170 | $spl172)),
  inference(unit_resulting_resolution,[],[f302,f296,f1641,f1636,f511])).
fof(f4722,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f514,f302,f354,f1636,f511])).
fof(f4729,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl172),
  inference(unit_resulting_resolution,[],[f302,f302,f1636,f518])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(duplicate_literal_removal,[],[f508])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f293])).
fof(f4353,plain,(
  ~sP5('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl301),
  inference(cnf_transformation,[],[f4353_D])).
fof(f4353_D,plain,(
  ~sP5('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl301),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl301])])).
fof(f4349,plain,(
  ~sP5('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl153),
  inference(unit_resulting_resolution,[],[f1616,f432])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f259])).
fof(f4343,plain,(
  ~sP5('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl299),
  inference(cnf_transformation,[],[f4343_D])).
fof(f4343_D,plain,(
  ~sP5('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl299),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl299])])).
fof(f4339,plain,(
  ~sP5('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl149),
  inference(unit_resulting_resolution,[],[f1612,f431])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f4333,plain,(
  ~sP5('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl297),
  inference(cnf_transformation,[],[f4333_D])).
fof(f4333_D,plain,(
  ~sP5('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl297),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl297])])).
fof(f4329,plain,(
  ~sP5('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl147),
  inference(unit_resulting_resolution,[],[f1610,f431])).
fof(f4306,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl294),
  inference(cnf_transformation,[],[f4306_D])).
fof(f4306_D,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f4301,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl132),
  inference(resolution,[],[f1596,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f271])).
fof(f4295,plain,(
  sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl132),
  inference(unit_resulting_resolution,[],[f1596,f453])).
fof(f4284,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl292),
  inference(cnf_transformation,[],[f4284_D])).
fof(f4284_D,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f4279,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl130),
  inference(resolution,[],[f1594,f453])).
fof(f4273,plain,(
  sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl130),
  inference(unit_resulting_resolution,[],[f1594,f453])).
fof(f4207,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl290),
  inference(cnf_transformation,[],[f4207_D])).
fof(f4207_D,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f4205,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl102),
  inference(resolution,[],[f1566,f489])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f286])).
fof(f4199,plain,(
  sP10('F_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl102),
  inference(unit_resulting_resolution,[],[f1566,f489])).
fof(f3901,plain,(
  sP5('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl288),
  inference(cnf_transformation,[],[f3901_D])).
fof(f3901_D,plain,(
  sP5('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f3897,plain,(
  sP5('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f302,f1257,f621,f1270,f436])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f1270,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f447])).
fof(f621,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f443])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f265])).
fof(f3898,plain,(
  sP5('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f302,f603,f621,f1270,f436])).
fof(f603,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f439])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f262])).
fof(f3890,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl286),
  inference(cnf_transformation,[],[f3890_D])).
fof(f3890_D,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f3881,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f1257,f1268,f1269,f434])).
fof(f3882,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f1257,f1268,f1269,f434])).
fof(f3883,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f1257,f1268,f1269,f434])).
fof(f3888,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl284),
  inference(cnf_transformation,[],[f3888_D])).
fof(f3888_D,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f3884,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f302,f603,f1268,f1269,f435])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f3885,plain,(
  sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl2)),
  inference(unit_resulting_resolution,[],[f302,f603,f621,f1269,f435])).
fof(f3850,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl283),
  inference(cnf_transformation,[],[f3850_D])).
fof(f3850_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl283),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl283])])).
fof(f3848,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl49),
  inference(unit_resulting_resolution,[],[f1117,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f259])).
fof(f3277,plain,(
  s(bool,'T_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | ~p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | p(s(bool,'T_0'))),
  introduced(theory_axiom,[])).
fof(f3269,plain,(
  ( ! [X229,X227,X230,X228] : (~p(s(bool,X229)) | sP11('const.bool.IN_2'(s(X227,X228),s(fun(X227,bool),'const.pred_set.UNIV_0')),X229,X230)) ) | $spl280),
  inference(cnf_transformation,[],[f3269_D])).
fof(f3269_D,plain,(
  ( ! [X229,X227,X230,X228] : (~p(s(bool,X229)) | sP11('const.bool.IN_2'(s(X227,X228),s(fun(X227,bool),'const.pred_set.UNIV_0')),X229,X230)) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f3139,plain,(
  ( ! [X229,X227,X230,X228] : (~p(s(bool,'T_0')) | ~p(s(bool,X229)) | sP11('const.bool.IN_2'(s(X227,X228),s(fun(X227,bool),'const.pred_set.UNIV_0')),X229,X230)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f500,f1641])).
fof(f3268,plain,(
  ( ! [X225,X223,X226,X224] : (p(s(bool,X225)) | sP11('const.bool.IN_2'(s(X223,X224),s(fun(X223,bool),'const.pred_set.UNIV_0')),X226,X225)) ) | $spl278),
  inference(cnf_transformation,[],[f3268_D])).
fof(f3268_D,plain,(
  ( ! [X225,X223,X226,X224] : (p(s(bool,X225)) | sP11('const.bool.IN_2'(s(X223,X224),s(fun(X223,bool),'const.pred_set.UNIV_0')),X226,X225)) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f3138,plain,(
  ( ! [X225,X223,X226,X224] : (~p(s(bool,'T_0')) | p(s(bool,X225)) | sP11('const.bool.IN_2'(s(X223,X224),s(fun(X223,bool),'const.pred_set.UNIV_0')),X226,X225)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f499,f1641])).
fof(f3267,plain,(
  ( ! [X222,X220,X219,X221] : (p(s(bool,X221)) | p(s(bool,X222)) | sP11(X222,X221,'const.bool.IN_2'(s(X219,X220),s(fun(X219,bool),'const.pred_set.UNIV_0')))) ) | $spl276),
  inference(cnf_transformation,[],[f3267_D])).
fof(f3267_D,plain,(
  ( ! [X222,X220,X219,X221] : (p(s(bool,X221)) | p(s(bool,X222)) | sP11(X222,X221,'const.bool.IN_2'(s(X219,X220),s(fun(X219,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f3137,plain,(
  ( ! [X222,X220,X219,X221] : (~p(s(bool,'T_0')) | p(s(bool,X221)) | sP11(X222,X221,'const.bool.IN_2'(s(X219,X220),s(fun(X219,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X222))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f498,f1641])).
fof(f3266,plain,(
  ( ! [X212,X213,X211,X214] : (sP10(X213,'const.bool.IN_2'(s(X211,X212),s(fun(X211,bool),'const.pred_set.UNIV_0')),X214)) ) | $spl274),
  inference(cnf_transformation,[],[f3266_D])).
fof(f3266_D,plain,(
  ( ! [X212,X213,X211,X214] : (sP10(X213,'const.bool.IN_2'(s(X211,X212),s(fun(X211,bool),'const.pred_set.UNIV_0')),X214)) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f3135,plain,(
  ( ! [X212,X213,X211,X214] : (~p(s(bool,'T_0')) | sP10(X213,'const.bool.IN_2'(s(X211,X212),s(fun(X211,bool),'const.pred_set.UNIV_0')),X214)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f496,f1641])).
fof(f3265,plain,(
  ( ! [X206,X204,X205,X203] : (~p(s(bool,X205)) | sP9('const.bool.IN_2'(s(X203,X204),s(fun(X203,bool),'const.pred_set.UNIV_0')),X205,X206)) ) | $spl272),
  inference(cnf_transformation,[],[f3265_D])).
fof(f3265_D,plain,(
  ( ! [X206,X204,X205,X203] : (~p(s(bool,X205)) | sP9('const.bool.IN_2'(s(X203,X204),s(fun(X203,bool),'const.pred_set.UNIV_0')),X205,X206)) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f3133,plain,(
  ( ! [X206,X204,X205,X203] : (~p(s(bool,'T_0')) | ~p(s(bool,X205)) | sP9('const.bool.IN_2'(s(X203,X204),s(fun(X203,bool),'const.pred_set.UNIV_0')),X205,X206)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f483,f1641])).
fof(f3264,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,X201)) | sP9('const.bool.IN_2'(s(X199,X200),s(fun(X199,bool),'const.pred_set.UNIV_0')),X202,X201)) ) | $spl270),
  inference(cnf_transformation,[],[f3264_D])).
fof(f3264_D,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,X201)) | sP9('const.bool.IN_2'(s(X199,X200),s(fun(X199,bool),'const.pred_set.UNIV_0')),X202,X201)) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f3132,plain,(
  ( ! [X202,X200,X199,X201] : (~p(s(bool,'T_0')) | ~p(s(bool,X201)) | sP9('const.bool.IN_2'(s(X199,X200),s(fun(X199,bool),'const.pred_set.UNIV_0')),X202,X201)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f482,f1641])).
fof(f3263,plain,(
  ( ! [X182,X180,X181,X179] : (sP8(X181,'const.bool.IN_2'(s(X179,X180),s(fun(X179,bool),'const.pred_set.UNIV_0')),X182)) ) | $spl268),
  inference(cnf_transformation,[],[f3263_D])).
fof(f3263_D,plain,(
  ( ! [X182,X180,X181,X179] : (sP8(X181,'const.bool.IN_2'(s(X179,X180),s(fun(X179,bool),'const.pred_set.UNIV_0')),X182)) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f3127,plain,(
  ( ! [X182,X180,X181,X179] : (~p(s(bool,'T_0')) | sP8(X181,'const.bool.IN_2'(s(X179,X180),s(fun(X179,bool),'const.pred_set.UNIV_0')),X182)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f479,f1641])).
fof(f3262,plain,(
  ( ! [X177,X175,X178,X176] : (sP8(X177,X178,'const.bool.IN_2'(s(X175,X176),s(fun(X175,bool),'const.pred_set.UNIV_0')))) ) | $spl266),
  inference(cnf_transformation,[],[f3262_D])).
fof(f3262_D,plain,(
  ( ! [X177,X175,X178,X176] : (sP8(X177,X178,'const.bool.IN_2'(s(X175,X176),s(fun(X175,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f3126,plain,(
  ( ! [X177,X175,X178,X176] : (~p(s(bool,'T_0')) | sP8(X177,X178,'const.bool.IN_2'(s(X175,X176),s(fun(X175,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f478,f1641])).
fof(f3261,plain,(
  ( ! [X173,X171,X174,X172] : (~p(s(bool,X173)) | sP7('const.bool.IN_2'(s(X171,X172),s(fun(X171,bool),'const.pred_set.UNIV_0')),X173,X174) | ~p(s(bool,X174))) ) | $spl264),
  inference(cnf_transformation,[],[f3261_D])).
fof(f3261_D,plain,(
  ( ! [X173,X171,X174,X172] : (~p(s(bool,X173)) | sP7('const.bool.IN_2'(s(X171,X172),s(fun(X171,bool),'const.pred_set.UNIV_0')),X173,X174) | ~p(s(bool,X174))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f3125,plain,(
  ( ! [X173,X171,X174,X172] : (~p(s(bool,'T_0')) | ~p(s(bool,X173)) | ~p(s(bool,X174)) | sP7('const.bool.IN_2'(s(X171,X172),s(fun(X171,bool),'const.pred_set.UNIV_0')),X173,X174)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f466,f1641])).
fof(f3260,plain,(
  ( ! [X146,X144,X145,X143] : (sP6(X145,X146,'const.bool.IN_2'(s(X143,X144),s(fun(X143,bool),'const.pred_set.UNIV_0')))) ) | $spl262),
  inference(cnf_transformation,[],[f3260_D])).
fof(f3260_D,plain,(
  ( ! [X146,X144,X145,X143] : (sP6(X145,X146,'const.bool.IN_2'(s(X143,X144),s(fun(X143,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f3118,plain,(
  ( ! [X146,X144,X145,X143] : (~p(s(bool,'T_0')) | sP6(X145,X146,'const.bool.IN_2'(s(X143,X144),s(fun(X143,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f461,f1641])).
fof(f3259,plain,(
  ( ! [X121,X122,X120,X119] : (sP2('const.bool.IN_2'(s(X119,X120),s(fun(X119,bool),'const.pred_set.UNIV_0')),X121,X122)) ) | $spl260),
  inference(cnf_transformation,[],[f3259_D])).
fof(f3259_D,plain,(
  ( ! [X121,X122,X120,X119] : (sP2('const.bool.IN_2'(s(X119,X120),s(fun(X119,bool),'const.pred_set.UNIV_0')),X121,X122)) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f3112,plain,(
  ( ! [X121,X122,X120,X119] : (~p(s(bool,'T_0')) | sP2('const.bool.IN_2'(s(X119,X120),s(fun(X119,bool),'const.pred_set.UNIV_0')),X121,X122)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f448,f1641])).
fof(f3258,plain,(
  ( ! [X118,X116,X117,X115] : (sP2(X117,'const.bool.IN_2'(s(X115,X116),s(fun(X115,bool),'const.pred_set.UNIV_0')),X118)) ) | $spl258),
  inference(cnf_transformation,[],[f3258_D])).
fof(f3258_D,plain,(
  ( ! [X118,X116,X117,X115] : (sP2(X117,'const.bool.IN_2'(s(X115,X116),s(fun(X115,bool),'const.pred_set.UNIV_0')),X118)) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f3111,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,'T_0')) | sP2(X117,'const.bool.IN_2'(s(X115,X116),s(fun(X115,bool),'const.pred_set.UNIV_0')),X118)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f447,f1641])).
fof(f3257,plain,(
  ( ! [X111,X114,X112,X113] : (sP2(X113,X114,'const.bool.IN_2'(s(X111,X112),s(fun(X111,bool),'const.pred_set.UNIV_0')))) ) | $spl256),
  inference(cnf_transformation,[],[f3257_D])).
fof(f3257_D,plain,(
  ( ! [X111,X114,X112,X113] : (sP2(X113,X114,'const.bool.IN_2'(s(X111,X112),s(fun(X111,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f3110,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,'T_0')) | sP2(X113,X114,'const.bool.IN_2'(s(X111,X112),s(fun(X111,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f446,f1641])).
fof(f3256,plain,(
  ( ! [X101,X99,X102,X100] : (sP3(X101,X102,'const.bool.IN_2'(s(X99,X100),s(fun(X99,bool),'const.pred_set.UNIV_0')))) ) | $spl254),
  inference(cnf_transformation,[],[f3256_D])).
fof(f3256_D,plain,(
  ( ! [X101,X99,X102,X100] : (sP3(X101,X102,'const.bool.IN_2'(s(X99,X100),s(fun(X99,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f3107,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,'T_0')) | sP3(X101,X102,'const.bool.IN_2'(s(X99,X100),s(fun(X99,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f442,f1641])).
fof(f3255,plain,(
  ( ! [X90,X88,X87,X89] : (sP4(X89,X90,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),'const.pred_set.UNIV_0')))) ) | $spl252),
  inference(cnf_transformation,[],[f3255_D])).
fof(f3255_D,plain,(
  ( ! [X90,X88,X87,X89] : (sP4(X89,X90,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f3104,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | sP4(X89,X90,'const.bool.IN_2'(s(X87,X88),s(fun(X87,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f438,f1641])).
fof(f3254,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X86)) | ~sP1('const.bool.IN_2'(s(X84,X85),s(fun(X84,bool),'const.pred_set.UNIV_0')),X86)) ) | $spl250),
  inference(cnf_transformation,[],[f3254_D])).
fof(f3254_D,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,X86)) | ~sP1('const.bool.IN_2'(s(X84,X85),s(fun(X84,bool),'const.pred_set.UNIV_0')),X86)) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f3103,plain,(
  ( ! [X85,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X86)) | ~sP1('const.bool.IN_2'(s(X84,X85),s(fun(X84,bool),'const.pred_set.UNIV_0')),X86)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f409,f1641])).
fof(f3252,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,X74)) | sP1('const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.pred_set.UNIV_0')),X74)) ) | $spl246),
  inference(cnf_transformation,[],[f3252_D])).
fof(f3252_D,plain,(
  ( ! [X74,X72,X73] : (p(s(bool,X74)) | sP1('const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.pred_set.UNIV_0')),X74)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f3102,plain,(
  ( ! [X83,X81,X82] : (~p(s(bool,'T_0')) | p(s(bool,X83)) | sP1('const.bool.IN_2'(s(X81,X82),s(fun(X81,bool),'const.pred_set.UNIV_0')),X83)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f408,f1641])).
fof(f3253,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,X77)) | sP1(X77,'const.bool.IN_2'(s(X75,X76),s(fun(X75,bool),'const.pred_set.UNIV_0')))) ) | $spl248),
  inference(cnf_transformation,[],[f3253_D])).
fof(f3253_D,plain,(
  ( ! [X76,X77,X75] : (p(s(bool,X77)) | sP1(X77,'const.bool.IN_2'(s(X75,X76),s(fun(X75,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f3101,plain,(
  ( ! [X80,X78,X79] : (~p(s(bool,'T_0')) | sP1(X80,'const.bool.IN_2'(s(X78,X79),s(fun(X78,bool),'const.pred_set.UNIV_0'))) | p(s(bool,X80))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f407,f1641])).
fof(f3100,plain,(
  ( ! [X76,X77,X75] : (~p(s(bool,'T_0')) | p(s(bool,X77)) | sP1(X77,'const.bool.IN_2'(s(X75,X76),s(fun(X75,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f406,f1641])).
fof(f3099,plain,(
  ( ! [X74,X72,X73] : (~p(s(bool,'T_0')) | p(s(bool,X74)) | sP1('const.bool.IN_2'(s(X72,X73),s(fun(X72,bool),'const.pred_set.UNIV_0')),X74)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f403,f1641])).
fof(f3251,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X71)) | ~sP1(X71,'const.bool.IN_2'(s(X69,X70),s(fun(X69,bool),'const.pred_set.UNIV_0')))) ) | $spl244),
  inference(cnf_transformation,[],[f3251_D])).
fof(f3251_D,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,X71)) | ~sP1(X71,'const.bool.IN_2'(s(X69,X70),s(fun(X69,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f3098,plain,(
  ( ! [X70,X71,X69] : (~p(s(bool,'T_0')) | ~p(s(bool,X71)) | ~sP1(X71,'const.bool.IN_2'(s(X69,X70),s(fun(X69,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f402,f1641])).
fof(f3249,plain,(
  ( ! [X61,X62,X60] : (sP0('const.bool.IN_2'(s(X60,X61),s(fun(X60,bool),'const.pred_set.UNIV_0')),X62)) ) | $spl242),
  inference(cnf_transformation,[],[f3249_D])).
fof(f3249_D,plain,(
  ( ! [X61,X62,X60] : (sP0('const.bool.IN_2'(s(X60,X61),s(fun(X60,bool),'const.pred_set.UNIV_0')),X62)) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f3097,plain,(
  ( ! [X68,X66,X67] : (~p(s(bool,'T_0')) | sP0('const.bool.IN_2'(s(X66,X67),s(fun(X66,bool),'const.pred_set.UNIV_0')),X68)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f393,f1641])).
fof(f3248,plain,(
  ( ! [X59,X57,X58] : (sP0(X59,'const.bool.IN_2'(s(X57,X58),s(fun(X57,bool),'const.pred_set.UNIV_0')))) ) | $spl240),
  inference(cnf_transformation,[],[f3248_D])).
fof(f3248_D,plain,(
  ( ! [X59,X57,X58] : (sP0(X59,'const.bool.IN_2'(s(X57,X58),s(fun(X57,bool),'const.pred_set.UNIV_0')))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f3250,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | sP0(X65,'const.bool.IN_2'(s(X63,X64),s(fun(X63,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(global_subsumption,[],[f295,f294,f522,f1492,f1490,f1483,f1472,f1462,f1461,f1460,f1459,f1458,f1456,f1454,f1447,f1441,f1440,f1439,f1438,f1437,f1436,f1435,f1434,f1433,f1432,f1431,f1430,f1429,f1428,f1427,f1423,f1422,f1419,f1418,f1417,f1416,f1415,f1414,f1413,f1412,f1411,f1410,f1409,f1408,f1407,f1405,f1404,f1403,f1402,f1401,f1400,f1396,f1395,f1392,f1391,f1390,f1386,f1385,f1382,f1381,f1380,f1377,f1376,f1375,f1374,f1373,f1372,f1371,f1370,f1369,f1368,f1367,f1366,f1365,f1364,f1363,f1362,f1361,f1360,f1359,f1358,f1355,f1354,f1353,f1352,f1351,f1350,f1348,f1347,f1346,f1345,f1343,f1342,f1341,f1340,f1338,f1337,f1336,f1335,f1334,f1333,f1330,f1329,f1328,f1327,f1326,f1325,f1323,f1322,f1321,f1320,f1318,f1317,f1316,f1315,f1313,f1312,f1311,f1309,f1308,f1305,f1304,f1303,f1302,f1301,f1300,f1298,f1297,f1296,f1295,f1293,f1292,f1291,f1290,f1289,f1288,f1287,f1286,f1285,f1284,f1283,f1282,f1281,f1280,f1275,f1274,f1273,f1272,f1271,f1270,f1269,f1268,f1267,f1266,f1265,f1264,f1263,f1262,f1261,f1260,f1259,f1258,f1257,f1256,f1255,f1254,f1253,f1252,f1251,f1250,f1249,f1248,f1247,f1246,f1245,f1244,f1243,f1242,f1241,f1240,f1239,f1238,f1237,f1236,f1235,f1234,f1233,f1232,f1231,f1230,f1229,f1228,f1227,f1226,f1225,f1224,f1223,f1222,f1220,f1219,f1218,f1217,f1215,f1214,f1213,f1212,f1211,f1207,f1206,f1205,f1204,f1201,f1200,f1199,f1193,f1192,f1191,f1190,f1188,f1187,f1494,f1495,f1496,f1497,f1502,f1507,f1508,f1509,f1512,f1519,f1522,f1523,f1525,f1526,f1527,f1528,f1639,f1634,f1641,f3025,f3024,f2972,f2971,f3070,f3094,f3095,f3096])).
fof(f3096,plain,(
  ( ! [X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,'F_0')) | sP0(X65,'const.bool.IN_2'(s(X63,X64),s(fun(X63,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f392,f1641])).
fof(f2971,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f296,f1641,f450])).
fof(f2972,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f302,f1641,f452])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X2)) | ~sP5(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f268])).
fof(f3024,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f296,f1641,f450])).
fof(f3025,plain,(
  ( ! [X0,X1] : (~sP5('F_0','T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(unit_resulting_resolution,[],[f302,f1641,f452])).
fof(f1528,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(superposition,[],[f522,f303])).
fof(f1527,plain,(
  ( ! [X39,X40] : (~p(s(bool,X39)) | sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X39,X40)) ) | $spl2),
  inference(resolution,[],[f522,f500])).
fof(f1526,plain,(
  ( ! [X37,X38] : (p(s(bool,X37)) | sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X38,X37)) ) | $spl2),
  inference(resolution,[],[f522,f499])).
fof(f1525,plain,(
  ( ! [X35,X36] : (p(s(bool,X35)) | sP11(X36,X35,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | p(s(bool,X36))) ) | $spl2),
  inference(resolution,[],[f522,f498])).
fof(f1523,plain,(
  ( ! [X31,X32] : (~p(s(bool,X31)) | sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X31,X32)) ) | $spl2),
  inference(resolution,[],[f522,f483])).
fof(f1522,plain,(
  ( ! [X30,X29] : (~p(s(bool,X29)) | sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X30,X29)) ) | $spl2),
  inference(resolution,[],[f522,f482])).
fof(f1519,plain,(
  ( ! [X24,X23] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X23,X24)) ) | $spl2),
  inference(resolution,[],[f522,f466])).
fof(f1512,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X10)) ) | $spl2),
  inference(resolution,[],[f522,f409])).
fof(f1509,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(resolution,[],[f522,f406])).
fof(f1508,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X6)) ) | $spl2),
  inference(resolution,[],[f522,f403])).
fof(f1507,plain,(
  ( ! [X5] : (~p(s(bool,X5)) | ~sP1(X5,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(resolution,[],[f522,f402])).
fof(f1502,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(resolution,[],[f522,f381])).
fof(f1211,plain,(
  ( ! [X0] : (sP0(X0,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f392])).
fof(f1212,plain,(
  ( ! [X0] : (sP0('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f393])).
fof(f1230,plain,(
  ( ! [X0,X1] : (~sP1('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f409])).
fof(f1235,plain,(
  ( ! [X0,X1] : (~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f409])).
fof(f1240,plain,(
  ( ! [X0,X1] : (~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f433])).
fof(f1245,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f433])).
fof(f1250,plain,(
  ( ! [X0,X1] : (~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f437])).
fof(f1255,plain,(
  ( ! [X0,X1] : (~sP4('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f437])).
fof(f1261,plain,(
  ( ! [X0,X1] : (~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f441])).
fof(f1266,plain,(
  ( ! [X0,X1] : (~sP3('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f441])).
fof(f1271,plain,(
  ( ! [X0,X1] : (sP2('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f448])).
fof(f1272,plain,(
  ( ! [X0] : (sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'const.combin.I_1'(s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f354,f522,f450])).
fof(f1273,plain,(
  ( ! [X0] : (sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,X0)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f356,f522,f450])).
fof(f1274,plain,(
  ( ! [X0] : (sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,X0)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f354,f522,f450])).
fof(f1275,plain,(
  ( ! [X0] : (sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),happ(s(fun(bool,bool),'const.combin.I_0'),s(bool,X0)),'const.combin.I_1'(s(bool,X0)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f356,f522,f450])).
fof(f1283,plain,(
  ( ! [X0,X1] : (~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f460])).
fof(f1288,plain,(
  ( ! [X0,X1] : (~sP6('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f460])).
fof(f1304,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f296,f522,f466])).
fof(f1308,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f355,f522,f466])).
fof(f1309,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f355,f522,f466])).
fof(f1329,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f296,f522,f466])).
fof(f1333,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f355,f522,f466])).
fof(f1334,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f355,f522,f466])).
fof(f1335,plain,(
  ( ! [X0,X1] : (sP7('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f355,f522,f466])).
fof(f1354,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f296,f522,f466])).
fof(f1358,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f355,f522,f466])).
fof(f1359,plain,(
  ( ! [X2,X0,X3,X1] : (sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.bool.IN_2'(s(X2,X3),s(fun(X2,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f355,f522,f466])).
fof(f1360,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f355,f522,f466])).
fof(f1364,plain,(
  ( ! [X0,X1] : (sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f522,f466])).
fof(f1366,plain,(
  ( ! [X0] : (~sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f467])).
fof(f1367,plain,(
  ( ! [X0] : (~sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f468])).
fof(f1371,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f355,f522,f469])).
fof(f1376,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f355,f522,f469])).
fof(f1381,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f478])).
fof(f1382,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f479])).
fof(f1385,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f482])).
fof(f1386,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f482])).
fof(f1390,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'T_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f482])).
fof(f1391,plain,(
  ( ! [X2,X0,X1] : (sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f482])).
fof(f1392,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f482])).
fof(f1395,plain,(
  ( ! [X0] : (sP9('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f483])).
fof(f1396,plain,(
  ( ! [X2,X0,X1] : (sP9('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f483])).
fof(f1400,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f483])).
fof(f1401,plain,(
  ( ! [X2,X0,X1] : (sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f483])).
fof(f1402,plain,(
  ( ! [X0] : (sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f483])).
fof(f1404,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f485])).
fof(f1405,plain,(
  ( ! [X0] : (~sP9('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f486])).
fof(f1410,plain,(
  ( ! [X0,X1] : (~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f355,f522,f494])).
fof(f1415,plain,(
  ( ! [X0,X1] : (~sP10('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f302,f522,f494])).
fof(f1417,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X1)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f496])).
fof(f1419,plain,(
  ( ! [X0] : (sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0,'F_0')) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f499])).
fof(f1422,plain,(
  ( ! [X0] : (sP11('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f500])).
fof(f1423,plain,(
  ( ! [X2,X0,X1] : (sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f500])).
fof(f1427,plain,(
  ( ! [X0] : (sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0',X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f500])).
fof(f1428,plain,(
  ( ! [X2,X0,X1] : (sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),X2)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f355,f522,f500])).
fof(f1429,plain,(
  ( ! [X0] : (sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f500])).
fof(f1433,plain,(
  ( ! [X0,X1] : (~sP11('const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f355,f522,f501])).
fof(f1438,plain,(
  ( ! [X0,X1] : (~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.UNIV_0')))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f355,f522,f501])).
fof(f1440,plain,(
  ( ! [X0] : (~sP11('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),X0)) ) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f503])).
fof(f3095,plain,(
  ( ! [X61,X62,X60] : (~p(s(bool,'T_0')) | sP0('const.bool.IN_2'(s(X60,X61),s(fun(X60,bool),'const.pred_set.UNIV_0')),X62)) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f388,f1641])).
fof(f388,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f165])).
fof(f3094,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,'T_0')) | sP0(X59,'const.bool.IN_2'(s(X57,X58),s(fun(X57,bool),'const.pred_set.UNIV_0')))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f387,f1641])).
fof(f387,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f236])).
fof(f3241,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),sK15('const.pred_set.UNIV_0',X0)),s('type.num.num',sK16(X1,'const.pred_set.UNIV_0',X0)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),'const.pred_set.UNIV_0'))))) ) | $spl236),
  inference(cnf_transformation,[],[f3241_D])).
fof(f3241_D,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),sK15('const.pred_set.UNIV_0',X0)),s('type.num.num',sK16(X1,'const.pred_set.UNIV_0',X0)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),'const.pred_set.UNIV_0'))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f3070,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,happ(s(fun('type.num.num',X0),sK15('const.pred_set.UNIV_0',X0)),s('type.num.num',sK16(X1,'const.pred_set.UNIV_0',X0)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),'const.pred_set.UNIV_0'))))) ) | ($spl2 | $spl170)),
  inference(superposition,[],[f359,f1641])).
fof(f359,plain,(
  ( ! [X6,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X6),s(fun(X0,bool),X1)))) | s(X0,X6) = s(X0,happ(s(fun('type.num.num',X0),sK15(X1,X0)),s('type.num.num',sK16(X6,X1,X0)))) | ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f215])).
fof(f2964,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl234),
  inference(cnf_transformation,[],[f2964_D])).
fof(f2964_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f2943,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','F_0') | ($spl1 | $spl7)),
  inference(unit_resulting_resolution,[],[f1036,f1013])).
fof(f1013,plain,(
  ( ! [X10,X11] : (sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X10,X11) | s(bool,X10) = s(bool,X11)) ) | $spl1),
  inference(resolution,[],[f520,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f268])).
fof(f2962,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl232),
  inference(cnf_transformation,[],[f2962_D])).
fof(f2962_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f2944,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl5)),
  inference(unit_resulting_resolution,[],[f1034,f1013])).
fof(f2960,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl230),
  inference(cnf_transformation,[],[f2960_D])).
fof(f2960_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f2945,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl1 | $spl7)),
  inference(unit_resulting_resolution,[],[f1036,f1013])).
fof(f2958,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl228),
  inference(cnf_transformation,[],[f2958_D])).
fof(f2958_D,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f2946,plain,(
  sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | ($spl1 | $spl5)),
  inference(unit_resulting_resolution,[],[f1034,f1013])).
fof(f2525,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl227),
  inference(cnf_transformation,[],[f2525_D])).
fof(f2525_D,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f2523,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl57),
  inference(unit_resulting_resolution,[],[f1125,f430])).
fof(f2517,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl225),
  inference(cnf_transformation,[],[f2517_D])).
fof(f2517_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f2515,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl53),
  inference(unit_resulting_resolution,[],[f1121,f430])).
fof(f2509,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl223),
  inference(cnf_transformation,[],[f2509_D])).
fof(f2509_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl223),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl223])])).
fof(f2507,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl51),
  inference(unit_resulting_resolution,[],[f1119,f430])).
fof(f2488,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl220),
  inference(cnf_transformation,[],[f2488_D])).
fof(f2488_D,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f2485,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl32),
  inference(resolution,[],[f1101,f472])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f178])).
fof(f2479,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl32),
  inference(unit_resulting_resolution,[],[f1101,f472])).
fof(f2330,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'F_0'))) | $spl219),
  inference(cnf_transformation,[],[f2330_D])).
fof(f2330_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'F_0'))) <=> ~$spl219),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl219])])).
fof(f2280,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'F_0'))) | ($spl1 | $spl84)),
  inference(superposition,[],[f637,f1153])).
fof(f637,plain,(
  ( ! [X0] : (~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X0,'const.combin.I_1'(s(bool,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f354,f520,f452])).
fof(f2328,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'F_0')),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl217),
  inference(cnf_transformation,[],[f2328_D])).
fof(f2328_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'F_0')),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl217),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl217])])).
fof(f2279,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,'F_0')),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ($spl1 | $spl84)),
  inference(superposition,[],[f635,f1153])).
fof(f635,plain,(
  ( ! [X0] : (~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.combin.I_1'(s(bool,X0)),X0)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f354,f520,f452])).
fof(f2316,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl214),
  inference(cnf_transformation,[],[f2316_D])).
fof(f2316_D,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f2168,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f514,f1153,f450])).
fof(f2169,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f297,f1153,f450])).
fof(f2170,plain,(
  sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl84),
  inference(unit_resulting_resolution,[],[f296,f1153,f450])).
fof(f2306,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl212),
  inference(cnf_transformation,[],[f2306_D])).
fof(f2306_D,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f2183,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl84),
  inference(unit_resulting_resolution,[],[f514,f1153,f450])).
fof(f2184,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl84),
  inference(unit_resulting_resolution,[],[f297,f1153,f450])).
fof(f2185,plain,(
  sP5('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl84),
  inference(unit_resulting_resolution,[],[f296,f1153,f450])).
fof(f2163,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl210),
  inference(cnf_transformation,[],[f2163_D])).
fof(f2163_D,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f2146,plain,(
  sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl5),
  inference(unit_resulting_resolution,[],[f302,f1034,f449])).
fof(f2160,plain,(
  sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl208),
  inference(cnf_transformation,[],[f2160_D])).
fof(f2160_D,plain,(
  sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f2151,plain,(
  sP5('F_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl5),
  inference(unit_resulting_resolution,[],[f302,f1034,f449])).
fof(f2121,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl207),
  inference(cnf_transformation,[],[f2121_D])).
fof(f2121_D,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl207),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl207])])).
fof(f2117,plain,(
  ~sP5('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl65),
  inference(unit_resulting_resolution,[],[f1133,f432])).
fof(f2112,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl205),
  inference(cnf_transformation,[],[f2112_D])).
fof(f2112_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f2108,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl63),
  inference(unit_resulting_resolution,[],[f1131,f431])).
fof(f2103,plain,(
  ~sP5('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl203),
  inference(cnf_transformation,[],[f2103_D])).
fof(f2103_D,plain,(
  ~sP5('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl203),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl203])])).
fof(f2101,plain,(
  ~sP5('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl61),
  inference(unit_resulting_resolution,[],[f1129,f430])).
fof(f2096,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl201),
  inference(cnf_transformation,[],[f2096_D])).
fof(f2096_D,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl201),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl201])])).
fof(f2094,plain,(
  ~sP5('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl59),
  inference(unit_resulting_resolution,[],[f1127,f430])).
fof(f2089,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl199),
  inference(cnf_transformation,[],[f2089_D])).
fof(f2089_D,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl199),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl199])])).
fof(f2087,plain,(
  ~sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl55),
  inference(unit_resulting_resolution,[],[f1123,f430])).
fof(f2050,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl196),
  inference(cnf_transformation,[],[f2050_D])).
fof(f2050_D,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f2048,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl36),
  inference(resolution,[],[f1105,f472])).
fof(f2042,plain,(
  sP8('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl36),
  inference(unit_resulting_resolution,[],[f1105,f472])).
fof(f2033,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl194),
  inference(cnf_transformation,[],[f2033_D])).
fof(f2033_D,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f2031,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl34),
  inference(resolution,[],[f1103,f472])).
fof(f2025,plain,(
  sP8('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl34),
  inference(unit_resulting_resolution,[],[f1103,f472])).
fof(f1993,plain,(
  sP4('F_0','T_0','F_0') | $spl188),
  inference(cnf_transformation,[],[f1993_D])).
fof(f1993_D,plain,(
  sP4('F_0','T_0','F_0') <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f1989,plain,(
  sP4('F_0','T_0','F_0') | $spl92),
  inference(resolution,[],[f1186,f432])).
fof(f1995,plain,(
  sP3('F_0','T_0','F_0') | $spl190),
  inference(cnf_transformation,[],[f1995_D])).
fof(f1995_D,plain,(
  sP3('F_0','T_0','F_0') <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f1988,plain,(
  sP3('F_0','T_0','F_0') | $spl92),
  inference(resolution,[],[f1186,f431])).
fof(f1997,plain,(
  sP2('T_0','F_0','F_0') | $spl192),
  inference(cnf_transformation,[],[f1997_D])).
fof(f1997_D,plain,(
  sP2('T_0','F_0','F_0') <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f1987,plain,(
  sP2('T_0','F_0','F_0') | $spl92),
  inference(resolution,[],[f1186,f430])).
fof(f1983,plain,(
  sP2('T_0','F_0','F_0') | $spl92),
  inference(unit_resulting_resolution,[],[f1186,f430])).
fof(f1984,plain,(
  sP3('F_0','T_0','F_0') | $spl92),
  inference(unit_resulting_resolution,[],[f1186,f431])).
fof(f1985,plain,(
  sP4('F_0','T_0','F_0') | $spl92),
  inference(unit_resulting_resolution,[],[f1186,f432])).
fof(f1969,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl186),
  inference(cnf_transformation,[],[f1969_D])).
fof(f1969_D,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f1967,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl14),
  inference(resolution,[],[f1083,f489])).
fof(f1961,plain,(
  sP10('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl14),
  inference(unit_resulting_resolution,[],[f1083,f489])).
fof(f1795,plain,(
  sP4('F_0','F_0','T_0') | $spl180),
  inference(cnf_transformation,[],[f1795_D])).
fof(f1795_D,plain,(
  sP4('F_0','F_0','T_0') <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f1792,plain,(
  sP4('F_0','F_0','T_0') | $spl88),
  inference(resolution,[],[f1182,f432])).
fof(f1797,plain,(
  sP3('T_0','F_0','F_0') | $spl182),
  inference(cnf_transformation,[],[f1797_D])).
fof(f1797_D,plain,(
  sP3('T_0','F_0','F_0') <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f1791,plain,(
  sP3('T_0','F_0','F_0') | $spl88),
  inference(resolution,[],[f1182,f431])).
fof(f1799,plain,(
  sP2('F_0','T_0','F_0') | $spl184),
  inference(cnf_transformation,[],[f1799_D])).
fof(f1799_D,plain,(
  sP2('F_0','T_0','F_0') <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f1790,plain,(
  sP2('F_0','T_0','F_0') | $spl88),
  inference(resolution,[],[f1182,f430])).
fof(f1785,plain,(
  sP2('F_0','T_0','F_0') | $spl88),
  inference(unit_resulting_resolution,[],[f1182,f430])).
fof(f1786,plain,(
  sP3('T_0','F_0','F_0') | $spl88),
  inference(unit_resulting_resolution,[],[f1182,f431])).
fof(f1787,plain,(
  sP4('F_0','F_0','T_0') | $spl88),
  inference(unit_resulting_resolution,[],[f1182,f432])).
fof(f1705,plain,(
  s(bool,'F_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~p(s(bool,'F_0')) | p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))))),
  introduced(theory_axiom,[])).
fof(f1702,plain,(
  ( ! [X6,X7] : (sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) | $spl176),
  inference(cnf_transformation,[],[f1702_D])).
fof(f1702_D,plain,(
  ( ! [X6,X7] : (sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f1704,plain,(
  p(s(bool,'F_0')) | $spl178),
  inference(cnf_transformation,[],[f1704_D])).
fof(f1704_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f1701,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) | ($spl1 | $spl84)),
  inference(forward_demodulation,[],[f1698,f1153])).
fof(f1698,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) | sP5('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X6,X7) | ~sP3(X6,X7,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | ~sP2(X7,X6,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) ) | $spl1),
  inference(resolution,[],[f604,f436])).
fof(f604,plain,(
  ( ! [X0,X1] : (sP4('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),X0,X1)) ) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f262])).
fof(f1634,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl170),
  inference(cnf_transformation,[],[f1634_D])).
fof(f1634_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f1639,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(global_subsumption,[],[f295,f294,f522,f1492,f1490,f1483,f1472,f1462,f1461,f1460,f1459,f1458,f1456,f1454,f1447,f1441,f1440,f1439,f1438,f1437,f1436,f1435,f1434,f1433,f1432,f1431,f1430,f1429,f1428,f1427,f1423,f1422,f1419,f1418,f1417,f1416,f1415,f1414,f1413,f1412,f1411,f1410,f1409,f1408,f1407,f1405,f1404,f1403,f1402,f1401,f1400,f1396,f1395,f1392,f1391,f1390,f1386,f1385,f1382,f1381,f1380,f1377,f1376,f1375,f1374,f1373,f1372,f1371,f1370,f1369,f1368,f1367,f1366,f1365,f1364,f1363,f1362,f1361,f1360,f1359,f1358,f1355,f1354,f1353,f1352,f1351,f1350,f1348,f1347,f1346,f1345,f1343,f1342,f1341,f1340,f1338,f1337,f1336,f1335,f1334,f1333,f1330,f1329,f1328,f1327,f1326,f1325,f1323,f1322,f1321,f1320,f1318,f1317,f1316,f1315,f1313,f1312,f1311,f1309,f1308,f1305,f1304,f1303,f1302,f1301,f1300,f1298,f1297,f1296,f1295,f1293,f1292,f1291,f1290,f1289,f1288,f1287,f1286,f1285,f1284,f1283,f1282,f1281,f1280,f1275,f1274,f1273,f1272,f1271,f1270,f1269,f1268,f1267,f1266,f1265,f1264,f1263,f1262,f1261,f1260,f1259,f1258,f1257,f1256,f1255,f1254,f1253,f1252,f1251,f1250,f1249,f1248,f1247,f1246,f1245,f1244,f1243,f1242,f1241,f1240,f1239,f1238,f1237,f1236,f1235,f1234,f1233,f1232,f1231,f1230,f1229,f1228,f1227,f1226,f1225,f1224,f1223,f1222,f1220,f1219,f1218,f1217,f1215,f1214,f1213,f1212,f1211,f1207,f1206,f1205,f1204,f1201,f1200,f1199,f1193,f1192,f1191,f1190,f1188,f1187,f1494,f1495,f1496,f1497,f1502,f1507,f1508,f1509,f1512,f1519,f1522,f1523,f1525,f1526,f1527,f1528])).
fof(f1497,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(resolution,[],[f522,f315])).
fof(f1496,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(resolution,[],[f522,f313])).
fof(f1636,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | $spl172),
  inference(cnf_transformation,[],[f1636_D])).
fof(f1636_D,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f1495,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | $spl2),
  inference(resolution,[],[f522,f415])).
fof(f415,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & (~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f249])).
fof(f249,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))))) & ((~p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | ~p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) | p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X16] : ! [X8,X7,X11] : (p(s(bool,'const.pred_set.BIJ_3'(s(fun(X6,X16),X8),s(fun(X6,bool),X7),s(fun(X16,bool),X11)))) <=> (p(s(bool,'const.pred_set.INJ_3'(s(fun(X6,X16),X8),s(fun(X6,bool),X7),s(fun(X16,bool),X11)))) & p(s(bool,'const.pred_set.SURJ_3'(s(fun(X6,X16),X8),s(fun(X6,bool),X7),s(fun(X16,bool),X11))))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.pred_set.BIJ_DEF')).
fof(f1638,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | $spl174),
  inference(cnf_transformation,[],[f1638_D])).
fof(f1638_D,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f1494,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | $spl2),
  inference(resolution,[],[f522,f414])).
fof(f414,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.BIJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4)))) | p(s(bool,'const.pred_set.INJ_3'(s(fun(X0,X1),X2),s(fun(X0,bool),X3),s(fun(X1,bool),X4))))) )),
  inference(cnf_transformation,[],[f250])).
fof(f1187,plain,(
  p(s(bool,'const.pred_set.INJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f414])).
fof(f1188,plain,(
  p(s(bool,'const.pred_set.SURJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f415])).
fof(f1190,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f313])).
fof(f1191,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f315])).
fof(f1531,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl95),
  inference(cnf_transformation,[],[f1531_D])).
fof(f1531_D,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f1192,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f316])).
fof(f1193,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f318])).
fof(f1199,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f381])).
fof(f1200,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f381])).
fof(f1201,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f381])).
fof(f1204,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f381])).
fof(f1205,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f381])).
fof(f1206,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f381])).
fof(f1626,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl163),
  inference(cnf_transformation,[],[f1626_D])).
fof(f1626_D,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f1213,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f402])).
fof(f1214,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f402])).
fof(f1215,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f402])).
fof(f1624,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl161),
  inference(cnf_transformation,[],[f1624_D])).
fof(f1624_D,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl161),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl161])])).
fof(f1217,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f402])).
fof(f1628,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl165),
  inference(cnf_transformation,[],[f1628_D])).
fof(f1628_D,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl165),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl165])])).
fof(f1218,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f402])).
fof(f1219,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f402])).
fof(f1220,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f402])).
fof(f1222,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f402])).
fof(f1630,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl166),
  inference(cnf_transformation,[],[f1630_D])).
fof(f1630_D,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f1223,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f403])).
fof(f1632,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl168),
  inference(cnf_transformation,[],[f1632_D])).
fof(f1632_D,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f1224,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f406])).
fof(f1225,plain,(
  sP1('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f407])).
fof(f1226,plain,(
  sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f408])).
fof(f1227,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f409])).
fof(f1228,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f409])).
fof(f1229,plain,(
  ~sP1('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f409])).
fof(f1231,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f409])).
fof(f1232,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f409])).
fof(f1233,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f409])).
fof(f1234,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f409])).
fof(f1236,plain,(
  ~sP1('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f409])).
fof(f1606,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl143),
  inference(cnf_transformation,[],[f1606_D])).
fof(f1606_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f1237,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f433])).
fof(f1238,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f433])).
fof(f1239,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f433])).
fof(f1620,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl157),
  inference(cnf_transformation,[],[f1620_D])).
fof(f1620_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl157),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl157])])).
fof(f1241,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f433])).
fof(f1622,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl159),
  inference(cnf_transformation,[],[f1622_D])).
fof(f1622_D,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f1242,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f433])).
fof(f1243,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f433])).
fof(f1244,plain,(
  ~sP5('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f433])).
fof(f1246,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f433])).
fof(f1618,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl155),
  inference(cnf_transformation,[],[f1618_D])).
fof(f1618_D,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f1247,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f437])).
fof(f1248,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f437])).
fof(f1249,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f437])).
fof(f1614,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl151),
  inference(cnf_transformation,[],[f1614_D])).
fof(f1614_D,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f1251,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f437])).
fof(f1616,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl153),
  inference(cnf_transformation,[],[f1616_D])).
fof(f1616_D,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl153),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl153])])).
fof(f1252,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f437])).
fof(f1253,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f437])).
fof(f1254,plain,(
  ~sP4('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f437])).
fof(f1256,plain,(
  ~sP4('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f437])).
fof(f1612,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl149),
  inference(cnf_transformation,[],[f1612_D])).
fof(f1612_D,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl149),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl149])])).
fof(f1258,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f441])).
fof(f1259,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f441])).
fof(f1260,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f441])).
fof(f1608,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl145),
  inference(cnf_transformation,[],[f1608_D])).
fof(f1608_D,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f1262,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f441])).
fof(f1610,plain,(
  ~sP3('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl147),
  inference(cnf_transformation,[],[f1610_D])).
fof(f1610_D,plain,(
  ~sP3('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f1263,plain,(
  ~sP3('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f441])).
fof(f1264,plain,(
  ~sP3('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f441])).
fof(f1265,plain,(
  ~sP3('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f441])).
fof(f1267,plain,(
  ~sP3('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f441])).
fof(f1276,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | ($spl2 | $spl7)),
  inference(unit_resulting_resolution,[],[f1036,f522,f451])).
fof(f1604,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl141),
  inference(cnf_transformation,[],[f1604_D])).
fof(f1604_D,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl141])])).
fof(f1277,plain,(
  ~sP5('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | ($spl2 | $spl7)),
  inference(unit_resulting_resolution,[],[f1036,f522,f451])).
fof(f1602,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl139),
  inference(cnf_transformation,[],[f1602_D])).
fof(f1602_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') <=> ~$spl139),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl139])])).
fof(f1280,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f460])).
fof(f1281,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f460])).
fof(f1282,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f460])).
fof(f1598,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl135),
  inference(cnf_transformation,[],[f1598_D])).
fof(f1598_D,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f1284,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f460])).
fof(f1600,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl137),
  inference(cnf_transformation,[],[f1600_D])).
fof(f1600_D,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f1285,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f460])).
fof(f1286,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f460])).
fof(f1287,plain,(
  ~sP6('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f460])).
fof(f1289,plain,(
  ~sP6('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f460])).
fof(f1596,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl132),
  inference(cnf_transformation,[],[f1596_D])).
fof(f1596_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f1291,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f514,f522,f466])).
fof(f1292,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f514,f522,f466])).
fof(f1293,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f514,f522,f466])).
fof(f1588,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl124),
  inference(cnf_transformation,[],[f1588_D])).
fof(f1588_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f1295,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f514,f522,f466])).
fof(f1296,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f297,f522,f466])).
fof(f1297,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f297,f522,f466])).
fof(f1298,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f297,f522,f466])).
fof(f1300,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f297,f522,f466])).
fof(f1301,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f296,f522,f466])).
fof(f1302,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f296,f522,f466])).
fof(f1303,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f296,f522,f466])).
fof(f1305,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f296,f522,f466])).
fof(f1592,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl128),
  inference(cnf_transformation,[],[f1592_D])).
fof(f1592_D,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f1311,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f522,f466])).
fof(f1312,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f522,f466])).
fof(f1313,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f522,f466])).
fof(f1584,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl120),
  inference(cnf_transformation,[],[f1584_D])).
fof(f1584_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f1315,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f522,f466])).
fof(f1594,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl130),
  inference(cnf_transformation,[],[f1594_D])).
fof(f1594_D,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f1316,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f514,f522,f466])).
fof(f1317,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f514,f522,f466])).
fof(f1318,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f514,f522,f466])).
fof(f1320,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f514,f522,f466])).
fof(f1321,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f297,f522,f466])).
fof(f1322,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f297,f522,f466])).
fof(f1323,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f297,f522,f466])).
fof(f1325,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f297,f522,f466])).
fof(f1326,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f296,f522,f466])).
fof(f1327,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f296,f522,f466])).
fof(f1328,plain,(
  sP7('T_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f296,f522,f466])).
fof(f1330,plain,(
  sP7('T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f296,f522,f466])).
fof(f1586,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl122),
  inference(cnf_transformation,[],[f1586_D])).
fof(f1586_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f1336,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f522,f466])).
fof(f1337,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f522,f466])).
fof(f1338,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f522,f466])).
fof(f1340,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f522,f466])).
fof(f1590,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl126),
  inference(cnf_transformation,[],[f1590_D])).
fof(f1590_D,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f1341,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f514,f522,f466])).
fof(f1342,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f514,f522,f466])).
fof(f1343,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f514,f522,f466])).
fof(f1345,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f514,f522,f466])).
fof(f1346,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f297,f522,f466])).
fof(f1347,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f297,f522,f466])).
fof(f1348,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f297,f522,f466])).
fof(f1350,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f297,f522,f466])).
fof(f1351,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f514,f296,f522,f466])).
fof(f1352,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f297,f296,f522,f466])).
fof(f1353,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f296,f296,f522,f466])).
fof(f1355,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f522,f296,f522,f466])).
fof(f1361,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f522,f522,f466])).
fof(f1362,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f522,f522,f466])).
fof(f1363,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f522,f522,f466])).
fof(f1365,plain,(
  sP7('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f522,f522,f466])).
fof(f1582,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl119),
  inference(cnf_transformation,[],[f1582_D])).
fof(f1582_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl119),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl119])])).
fof(f1368,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f514,f522,f469])).
fof(f1369,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f297,f522,f469])).
fof(f1370,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f296,f522,f469])).
fof(f1578,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl115),
  inference(cnf_transformation,[],[f1578_D])).
fof(f1578_D,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl115),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl115])])).
fof(f1372,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f522,f469])).
fof(f1580,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl117),
  inference(cnf_transformation,[],[f1580_D])).
fof(f1580_D,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f1373,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f514,f522,f469])).
fof(f1374,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f297,f522,f469])).
fof(f1375,plain,(
  ~sP7('F_0','T_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f296,f522,f469])).
fof(f1377,plain,(
  ~sP7('F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f522,f469])).
fof(f1576,plain,(
  ~sP8('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl113),
  inference(cnf_transformation,[],[f1576_D])).
fof(f1576_D,plain,(
  ~sP8('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f1380,plain,(
  ~sP8('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f522,f477])).
fof(f1574,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl111),
  inference(cnf_transformation,[],[f1574_D])).
fof(f1574_D,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl111),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl111])])).
fof(f1403,plain,(
  ~sP9('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f522,f484])).
fof(f1572,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl109),
  inference(cnf_transformation,[],[f1572_D])).
fof(f1572_D,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl109),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl109])])).
fof(f1407,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f514,f522,f494])).
fof(f1408,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f297,f522,f494])).
fof(f1409,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f296,f522,f494])).
fof(f1568,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl105),
  inference(cnf_transformation,[],[f1568_D])).
fof(f1568_D,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl105])])).
fof(f1411,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f522,f494])).
fof(f1570,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl107),
  inference(cnf_transformation,[],[f1570_D])).
fof(f1570_D,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl107),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl107])])).
fof(f1412,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f522,f494])).
fof(f1413,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f297,f302,f522,f494])).
fof(f1414,plain,(
  ~sP10('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f296,f302,f522,f494])).
fof(f1416,plain,(
  ~sP10('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f522,f494])).
fof(f1566,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl102),
  inference(cnf_transformation,[],[f1566_D])).
fof(f1566_D,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f1418,plain,(
  sP11('F_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f522,f498])).
fof(f1564,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl101),
  inference(cnf_transformation,[],[f1564_D])).
fof(f1564_D,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl101])])).
fof(f1430,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f514,f522,f501])).
fof(f1431,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f297,f522,f501])).
fof(f1432,plain,(
  ~sP11('T_0','F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f296,f522,f501])).
fof(f1560,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl97),
  inference(cnf_transformation,[],[f1560_D])).
fof(f1560_D,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f1434,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f522,f501])).
fof(f1562,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl99),
  inference(cnf_transformation,[],[f1562_D])).
fof(f1562_D,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f1435,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f514,f522,f501])).
fof(f1436,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f297,f522,f501])).
fof(f1437,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl2),
  inference(unit_resulting_resolution,[],[f302,f296,f522,f501])).
fof(f1439,plain,(
  ~sP11('const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)),'F_0','const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f522,f522,f501])).
fof(f1441,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f354,f522,f505])).
fof(f1447,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f354,f522,f507])).
fof(f1454,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f354,f522,f507])).
fof(f1456,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f354,f522,f509])).
fof(f1458,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f514,f522,f510])).
fof(f1459,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f297,f522,f510])).
fof(f1460,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f296,f522,f510])).
fof(f1461,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f355,f522,f510])).
fof(f1462,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f522,f522,f510])).
fof(f1472,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f522,f522,f510])).
fof(f1483,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f522,f302,f354,f522,f511])).
fof(f1490,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f514,f302,f354,f522,f511])).
fof(f1492,plain,(
  s(bool,'F_0') != s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))) | $spl2),
  inference(unit_resulting_resolution,[],[f302,f302,f522,f518])).
fof(f1186,plain,(
  sP5('F_0','F_0','T_0') | $spl92),
  inference(cnf_transformation,[],[f1186_D])).
fof(f1186_D,plain,(
  sP5('F_0','F_0','T_0') <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f1160,plain,(
  sP5('F_0','F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f302,f1036,f449])).
fof(f1184,plain,(
  ~sP5('T_0','F_0','T_0') | $spl91),
  inference(cnf_transformation,[],[f1184_D])).
fof(f1184_D,plain,(
  ~sP5('T_0','F_0','T_0') <=> ~$spl91),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl91])])).
fof(f1161,plain,(
  ~sP5('T_0','F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f514,f1036,f451])).
fof(f1162,plain,(
  ~sP5('T_0','F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f297,f1036,f451])).
fof(f1163,plain,(
  ~sP5('T_0','F_0','T_0') | $spl7),
  inference(unit_resulting_resolution,[],[f296,f1036,f451])).
fof(f1182,plain,(
  sP5('F_0','T_0','F_0') | $spl88),
  inference(cnf_transformation,[],[f1182_D])).
fof(f1182_D,plain,(
  sP5('F_0','T_0','F_0') <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f1170,plain,(
  sP5('F_0','T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f302,f1036,f449])).
fof(f1180,plain,(
  ~sP5('T_0','T_0','F_0') | $spl87),
  inference(cnf_transformation,[],[f1180_D])).
fof(f1180_D,plain,(
  ~sP5('T_0','T_0','F_0') <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f1171,plain,(
  ~sP5('T_0','T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f514,f1036,f451])).
fof(f1172,plain,(
  ~sP5('T_0','T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f297,f1036,f451])).
fof(f1173,plain,(
  ~sP5('T_0','T_0','F_0') | $spl7),
  inference(unit_resulting_resolution,[],[f296,f1036,f451])).
fof(f1153,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl84),
  inference(cnf_transformation,[],[f1153_D])).
fof(f1153_D,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f1003,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(resolution,[],[f520,f319])).
fof(f319,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f189])).
fof(f1002,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(resolution,[],[f520,f317])).
fof(f317,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f188])).
fof(f1034,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl5),
  inference(cnf_transformation,[],[f1034_D])).
fof(f1034_D,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f525,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f312])).
fof(f312,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f186])).
fof(f526,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f314])).
fof(f314,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f527,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f317])).
fof(f528,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f319])).
fof(f535,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f378])).
fof(f378,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f153])).
fof(f537,plain,(
  s(bool,'F_0') = s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f378])).
fof(f1151,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl83),
  inference(cnf_transformation,[],[f1151_D])).
fof(f1151_D,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f539,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f386])).
fof(f386,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f236])).
fof(f1147,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl79),
  inference(cnf_transformation,[],[f1147_D])).
fof(f1147_D,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f540,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f386])).
fof(f1149,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl81),
  inference(cnf_transformation,[],[f1149_D])).
fof(f1149_D,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f541,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f386])).
fof(f542,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f386])).
fof(f543,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f390])).
fof(f390,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f238])).
fof(f544,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f390])).
fof(f545,plain,(
  ~sP0('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f390])).
fof(f546,plain,(
  ~sP0('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f390])).
fof(f1139,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl71),
  inference(cnf_transformation,[],[f1139_D])).
fof(f1139_D,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f547,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f401])).
fof(f401,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f243])).
fof(f1137,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl69),
  inference(cnf_transformation,[],[f1137_D])).
fof(f1137_D,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f548,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f401])).
fof(f1141,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl73),
  inference(cnf_transformation,[],[f1141_D])).
fof(f1141_D,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f549,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f401])).
fof(f550,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f401])).
fof(f1143,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl74),
  inference(cnf_transformation,[],[f1143_D])).
fof(f1143_D,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f551,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f403])).
fof(f552,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f403])).
fof(f553,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f403])).
fof(f1145,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl76),
  inference(cnf_transformation,[],[f1145_D])).
fof(f1145_D,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f555,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f406])).
fof(f556,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f406])).
fof(f557,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f406])).
fof(f559,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f407])).
fof(f560,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f407])).
fof(f561,plain,(
  sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f407])).
fof(f563,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f408])).
fof(f564,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f408])).
fof(f565,plain,(
  sP1('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f408])).
fof(f567,plain,(
  ~sP1('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f410])).
fof(f410,plain,(
  ( ! [X0,X1] : (~sP1(X1,X0) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f244])).
fof(f568,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f410])).
fof(f569,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f410])).
fof(f570,plain,(
  ~sP1('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f410])).
fof(f1135,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl67),
  inference(cnf_transformation,[],[f1135_D])).
fof(f1135_D,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f571,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f520,f433])).
fof(f572,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f520,f433])).
fof(f573,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f520,f433])).
fof(f575,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f520,f433])).
fof(f576,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f520,f433])).
fof(f577,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f520,f433])).
fof(f579,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f520,f433])).
fof(f580,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f520,f433])).
fof(f581,plain,(
  ~sP5('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f520,f433])).
fof(f1133,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl65),
  inference(cnf_transformation,[],[f1133_D])).
fof(f1133_D,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f587,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f520,f437])).
fof(f588,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f520,f437])).
fof(f589,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f520,f437])).
fof(f591,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f520,f437])).
fof(f592,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f520,f437])).
fof(f593,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f520,f437])).
fof(f595,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f520,f437])).
fof(f596,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f520,f437])).
fof(f597,plain,(
  ~sP4('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f520,f437])).
fof(f1131,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl63),
  inference(cnf_transformation,[],[f1131_D])).
fof(f1131_D,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f605,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f520,f441])).
fof(f606,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f520,f441])).
fof(f607,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f520,f441])).
fof(f609,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f520,f441])).
fof(f610,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f520,f441])).
fof(f611,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f520,f441])).
fof(f613,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f520,f441])).
fof(f614,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f520,f441])).
fof(f615,plain,(
  ~sP3('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f520,f441])).
fof(f1129,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl61),
  inference(cnf_transformation,[],[f1129_D])).
fof(f1129_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f623,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f445])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f267])).
fof(f1125,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl57),
  inference(cnf_transformation,[],[f1125_D])).
fof(f1125_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f624,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f445])).
fof(f1121,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl53),
  inference(cnf_transformation,[],[f1121_D])).
fof(f1121_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f625,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f520,f445])).
fof(f1117,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl49),
  inference(cnf_transformation,[],[f1117_D])).
fof(f1117_D,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f626,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f520,f445])).
fof(f1127,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl59),
  inference(cnf_transformation,[],[f1127_D])).
fof(f1127_D,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f627,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f445])).
fof(f628,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f445])).
fof(f1119,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl51),
  inference(cnf_transformation,[],[f1119_D])).
fof(f1119_D,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f629,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f520,f445])).
fof(f630,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f520,f445])).
fof(f1123,plain,(
  ~sP2('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl55),
  inference(cnf_transformation,[],[f1123_D])).
fof(f1123_D,plain,(
  ~sP2('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f631,plain,(
  ~sP2('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f445])).
fof(f632,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f445])).
fof(f633,plain,(
  ~sP2('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f520,f445])).
fof(f634,plain,(
  ~sP2('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f520,f445])).
fof(f1115,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl47),
  inference(cnf_transformation,[],[f1115_D])).
fof(f1115_D,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f647,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f520,f460])).
fof(f648,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f520,f460])).
fof(f649,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f520,f460])).
fof(f651,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f520,f460])).
fof(f652,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f520,f460])).
fof(f653,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f520,f460])).
fof(f655,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f520,f460])).
fof(f656,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f520,f460])).
fof(f657,plain,(
  ~sP6('T_0','T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f520,f460])).
fof(f1113,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl45),
  inference(cnf_transformation,[],[f1113_D])).
fof(f1113_D,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f681,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f520,f469])).
fof(f682,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f520,f469])).
fof(f683,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f520,f469])).
fof(f685,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f520,f469])).
fof(f686,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f520,f469])).
fof(f687,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f520,f469])).
fof(f689,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f520,f469])).
fof(f690,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f520,f469])).
fof(f691,plain,(
  ~sP7('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f520,f469])).
fof(f1111,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl43),
  inference(cnf_transformation,[],[f1111_D])).
fof(f1111_D,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f705,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f520,f477])).
fof(f706,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f520,f477])).
fof(f707,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f520,f477])).
fof(f1107,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl39),
  inference(cnf_transformation,[],[f1107_D])).
fof(f1107_D,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f709,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f520,f477])).
fof(f710,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f520,f477])).
fof(f711,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f520,f477])).
fof(f1109,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl41),
  inference(cnf_transformation,[],[f1109_D])).
fof(f1109_D,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f713,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f520,f477])).
fof(f714,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f520,f477])).
fof(f715,plain,(
  ~sP8('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f520,f477])).
fof(f717,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f520,f477])).
fof(f718,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f520,f477])).
fof(f719,plain,(
  ~sP8('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f520,f477])).
fof(f1105,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl36),
  inference(cnf_transformation,[],[f1105_D])).
fof(f1105_D,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f722,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f283])).
fof(f1097,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl28),
  inference(cnf_transformation,[],[f1097_D])).
fof(f1097_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f723,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f481])).
fof(f1101,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl32),
  inference(cnf_transformation,[],[f1101_D])).
fof(f1101_D,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f724,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f520,f481])).
fof(f1093,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl24),
  inference(cnf_transformation,[],[f1093_D])).
fof(f1093_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f725,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f520,f481])).
fof(f1103,plain,(
  sP9('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl34),
  inference(cnf_transformation,[],[f1103_D])).
fof(f1103_D,plain,(
  sP9('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f726,plain,(
  sP9('F_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f481])).
fof(f1095,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl26),
  inference(cnf_transformation,[],[f1095_D])).
fof(f1095_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f727,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f481])).
fof(f728,plain,(
  sP9('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f520,f481])).
fof(f729,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f520,f481])).
fof(f1099,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl30),
  inference(cnf_transformation,[],[f1099_D])).
fof(f1099_D,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f730,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f302,f520,f481])).
fof(f731,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f520,f302,f520,f481])).
fof(f732,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f520,f520,f481])).
fof(f733,plain,(
  sP9('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f520,f520,f481])).
fof(f1091,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl23),
  inference(cnf_transformation,[],[f1091_D])).
fof(f1091_D,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f734,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f520,f484])).
fof(f1087,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl19),
  inference(cnf_transformation,[],[f1087_D])).
fof(f1087_D,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f735,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f514,f520,f484])).
fof(f736,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f520,f484])).
fof(f737,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f297,f520,f484])).
fof(f738,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f520,f484])).
fof(f739,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f296,f520,f484])).
fof(f1089,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl21),
  inference(cnf_transformation,[],[f1089_D])).
fof(f1089_D,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f742,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f520,f484])).
fof(f743,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f514,f520,f484])).
fof(f744,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f520,f484])).
fof(f745,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f297,f520,f484])).
fof(f746,plain,(
  ~sP9('T_0','F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f520,f484])).
fof(f747,plain,(
  ~sP9('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f296,f520,f484])).
fof(f1085,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl17),
  inference(cnf_transformation,[],[f1085_D])).
fof(f1085_D,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f766,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f520,f494])).
fof(f767,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f520,f494])).
fof(f768,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f520,f494])).
fof(f770,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f520,f494])).
fof(f771,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f520,f494])).
fof(f772,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f520,f494])).
fof(f774,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f520,f494])).
fof(f775,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f520,f494])).
fof(f776,plain,(
  ~sP10('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f520,f494])).
fof(f1083,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl14),
  inference(cnf_transformation,[],[f1083_D])).
fof(f1083_D,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f784,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f520,f498])).
fof(f1079,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl10),
  inference(cnf_transformation,[],[f1079_D])).
fof(f1079_D,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f785,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f520,f514,f520,f498])).
fof(f786,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f520,f498])).
fof(f787,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f520,f297,f520,f498])).
fof(f788,plain,(
  sP11('F_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f520,f498])).
fof(f789,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f520,f296,f520,f498])).
fof(f1081,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl12),
  inference(cnf_transformation,[],[f1081_D])).
fof(f1081_D,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f792,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f520,f498])).
fof(f793,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f520,f498])).
fof(f794,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'F_0','T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f520,f498])).
fof(f796,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f520,f498])).
fof(f797,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f520,f498])).
fof(f798,plain,(
  sP11('const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f520,f498])).
fof(f1077,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl9),
  inference(cnf_transformation,[],[f1077_D])).
fof(f1077_D,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f804,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f520,f501])).
fof(f805,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f520,f501])).
fof(f806,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f520,f501])).
fof(f808,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f520,f501])).
fof(f809,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f520,f501])).
fof(f810,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f520,f501])).
fof(f812,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f520,f501])).
fof(f813,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f520,f501])).
fof(f814,plain,(
  ~sP11('T_0','const.util_prob.countable_1'(s(fun(sK12,bool),sK13)),'T_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f520,f501])).
fof(f828,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f354,f520,f505])).
fof(f829,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f354,f520,f505])).
fof(f830,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f354,f520,f505])).
fof(f832,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f354,f520,f505])).
fof(f833,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f354,f520,f505])).
fof(f834,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f354,f520,f505])).
fof(f1036,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl7),
  inference(cnf_transformation,[],[f1036_D])).
fof(f1036_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f837,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f354,f520,f505])).
fof(f838,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f514,f354,f520,f505])).
fof(f839,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f354,f520,f505])).
fof(f840,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f297,f354,f520,f505])).
fof(f841,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f354,f520,f505])).
fof(f842,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f296,f354,f520,f505])).
fof(f851,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f354,f520,f507])).
fof(f852,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f354,f520,f507])).
fof(f853,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f354,f520,f507])).
fof(f855,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f354,f520,f507])).
fof(f856,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f354,f520,f507])).
fof(f857,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f354,f520,f507])).
fof(f859,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f354,f520,f507])).
fof(f860,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f354,f520,f507])).
fof(f861,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f354,f520,f507])).
fof(f863,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f355,f354,f520,f507])).
fof(f864,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f355,f354,f520,f507])).
fof(f865,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f355,f354,f520,f507])).
fof(f868,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f354,f520,f509])).
fof(f869,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f354,f520,f509])).
fof(f870,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f354,f520,f509])).
fof(f872,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f354,f520,f509])).
fof(f873,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f354,f520,f509])).
fof(f874,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f354,f520,f509])).
fof(f877,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f354,f520,f509])).
fof(f878,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f514,f354,f520,f509])).
fof(f879,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f354,f520,f509])).
fof(f880,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f297,f354,f520,f509])).
fof(f881,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f354,f520,f509])).
fof(f882,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f520,f296,f354,f520,f509])).
fof(f891,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f514,f520,f510])).
fof(f892,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f514,f520,f510])).
fof(f893,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f514,f520,f510])).
fof(f896,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f297,f520,f510])).
fof(f897,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f297,f520,f510])).
fof(f898,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f297,f520,f510])).
fof(f901,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f296,f520,f510])).
fof(f902,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f296,f520,f510])).
fof(f903,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f296,f520,f510])).
fof(f906,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f355,f520,f510])).
fof(f907,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f355,f520,f510])).
fof(f908,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f355,f520,f510])).
fof(f915,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f514,f520,f510])).
fof(f916,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f514,f520,f510])).
fof(f917,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f514,f520,f510])).
fof(f919,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f514,f520,f510])).
fof(f920,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f514,f520,f510])).
fof(f921,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f514,f520,f510])).
fof(f923,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f297,f520,f510])).
fof(f924,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f297,f520,f510])).
fof(f925,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f297,f520,f510])).
fof(f927,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f297,f520,f510])).
fof(f928,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f297,f520,f510])).
fof(f929,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f297,f520,f510])).
fof(f931,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f296,f520,f510])).
fof(f932,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f296,f520,f510])).
fof(f933,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f296,f520,f510])).
fof(f935,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f296,f520,f510])).
fof(f936,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f296,f520,f510])).
fof(f937,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f296,f520,f510])).
fof(f939,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f355,f520,f510])).
fof(f940,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f355,f520,f510])).
fof(f941,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f355,f520,f510])).
fof(f943,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f355,f520,f510])).
fof(f944,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f355,f520,f510])).
fof(f945,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f355,f520,f510])).
fof(f965,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f514,f354,f520,f511])).
fof(f966,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f514,f354,f520,f511])).
fof(f967,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f514,f354,f520,f511])).
fof(f969,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f297,f354,f520,f511])).
fof(f970,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f297,f354,f520,f511])).
fof(f971,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f297,f354,f520,f511])).
fof(f973,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f296,f354,f520,f511])).
fof(f974,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f296,f354,f520,f511])).
fof(f975,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f296,f354,f520,f511])).
fof(f977,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f355,f354,f520,f511])).
fof(f978,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f355,f354,f520,f511])).
fof(f979,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f355,f354,f520,f511])).
fof(f982,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f514,f520,f518])).
fof(f983,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f297,f520,f518])).
fof(f984,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f302,f296,f520,f518])).
fof(f988,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f514,f302,f520,f518])).
fof(f989,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f297,f302,f520,f518])).
fof(f990,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl1),
  inference(unit_resulting_resolution,[],[f296,f302,f520,f518])).
fof(f992,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f514,f520,f520,f518])).
fof(f993,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f297,f520,f520,f518])).
fof(f994,plain,(
  s(bool,'T_0') != s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))) | $spl1),
  inference(unit_resulting_resolution,[],[f296,f520,f520,f518])).
fof(f522,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) | $spl2),
  inference(cnf_transformation,[],[f522_D])).
fof(f522_D,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f294,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13))))),
  inference(cnf_transformation,[],[f183])).
fof(f183,plain,(
  p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',sK12),sK14),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(sK12,bool),sK13)))) & ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f136])).
fof(f136,plain,(
  ? [X0,X1] : (? [X2] : p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',X0),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X1)))) & ~p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ~! [X0,X1] : (? [X2] : p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',X0),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X1)))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))))),
  inference(flattening,[],[f55])).
fof(f55,plain,(
  ~! [X0] : ! [X1] : (? [X2] : p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',X0),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(X0,bool),X1)))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X0,bool),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (? [X8] : p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',X6),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(X6,bool),X7)))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (? [X8] : p(s(bool,'const.pred_set.BIJ_3'(s(fun('type.num.num',X6),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(X6,bool),X7)))) => p(s(bool,'const.util_prob.countable_1'(s(fun(X6,bool),X7)))))),
  file('thm.util_prob.BIJ_NUM_COUNTABLE','thm.util_prob.BIJ_NUM_COUNTABLE')).
fof(f520,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) | $spl1),
  inference(cnf_transformation,[],[f520_D])).
fof(f520_D,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f295,plain,(
  ~p(s(bool,'const.util_prob.countable_1'(s(fun(sK12,bool),sK13))))),
  inference(cnf_transformation,[],[f183])).
% SZS output end Proof for thm.util_prob
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 329
% Passive clauses: 487
% Generated clauses: 4208
% Final active clauses: 325
% Final passive clauses: 304
% Input formulas: 53
% Initial clauses: 218
% 
% Duplicate literals: 40
% Trivial inequalities: 5
% Fw demodulations: 811
% Bw demodulations: 11
% Global subsumptions: 5
% 
% Simple tautologies: 166
% Equational tautologies: 15
% Forward subsumptions: 2063
% Backward subsumptions: 2
% 
% Binary resolution: 613
% Unit resulting resolution: 2201
% Forward superposition: 191
% Backward superposition: 185
% Self superposition: 1
% 
% Split clauses: 28
% Split components: 56
% SAT solver clauses: 1866
% SAT solver unit clauses: 4
% SAT solver binary clauses: 1623
% 
% Sat splits: 28
% 
% Memory used [KB]: 4221
% Time elapsed: 0.104 s
% ------------------------------
% ------------------------------
% Success in time 18.139 s
