Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.real
% (101289)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1534
% Passive clauses: 9394
% Generated clauses: 50312
% Final active clauses: 1235
% Final passive clauses: 7072
% Input formulas: 517
% Initial clauses: 818
% 
% Duplicate literals: 234
% Trivial inequalities: 114
% Fw subsumption resolutions: 2717
% Fw demodulations: 4406
% Bw demodulations: 293
% 
% Simple tautologies: 299
% Equational tautologies: 38
% Forward subsumptions: 33087
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 9018
% Factoring: 834
% Forward superposition: 15038
% Backward superposition: 16511
% Self superposition: 96
% Equality factoring: 2
% Equality resolution: 25
% 
% Split clauses: 382
% Split components: 766
% SAT solver clauses: 732
% SAT solver unit clauses: 34
% SAT solver binary clauses: 146
% 
% Sat splits: 382
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 732
% TWLsolver calls for satisfiability: 593
% 
% Memory used [KB]: 29295
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.real
% (101303)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 764
% Passive clauses: 3427
% Generated clauses: 29556
% Final active clauses: 628
% Final passive clauses: 2250
% Input formulas: 517
% Initial clauses: 818
% 
% Duplicate literals: 103
% Trivial inequalities: 126
% Fw subsumption resolutions: 806
% Fw demodulations: 3853
% Bw demodulations: 543
% Condensations: 85
% Global subsumptions: 37
% 
% Simple tautologies: 277
% Equational tautologies: 19
% Forward subsumptions: 16636
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 2721
% Factoring: 163
% Forward superposition: 9362
% Backward superposition: 7293
% Self superposition: 100
% Equality factoring: 2
% Equality resolution: 26
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 5083
% SAT solver unit clauses: 1118
% SAT solver binary clauses: 256
% SAT solver learnt clauses: 74
% SAT solver learnt literals: 12
% 
% Sat splits: 19
% 
% TWLsolver clauses: 4935
% TWLsolver calls for satisfiability: 4926
% 
% Memory used [KB]: 22899
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.real
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.real
% SZS output start Proof for thm.real
fof(f27342,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f27341])).
fof(f27341,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0')))),
  inference(forward_demodulation,[],[f27340,f2520])).
fof(f2520,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_|49|_0') = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_|49|_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f2519,f2448])).
fof(f2448,plain,(
  s('type.realax.real','const.realax.real_|49|_0') = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f1613,f1610])).
fof(f1610,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f72])).
fof(f72,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.real.POW_M|49|','thm.arithmetic.ONE')).
fof(f1613,plain,(
  s('type.realax.real','const.realax.real_|49|_0') = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f26])).
fof(f26,axiom,(
  s('type.realax.real','const.realax.real_|49|_0') = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.real.POW_M|49|','thm.real.REAL_|49|')).
fof(f2519,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f1740,f1610])).
fof(f1740,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f636])).
fof(f636,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num',X6)))),
  file('thm.real.POW_M|49|','thm.real.POW_ONE')).
fof(f27340,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_|49|_0'),s('type.num.num',sK15)))))),
  inference(forward_demodulation,[],[f27339,f2461])).
fof(f2461,plain,(
  s('type.realax.real','const.realax.real_|49|_0') = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_|49|_0')))),
  inference(forward_demodulation,[],[f2460,f2448])).
fof(f2460,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(forward_demodulation,[],[f1620,f1610])).
fof(f1620,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(cnf_transformation,[],[f24])).
fof(f24,axiom,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.real.POW_M|49|','thm.real.ABS_|49|')).
fof(f27339,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_|49|_0'))),s('type.num.num',sK15)))))),
  inference(forward_demodulation,[],[f27338,f1672])).
fof(f1672,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f568])).
fof(f568,plain,(
  ! [X0] : s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X8] : s('type.realax.real','const.real.abs_1'(s('type.realax.real',X8))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8)))))),
  file('thm.real.POW_M|49|','thm.real.ABS_NEG')).
fof(f27338,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))))),s('type.num.num',sK15)))))),
  inference(forward_demodulation,[],[f27229,f2003])).
fof(f2003,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))),s('type.num.num',X1))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f871])).
fof(f871,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X0))),s('type.num.num',X1))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X13,X6] : s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.abs_1'(s('type.realax.real',X13))),s('type.num.num',X6))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X13),s('type.num.num',X6)))))),
  file('thm.real.POW_M|49|','thm.real.POW_ABS')).
fof(f27229,plain,(
  s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))),s('type.num.num',sK15)))))))),
  inference(unit_resulting_resolution,[],[f2449,f2130])).
fof(f2130,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f1492])).
fof(f1492,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(nnf_transformation,[],[f988])).
fof(f988,plain,(
  ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f167])).
fof(f167,axiom,(
  ! [X8,X12] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))) <=> s('type.realax.real',X8) = s('type.realax.real',X12))),
  file('thm.real.POW_M|49|','thm.real.REAL_EQ_NEG')).
fof(f2449,plain,(
  s('type.realax.real','const.realax.real_|49|_0') != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_|49|_0'))),s('type.num.num',sK15)))))),
  inference(backward_demodulation,[],[f2448,f2445])).
fof(f2445,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))),s('type.num.num',sK15)))))),
  inference(backward_demodulation,[],[f1610,f1601])).
fof(f1601,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',sK15)))))),
  inference(cnf_transformation,[],[f1372])).
fof(f1372,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',sK15)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f1145])).
fof(f1145,plain,(
  ? [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) != s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',X0)))))),
  inference(ennf_transformation,[],[f519])).
fof(f519,plain,(
  ~! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',X0)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',X6)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.abs_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),s('type.num.num',X6)))))),
  file('thm.real.POW_M|49|','thm.real.POW_M|49|')).
% SZS output end Proof for thm.real
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 674
% Passive clauses: 2970
% Generated clauses: 25726
% Final active clauses: 635
% Final passive clauses: 2202
% Input formulas: 517
% Initial clauses: 818
% 
% Duplicate literals: 8
% Trivial inequalities: 46
% Fw subsumption resolutions: 159
% Fw demodulations: 4691
% Bw demodulations: 133
% Condensations: 19
% 
% Simple tautologies: 114
% Equational tautologies: 25
% Forward subsumptions: 17575
% Fw demodulations to eq. taut.: 4
% 
% Unit resulting resolution: 5881
% Factoring: 1
% Forward superposition: 6019
% Backward superposition: 7842
% Self superposition: 95
% Equality factoring: 2
% Equality resolution: 7
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 2989
% SAT solver unit clauses: 1777
% SAT solver binary clauses: 941
% 
% Sat splits: 8
% Sat splitting refutations: 1
% 
% Memory used [KB]: 15735
% Time elapsed: 0.404 s
% ------------------------------
% ------------------------------
% Success in time 1.702 s
