Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.gcd
% (102862)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1076
% Passive clauses: 7707
% Generated clauses: 48529
% Final active clauses: 1061
% Final passive clauses: 6311
% Input formulas: 515
% Initial clauses: 851
% 
% Duplicate literals: 150
% Trivial inequalities: 175
% Fw subsumption resolutions: 1720
% Fw demodulations: 10418
% Bw demodulations: 14
% 
% Simple tautologies: 347
% Equational tautologies: 24
% Forward subsumptions: 27671
% Fw demodulations to eq. taut.: 6
% 
% Binary resolution: 6496
% Factoring: 419
% Forward superposition: 16756
% Backward superposition: 11270
% Self superposition: 49
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 519
% Split components: 1040
% SAT solver clauses: 628
% SAT solver unit clauses: 27
% SAT solver binary clauses: 61
% 
% Sat splits: 518
% 
% TWLsolver clauses: 628
% TWLsolver calls for satisfiability: 400
% 
% Memory used [KB]: 23539
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.gcd
% (102873)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 656
% Passive clauses: 3395
% Generated clauses: 28505
% Final active clauses: 633
% Final passive clauses: 2550
% Input formulas: 515
% Initial clauses: 851
% 
% Duplicate literals: 122
% Trivial inequalities: 33
% Fw subsumption resolutions: 991
% Fw demodulations: 3375
% Bw demodulations: 205
% Condensations: 118
% Global subsumptions: 28
% 
% Simple tautologies: 328
% Equational tautologies: 14
% Forward subsumptions: 15964
% Fw demodulations to eq. taut.: 8
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 3324
% Factoring: 306
% Forward superposition: 9197
% Backward superposition: 6460
% Self superposition: 90
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 97
% Split components: 194
% SAT solver clauses: 5009
% SAT solver unit clauses: 1157
% SAT solver binary clauses: 238
% SAT solver learnt clauses: 56
% SAT solver learnt literals: 16
% 
% Sat splits: 97
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 4897
% TWLsolver calls for satisfiability: 4456
% 
% Memory used [KB]: 24562
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.gcd
% (102881)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1134
% Passive clauses: 14074
% Generated clauses: 120271
% Final active clauses: 878
% Final passive clauses: 2593
% Input formulas: 515
% Initial clauses: 851
% Discarded non-redundant clauses: 4586
% 
% Duplicate literals: 13
% Trivial inequalities: 135
% Fw subsumption resolutions: 1884
% Fw demodulations: 24561
% Bw demodulations: 5919
% Condensations: 71
% Global subsumptions: 1165
% 
% Simple tautologies: 190
% Equational tautologies: 89
% Forward subsumptions: 67071
% Fw demodulations to eq. taut.: 708
% Bw demodulations to eq. taut.: 29
% 
% Unit resulting resolution: 33935
% Factoring: 1
% Forward superposition: 30535
% Backward superposition: 20226
% Self superposition: 158
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 19
% Split components: 38
% SAT solver clauses: 18223
% SAT solver unit clauses: 9148
% SAT solver binary clauses: 8185
% 
% Sat splits: 19
% Sat splitting refutations: 2943
% 
% Memory used [KB]: 234750
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.gcd
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.gcd
% SZS output start Proof for thm.gcd
fof(f10975,plain,(
  $false),
  inference(subsumption_resolution,[],[f10710,f1193])).
fof(f1193,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))) != s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))),s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f869,f921])).
fof(f921,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f534])).
fof(f534,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6)))),
  file('thm.gcd.gcd_def_compute|split|2','thm.arithmetic.NUMERAL_DEF')).
fof(f869,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))))) != s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))))),s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f720])).
fof(f720,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))))) != s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))))),s('type.num.num','const.num.|48|_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f670])).
fof(f670,plain,(
  ? [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) != s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))),
  inference(ennf_transformation,[],[f517])).
fof(f517,plain,(
  ~! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.num.num','const.num.|48|_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X6))))),s('type.num.num','const.num.|48|_0')))),
  file('thm.gcd.gcd_def_compute|split|2','thm.gcd.gcd_def_compute|split|2')).
fof(f10710,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f924,f1310])).
fof(f1310,plain,(
  s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.arithmetic.BIT2_1'(s('type.num.num',sK13))))))),
  inference(subsumption_resolution,[],[f1223,f925])).
fof(f925,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f538])).
fof(f538,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.gcd.gcd_def_compute|split|2','thm.gcd.gcd_def|split|0')).
fof(f1223,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',sK13))) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16('const.arithmetic.BIT2_1'(s('type.num.num',sK13))))))),
  inference(superposition,[],[f1193,f956])).
fof(f956,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f756])).
fof(f756,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',sK16(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f563])).
fof(f563,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) | ? [X1] : s('type.num.num',X0) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f481])).
fof(f481,axiom,(
  ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10) | ? [X11] : s('type.num.num',X10) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))))),
  file('thm.gcd.gcd_def_compute|split|2','thm.arithmetic.num_CASES')).
fof(f924,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f537])).
fof(f537,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))) = s('type.num.num','const.gcd.gcd_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.num.|48|_0')))),
  file('thm.gcd.gcd_def_compute|split|2','thm.gcd.gcd_def|split|1')).
% SZS output end Proof for thm.gcd
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 410
% Passive clauses: 3406
% Generated clauses: 10075
% Final active clauses: 387
% Final passive clauses: 2351
% Input formulas: 515
% Initial clauses: 316
% 
% Selected by SInE selection: 116
% SInE iterations: 1
% 
% Duplicate literals: 130
% Trivial inequalities: 188
% Fw subsumption resolutions: 486
% Fw demodulations: 1385
% Bw demodulations: 6
% 
% Simple tautologies: 306
% Equational tautologies: 2
% Forward subsumptions: 3425
% 
% Binary resolution: 3196
% Forward superposition: 3539
% Backward superposition: 744
% Equality resolution: 2
% 
% Split clauses: 320
% Split components: 657
% SAT solver clauses: 347
% SAT solver unit clauses: 20
% SAT solver binary clauses: 205
% 
% Sat splits: 320
% Sat splitting refutations: 26
% 
% Memory used [KB]: 5373
% Time elapsed: 0.155 s
% ------------------------------
% ------------------------------
% Success in time 11.917 s
