Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.quotient
% (107895)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 122
% Passive clauses: 2014
% Generated clauses: 20243
% Final active clauses: 99
% Final passive clauses: 1693
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 1208
% Trivial inequalities: 216
% Fw subsumption resolutions: 3061
% Fw demodulations: 1902
% Bw demodulations: 19
% 
% Simple tautologies: 345
% Equational tautologies: 62
% Forward subsumptions: 9593
% 
% Binary resolution: 339
% Factoring: 2
% Forward superposition: 7487
% Backward superposition: 5280
% Self superposition: 270
% Equality factoring: 184
% Equality resolution: 13
% 
% Split clauses: 1279
% Split components: 2603
% SAT solver clauses: 2004
% SAT solver unit clauses: 47
% SAT solver binary clauses: 222
% 
% Sat splits: 1279
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1963
% TWLsolver calls for satisfiability: 191
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% (107910)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 94
% Passive clauses: 632
% Generated clauses: 11660
% Final active clauses: 79
% Final passive clauses: 391
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 513
% Trivial inequalities: 70
% Fw subsumption resolutions: 989
% Fw demodulations: 1922
% Bw demodulations: 109
% Condensations: 554
% 
% Simple tautologies: 111
% Equational tautologies: 30
% Forward subsumptions: 5075
% Bw demodulations to eq. taut.: 51
% 
% Binary resolution: 299
% Factoring: 8
% Forward superposition: 3440
% Backward superposition: 2807
% Self superposition: 92
% Equality factoring: 66
% Equality resolution: 7
% 
% Split clauses: 227
% Split components: 478
% SAT solver clauses: 1494
% SAT solver unit clauses: 18
% SAT solver binary clauses: 25
% 
% Sat splits: 227
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 1494
% TWLsolver calls for satisfiability: 149
% 
% Memory used [KB]: 9210
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.quotient
% (107930)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 448
% Passive clauses: 3478
% Generated clauses: 131194
% Final active clauses: 63
% Final passive clauses: 507
% Input formulas: 8
% Initial clauses: 14
% 
% Duplicate literals: 8145
% Trivial inequalities: 941
% Fw subsumption resolutions: 10510
% Fw demodulations: 13076
% Bw demodulations: 2033
% Condensations: 15835
% Global subsumptions: 702
% 
% Simple tautologies: 1068
% Equational tautologies: 2852
% Forward subsumptions: 73434
% Fw demodulations to eq. taut.: 7
% Bw demodulations to eq. taut.: 82
% 
% Unit resulting resolution: 9597
% Factoring: 1
% Forward superposition: 35190
% Backward superposition: 21663
% Self superposition: 657
% Equality factoring: 417
% Equality resolution: 71
% 
% Split clauses: 440
% Split components: 889
% SAT solver clauses: 5797
% SAT solver unit clauses: 39
% SAT solver binary clauses: 485
% 
% Sat splits: 440
% Sat splitting refutations: 735
% 
% Memory used [KB]: 45670
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f9081,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f63,f64,f65,f86,f87,f88,f108,f127,f128,f147,f148,f227,f228,f229,f230,f231,f232,f284,f289,f343,f348,f487,f508,f585,f620,f751,f752,f772,f773,f788,f789,f797,f800,f816,f817,f825,f828,f852,f853,f1061,f1064,f881,f882,f1069,f1070,f904,f905,f919,f922,f947,f948,f962,f965,f697,f1132,f1258,f1259,f1260,f1261,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1320,f82,f1321,f104,f1491,f1493,f1501,f1504,f1509,f1511,f1607,f1608,f1609,f1651,f1653,f1652,f1681,f1682,f1063,f1765,f1766,f1066,f1843,f1846,f1845,f1853,f1854,f1855,f1856,f1857,f1858,f1859,f1860,f1861,f1862,f1863,f1864,f1865,f1866,f1893,f1896,f1895,f1927,f1928,f2138,f2142,f2123,f2124,f2143,f2139,f2144,f2145,f2151,f2314,f2316,f2347,f2348,f2359,f2360,f2361,f2362,f2363,f2364,f2365,f2366,f2367,f2368,f2397,f2398,f2533,f2534,f2559,f2560,f235,f2597,f2599,f2598,f2656,f2658,f2315,f2722,f2724,f2921,f2922,f2934,f2935,f2940,f2941,f2951,f2952,f2953,f2954,f2964,f2965,f2966,f2967,f3235,f3237,f3246,f84,f3247,f106,f3251,f3253,f3260,f83,f3261,f105,f3265,f3267,f3272,f3274,f2358,f2357,f244,f241,f2356,f2355,f243,f242,f3325,f3328,f3327,f3326,f3329,f3331,f3371,f3372,f3373,f3374,f3434,f3435,f3330,f3442,f3443,f3475,f3476,f3595,f3597,f3596,f3685,f3686,f2911,f3891,f3226,f3892,f3224,f3893,f2909,f3920,f3921,f3954,f3957,f3956,f3960,f3995,f3998,f3997,f4073,f4191,f3471,f4324,f143,f4325,f144,f4434,f163,f4435,f164,f4515,f4518,f4516,f4610,f3950,f4611,f3951,f4726,f3990,f4800,f4801,f4858,f4859,f4917,f4919,f4977,f4979,f5034,f5035,f5036,f5037,f5038,f5039,f5040,f5041,f5113,f5115,f5202,f5204,f5265,f5267,f5327,f5329,f5389,f5391,f5453,f5454,f5514,f5516,f5576,f5577,f5640,f5641,f5711,f5712,f1048,f6840,f6836,f6835,f6831,f6830,f6827,f6824,f6820,f6819,f6815,f6814,f6813,f6812,f6811,f6810,f6809,f6805,f6802,f6799,f6981,f6983,f6993,f6793,f6995,f6998,f7004,f3273,f7005,f1510,f7010,f3266,f7011,f3252,f7012,f1503,f7022,f7027,f7028,f7029,f7031,f7032,f7034,f7035,f4517,f7036,f7037,f7038,f7040,f7041,f7042,f7043,f7044,f6785,f6784,f5266,f6783,f7051,f7052,f7053,f7054,f7055,f6782,f7056,f6781,f7057,f7058,f7059,f7060,f7061,f7062,f6779,f7075,f7079,f7138,f7141,f7142,f7145,f7147,f6872,f7149,f5328,f7150,f7151,f7152,f7155,f7157,f2141,f7158,f7159,f7160,f7161,f7162,f6756,f6755,f6754,f6753,f6752,f6751,f695,f6750,f6911,f6749,f6912,f6748,f6913,f6747,f6914,f6746,f6915,f6745,f6916,f7165,f7166,f7167,f7168,f7185,f5203,f7194,f4918,f7195,f6740,f5114,f7196,f7197,f5515,f7200,f5390,f7216,f3236,f7217,f1492,f7299,f7343,f7394,f7410,f7412,f7413,f7411,f7414,f7415,f7485,f7492,f7496,f7497,f7499,f7500,f7501,f4978,f7581,f7030,f7608,f6854,f7580,f7156,f7614,f7615,f2657,f6876,f7616,f7661,f7617,f7662,f7667,f7669,f7670,f7668,f7671,f7672,f7719,f7722,f7726,f7728,f7721,f7831,f7828,f7727,f7883,f7929,f7884,f7930,f7938,f7940,f7941,f7939,f7942,f7943,f7992,f7995,f7999,f8001,f7994,f8003,f8004,f8005,f8051,f8054,f8058,f8060,f8063,f8107,f8064,f8108,f8110,f8113,f8117,f8118,f8119,f8112,f8223,f8220,f8000,f8276,f8321,f8277,f8322,f8327,f8329,f8330,f8328,f8331,f8332,f8379,f8382,f8386,f8388,f8381,f8492,f8489,f8387,f8544,f8591,f8545,f8592,f8594,f4612,f3922,f2140,f8605,f8607,f8604,f8603,f8602,f8608,f8611,f3227,f8612,f3225,f8613,f2912,f8655,f8656,f8660,f8661,f8662,f8705,f8663,f2908,f8709,f8710,f8711,f239,f8712,f8713,f8715,f8758,f8762,f8765,f8766,f6982,f8764,f8767,f1502,f8768,f8772,f8815,f8773,f8816,f8818,f8821,f8825,f8826,f8827,f8820,f8829,f8872,f8830,f8873,f8831,f8875,f8877,f8880,f8884,f8885,f8886,f8879,f8895,f8896,f8897,f8898,f8899,f8900,f8901,f8902,f8903,f8904,f8905,f240,f8913,f8906,f8914,f234,f8907,f1275,f1272,f8908,f236,f8915,f8909,f1273,f1274,f8910,f8916,f238,f8964,f8967,f8971,f8972,f8973,f8966,f8975,f8046,f8976,f8047,f8977,f8049,f80,f9019,f9020,f9024,f8059,f288,f9025,f7186,f8053,f9072])).
fof(f9072,plain,(
  $false | ($spl4 | $spl10 | $spl30 | $spl85)),
  inference(subsumption_resolution,[],[f9026,f2908])).
fof(f9026,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) | ($spl4 | $spl10 | $spl30)),
  inference(resolution,[],[f8828,f105])).
fof(f8828,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK5)))) | ($spl4 | $spl30)),
  inference(resolution,[],[f8609,f44])).
fof(f44,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1))))) )),
  inference(resolution,[],[f30,f41])).
fof(f41,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5))))))),
  inference(ennf_transformation,[],[f18])).
fof(f18,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5))))))),
  inference(flattening,[],[f17])).
fof(f17,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X5))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(X7,X6),X9),s(fun(X6,X7),X10)))) => ! [X17,X18] : (p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X17))),s(X7,X18)))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X18))),s(X7,X17))))))),
  file('thm.quotient.EQUALS_RSP','thm.quotient.QUOTIENT_SYM')).
fof(f30,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))))),
  inference(cnf_transformation,[],[f28])).
fof(f28,plain,(
  p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4)))) & (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6)))) & p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8)))) & s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3,sK4,sK5,sK6,sK7,sK8])],[f21])).
fof(f21,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) & ? [X5,X6,X7,X8] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7))),s(X1,X8)))) & s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))) != s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X8)))))),
  inference(flattening,[],[f20])).
fof(f20,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) & ? [X5,X6,X7,X8] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7))),s(X1,X8))))) & s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))) != s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X8)))))),
  inference(ennf_transformation,[],[f11])).
fof(f11,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6,X7,X8] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7))),s(X1,X8))))) => s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X8)))))),
  inference(flattening,[],[f10])).
fof(f10,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6,X7,X8] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X7))),s(X1,X8))))) => s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X8)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(X7,X6),X9),s(fun(X6,X7),X10)))) => ! [X11,X12,X13,X14] : ((p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X12)))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X13))),s(X7,X14))))) => s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X13))) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X12))),s(X7,X14)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(X7,X6),X9),s(fun(X6,X7),X10)))) => ! [X11,X12,X13,X14] : ((p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X12)))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X13))),s(X7,X14))))) => s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X11))),s(X7,X13))) = s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X12))),s(X7,X14)))))),
  file('thm.quotient.EQUALS_RSP','thm.quotient.EQUALS_RSP')).
fof(f8609,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8)))) | ($spl4 | $spl30)),
  inference(resolution,[],[f242,f83])).
fof(f8053,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | $spl242),
  inference(cnf_transformation,[],[f8053_D])).
fof(f8053_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f7186,plain,(
  ( ! [X110,X113] : (p(s(bool,X110)) | s(bool,X110) = s(bool,X113)) ) | $spl200),
  inference(cnf_transformation,[],[f7186_D])).
fof(f7186_D,plain,(
  ( ! [X110,X113] : (p(s(bool,X110)) | s(bool,X110) = s(bool,X113)) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f9025,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl10 | $spl76 | $spl86)),
  inference(subsumption_resolution,[],[f9015,f2404])).
fof(f2404,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') != s(bool,X4)) ) | $spl76),
  inference(resolution,[],[f2143,f35])).
fof(f35,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f24])).
fof(f24,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f23])).
fof(f23,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f14])).
fof(f14,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X15,X16] : ((p(s(bool,X15)) => p(s(bool,X16))) => ((p(s(bool,X16)) => p(s(bool,X15))) => s(bool,X15) = s(bool,X16)))),
  file('thm.quotient.EQUALS_RSP','thm.bool.IMP_ANTISYM_AX')).
fof(f9015,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl10 | $spl86)),
  inference(superposition,[],[f8770,f401])).
fof(f401,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) = s(bool,X3) | s(bool,'F_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(resolution,[],[f385,f245])).
fof(f245,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(resolution,[],[f210,f35])).
fof(f210,plain,(
  ( ! [X11] : (~p(s(bool,X11)) | s(bool,'T_0') = s(bool,X11)) )),
  inference(forward_demodulation,[],[f209,f62])).
fof(f62,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4)))),
  inference(subsumption_resolution,[],[f55,f42])).
fof(f42,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f19])).
fof(f19,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.quotient.EQUALS_RSP',aHL_FALSITY)).
fof(f55,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4)))),
  inference(superposition,[],[f30,f39])).
fof(f39,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.quotient.EQUALS_RSP',aHL_BOOL_CASES)).
fof(f209,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | ~p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f186,f43])).
fof(f43,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.EQUALS_RSP',aHL_TRUTH)).
fof(f186,plain,(
  ( ! [X11] : (~p(s(bool,'T_0')) | s(bool,X11) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f38,f62])).
fof(f38,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f24])).
fof(f385,plain,(
  ( ! [X2,X1] : (~p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s(bool,'F_0') = s(bool,X1)) )),
  inference(resolution,[],[f357,f38])).
fof(f357,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,'F_0') = s(bool,X1)) )),
  inference(resolution,[],[f42,f35])).
fof(f8770,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) | ($spl10 | $spl86)),
  inference(resolution,[],[f2912,f105])).
fof(f288,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17)) ) | $spl36),
  inference(cnf_transformation,[],[f288_D])).
fof(f288_D,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f8059,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21)) ) | $spl244),
  inference(cnf_transformation,[],[f8059_D])).
fof(f8059_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f9024,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl10 | $spl86)),
  inference(subsumption_resolution,[],[f9014,f210])).
fof(f9014,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl10 | $spl86)),
  inference(superposition,[],[f8770,f401])).
fof(f9020,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl10 | $spl86)),
  inference(subsumption_resolution,[],[f9004,f210])).
fof(f9004,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl10 | $spl86)),
  inference(superposition,[],[f8770,f266])).
fof(f266,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X1)) )),
  inference(resolution,[],[f245,f210])).
fof(f9019,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl10 | $spl86)),
  inference(subsumption_resolution,[],[f9003,f210])).
fof(f9003,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl10 | $spl86)),
  inference(superposition,[],[f8770,f266])).
fof(f80,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X0),X1),s(fun(X0,sK1),X2))))) ) | $spl0),
  inference(cnf_transformation,[],[f80_D])).
fof(f80_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X0),X1),s(fun(X0,sK1),X2))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f8049,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) | $spl240),
  inference(cnf_transformation,[],[f8049_D])).
fof(f8049_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f8977,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) ) | ($spl10 | $spl86)),
  inference(resolution,[],[f8770,f41])).
fof(f8047,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7))))) ) | $spl238),
  inference(cnf_transformation,[],[f8047_D])).
fof(f8047_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7))))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f8976,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | ($spl10 | $spl86)),
  inference(resolution,[],[f8770,f40])).
fof(f40,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4))))) )),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6,X7] : (~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))))))),
  inference(flattening,[],[f25])).
fof(f25,plain,(
  ! [X0,X1,X2,X3,X4] : (~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) | ! [X5,X6,X7] : ((~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7))))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))))))),
  inference(ennf_transformation,[],[f16])).
fof(f16,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))))))),
  inference(flattening,[],[f15])).
fof(f15,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X1,fun(X1,bool)),X2),s(fun(X1,X0),X3),s(fun(X0,X1),X4)))) => ! [X5,X6,X7] : ((p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X6)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X6))),s(X1,X7))))) => p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X1,fun(X1,bool)),X2),s(X1,X5))),s(X1,X7))))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X7,fun(X7,bool)),X8),s(fun(X7,X6),X9),s(fun(X6,X7),X10)))) => ! [X17,X18,X19] : ((p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X17))),s(X7,X18)))) & p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X18))),s(X7,X19))))) => p(s(bool,happ(s(fun(X7,bool),happ(s(fun(X7,fun(X7,bool)),X8),s(X7,X17))),s(X7,X19))))))),
  file('thm.quotient.EQUALS_RSP','thm.quotient.QUOTIENT_TRANS')).
fof(f8046,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0))))) ) | $spl236),
  inference(cnf_transformation,[],[f8046_D])).
fof(f8046_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0))))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f8975,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | ($spl10 | $spl86)),
  inference(resolution,[],[f8770,f40])).
fof(f8966,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) | $spl310),
  inference(cnf_transformation,[],[f8966_D])).
fof(f8966_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f8973,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8)))) ) | ($spl8 | $spl12 | $spl76)),
  inference(subsumption_resolution,[],[f8959,f2404])).
fof(f8959,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8)))) ) | ($spl8 | $spl12)),
  inference(superposition,[],[f8598,f401])).
fof(f8598,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8)))) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f8568,f43])).
fof(f8568,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8)))) | ($spl8 | $spl12)),
  inference(superposition,[],[f106,f165])).
fof(f165,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | $spl8),
  inference(subsumption_resolution,[],[f156,f42])).
fof(f156,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | $spl8),
  inference(superposition,[],[f104,f39])).
fof(f8972,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21)) ) | $spl312),
  inference(cnf_transformation,[],[f8972_D])).
fof(f8972_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f8971,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f8958,f210])).
fof(f8958,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl8 | $spl12)),
  inference(superposition,[],[f8598,f401])).
fof(f8967,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8)))) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f8948,f210])).
fof(f8948,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8)))) ) | ($spl8 | $spl12)),
  inference(superposition,[],[f8598,f266])).
fof(f8964,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f8947,f210])).
fof(f8947,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl8 | $spl12)),
  inference(superposition,[],[f8598,f266])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | $spl26),
  inference(cnf_transformation,[],[f238_D])).
fof(f238_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f8916,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) = s(bool,X23)) ) | $spl308),
  inference(cnf_transformation,[],[f8916_D])).
fof(f8916_D,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X23) | s(bool,X22) = s(bool,X23)) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f8910,plain,(
  ( ! [X23,X22] : (s(bool,X22) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f401])).
fof(f33,plain,(
  s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(cnf_transformation,[],[f28])).
fof(f1274,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X8)) ) | $spl54),
  inference(cnf_transformation,[],[f1274_D])).
fof(f1274_D,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X8)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f1273,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl52),
  inference(cnf_transformation,[],[f1273_D])).
fof(f1273_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f8909,plain,(
  ( ! [X21,X20] : (s(bool,X20) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f33,f401])).
fof(f8915,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X18)) ) | $spl306),
  inference(cnf_transformation,[],[f8915_D])).
fof(f8915_D,plain,(
  ( ! [X19,X18] : (s(bool,X18) = s(bool,X19) | s(bool,X19) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X18)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f236,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | $spl25),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f8908,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X18) = s(bool,X19) | s(bool,'F_0') = s(bool,X18) | s(bool,X19) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f401])).
fof(f1272,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X4)) ) | $spl50),
  inference(cnf_transformation,[],[f1272_D])).
fof(f1272_D,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X4)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f1275,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X9)) ) | $spl56),
  inference(cnf_transformation,[],[f1275_D])).
fof(f1275_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X9)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f8907,plain,(
  ( ! [X17,X16] : (s(bool,X16) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f33,f401])).
fof(f234,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | $spl22),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f8914,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X15)) ) | $spl304),
  inference(cnf_transformation,[],[f8914_D])).
fof(f8914_D,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X15) | s(bool,X14) = s(bool,X15)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f8906,plain,(
  ( ! [X14,X15] : (s(bool,X14) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f33,f401])).
fof(f8913,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X13) | s(bool,X12) = s(bool,X13)) ) | $spl302),
  inference(cnf_transformation,[],[f8913_D])).
fof(f8913_D,plain,(
  ( ! [X12,X13] : (s(bool,X12) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X13) | s(bool,X12) = s(bool,X13)) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f240,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | $spl29),
  inference(cnf_transformation,[],[f240_D])).
fof(f240_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f8905,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X12) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f33,f401])).
fof(f8904,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X11) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f33,f400])).
fof(f400,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X1)) )),
  inference(resolution,[],[f385,f357])).
fof(f8903,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f400])).
fof(f8902,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f33,f400])).
fof(f8901,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X8) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f33,f400])).
fof(f8900,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X7) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f33,f266])).
fof(f8899,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f266])).
fof(f8898,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f33,f266])).
fof(f8897,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X4) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f33,f266])).
fof(f8896,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f8895,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f8879,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) | $spl298),
  inference(cnf_transformation,[],[f8879_D])).
fof(f8879_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f8886,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8)))) ) | ($spl4 | $spl30 | $spl76)),
  inference(subsumption_resolution,[],[f8869,f2404])).
fof(f8869,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8)))) ) | ($spl4 | $spl30)),
  inference(superposition,[],[f8609,f401])).
fof(f8885,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21)) ) | $spl300),
  inference(cnf_transformation,[],[f8885_D])).
fof(f8885_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f8884,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f8868,f210])).
fof(f8868,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl4 | $spl30)),
  inference(superposition,[],[f8609,f401])).
fof(f8880,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8)))) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f8858,f210])).
fof(f8858,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8)))) ) | ($spl4 | $spl30)),
  inference(superposition,[],[f8609,f266])).
fof(f8877,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl30)),
  inference(subsumption_resolution,[],[f8857,f210])).
fof(f8857,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl4 | $spl30)),
  inference(superposition,[],[f8609,f266])).
fof(f8875,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK5)))) | $spl296),
  inference(cnf_transformation,[],[f8875_D])).
fof(f8875_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK5)))) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f8831,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) ) | ($spl4 | $spl30)),
  inference(resolution,[],[f8609,f41])).
fof(f8873,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5))))) ) | $spl294),
  inference(cnf_transformation,[],[f8873_D])).
fof(f8873_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5))))) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f8830,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | ($spl4 | $spl30)),
  inference(resolution,[],[f8609,f40])).
fof(f8872,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0))))) ) | $spl292),
  inference(cnf_transformation,[],[f8872_D])).
fof(f8872_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0))))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f8829,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | ($spl4 | $spl30)),
  inference(resolution,[],[f8609,f40])).
fof(f8820,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | $spl288),
  inference(cnf_transformation,[],[f8820_D])).
fof(f8820_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f8827,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) ) | ($spl76 | $spl86)),
  inference(subsumption_resolution,[],[f8812,f2404])).
fof(f8812,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) ) | $spl86),
  inference(superposition,[],[f2912,f401])).
fof(f8826,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21)) ) | $spl290),
  inference(cnf_transformation,[],[f8826_D])).
fof(f8826_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f8825,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl86),
  inference(subsumption_resolution,[],[f8811,f210])).
fof(f8811,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl86),
  inference(superposition,[],[f2912,f401])).
fof(f8821,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) ) | $spl86),
  inference(subsumption_resolution,[],[f8801,f210])).
fof(f8801,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) ) | $spl86),
  inference(superposition,[],[f2912,f266])).
fof(f8818,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | $spl86),
  inference(subsumption_resolution,[],[f8800,f210])).
fof(f8800,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | $spl86),
  inference(superposition,[],[f2912,f266])).
fof(f8816,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8))))) ) | $spl286),
  inference(cnf_transformation,[],[f8816_D])).
fof(f8816_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8))))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f8773,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl86),
  inference(resolution,[],[f2912,f40])).
fof(f8815,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0))))) ) | $spl284),
  inference(cnf_transformation,[],[f8815_D])).
fof(f8815_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0))))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f8772,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl86),
  inference(resolution,[],[f2912,f40])).
fof(f8768,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X17)) ) | $spl282),
  inference(cnf_transformation,[],[f8768_D])).
fof(f8768_D,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f1502,plain,(
  ( ! [X90] : (s(bool,'F_0') = s(bool,X90)) ) | $spl60),
  inference(cnf_transformation,[],[f1502_D])).
fof(f1502_D,plain,(
  ( ! [X90] : (s(bool,'F_0') = s(bool,X90)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8767,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X17)) ) | $spl85),
  inference(subsumption_resolution,[],[f8752,f357])).
fof(f8752,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X17)) ) | $spl85),
  inference(superposition,[],[f2908,f401])).
fof(f8764,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | $spl280),
  inference(cnf_transformation,[],[f8764_D])).
fof(f8764_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f6982,plain,(
  ( ! [X118,X115] : (~p(s(bool,X115)) | s(bool,X115) = s(bool,X118)) ) | $spl162),
  inference(cnf_transformation,[],[f6982_D])).
fof(f6982_D,plain,(
  ( ! [X118,X115] : (~p(s(bool,X115)) | s(bool,X115) = s(bool,X118)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f8766,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,X14) = s(bool,X15)) ) | ($spl43 | $spl85)),
  inference(subsumption_resolution,[],[f8751,f1223])).
fof(f1223,plain,(
  ( ! [X2,X1] : (~p(s(bool,X2)) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') != s(bool,X1)) ) | $spl43),
  inference(resolution,[],[f1190,f38])).
fof(f1190,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') != s(bool,X5)) ) | $spl43),
  inference(subsumption_resolution,[],[f1172,f42])).
fof(f1172,plain,(
  ( ! [X5] : (s(bool,'T_0') != s(bool,X5) | p(s(bool,X5)) | p(s(bool,'F_0'))) ) | $spl43),
  inference(superposition,[],[f696,f35])).
fof(f696,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl43),
  inference(cnf_transformation,[],[f696_D])).
fof(f696_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f8751,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl85),
  inference(superposition,[],[f2908,f401])).
fof(f8765,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) ) | $spl85),
  inference(subsumption_resolution,[],[f8748,f357])).
fof(f8748,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) ) | $spl85),
  inference(superposition,[],[f2908,f400])).
fof(f8762,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X9)) ) | $spl85),
  inference(subsumption_resolution,[],[f8747,f357])).
fof(f8747,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X9)) ) | $spl85),
  inference(superposition,[],[f2908,f400])).
fof(f8758,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0))))) ) | $spl278),
  inference(cnf_transformation,[],[f8758_D])).
fof(f8758_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0))))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f8715,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl85),
  inference(resolution,[],[f2908,f40])).
fof(f8713,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X17)) ) | $spl276),
  inference(cnf_transformation,[],[f8713_D])).
fof(f8713_D,plain,(
  ( ! [X17] : (s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f8712,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X17)) ) | $spl73),
  inference(subsumption_resolution,[],[f8699,f357])).
fof(f8699,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X17)) ) | $spl73),
  inference(superposition,[],[f2139,f401])).
fof(f239,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | $spl28),
  inference(cnf_transformation,[],[f239_D])).
fof(f239_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f8711,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X14) = s(bool,X15)) ) | ($spl43 | $spl73)),
  inference(subsumption_resolution,[],[f8698,f1223])).
fof(f8698,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X14) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) ) | $spl73),
  inference(superposition,[],[f2139,f401])).
fof(f8710,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl73),
  inference(subsumption_resolution,[],[f8695,f357])).
fof(f8695,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl73),
  inference(superposition,[],[f2139,f400])).
fof(f8709,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X9)) ) | $spl73),
  inference(subsumption_resolution,[],[f8694,f357])).
fof(f8694,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X9)) ) | $spl73),
  inference(superposition,[],[f2139,f400])).
fof(f2908,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) | $spl85),
  inference(cnf_transformation,[],[f2908_D])).
fof(f2908_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f8663,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X4),X5),s(fun(X4,sK1),X6))))) ) | $spl73),
  inference(resolution,[],[f2139,f41])).
fof(f8705,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0))))) ) | $spl274),
  inference(cnf_transformation,[],[f8705_D])).
fof(f8705_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0))))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f8662,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl73),
  inference(resolution,[],[f2139,f40])).
fof(f8661,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | ($spl30 | $spl76)),
  inference(subsumption_resolution,[],[f8651,f2404])).
fof(f8651,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl30),
  inference(superposition,[],[f242,f401])).
fof(f8660,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl30),
  inference(subsumption_resolution,[],[f8650,f210])).
fof(f8650,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | $spl30),
  inference(superposition,[],[f242,f401])).
fof(f8656,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl30),
  inference(subsumption_resolution,[],[f8640,f210])).
fof(f8640,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl30),
  inference(superposition,[],[f242,f266])).
fof(f8655,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | $spl30),
  inference(subsumption_resolution,[],[f8639,f210])).
fof(f8639,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | $spl30),
  inference(superposition,[],[f242,f266])).
fof(f2912,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) | $spl86),
  inference(cnf_transformation,[],[f2912_D])).
fof(f2912_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f8613,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) ) | $spl30),
  inference(resolution,[],[f242,f41])).
fof(f3225,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK6))))) ) | $spl90),
  inference(cnf_transformation,[],[f3225_D])).
fof(f3225_D,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK6))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8612,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl30),
  inference(resolution,[],[f242,f40])).
fof(f3227,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X8))))) ) | $spl94),
  inference(cnf_transformation,[],[f3227_D])).
fof(f3227_D,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X8))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f8611,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl30),
  inference(resolution,[],[f242,f40])).
fof(f8608,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK7)))) ) | $spl272),
  inference(cnf_transformation,[],[f8608_D])).
fof(f8608_D,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK7)))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f8602,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21)) ) | $spl12),
  inference(subsumption_resolution,[],[f8580,f357])).
fof(f8580,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21)) ) | $spl12),
  inference(superposition,[],[f106,f400])).
fof(f8603,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK7)))) ) | $spl12),
  inference(subsumption_resolution,[],[f8581,f357])).
fof(f8581,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK7)))) ) | $spl12),
  inference(superposition,[],[f106,f400])).
fof(f8604,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK7))) | s(bool,X30) = s(bool,X31)) ) | ($spl12 | $spl43)),
  inference(subsumption_resolution,[],[f8584,f1223])).
fof(f8584,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK7))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl12),
  inference(superposition,[],[f106,f401])).
fof(f8607,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK8)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK7)))) ) | $spl270),
  inference(cnf_transformation,[],[f8607_D])).
fof(f8607_D,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK8)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK7)))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f8605,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X34)) ) | $spl12),
  inference(subsumption_resolution,[],[f8585,f357])).
fof(f8585,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X33)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK8)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X34)) ) | $spl12),
  inference(superposition,[],[f106,f401])).
fof(f2140,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | $spl72),
  inference(cnf_transformation,[],[f2140_D])).
fof(f2140_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3922,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) ) | $spl84),
  inference(resolution,[],[f2909,f41])).
fof(f4612,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X12),X13),s(fun(X12,sK1),X14))))) ) | $spl84),
  inference(subsumption_resolution,[],[f4548,f43])).
fof(f4548,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X12),X13),s(fun(X12,sK1),X14))))) ) | $spl84),
  inference(superposition,[],[f41,f3925])).
fof(f3925,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | $spl84),
  inference(resolution,[],[f2909,f210])).
fof(f8594,plain,(
  $false | ($spl6 | $spl12 | $spl22 | $spl84)),
  inference(subsumption_resolution,[],[f8593,f42])).
fof(f8593,plain,(
  p(s(bool,'F_0')) | ($spl6 | $spl12 | $spl22 | $spl84)),
  inference(forward_demodulation,[],[f8546,f234])).
fof(f8546,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | ($spl6 | $spl12 | $spl84)),
  inference(resolution,[],[f106,f8002])).
fof(f8002,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) | ($spl6 | $spl84)),
  inference(resolution,[],[f7934,f44])).
fof(f7934,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) | ($spl6 | $spl84)),
  inference(subsumption_resolution,[],[f7906,f43])).
fof(f7906,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) | ($spl6 | $spl84)),
  inference(superposition,[],[f84,f3925])).
fof(f8592,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X5))))) ) | $spl268),
  inference(cnf_transformation,[],[f8592_D])).
fof(f8592_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X5))))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f8545,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X6),X7),s(fun(X6,sK1),X8))))) ) | $spl12),
  inference(resolution,[],[f106,f41])).
fof(f8591,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK7))))) ) | $spl266),
  inference(cnf_transformation,[],[f8591_D])).
fof(f8591_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK7))))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f8544,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X2),X3),s(fun(X2,sK1),X4))))) ) | $spl12),
  inference(resolution,[],[f106,f40])).
fof(f8387,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21)) ) | $spl264),
  inference(cnf_transformation,[],[f8387_D])).
fof(f8387_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f8489,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X43)) ) | ($spl8 | $spl10)),
  inference(duplicate_literal_removal,[],[f8440])).
fof(f8440,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f401,f8340])).
fof(f8340,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | ($spl8 | $spl10)),
  inference(resolution,[],[f8323,f210])).
fof(f8323,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7)))) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8297,f43])).
fof(f8297,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7)))) | ($spl8 | $spl10)),
  inference(superposition,[],[f105,f165])).
fof(f8492,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X17)) ) | ($spl8 | $spl10)),
  inference(duplicate_literal_removal,[],[f8409])).
fof(f8409,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8340,f401])).
fof(f8381,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | $spl262),
  inference(cnf_transformation,[],[f8381_D])).
fof(f8381_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f8388,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7)))) ) | ($spl8 | $spl10 | $spl76)),
  inference(subsumption_resolution,[],[f8375,f2404])).
fof(f8375,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7)))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8323,f401])).
fof(f8386,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8374,f210])).
fof(f8374,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8323,f401])).
fof(f8382,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7)))) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8364,f210])).
fof(f8364,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7)))) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8323,f266])).
fof(f8379,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl8 | $spl10)),
  inference(subsumption_resolution,[],[f8363,f210])).
fof(f8363,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl8 | $spl10)),
  inference(superposition,[],[f8323,f266])).
fof(f8332,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X32)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X32)))) ) | $spl260),
  inference(cnf_transformation,[],[f8332_D])).
fof(f8332_D,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X32)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X32)))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f8331,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X32)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X32))) | s(bool,'T_0') = s(bool,X34)) ) | $spl10),
  inference(subsumption_resolution,[],[f8315,f357])).
fof(f8315,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X33)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X32)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X32))) | s(bool,'T_0') = s(bool,X34)) ) | $spl10),
  inference(superposition,[],[f105,f401])).
fof(f8328,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X20)))) ) | $spl258),
  inference(cnf_transformation,[],[f8328_D])).
fof(f8328_D,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X20)))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f8330,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X29)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X29))) | s(bool,X30) = s(bool,X31)) ) | ($spl10 | $spl43)),
  inference(subsumption_resolution,[],[f8314,f1223])).
fof(f8314,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X29)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X29))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl10),
  inference(superposition,[],[f105,f401])).
fof(f8329,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X22)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X22)))) ) | $spl10),
  inference(subsumption_resolution,[],[f8311,f357])).
fof(f8311,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X22)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X22)))) ) | $spl10),
  inference(superposition,[],[f105,f400])).
fof(f8327,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X20))) | s(bool,'F_0') = s(bool,X21)) ) | $spl10),
  inference(subsumption_resolution,[],[f8310,f357])).
fof(f8310,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X20))) | s(bool,'F_0') = s(bool,X21)) ) | $spl10),
  inference(superposition,[],[f105,f400])).
fof(f8322,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK8))))) ) | $spl256),
  inference(cnf_transformation,[],[f8322_D])).
fof(f8322_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK8))))) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f8277,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X6),X7),s(fun(X6,sK1),X8))))) ) | $spl10),
  inference(resolution,[],[f105,f41])).
fof(f8321,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0))))) ) | $spl254),
  inference(cnf_transformation,[],[f8321_D])).
fof(f8321_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0))))) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f8276,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X1)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X2),X3),s(fun(X2,sK1),X4))))) ) | $spl10),
  inference(resolution,[],[f105,f40])).
fof(f8000,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21)) ) | $spl234),
  inference(cnf_transformation,[],[f8000_D])).
fof(f8000_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f8220,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X43)) ) | ($spl2 | $spl6)),
  inference(duplicate_literal_removal,[],[f8171])).
fof(f8171,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f401,f7952])).
fof(f7952,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | ($spl2 | $spl6)),
  inference(resolution,[],[f7933,f210])).
fof(f7933,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6)))) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f7905,f43])).
fof(f7905,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6)))) | ($spl2 | $spl6)),
  inference(superposition,[],[f84,f145])).
fof(f145,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | $spl2),
  inference(subsumption_resolution,[],[f136,f42])).
fof(f136,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | $spl2),
  inference(superposition,[],[f82,f39])).
fof(f8223,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X17)) ) | ($spl2 | $spl6)),
  inference(duplicate_literal_removal,[],[f8139])).
fof(f8139,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f7952,f401])).
fof(f8112,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) | $spl250),
  inference(cnf_transformation,[],[f8112_D])).
fof(f8112_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f8119,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) ) | ($spl6 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f8103,f2404])).
fof(f8103,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f8002,f401])).
fof(f8118,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21)) ) | $spl252),
  inference(cnf_transformation,[],[f8118_D])).
fof(f8118_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f8117,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl6 | $spl84)),
  inference(subsumption_resolution,[],[f8102,f210])).
fof(f8102,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f8002,f401])).
fof(f8113,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) ) | ($spl6 | $spl84)),
  inference(subsumption_resolution,[],[f8092,f210])).
fof(f8092,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f8002,f266])).
fof(f8110,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl6 | $spl84)),
  inference(subsumption_resolution,[],[f8091,f210])).
fof(f8091,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f8002,f266])).
fof(f8108,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6))))) ) | $spl248),
  inference(cnf_transformation,[],[f8108_D])).
fof(f8108_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f8064,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | ($spl6 | $spl84)),
  inference(resolution,[],[f8002,f40])).
fof(f8107,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0))))) ) | $spl246),
  inference(cnf_transformation,[],[f8107_D])).
fof(f8107_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0))))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f8063,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | ($spl6 | $spl84)),
  inference(resolution,[],[f8002,f40])).
fof(f8060,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl6 | $spl76 | $spl84)),
  inference(subsumption_resolution,[],[f8043,f2404])).
fof(f8043,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f7934,f401])).
fof(f8058,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl6 | $spl84)),
  inference(subsumption_resolution,[],[f8042,f210])).
fof(f8042,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f7934,f401])).
fof(f8054,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl6 | $spl84)),
  inference(subsumption_resolution,[],[f8032,f210])).
fof(f8032,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6)))) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f7934,f266])).
fof(f8051,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl6 | $spl84)),
  inference(subsumption_resolution,[],[f8031,f210])).
fof(f8031,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl6 | $spl84)),
  inference(superposition,[],[f7934,f266])).
fof(f8005,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) ) | ($spl6 | $spl84)),
  inference(resolution,[],[f7934,f41])).
fof(f8004,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | ($spl6 | $spl84)),
  inference(resolution,[],[f7934,f40])).
fof(f8003,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | ($spl6 | $spl84)),
  inference(resolution,[],[f7934,f40])).
fof(f7994,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | $spl232),
  inference(cnf_transformation,[],[f7994_D])).
fof(f7994_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f8001,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6)))) ) | ($spl2 | $spl6 | $spl76)),
  inference(subsumption_resolution,[],[f7987,f2404])).
fof(f7987,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6)))) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f7933,f401])).
fof(f7999,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f7986,f210])).
fof(f7986,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f7933,f401])).
fof(f7995,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6)))) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f7976,f210])).
fof(f7976,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6)))) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f7933,f266])).
fof(f7992,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f7975,f210])).
fof(f7975,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl2 | $spl6)),
  inference(superposition,[],[f7933,f266])).
fof(f7943,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK6)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK5)))) ) | $spl230),
  inference(cnf_transformation,[],[f7943_D])).
fof(f7943_D,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK6)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK5)))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f7942,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X34)) ) | $spl6),
  inference(subsumption_resolution,[],[f7923,f357])).
fof(f7923,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X33)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X34)) ) | $spl6),
  inference(superposition,[],[f84,f401])).
fof(f7939,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK5)))) ) | $spl228),
  inference(cnf_transformation,[],[f7939_D])).
fof(f7939_D,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK5)))) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f7941,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK5))) | s(bool,X30) = s(bool,X31)) ) | ($spl6 | $spl43)),
  inference(subsumption_resolution,[],[f7922,f1223])).
fof(f7922,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X29))),s(sK1,sK5))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl6),
  inference(superposition,[],[f84,f401])).
fof(f7940,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK5)))) ) | $spl6),
  inference(subsumption_resolution,[],[f7919,f357])).
fof(f7919,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X22))),s(sK1,sK5)))) ) | $spl6),
  inference(superposition,[],[f84,f400])).
fof(f7938,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X21)) ) | $spl6),
  inference(subsumption_resolution,[],[f7918,f357])).
fof(f7918,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK6)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X20))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X21)) ) | $spl6),
  inference(superposition,[],[f84,f400])).
fof(f7930,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X5))))) ) | $spl226),
  inference(cnf_transformation,[],[f7930_D])).
fof(f7930_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X5))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f7884,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X6),X7),s(fun(X6,sK1),X8))))) ) | $spl6),
  inference(resolution,[],[f84,f41])).
fof(f7929,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK5))))) ) | $spl224),
  inference(cnf_transformation,[],[f7929_D])).
fof(f7929_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK5))))) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f7883,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X2),X3),s(fun(X2,sK1),X4))))) ) | $spl6),
  inference(resolution,[],[f84,f40])).
fof(f7727,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X21)) ) | $spl222),
  inference(cnf_transformation,[],[f7727_D])).
fof(f7727_D,plain,(
  ( ! [X21] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X21)) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f7828,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X43)) ) | ($spl2 | $spl4)),
  inference(duplicate_literal_removal,[],[f7779])).
fof(f7779,plain,(
  ( ! [X43,X42] : (s(bool,'T_0') = s(bool,X42) | s(bool,X43) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X43) | s(bool,'T_0') = s(bool,X42)) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f401,f7680])).
fof(f7680,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | ($spl2 | $spl4)),
  inference(resolution,[],[f7663,f210])).
fof(f7663,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5)))) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7638,f43])).
fof(f7638,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5)))) | ($spl2 | $spl4)),
  inference(superposition,[],[f83,f145])).
fof(f7831,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X17)) ) | ($spl2 | $spl4)),
  inference(duplicate_literal_removal,[],[f7748])).
fof(f7748,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X16) | s(bool,X17) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X16)) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f7680,f401])).
fof(f7721,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | $spl220),
  inference(cnf_transformation,[],[f7721_D])).
fof(f7721_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f7728,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5)))) ) | ($spl2 | $spl4 | $spl76)),
  inference(subsumption_resolution,[],[f7715,f2404])).
fof(f7715,plain,(
  ( ! [X23,X22] : (p(s(bool,X22)) | s(bool,X22) = s(bool,X23) | s(bool,'F_0') = s(bool,X23) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5)))) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f7663,f401])).
fof(f7726,plain,(
  ( ! [X21,X20] : (s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7714,f210])).
fof(f7714,plain,(
  ( ! [X21,X20] : (p(s(bool,X20)) | s(bool,X21) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X21) | s(bool,'T_0') = s(bool,X20)) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f7663,f401])).
fof(f7722,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5)))) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7704,f210])).
fof(f7704,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5)))) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f7663,f266])).
fof(f7719,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl2 | $spl4)),
  inference(subsumption_resolution,[],[f7703,f210])).
fof(f7703,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X5)) ) | ($spl2 | $spl4)),
  inference(superposition,[],[f7663,f266])).
fof(f7672,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X32)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X32)))) ) | $spl218),
  inference(cnf_transformation,[],[f7672_D])).
fof(f7672_D,plain,(
  ( ! [X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X32)))) | s(bool,'T_0') = s(bool,X34) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X32)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f7671,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X32)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X32))) | s(bool,'T_0') = s(bool,X34)) ) | $spl4),
  inference(subsumption_resolution,[],[f7655,f357])).
fof(f7655,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X33)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X32)))) | s(bool,'F_0') = s(bool,X33) | s(bool,X34) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X32))) | s(bool,'T_0') = s(bool,X34)) ) | $spl4),
  inference(superposition,[],[f83,f401])).
fof(f7668,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X20)))) ) | $spl216),
  inference(cnf_transformation,[],[f7668_D])).
fof(f7668_D,plain,(
  ( ! [X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X20)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f7670,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X29)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X29))) | s(bool,X30) = s(bool,X31)) ) | ($spl4 | $spl43)),
  inference(subsumption_resolution,[],[f7654,f1223])).
fof(f7654,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,X30)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X29)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X29))) | s(bool,X30) = s(bool,X31) | s(bool,'T_0') = s(bool,X31)) ) | $spl4),
  inference(superposition,[],[f83,f401])).
fof(f7669,plain,(
  ( ! [X23,X22] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X22)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X22)))) ) | $spl4),
  inference(subsumption_resolution,[],[f7651,f357])).
fof(f7651,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X22)))) | s(bool,'F_0') = s(bool,X23) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X22)))) ) | $spl4),
  inference(superposition,[],[f83,f400])).
fof(f7667,plain,(
  ( ! [X21,X20] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X20))) | s(bool,'F_0') = s(bool,X21)) ) | $spl4),
  inference(subsumption_resolution,[],[f7650,f357])).
fof(f7650,plain,(
  ( ! [X21,X20] : (~p(s(bool,X21)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X20)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X20))) | s(bool,'F_0') = s(bool,X21)) ) | $spl4),
  inference(superposition,[],[f83,f400])).
fof(f7662,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK6))))) ) | $spl214),
  inference(cnf_transformation,[],[f7662_D])).
fof(f7662_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK6))))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f7617,plain,(
  ( ! [X6,X8,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X5))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X6),X7),s(fun(X6,sK1),X8))))) ) | $spl4),
  inference(resolution,[],[f83,f41])).
fof(f7661,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0))))) ) | $spl212),
  inference(cnf_transformation,[],[f7661_D])).
fof(f7661_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f7616,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X1)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X2),X3),s(fun(X2,sK1),X4))))) ) | $spl4),
  inference(resolution,[],[f83,f40])).
fof(f6876,plain,(
  ( ! [X133] : (s(bool,X133) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X133)) ) | $spl148),
  inference(cnf_transformation,[],[f6876_D])).
fof(f6876_D,plain,(
  ( ! [X133] : (s(bool,X133) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X133)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f2657,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20)) ) | $spl82),
  inference(cnf_transformation,[],[f2657_D])).
fof(f2657_D,plain,(
  ( ! [X20] : (s(bool,'F_0') != s(bool,X20)) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f7615,plain,(
  ( ! [X132,X133] : (s(bool,'F_0') != s(bool,X132) | s(bool,X133) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X133)) ) | ($spl43 | $spl76 | $spl87)),
  inference(subsumption_resolution,[],[f7575,f2401])).
fof(f2401,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | s(bool,'T_0') != s(bool,X0)) ) | ($spl43 | $spl76)),
  inference(resolution,[],[f2143,f1190])).
fof(f7575,plain,(
  ( ! [X132,X133] : (s(bool,'F_0') != s(bool,X132) | s(bool,X133) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X133) | s(bool,'T_0') = s(bool,X132)) ) | ($spl76 | $spl87)),
  inference(superposition,[],[f2661,f3965])).
fof(f3965,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | $spl87),
  inference(resolution,[],[f2911,f357])).
fof(f2661,plain,(
  ( ! [X10,X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) ) | $spl76),
  inference(superposition,[],[f2503,f39])).
fof(f2503,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') != s(bool,X8) | s(bool,X8) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl76),
  inference(resolution,[],[f2404,f210])).
fof(f7614,plain,(
  ( ! [X123,X122] : (s(bool,'F_0') != s(bool,X122) | s(bool,X123) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X123)) ) | ($spl22 | $spl43 | $spl76)),
  inference(subsumption_resolution,[],[f7570,f2401])).
fof(f7570,plain,(
  ( ! [X123,X122] : (s(bool,'F_0') != s(bool,X122) | s(bool,X123) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X123) | s(bool,'T_0') = s(bool,X122)) ) | ($spl22 | $spl76)),
  inference(superposition,[],[f2661,f234])).
fof(f7156,plain,(
  ( ! [X198,X196,X200,X199] : (s(bool,'T_0') = s(bool,X198) | s(bool,'F_0') = s(bool,X200) | s(bool,X196) = s(bool,X200) | s(bool,X198) != s(bool,X199) | s(bool,X196) = s(bool,X199)) ) | $spl190),
  inference(cnf_transformation,[],[f7156_D])).
fof(f7156_D,plain,(
  ( ! [X198,X196,X200,X199] : (s(bool,'T_0') = s(bool,X198) | s(bool,'F_0') = s(bool,X200) | s(bool,X196) = s(bool,X200) | s(bool,X198) != s(bool,X199) | s(bool,X196) = s(bool,X199)) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f7580,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X85) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,X84) = s(bool,X88)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f7559])).
fof(f7559,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,X84) != s(bool,X85) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X85) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,X84) = s(bool,X88)) ) | $spl76),
  inference(superposition,[],[f2661,f401])).
fof(f6854,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X63) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,X62)) ) | $spl144),
  inference(cnf_transformation,[],[f6854_D])).
fof(f6854_D,plain,(
  ( ! [X61,X62,X60,X63] : (s(bool,X60) = s(bool,X61) | s(bool,'T_0') = s(bool,X63) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X62) | s(bool,X61) = s(bool,X62)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f7608,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') != s(bool,X55) | s(bool,X54) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X54) = s(bool,X57) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f7553,f2401])).
fof(f7553,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') != s(bool,X55) | s(bool,X54) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,X55) | s(bool,X54) = s(bool,X57) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) ) | $spl76),
  inference(superposition,[],[f2661,f401])).
fof(f7030,plain,(
  ( ! [X118,X116,X114,X117] : (s(bool,X116) = s(bool,X117) | s(bool,'T_0') = s(bool,X118) | s(bool,X114) = s(bool,X118) | s(bool,X114) != s(bool,X116) | s(bool,'F_0') = s(bool,X117)) ) | $spl164),
  inference(cnf_transformation,[],[f7030_D])).
fof(f7030_D,plain,(
  ( ! [X118,X116,X114,X117] : (s(bool,X116) = s(bool,X117) | s(bool,'T_0') = s(bool,X118) | s(bool,X114) = s(bool,X118) | s(bool,X114) != s(bool,X116) | s(bool,'F_0') = s(bool,X117)) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f7581,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X79) | s(bool,X80) = s(bool,X83) | s(bool,'F_0') = s(bool,X83)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f7521])).
fof(f7521,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X80) != s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X79) | s(bool,X80) = s(bool,X83) | s(bool,'F_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X79)) ) | $spl76),
  inference(superposition,[],[f2661,f401])).
fof(f4978,plain,(
  ( ! [X28,X26,X29] : (s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') = s(bool,X29)) ) | $spl130),
  inference(cnf_transformation,[],[f4978_D])).
fof(f4978_D,plain,(
  ( ! [X28,X26,X29] : (s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f7501,plain,(
  ( ! [X80,X83] : (s(bool,'F_0') != s(bool,X80) | s(bool,X80) = s(bool,X83)) ) | $spl210),
  inference(cnf_transformation,[],[f7501_D])).
fof(f7501_D,plain,(
  ( ! [X80,X83] : (s(bool,'F_0') != s(bool,X80) | s(bool,X80) = s(bool,X83)) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f7500,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'F_0') != s(bool,X80) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X81) | s(bool,X80) = s(bool,X83) | s(bool,'T_0') = s(bool,X79)) ) | $spl76),
  inference(subsumption_resolution,[],[f7460,f2500])).
fof(f2500,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,'F_0') != s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl76),
  inference(resolution,[],[f2404,f2143])).
fof(f7460,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,'F_0') != s(bool,X80) | s(bool,X81) != s(bool,X82) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X81) | s(bool,X80) = s(bool,X83) | s(bool,'F_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X79)) ) | $spl76),
  inference(superposition,[],[f2607,f401])).
fof(f2607,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,X10) | s(bool,X8) != s(bool,X9) | s(bool,X9) = s(bool,X10) | s(bool,'T_0') = s(bool,X8)) ) | $spl76),
  inference(superposition,[],[f2500,f39])).
fof(f7499,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'F_0') != s(bool,X75) | s(bool,X76) != s(bool,X77) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X76) | s(bool,X74) = s(bool,X78) | s(bool,'F_0') = s(bool,X78)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f7459,f2401])).
fof(f7459,plain,(
  ( ! [X78,X76,X74,X77,X75] : (s(bool,'F_0') != s(bool,X75) | s(bool,X76) != s(bool,X77) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X76) | s(bool,X74) = s(bool,X78) | s(bool,'F_0') = s(bool,X78) | s(bool,'T_0') = s(bool,X75)) ) | $spl76),
  inference(superposition,[],[f2607,f401])).
fof(f7497,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X32) != s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f7449,f2401])).
fof(f7449,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X32) != s(bool,X33) | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl76),
  inference(superposition,[],[f2607,f266])).
fof(f7496,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,X27) | s(bool,X28) != s(bool,X29) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X26)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f7448,f2401])).
fof(f7448,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,X27) | s(bool,X28) != s(bool,X29) | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl76),
  inference(superposition,[],[f2607,f266])).
fof(f7492,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X93) != s(bool,X94) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,X95) != s(bool,X96)) ) | $spl208),
  inference(cnf_transformation,[],[f7492_D])).
fof(f7492_D,plain,(
  ( ! [X94,X95,X93,X96] : (s(bool,X93) != s(bool,X94) | s(bool,'T_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X95) | s(bool,X94) = s(bool,X96) | s(bool,X95) != s(bool,X96)) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f7485,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,X93) != s(bool,X94) | s(bool,X95) != s(bool,X96) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X97) | s(bool,'T_0') = s(bool,X93)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f7439])).
fof(f7439,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,X93) != s(bool,X94) | s(bool,X95) != s(bool,X96) | s(bool,X94) = s(bool,X96) | s(bool,'T_0') = s(bool,X95) | s(bool,'F_0') = s(bool,X97) | s(bool,'F_0') = s(bool,X97) | s(bool,'T_0') = s(bool,X93)) ) | $spl76),
  inference(superposition,[],[f2607,f401])).
fof(f7415,plain,(
  ( ! [X48,X51,X49] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X49))),s(sK1,X48)))) | s(bool,'T_0') = s(bool,X51) | s(bool,X51) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X48))),s(sK1,X49)))) ) | $spl206),
  inference(cnf_transformation,[],[f7415_D])).
fof(f7415_D,plain,(
  ( ! [X48,X51,X49] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X49))),s(sK1,X48)))) | s(bool,'T_0') = s(bool,X51) | s(bool,X51) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X48))),s(sK1,X49)))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f7414,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X49))),s(sK1,X48)))) | s(bool,'F_0') = s(bool,X50) | s(bool,X51) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X48))),s(sK1,X49))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(subsumption_resolution,[],[f7388,f357])).
fof(f7388,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X50)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X49))),s(sK1,X48)))) | s(bool,'F_0') = s(bool,X50) | s(bool,X51) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X48))),s(sK1,X49))) | s(bool,'T_0') = s(bool,X51)) )),
  inference(superposition,[],[f44,f401])).
fof(f7411,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,X31)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X31))),s(sK1,X32)))) ) | $spl204),
  inference(cnf_transformation,[],[f7411_D])).
fof(f7411_D,plain,(
  ( ! [X31,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,X31)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X31))),s(sK1,X32)))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f7413,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X46)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X45))),s(sK1,X44)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X44))),s(sK1,X45))) | s(bool,X46) = s(bool,X47)) ) | $spl43),
  inference(subsumption_resolution,[],[f7387,f1223])).
fof(f7387,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X46)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X45))),s(sK1,X44)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X44))),s(sK1,X45))) | s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f44,f401])).
fof(f7412,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X35))),s(sK1,X34)))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X34))),s(sK1,X35)))) )),
  inference(subsumption_resolution,[],[f7384,f357])).
fof(f7384,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X36)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X35))),s(sK1,X34)))) | s(bool,'F_0') = s(bool,X36) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X34))),s(sK1,X35)))) )),
  inference(superposition,[],[f44,f400])).
fof(f7410,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,X31)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X31))),s(sK1,X32))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f7383,f357])).
fof(f7383,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X33)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X32))),s(sK1,X31)))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X31))),s(sK1,X32))) | s(bool,'F_0') = s(bool,X33)) )),
  inference(superposition,[],[f44,f400])).
fof(f7394,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X0))))) ) | $spl202),
  inference(cnf_transformation,[],[f7394_D])).
fof(f7394_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X0))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f7343,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,X1)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X2))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X1))),s(sK1,X2)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X3),X4),s(fun(X3,sK1),X5))))) )),
  inference(resolution,[],[f44,f40])).
fof(f7299,plain,(
  $false | $spl126),
  inference(resolution,[],[f4517,f43])).
fof(f1492,plain,(
  ( ! [X61,X59,X57,X54,X60,X56,X55] : (~p(s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X57))),s(X54,X56)))) | s(bool,'T_0') = s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X56))),s(X54,X57))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X54,fun(X54,bool)),X55),s(fun(X54,X59),X60),s(fun(X59,X54),X61))))) ) | $spl58),
  inference(cnf_transformation,[],[f1492_D])).
fof(f1492_D,plain,(
  ( ! [X61,X59,X57,X54,X60,X56,X55] : (~p(s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X57))),s(X54,X56)))) | s(bool,'T_0') = s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X56))),s(X54,X57))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X54,fun(X54,bool)),X55),s(fun(X54,X59),X60),s(fun(X59,X54),X61))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7217,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (p(s(bool,X375)) | ~p(s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X374))),s(X371,X373)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X371,fun(X371,bool)),X372),s(fun(X371,X376),X377),s(fun(X376,X371),X378)))) | s(bool,X375) = s(bool,X379) | s(bool,'T_0') = s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X373))),s(X371,X374)))) ) | $spl76),
  inference(subsumption_resolution,[],[f6507,f2404])).
fof(f6507,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (p(s(bool,X375)) | ~p(s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X374))),s(X371,X373)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X371,fun(X371,bool)),X372),s(fun(X371,X376),X377),s(fun(X376,X371),X378)))) | s(bool,X375) = s(bool,X379) | s(bool,'F_0') = s(bool,X379) | s(bool,'T_0') = s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X373))),s(X371,X374)))) )),
  inference(superposition,[],[f41,f401])).
fof(f3236,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75] : (~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X74))),s(X69,X72)))) | s(bool,'T_0') = s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X72))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X69,fun(X69,bool)),X70),s(fun(X69,X75),X76),s(fun(X75,X69),X77)))) | ~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X74))))) ) | $spl96),
  inference(cnf_transformation,[],[f3236_D])).
fof(f3236_D,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75] : (~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X74))),s(X69,X72)))) | s(bool,'T_0') = s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X72))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X69,fun(X69,bool)),X70),s(fun(X69,X75),X76),s(fun(X75,X69),X77)))) | ~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X74))))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f7216,plain,(
  ( ! [X356,X358,X352,X354,X360,X357,X359,X353,X355,X361] : (p(s(bool,X356)) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X357))),s(X352,X355)))) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X357)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X352,fun(X352,bool)),X353),s(fun(X352,X358),X359),s(fun(X358,X352),X360)))) | s(bool,X356) = s(bool,X361) | s(bool,'T_0') = s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X355)))) ) | $spl76),
  inference(subsumption_resolution,[],[f6505,f2404])).
fof(f6505,plain,(
  ( ! [X356,X358,X352,X354,X360,X357,X359,X353,X355,X361] : (p(s(bool,X356)) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X357))),s(X352,X355)))) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X357)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X352,fun(X352,bool)),X353),s(fun(X352,X358),X359),s(fun(X358,X352),X360)))) | s(bool,X356) = s(bool,X361) | s(bool,'F_0') = s(bool,X361) | s(bool,'T_0') = s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X355)))) )),
  inference(superposition,[],[f40,f401])).
fof(f5390,plain,(
  ( ! [X45,X44,X42] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') = s(bool,X45)) ) | $spl140),
  inference(cnf_transformation,[],[f5390_D])).
fof(f5390_D,plain,(
  ( ! [X45,X44,X42] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X44) | s(bool,'T_0') = s(bool,X45)) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f7200,plain,(
  ( ! [X222,X220,X219,X223,X221] : (p(s(bool,X220)) | s(bool,'F_0') = s(bool,X221) | s(bool,X221) != s(bool,X222) | s(bool,X219) != s(bool,X222) | s(bool,X220) = s(bool,X223) | s(bool,'T_0') = s(bool,X219)) ) | ($spl76 | $spl112)),
  inference(subsumption_resolution,[],[f6471,f2404])).
fof(f6471,plain,(
  ( ! [X222,X220,X219,X223,X221] : (p(s(bool,X220)) | s(bool,'F_0') = s(bool,X221) | s(bool,X221) != s(bool,X222) | s(bool,X219) != s(bool,X222) | s(bool,X220) = s(bool,X223) | s(bool,'F_0') = s(bool,X223) | s(bool,'T_0') = s(bool,X219)) ) | ($spl76 | $spl112)),
  inference(superposition,[],[f3800,f401])).
fof(f3800,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) != s(bool,X13) | s(bool,X11) != s(bool,X12)) ) | ($spl76 | $spl112)),
  inference(subsumption_resolution,[],[f3758,f42])).
fof(f3758,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,'F_0') = s(bool,X13) | s(bool,X12) != s(bool,X13) | p(s(bool,X11)) | p(s(bool,'F_0'))) ) | ($spl76 | $spl112)),
  inference(superposition,[],[f3716,f35])).
fof(f3716,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X14) | s(bool,'F_0') = s(bool,X15) | s(bool,X14) != s(bool,X15)) ) | ($spl76 | $spl112)),
  inference(resolution,[],[f3596,f2143])).
fof(f5515,plain,(
  ( ! [X28,X26,X29] : (s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') != s(bool,X29)) ) | $spl142),
  inference(cnf_transformation,[],[f5515_D])).
fof(f5515_D,plain,(
  ( ! [X28,X26,X29] : (s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') != s(bool,X29)) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f7197,plain,(
  ( ! [X204,X202,X205,X203,X201] : (p(s(bool,X202)) | s(bool,X203) != s(bool,X204) | s(bool,'T_0') != s(bool,X204) | s(bool,X201) != s(bool,X203) | s(bool,X202) = s(bool,X205) | s(bool,'T_0') = s(bool,X201)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6467,f2404])).
fof(f6467,plain,(
  ( ! [X204,X202,X205,X203,X201] : (p(s(bool,X202)) | s(bool,X203) != s(bool,X204) | s(bool,'T_0') != s(bool,X204) | s(bool,X201) != s(bool,X203) | s(bool,X202) = s(bool,X205) | s(bool,'F_0') = s(bool,X205) | s(bool,'T_0') = s(bool,X201)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3523,f401])).
fof(f3523,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X13) | s(bool,X11) != s(bool,X12)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3480,f42])).
fof(f3480,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X12) != s(bool,X13) | s(bool,'T_0') != s(bool,X13) | p(s(bool,X11)) | p(s(bool,'F_0'))) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2561,f35])).
fof(f2561,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') != s(bool,X1) | s(bool,X0) != s(bool,X1) | s(bool,'T_0') != s(bool,X0)) ) | ($spl43 | $spl76)),
  inference(resolution,[],[f2494,f2143])).
fof(f2494,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | s(bool,'T_0') != s(bool,X9) | s(bool,X8) != s(bool,X9)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f2450,f42])).
fof(f2450,plain,(
  ( ! [X8,X9] : (s(bool,X8) != s(bool,X9) | s(bool,'T_0') != s(bool,X9) | p(s(bool,X8)) | p(s(bool,'F_0'))) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2401,f35])).
fof(f7196,plain,(
  ( ! [X198,X196,X200,X199,X197] : (p(s(bool,X197)) | s(bool,'T_0') = s(bool,X198) | s(bool,X196) = s(bool,X199) | s(bool,X198) != s(bool,X199) | s(bool,X197) = s(bool,X200) | s(bool,'T_0') = s(bool,X196)) ) | ($spl76 | $spl104)),
  inference(subsumption_resolution,[],[f6466,f2404])).
fof(f6466,plain,(
  ( ! [X198,X196,X200,X199,X197] : (p(s(bool,X197)) | s(bool,'T_0') = s(bool,X198) | s(bool,X196) = s(bool,X199) | s(bool,X198) != s(bool,X199) | s(bool,X197) = s(bool,X200) | s(bool,'F_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X196)) ) | $spl104),
  inference(superposition,[],[f3393,f401])).
fof(f3393,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,'T_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,X0) != s(bool,X1)) ) | $spl104),
  inference(resolution,[],[f3326,f35])).
fof(f5114,plain,(
  ( ! [X45,X44,X42] : (s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X45)) ) | $spl132),
  inference(cnf_transformation,[],[f5114_D])).
fof(f5114_D,plain,(
  ( ! [X45,X44,X42] : (s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X45)) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f6740,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,X184) != s(bool,X185) | s(bool,'T_0') = s(bool,X183) | s(bool,'F_0') = s(bool,X185) | s(bool,X184) = s(bool,X186) | s(bool,'F_0') = s(bool,X186)) ) | ($spl43 | $spl76)),
  inference(duplicate_literal_removal,[],[f6463])).
fof(f6463,plain,(
  ( ! [X185,X186,X184,X183] : (s(bool,X184) != s(bool,X185) | s(bool,'T_0') = s(bool,X183) | s(bool,'F_0') = s(bool,X185) | s(bool,X184) = s(bool,X186) | s(bool,'F_0') = s(bool,X186) | s(bool,'T_0') = s(bool,X183)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3318,f401])).
fof(f3318,plain,(
  ( ! [X12,X13] : (s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | ($spl43 | $spl76)),
  inference(trivial_inequality_removal,[],[f3293])).
fof(f3293,plain,(
  ( ! [X12,X13] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,X12) != s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,'F_0') = s(bool,X12)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f39])).
fof(f2449,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,X6) != s(bool,X7) | s(bool,'T_0') = s(bool,X6)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2401,f39])).
fof(f7195,plain,(
  ( ! [X177,X175,X178,X176,X174] : (p(s(bool,X175)) | s(bool,X176) = s(bool,X177) | s(bool,'T_0') = s(bool,X177) | s(bool,X174) != s(bool,X176) | s(bool,X175) = s(bool,X178) | s(bool,'T_0') = s(bool,X174)) ) | $spl76),
  inference(subsumption_resolution,[],[f6461,f2404])).
fof(f6461,plain,(
  ( ! [X177,X175,X178,X176,X174] : (p(s(bool,X175)) | s(bool,X176) = s(bool,X177) | s(bool,'T_0') = s(bool,X177) | s(bool,X174) != s(bool,X176) | s(bool,X175) = s(bool,X178) | s(bool,'F_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X174)) ) | $spl76),
  inference(superposition,[],[f2705,f401])).
fof(f2705,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | s(bool,X11) != s(bool,X12)) ) | $spl76),
  inference(subsumption_resolution,[],[f2662,f42])).
fof(f2662,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,X12) = s(bool,X13) | s(bool,'T_0') = s(bool,X13) | p(s(bool,X11)) | p(s(bool,'F_0'))) ) | $spl76),
  inference(superposition,[],[f2503,f35])).
fof(f4918,plain,(
  ( ! [X28,X26,X29] : (s(bool,'F_0') != s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) != s(bool,X29) | s(bool,X28) = s(bool,X29)) ) | $spl128),
  inference(cnf_transformation,[],[f4918_D])).
fof(f4918_D,plain,(
  ( ! [X28,X26,X29] : (s(bool,'F_0') != s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,X26) != s(bool,X29) | s(bool,X28) = s(bool,X29)) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f7194,plain,(
  ( ! [X173,X171,X169,X172,X170] : (p(s(bool,X170)) | s(bool,'F_0') != s(bool,X171) | s(bool,X171) = s(bool,X172) | s(bool,X169) != s(bool,X172) | s(bool,X170) = s(bool,X173) | s(bool,'T_0') = s(bool,X169)) ) | $spl76),
  inference(subsumption_resolution,[],[f6460,f2404])).
fof(f6460,plain,(
  ( ! [X173,X171,X169,X172,X170] : (p(s(bool,X170)) | s(bool,'F_0') != s(bool,X171) | s(bool,X171) = s(bool,X172) | s(bool,X169) != s(bool,X172) | s(bool,X170) = s(bool,X173) | s(bool,'F_0') = s(bool,X173) | s(bool,'T_0') = s(bool,X169)) ) | $spl76),
  inference(superposition,[],[f2648,f401])).
fof(f2648,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | s(bool,'F_0') != s(bool,X13) | s(bool,X12) = s(bool,X13) | s(bool,X11) != s(bool,X12)) ) | $spl76),
  inference(subsumption_resolution,[],[f2608,f42])).
fof(f2608,plain,(
  ( ! [X12,X13,X11] : (s(bool,X11) != s(bool,X12) | s(bool,'F_0') != s(bool,X13) | s(bool,X12) = s(bool,X13) | p(s(bool,X11)) | p(s(bool,'F_0'))) ) | $spl76),
  inference(superposition,[],[f2500,f35])).
fof(f5203,plain,(
  ( ! [X28,X26,X29] : (s(bool,'T_0') != s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,X28) = s(bool,X29) | s(bool,X26) = s(bool,X29)) ) | $spl134),
  inference(cnf_transformation,[],[f5203_D])).
fof(f5203_D,plain,(
  ( ! [X28,X26,X29] : (s(bool,'T_0') != s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,X28) = s(bool,X29) | s(bool,X26) = s(bool,X29)) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f7185,plain,(
  ( ! [X111,X109,X112,X110,X113] : (p(s(bool,X110)) | s(bool,'T_0') != s(bool,X111) | s(bool,X109) = s(bool,X112) | s(bool,X111) = s(bool,X112) | s(bool,X110) = s(bool,X113) | s(bool,'T_0') = s(bool,X109)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6446,f2404])).
fof(f6446,plain,(
  ( ! [X111,X109,X112,X110,X113] : (p(s(bool,X110)) | s(bool,'T_0') != s(bool,X111) | s(bool,X109) = s(bool,X112) | s(bool,X111) = s(bool,X112) | s(bool,X110) = s(bool,X113) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,X109)) ) | $spl43),
  inference(superposition,[],[f1622,f401])).
fof(f1622,plain,(
  ( ! [X8,X7,X9] : (p(s(bool,X9)) | s(bool,'T_0') != s(bool,X7) | s(bool,X8) = s(bool,X9) | s(bool,X7) = s(bool,X8)) ) | $spl43),
  inference(resolution,[],[f1223,f35])).
fof(f7168,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X371] : (~p(s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X374))),s(X371,X373)))) | s(bool,'F_0') = s(bool,X379) | s(bool,X379) = s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X373))),s(X371,X374))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X371,fun(X371,bool)),X372),s(fun(X371,X376),X377),s(fun(X376,X371),X378))))) ) | $spl198),
  inference(cnf_transformation,[],[f7168_D])).
fof(f7168_D,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X371] : (~p(s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X374))),s(X371,X373)))) | s(bool,'F_0') = s(bool,X379) | s(bool,X379) = s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X373))),s(X371,X374))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X371,fun(X371,bool)),X372),s(fun(X371,X376),X377),s(fun(X376,X371),X378))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f7167,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (~p(s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X374))),s(X371,X373)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X371,fun(X371,bool)),X372),s(fun(X371,X376),X377),s(fun(X376,X371),X378)))) | s(bool,X379) = s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X373))),s(X371,X374))) | s(bool,'F_0') = s(bool,X379) | s(bool,'T_0') = s(bool,X375)) )),
  inference(subsumption_resolution,[],[f6405,f210])).
fof(f6405,plain,(
  ( ! [X377,X379,X372,X374,X376,X378,X373,X375,X371] : (p(s(bool,X375)) | ~p(s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X374))),s(X371,X373)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X371,fun(X371,bool)),X372),s(fun(X371,X376),X377),s(fun(X376,X371),X378)))) | s(bool,X379) = s(bool,happ(s(fun(X371,bool),happ(s(fun(X371,fun(X371,bool)),X372),s(X371,X373))),s(X371,X374))) | s(bool,'F_0') = s(bool,X379) | s(bool,'T_0') = s(bool,X375)) )),
  inference(superposition,[],[f41,f401])).
fof(f7166,plain,(
  ( ! [X358,X352,X354,X360,X357,X359,X353,X355,X361] : (~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X357))),s(X352,X355)))) | s(bool,'F_0') = s(bool,X361) | s(bool,X361) = s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X355))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X352,fun(X352,bool)),X353),s(fun(X352,X358),X359),s(fun(X358,X352),X360)))) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X357))))) ) | $spl196),
  inference(cnf_transformation,[],[f7166_D])).
fof(f7166_D,plain,(
  ( ! [X358,X352,X354,X360,X357,X359,X353,X355,X361] : (~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X357))),s(X352,X355)))) | s(bool,'F_0') = s(bool,X361) | s(bool,X361) = s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X355))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X352,fun(X352,bool)),X353),s(fun(X352,X358),X359),s(fun(X358,X352),X360)))) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X357))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f7165,plain,(
  ( ! [X356,X358,X352,X354,X360,X357,X359,X353,X355,X361] : (~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X357))),s(X352,X355)))) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X357)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X352,fun(X352,bool)),X353),s(fun(X352,X358),X359),s(fun(X358,X352),X360)))) | s(bool,X361) = s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X355))) | s(bool,'F_0') = s(bool,X361) | s(bool,'T_0') = s(bool,X356)) )),
  inference(subsumption_resolution,[],[f6403,f210])).
fof(f6403,plain,(
  ( ! [X356,X358,X352,X354,X360,X357,X359,X353,X355,X361] : (p(s(bool,X356)) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X357))),s(X352,X355)))) | ~p(s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X357)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X352,fun(X352,bool)),X353),s(fun(X352,X358),X359),s(fun(X358,X352),X360)))) | s(bool,X361) = s(bool,happ(s(fun(X352,bool),happ(s(fun(X352,fun(X352,bool)),X353),s(X352,X354))),s(X352,X355))) | s(bool,'F_0') = s(bool,X361) | s(bool,'T_0') = s(bool,X356)) )),
  inference(superposition,[],[f40,f401])).
fof(f6916,plain,(
  ( ! [X131] : (s(bool,X131) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X131)) ) | $spl160),
  inference(cnf_transformation,[],[f6916_D])).
fof(f6916_D,plain,(
  ( ! [X131] : (s(bool,X131) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X131)) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6745,plain,(
  ( ! [X327,X326] : (s(bool,'T_0') = s(bool,X326) | s(bool,X327) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X327)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f6398])).
fof(f6398,plain,(
  ( ! [X327,X326] : (s(bool,'T_0') = s(bool,X326) | s(bool,X327) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X327) | s(bool,'T_0') = s(bool,X326)) ) | $spl8),
  inference(superposition,[],[f165,f401])).
fof(f6915,plain,(
  ( ! [X129] : (s(bool,X129) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X129)) ) | $spl158),
  inference(cnf_transformation,[],[f6915_D])).
fof(f6915_D,plain,(
  ( ! [X129] : (s(bool,X129) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X129)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f6746,plain,(
  ( ! [X325,X324] : (s(bool,'T_0') = s(bool,X324) | s(bool,X325) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X325)) ) | $spl84),
  inference(duplicate_literal_removal,[],[f6397])).
fof(f6397,plain,(
  ( ! [X325,X324] : (s(bool,'T_0') = s(bool,X324) | s(bool,X325) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X325) | s(bool,'T_0') = s(bool,X324)) ) | $spl84),
  inference(superposition,[],[f3925,f401])).
fof(f6914,plain,(
  ( ! [X127] : (s(bool,X127) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X127)) ) | $spl156),
  inference(cnf_transformation,[],[f6914_D])).
fof(f6914_D,plain,(
  ( ! [X127] : (s(bool,X127) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X127)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f6747,plain,(
  ( ! [X323,X322] : (s(bool,'T_0') = s(bool,X322) | s(bool,X323) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X323)) )),
  inference(duplicate_literal_removal,[],[f6396])).
fof(f6396,plain,(
  ( ! [X323,X322] : (s(bool,'T_0') = s(bool,X322) | s(bool,X323) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X323) | s(bool,'T_0') = s(bool,X322)) )),
  inference(superposition,[],[f107,f401])).
fof(f107,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8)))),
  inference(subsumption_resolution,[],[f96,f42])).
fof(f96,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8)))),
  inference(superposition,[],[f32,f39])).
fof(f32,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))))),
  inference(cnf_transformation,[],[f28])).
fof(f6913,plain,(
  ( ! [X125] : (s(bool,X125) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X125)) ) | $spl154),
  inference(cnf_transformation,[],[f6913_D])).
fof(f6913_D,plain,(
  ( ! [X125] : (s(bool,X125) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X125)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f6748,plain,(
  ( ! [X321,X320] : (s(bool,'T_0') = s(bool,X320) | s(bool,X321) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X321)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f6395])).
fof(f6395,plain,(
  ( ! [X321,X320] : (s(bool,'T_0') = s(bool,X320) | s(bool,X321) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X321) | s(bool,'T_0') = s(bool,X320)) ) | $spl2),
  inference(superposition,[],[f145,f401])).
fof(f6912,plain,(
  ( ! [X121] : (s(bool,X121) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X121)) ) | $spl152),
  inference(cnf_transformation,[],[f6912_D])).
fof(f6912_D,plain,(
  ( ! [X121] : (s(bool,X121) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X121)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f6749,plain,(
  ( ! [X315,X314] : (s(bool,'T_0') = s(bool,X314) | s(bool,X315) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X315)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f6392])).
fof(f6392,plain,(
  ( ! [X315,X314] : (s(bool,'T_0') = s(bool,X314) | s(bool,X315) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X315) | s(bool,'T_0') = s(bool,X314)) ) | $spl24),
  inference(superposition,[],[f235,f401])).
fof(f6911,plain,(
  ( ! [X119] : (s(bool,X119) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X119)) ) | $spl150),
  inference(cnf_transformation,[],[f6911_D])).
fof(f6911_D,plain,(
  ( ! [X119] : (s(bool,X119) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X119)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f6750,plain,(
  ( ! [X313,X312] : (s(bool,'T_0') = s(bool,X312) | s(bool,X313) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X313)) )),
  inference(duplicate_literal_removal,[],[f6391])).
fof(f6391,plain,(
  ( ! [X313,X312] : (s(bool,'T_0') = s(bool,X312) | s(bool,X313) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X313) | s(bool,'T_0') = s(bool,X312)) )),
  inference(superposition,[],[f85,f401])).
fof(f85,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6)))),
  inference(subsumption_resolution,[],[f73,f42])).
fof(f73,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6)))),
  inference(superposition,[],[f31,f39])).
fof(f31,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))))),
  inference(cnf_transformation,[],[f28])).
fof(f695,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X27)) ) | $spl40),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  ( ! [X27] : (s(bool,X27) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X27)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f6751,plain,(
  ( ! [X292,X291] : (s(bool,'T_0') = s(bool,X291) | s(bool,X292) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X292)) )),
  inference(duplicate_literal_removal,[],[f6388])).
fof(f6388,plain,(
  ( ! [X292,X291] : (s(bool,'T_0') = s(bool,X291) | s(bool,X292) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X292) | s(bool,'T_0') = s(bool,X291)) )),
  inference(superposition,[],[f62,f401])).
fof(f6752,plain,(
  ( ! [X288,X290,X287,X289] : (s(bool,X287) != s(bool,X288) | s(bool,'F_0') = s(bool,X289) | s(bool,X288) != s(bool,X289) | s(bool,'F_0') = s(bool,X290) | s(bool,'T_0') = s(bool,X287)) ) | ($spl76 | $spl112)),
  inference(duplicate_literal_removal,[],[f6387])).
fof(f6387,plain,(
  ( ! [X288,X290,X287,X289] : (s(bool,X287) != s(bool,X288) | s(bool,'F_0') = s(bool,X289) | s(bool,X288) != s(bool,X289) | s(bool,'F_0') = s(bool,X290) | s(bool,'F_0') = s(bool,X290) | s(bool,'T_0') = s(bool,X287)) ) | ($spl76 | $spl112)),
  inference(superposition,[],[f3716,f401])).
fof(f6753,plain,(
  ( ! [X285,X283,X284,X286] : (s(bool,X283) != s(bool,X284) | s(bool,X284) != s(bool,X285) | s(bool,'T_0') != s(bool,X285) | s(bool,'F_0') = s(bool,X286) | s(bool,'T_0') = s(bool,X283)) ) | ($spl43 | $spl76)),
  inference(duplicate_literal_removal,[],[f6386])).
fof(f6386,plain,(
  ( ! [X285,X283,X284,X286] : (s(bool,X283) != s(bool,X284) | s(bool,X284) != s(bool,X285) | s(bool,'T_0') != s(bool,X285) | s(bool,'F_0') = s(bool,X286) | s(bool,'F_0') = s(bool,X286) | s(bool,'T_0') = s(bool,X283)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2561,f401])).
fof(f6754,plain,(
  ( ! [X280,X282,X279,X281] : (s(bool,X279) != s(bool,X280) | s(bool,X280) = s(bool,X281) | s(bool,'T_0') = s(bool,X281) | s(bool,'F_0') = s(bool,X282) | s(bool,'T_0') = s(bool,X279)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f6385])).
fof(f6385,plain,(
  ( ! [X280,X282,X279,X281] : (s(bool,X279) != s(bool,X280) | s(bool,X280) = s(bool,X281) | s(bool,'T_0') = s(bool,X281) | s(bool,'F_0') = s(bool,X282) | s(bool,'F_0') = s(bool,X282) | s(bool,'T_0') = s(bool,X279)) ) | $spl76),
  inference(superposition,[],[f2503,f401])).
fof(f6755,plain,(
  ( ! [X277,X275,X276,X278,X274] : (s(bool,X274) != s(bool,X275) | s(bool,X275) = s(bool,X276) | s(bool,X276) = s(bool,X277) | s(bool,'F_0') = s(bool,X277) | s(bool,'F_0') = s(bool,X278) | s(bool,'T_0') = s(bool,X274)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f6384])).
fof(f6384,plain,(
  ( ! [X277,X275,X276,X278,X274] : (s(bool,X274) != s(bool,X275) | s(bool,X275) = s(bool,X276) | s(bool,X276) = s(bool,X277) | s(bool,'F_0') = s(bool,X277) | s(bool,'F_0') = s(bool,X278) | s(bool,'F_0') = s(bool,X278) | s(bool,'T_0') = s(bool,X274)) ) | $spl76),
  inference(superposition,[],[f2502,f401])).
fof(f2502,plain,(
  ( ! [X6,X7,X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,X5) = s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'F_0') = s(bool,X7)) ) | $spl76),
  inference(resolution,[],[f2404,f385])).
fof(f6756,plain,(
  ( ! [X271,X273,X270,X272] : (s(bool,X270) != s(bool,X271) | s(bool,'F_0') != s(bool,X272) | s(bool,X271) = s(bool,X272) | s(bool,'F_0') = s(bool,X273) | s(bool,'T_0') = s(bool,X270)) ) | $spl76),
  inference(duplicate_literal_removal,[],[f6383])).
fof(f6383,plain,(
  ( ! [X271,X273,X270,X272] : (s(bool,X270) != s(bool,X271) | s(bool,'F_0') != s(bool,X272) | s(bool,X271) = s(bool,X272) | s(bool,'F_0') = s(bool,X273) | s(bool,'F_0') = s(bool,X273) | s(bool,'T_0') = s(bool,X270)) ) | $spl76),
  inference(superposition,[],[f2500,f401])).
fof(f7162,plain,(
  ( ! [X222,X219,X223,X221] : (s(bool,'F_0') = s(bool,X221) | s(bool,'F_0') = s(bool,X223) | s(bool,X219) = s(bool,X223) | s(bool,X219) != s(bool,X222) | s(bool,X221) != s(bool,X222)) ) | $spl194),
  inference(cnf_transformation,[],[f7162_D])).
fof(f7162_D,plain,(
  ( ! [X222,X219,X223,X221] : (s(bool,'F_0') = s(bool,X221) | s(bool,'F_0') = s(bool,X223) | s(bool,X219) = s(bool,X223) | s(bool,X219) != s(bool,X222) | s(bool,X221) != s(bool,X222)) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f7161,plain,(
  ( ! [X222,X220,X219,X223,X221] : (s(bool,'F_0') = s(bool,X221) | s(bool,X221) != s(bool,X222) | s(bool,X219) != s(bool,X222) | s(bool,X219) = s(bool,X223) | s(bool,'F_0') = s(bool,X223) | s(bool,'T_0') = s(bool,X220)) ) | ($spl76 | $spl112)),
  inference(subsumption_resolution,[],[f6369,f210])).
fof(f6369,plain,(
  ( ! [X222,X220,X219,X223,X221] : (p(s(bool,X220)) | s(bool,'F_0') = s(bool,X221) | s(bool,X221) != s(bool,X222) | s(bool,X219) != s(bool,X222) | s(bool,X219) = s(bool,X223) | s(bool,'F_0') = s(bool,X223) | s(bool,'T_0') = s(bool,X220)) ) | ($spl76 | $spl112)),
  inference(superposition,[],[f3800,f401])).
fof(f7160,plain,(
  ( ! [X218,X216,X215,X217] : (s(bool,'F_0') != s(bool,X216) | s(bool,'F_0') = s(bool,X217) | s(bool,X215) != s(bool,X217) | s(bool,X215) = s(bool,X218) | s(bool,'F_0') = s(bool,X218)) ) | ($spl43 | $spl76 | $spl112)),
  inference(subsumption_resolution,[],[f6368,f2401])).
fof(f6368,plain,(
  ( ! [X218,X216,X215,X217] : (s(bool,'F_0') != s(bool,X216) | s(bool,'F_0') = s(bool,X217) | s(bool,X215) != s(bool,X217) | s(bool,X215) = s(bool,X218) | s(bool,'F_0') = s(bool,X218) | s(bool,'T_0') = s(bool,X216)) ) | ($spl76 | $spl112)),
  inference(superposition,[],[f3716,f401])).
fof(f7159,plain,(
  ( ! [X208,X206,X209,X207] : (s(bool,X206) != s(bool,X208) | s(bool,'F_0') = s(bool,X208) | s(bool,X206) = s(bool,X209) | s(bool,'F_0') = s(bool,X209) | s(bool,'T_0') = s(bool,X207)) ) | $spl112),
  inference(subsumption_resolution,[],[f6366,f210])).
fof(f6366,plain,(
  ( ! [X208,X206,X209,X207] : (p(s(bool,X207)) | s(bool,X206) != s(bool,X208) | s(bool,'F_0') = s(bool,X208) | s(bool,X206) = s(bool,X209) | s(bool,'F_0') = s(bool,X209) | s(bool,'T_0') = s(bool,X207)) ) | $spl112),
  inference(superposition,[],[f3596,f401])).
fof(f7158,plain,(
  ( ! [X204,X205,X203,X201] : (s(bool,X203) != s(bool,X204) | s(bool,'F_0') = s(bool,X205) | s(bool,X201) = s(bool,X205) | s(bool,X201) != s(bool,X203) | s(bool,'T_0') != s(bool,X204)) ) | $spl192),
  inference(cnf_transformation,[],[f7158_D])).
fof(f7158_D,plain,(
  ( ! [X204,X205,X203,X201] : (s(bool,X203) != s(bool,X204) | s(bool,'F_0') = s(bool,X205) | s(bool,X201) = s(bool,X205) | s(bool,X201) != s(bool,X203) | s(bool,'T_0') != s(bool,X204)) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f2141,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl74),
  inference(cnf_transformation,[],[f2141_D])).
fof(f2141_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f7157,plain,(
  ( ! [X204,X202,X205,X203,X201] : (p(s(bool,X202)) | s(bool,X203) != s(bool,X204) | s(bool,'T_0') != s(bool,X204) | s(bool,X201) != s(bool,X203) | s(bool,X201) = s(bool,X205) | s(bool,'F_0') = s(bool,X205)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6365,f1190])).
fof(f6365,plain,(
  ( ! [X204,X202,X205,X203,X201] : (p(s(bool,X202)) | s(bool,X203) != s(bool,X204) | s(bool,'T_0') != s(bool,X204) | s(bool,X201) != s(bool,X203) | s(bool,X201) = s(bool,X205) | s(bool,'F_0') = s(bool,X205) | s(bool,'T_0') = s(bool,X202)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3523,f401])).
fof(f7155,plain,(
  ( ! [X198,X196,X200,X199,X197] : (s(bool,'T_0') = s(bool,X198) | s(bool,X196) = s(bool,X199) | s(bool,X198) != s(bool,X199) | s(bool,X196) = s(bool,X200) | s(bool,'F_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X197)) ) | $spl104),
  inference(subsumption_resolution,[],[f6364,f210])).
fof(f6364,plain,(
  ( ! [X198,X196,X200,X199,X197] : (p(s(bool,X197)) | s(bool,'T_0') = s(bool,X198) | s(bool,X196) = s(bool,X199) | s(bool,X198) != s(bool,X199) | s(bool,X196) = s(bool,X200) | s(bool,'F_0') = s(bool,X200) | s(bool,'T_0') = s(bool,X197)) ) | $spl104),
  inference(superposition,[],[f3393,f401])).
fof(f7152,plain,(
  ( ! [X177,X175,X178,X176,X174] : (s(bool,X176) = s(bool,X177) | s(bool,'T_0') = s(bool,X177) | s(bool,X174) != s(bool,X176) | s(bool,X174) = s(bool,X178) | s(bool,'F_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X175)) ) | $spl76),
  inference(subsumption_resolution,[],[f6359,f210])).
fof(f6359,plain,(
  ( ! [X177,X175,X178,X176,X174] : (p(s(bool,X175)) | s(bool,X176) = s(bool,X177) | s(bool,'T_0') = s(bool,X177) | s(bool,X174) != s(bool,X176) | s(bool,X174) = s(bool,X178) | s(bool,'F_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X175)) ) | $spl76),
  inference(superposition,[],[f2705,f401])).
fof(f7151,plain,(
  ( ! [X173,X171,X169,X172] : (s(bool,'F_0') != s(bool,X171) | s(bool,'F_0') = s(bool,X173) | s(bool,X169) = s(bool,X173) | s(bool,X169) != s(bool,X172) | s(bool,X171) = s(bool,X172)) ) | $spl188),
  inference(cnf_transformation,[],[f7151_D])).
fof(f7151_D,plain,(
  ( ! [X173,X171,X169,X172] : (s(bool,'F_0') != s(bool,X171) | s(bool,'F_0') = s(bool,X173) | s(bool,X169) = s(bool,X173) | s(bool,X169) != s(bool,X172) | s(bool,X171) = s(bool,X172)) ) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f7150,plain,(
  ( ! [X173,X171,X169,X172,X170] : (s(bool,'F_0') != s(bool,X171) | s(bool,X171) = s(bool,X172) | s(bool,X169) != s(bool,X172) | s(bool,X169) = s(bool,X173) | s(bool,'F_0') = s(bool,X173) | s(bool,'T_0') = s(bool,X170)) ) | $spl76),
  inference(subsumption_resolution,[],[f6358,f210])).
fof(f6358,plain,(
  ( ! [X173,X171,X169,X172,X170] : (p(s(bool,X170)) | s(bool,'F_0') != s(bool,X171) | s(bool,X171) = s(bool,X172) | s(bool,X169) != s(bool,X172) | s(bool,X169) = s(bool,X173) | s(bool,'F_0') = s(bool,X173) | s(bool,'T_0') = s(bool,X170)) ) | $spl76),
  inference(superposition,[],[f2648,f401])).
fof(f5328,plain,(
  ( ! [X45,X44,X42] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') != s(bool,X44) | s(bool,X42) = s(bool,X45)) ) | $spl138),
  inference(cnf_transformation,[],[f5328_D])).
fof(f5328_D,plain,(
  ( ! [X45,X44,X42] : (s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'T_0') != s(bool,X44) | s(bool,X42) = s(bool,X45)) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7149,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,'F_0') != s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') != s(bool,X162) | s(bool,X160) = s(bool,X163) | s(bool,'F_0') = s(bool,X163)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6356,f2401])).
fof(f6356,plain,(
  ( ! [X163,X161,X162,X160] : (s(bool,'F_0') != s(bool,X161) | s(bool,X160) != s(bool,X162) | s(bool,'T_0') != s(bool,X162) | s(bool,X160) = s(bool,X163) | s(bool,'F_0') = s(bool,X163) | s(bool,'T_0') = s(bool,X161)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2561,f401])).
fof(f6872,plain,(
  ( ! [X88,X87,X85,X84] : (s(bool,X84) = s(bool,X85) | s(bool,X84) = s(bool,X88) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,'F_0') = s(bool,X85)) ) | $spl146),
  inference(cnf_transformation,[],[f6872_D])).
fof(f6872_D,plain,(
  ( ! [X88,X87,X85,X84] : (s(bool,X84) = s(bool,X85) | s(bool,X84) = s(bool,X88) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,'F_0') = s(bool,X85)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f7147,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'F_0') != s(bool,X147) | s(bool,X146) = s(bool,X148) | s(bool,X148) = s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X146) = s(bool,X150) | s(bool,'F_0') = s(bool,X150)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6353,f2401])).
fof(f6353,plain,(
  ( ! [X146,X149,X147,X150,X148] : (s(bool,'F_0') != s(bool,X147) | s(bool,X146) = s(bool,X148) | s(bool,X148) = s(bool,X149) | s(bool,'F_0') = s(bool,X149) | s(bool,X146) = s(bool,X150) | s(bool,'F_0') = s(bool,X150) | s(bool,'T_0') = s(bool,X147)) ) | $spl76),
  inference(superposition,[],[f2502,f401])).
fof(f7145,plain,(
  ( ! [X140,X138,X141,X139] : (p(s(bool,X139)) | s(bool,'T_0') != s(bool,X140) | s(bool,X138) != s(bool,X140) | s(bool,X138) = s(bool,X141) | s(bool,'F_0') = s(bool,X141)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6351,f1190])).
fof(f6351,plain,(
  ( ! [X140,X138,X141,X139] : (p(s(bool,X139)) | s(bool,'T_0') != s(bool,X140) | s(bool,X138) != s(bool,X140) | s(bool,X138) = s(bool,X141) | s(bool,'F_0') = s(bool,X141) | s(bool,'T_0') = s(bool,X139)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2494,f401])).
fof(f7142,plain,(
  ( ! [X111,X109,X112,X113] : (s(bool,'T_0') != s(bool,X111) | s(bool,'F_0') = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X111) = s(bool,X112) | s(bool,X109) = s(bool,X112)) ) | $spl186),
  inference(cnf_transformation,[],[f7142_D])).
fof(f7142_D,plain,(
  ( ! [X111,X109,X112,X113] : (s(bool,'T_0') != s(bool,X111) | s(bool,'F_0') = s(bool,X113) | s(bool,X109) = s(bool,X113) | s(bool,X111) = s(bool,X112) | s(bool,X109) = s(bool,X112)) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f7141,plain,(
  ( ! [X111,X109,X112,X110,X113] : (p(s(bool,X110)) | s(bool,'T_0') != s(bool,X111) | s(bool,X109) = s(bool,X112) | s(bool,X111) = s(bool,X112) | s(bool,X109) = s(bool,X113) | s(bool,'F_0') = s(bool,X113)) ) | $spl43),
  inference(subsumption_resolution,[],[f6344,f1190])).
fof(f6344,plain,(
  ( ! [X111,X109,X112,X110,X113] : (p(s(bool,X110)) | s(bool,'T_0') != s(bool,X111) | s(bool,X109) = s(bool,X112) | s(bool,X111) = s(bool,X112) | s(bool,X109) = s(bool,X113) | s(bool,'F_0') = s(bool,X113) | s(bool,'T_0') = s(bool,X110)) ) | $spl43),
  inference(superposition,[],[f1622,f401])).
fof(f7138,plain,(
  ( ! [X92,X90,X93,X91,X89] : (s(bool,'F_0') = s(bool,X91) | s(bool,X89) = s(bool,X92) | s(bool,X91) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f6339,f210])).
fof(f6339,plain,(
  ( ! [X92,X90,X93,X91,X89] : (p(s(bool,X90)) | s(bool,'F_0') = s(bool,X91) | s(bool,X89) = s(bool,X92) | s(bool,X91) = s(bool,X92) | s(bool,X89) = s(bool,X93) | s(bool,'F_0') = s(bool,X93) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f402,f401])).
fof(f402,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X7)) | s(bool,'F_0') = s(bool,X5) | s(bool,X6) = s(bool,X7) | s(bool,X5) = s(bool,X6)) )),
  inference(resolution,[],[f385,f35])).
fof(f7079,plain,(
  ( ! [X171,X169,X170,X168] : (s(bool,'T_0') != s(bool,X168) | s(bool,'F_0') != s(bool,X169) | s(bool,X168) = s(bool,X170) | s(bool,X170) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6196,f2401])).
fof(f6196,plain,(
  ( ! [X171,X169,X170,X168] : (s(bool,'T_0') != s(bool,X168) | s(bool,'F_0') != s(bool,X169) | s(bool,'T_0') = s(bool,X169) | s(bool,X168) = s(bool,X170) | s(bool,X170) = s(bool,X171) | s(bool,'T_0') = s(bool,X171)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f401])).
fof(f7075,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,'F_0') = s(bool,X140) | s(bool,X139) = s(bool,X141) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | $spl43),
  inference(subsumption_resolution,[],[f6189,f39])).
fof(f6189,plain,(
  ( ! [X142,X140,X141,X139] : (s(bool,'T_0') != s(bool,X139) | s(bool,'T_0') != s(bool,X140) | s(bool,'F_0') = s(bool,X140) | s(bool,X139) = s(bool,X141) | s(bool,X141) = s(bool,X142) | s(bool,'T_0') = s(bool,X142)) ) | $spl43),
  inference(superposition,[],[f1619,f401])).
fof(f1619,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X1) | s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,X1)) ) | $spl43),
  inference(resolution,[],[f1223,f1190])).
fof(f6779,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') != s(bool,X128) | s(bool,'F_0') = s(bool,X129) | s(bool,X128) = s(bool,X130) | s(bool,X130) = s(bool,X131) | s(bool,'T_0') = s(bool,X131)) ) | $spl43),
  inference(duplicate_literal_removal,[],[f6186])).
fof(f6186,plain,(
  ( ! [X130,X128,X131,X129] : (s(bool,'T_0') != s(bool,X128) | s(bool,'F_0') = s(bool,X129) | s(bool,'F_0') = s(bool,X129) | s(bool,X128) = s(bool,X130) | s(bool,X130) = s(bool,X131) | s(bool,'T_0') = s(bool,X131)) ) | $spl43),
  inference(superposition,[],[f1169,f401])).
fof(f1169,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,X0) | s(bool,X0) = s(bool,X1) | s(bool,'F_0') = s(bool,X1)) ) | $spl43),
  inference(superposition,[],[f696,f400])).
fof(f7062,plain,(
  ( ! [X364,X362,X368,X370,X365,X367,X363,X369] : (p(s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X365))),s(X362,X364)))) | s(bool,'T_0') = s(bool,X370) | s(bool,X370) = s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X364))),s(X362,X365))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X362,fun(X362,bool)),X363),s(fun(X362,X367),X368),s(fun(X367,X362),X369))))) ) | $spl184),
  inference(cnf_transformation,[],[f7062_D])).
fof(f7062_D,plain,(
  ( ! [X364,X362,X368,X370,X365,X367,X363,X369] : (p(s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X365))),s(X362,X364)))) | s(bool,'T_0') = s(bool,X370) | s(bool,X370) = s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X364))),s(X362,X365))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X362,fun(X362,bool)),X363),s(fun(X362,X367),X368),s(fun(X367,X362),X369))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f7061,plain,(
  ( ! [X364,X366,X362,X368,X370,X365,X367,X363,X369] : (p(s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X365))),s(X362,X364)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X362,fun(X362,bool)),X363),s(fun(X362,X367),X368),s(fun(X367,X362),X369)))) | s(bool,'F_0') = s(bool,X366) | s(bool,X370) = s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X364))),s(X362,X365))) | s(bool,'T_0') = s(bool,X370)) )),
  inference(subsumption_resolution,[],[f6140,f357])).
fof(f6140,plain,(
  ( ! [X364,X366,X362,X368,X370,X365,X367,X363,X369] : (~p(s(bool,X366)) | p(s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X365))),s(X362,X364)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X362,fun(X362,bool)),X363),s(fun(X362,X367),X368),s(fun(X367,X362),X369)))) | s(bool,'F_0') = s(bool,X366) | s(bool,X370) = s(bool,happ(s(fun(X362,bool),happ(s(fun(X362,fun(X362,bool)),X363),s(X362,X364))),s(X362,X365))) | s(bool,'T_0') = s(bool,X370)) )),
  inference(superposition,[],[f41,f401])).
fof(f7060,plain,(
  ( ! [X350,X344,X343,X349,X351,X345,X347,X342,X348] : (p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X345)))) | s(bool,'T_0') = s(bool,X351) | s(bool,X351) = s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X344))),s(X342,X345))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X342,fun(X342,bool)),X343),s(fun(X342,X348),X349),s(fun(X348,X342),X350)))) | ~p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X344))))) ) | $spl182),
  inference(cnf_transformation,[],[f7060_D])).
fof(f7060_D,plain,(
  ( ! [X350,X344,X343,X349,X351,X345,X347,X342,X348] : (p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X345)))) | s(bool,'T_0') = s(bool,X351) | s(bool,X351) = s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X344))),s(X342,X345))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X342,fun(X342,bool)),X343),s(fun(X342,X348),X349),s(fun(X348,X342),X350)))) | ~p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X344))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f7059,plain,(
  ( ! [X350,X344,X346,X343,X349,X351,X345,X347,X342,X348] : (p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X345)))) | ~p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X344)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X342,fun(X342,bool)),X343),s(fun(X342,X348),X349),s(fun(X348,X342),X350)))) | s(bool,'F_0') = s(bool,X346) | s(bool,X351) = s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X344))),s(X342,X345))) | s(bool,'T_0') = s(bool,X351)) )),
  inference(subsumption_resolution,[],[f6138,f357])).
fof(f6138,plain,(
  ( ! [X350,X344,X346,X343,X349,X351,X345,X347,X342,X348] : (~p(s(bool,X346)) | p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X345)))) | ~p(s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X347))),s(X342,X344)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X342,fun(X342,bool)),X343),s(fun(X342,X348),X349),s(fun(X348,X342),X350)))) | s(bool,'F_0') = s(bool,X346) | s(bool,X351) = s(bool,happ(s(fun(X342,bool),happ(s(fun(X342,fun(X342,bool)),X343),s(X342,X344))),s(X342,X345))) | s(bool,'T_0') = s(bool,X351)) )),
  inference(superposition,[],[f40,f401])).
fof(f7058,plain,(
  ( ! [X333,X335,X341,X337,X339,X332,X334,X340,X338] : (p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X337)))) | s(bool,'T_0') = s(bool,X341) | s(bool,X341) = s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X335))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X332,fun(X332,bool)),X333),s(fun(X332,X338),X339),s(fun(X338,X332),X340)))) | ~p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X335))),s(X332,X337))))) ) | $spl180),
  inference(cnf_transformation,[],[f7058_D])).
fof(f7058_D,plain,(
  ( ! [X333,X335,X341,X337,X339,X332,X334,X340,X338] : (p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X337)))) | s(bool,'T_0') = s(bool,X341) | s(bool,X341) = s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X335))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X332,fun(X332,bool)),X333),s(fun(X332,X338),X339),s(fun(X338,X332),X340)))) | ~p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X335))),s(X332,X337))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f7057,plain,(
  ( ! [X333,X335,X341,X337,X339,X332,X334,X340,X336,X338] : (p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X337)))) | ~p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X335))),s(X332,X337)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X332,fun(X332,bool)),X333),s(fun(X332,X338),X339),s(fun(X338,X332),X340)))) | s(bool,'F_0') = s(bool,X336) | s(bool,X341) = s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X335))) | s(bool,'T_0') = s(bool,X341)) )),
  inference(subsumption_resolution,[],[f6137,f357])).
fof(f6137,plain,(
  ( ! [X333,X335,X341,X337,X339,X332,X334,X340,X336,X338] : (~p(s(bool,X336)) | p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X337)))) | ~p(s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X335))),s(X332,X337)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X332,fun(X332,bool)),X333),s(fun(X332,X338),X339),s(fun(X338,X332),X340)))) | s(bool,'F_0') = s(bool,X336) | s(bool,X341) = s(bool,happ(s(fun(X332,bool),happ(s(fun(X332,fun(X332,bool)),X333),s(X332,X334))),s(X332,X335))) | s(bool,'T_0') = s(bool,X341)) )),
  inference(superposition,[],[f40,f401])).
fof(f6781,plain,(
  ( ! [X331,X330] : (s(bool,'F_0') = s(bool,X330) | s(bool,X331) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X331)) ) | $spl87),
  inference(duplicate_literal_removal,[],[f6136])).
fof(f6136,plain,(
  ( ! [X331,X330] : (s(bool,'F_0') = s(bool,X330) | s(bool,'F_0') = s(bool,X330) | s(bool,X331) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X331)) ) | $spl87),
  inference(superposition,[],[f3965,f401])).
fof(f7056,plain,(
  ( ! [X329,X328] : (s(bool,'F_0') = s(bool,X328) | s(bool,X329) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X329)) ) | ($spl43 | $spl87)),
  inference(subsumption_resolution,[],[f6135,f39])).
fof(f6135,plain,(
  ( ! [X329,X328] : (s(bool,'T_0') != s(bool,X328) | s(bool,'F_0') = s(bool,X328) | s(bool,X329) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X329)) ) | ($spl43 | $spl87)),
  inference(superposition,[],[f3966,f401])).
fof(f3966,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | ($spl43 | $spl87)),
  inference(resolution,[],[f2911,f1190])).
fof(f6782,plain,(
  ( ! [X318,X319] : (s(bool,'F_0') = s(bool,X318) | s(bool,X319) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X319)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f6130])).
fof(f6130,plain,(
  ( ! [X318,X319] : (s(bool,'F_0') = s(bool,X318) | s(bool,'F_0') = s(bool,X318) | s(bool,X319) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X319)) ) | $spl22),
  inference(superposition,[],[f234,f401])).
fof(f7055,plain,(
  ( ! [X316,X317] : (s(bool,'F_0') = s(bool,X316) | s(bool,X317) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X317)) ) | $spl27),
  inference(subsumption_resolution,[],[f6129,f39])).
fof(f6129,plain,(
  ( ! [X316,X317] : (s(bool,'T_0') != s(bool,X316) | s(bool,'F_0') = s(bool,X316) | s(bool,X317) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X317)) ) | $spl27),
  inference(superposition,[],[f237,f401])).
fof(f237,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | $spl27),
  inference(cnf_transformation,[],[f237_D])).
fof(f237_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f7054,plain,(
  ( ! [X310,X304,X306,X303,X309,X311,X305,X307] : (p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X309))),s(X303,X310)))) | s(bool,'T_0') = s(bool,X311) | s(bool,X311) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X303,fun(X303,bool)),X304),s(fun(X303,X305),X306),s(fun(X305,X303),X307))) | ~p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X310))),s(X303,X309))))) ) | $spl178),
  inference(cnf_transformation,[],[f7054_D])).
fof(f7054_D,plain,(
  ( ! [X310,X304,X306,X303,X309,X311,X305,X307] : (p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X309))),s(X303,X310)))) | s(bool,'T_0') = s(bool,X311) | s(bool,X311) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X303,fun(X303,bool)),X304),s(fun(X303,X305),X306),s(fun(X305,X303),X307))) | ~p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X310))),s(X303,X309))))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f7053,plain,(
  ( ! [X308,X310,X304,X306,X303,X309,X311,X305,X307] : (p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X309))),s(X303,X310)))) | ~p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X310))),s(X303,X309)))) | s(bool,'F_0') = s(bool,X308) | s(bool,X311) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X303,fun(X303,bool)),X304),s(fun(X303,X305),X306),s(fun(X305,X303),X307))) | s(bool,'T_0') = s(bool,X311)) )),
  inference(subsumption_resolution,[],[f6126,f357])).
fof(f6126,plain,(
  ( ! [X308,X310,X304,X306,X303,X309,X311,X305,X307] : (~p(s(bool,X308)) | p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X309))),s(X303,X310)))) | ~p(s(bool,happ(s(fun(X303,bool),happ(s(fun(X303,fun(X303,bool)),X304),s(X303,X310))),s(X303,X309)))) | s(bool,'F_0') = s(bool,X308) | s(bool,X311) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X303,fun(X303,bool)),X304),s(fun(X303,X305),X306),s(fun(X305,X303),X307))) | s(bool,'T_0') = s(bool,X311)) )),
  inference(superposition,[],[f41,f401])).
fof(f7052,plain,(
  ( ! [X294,X300,X302,X296,X293,X295,X301,X297,X299] : (p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X300)))) | s(bool,'T_0') = s(bool,X302) | s(bool,X302) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X293,fun(X293,bool)),X294),s(fun(X293,X295),X296),s(fun(X295,X293),X297))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X301)))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X301))),s(X293,X300))))) ) | $spl176),
  inference(cnf_transformation,[],[f7052_D])).
fof(f7052_D,plain,(
  ( ! [X294,X300,X302,X296,X293,X295,X301,X297,X299] : (p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X300)))) | s(bool,'T_0') = s(bool,X302) | s(bool,X302) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X293,fun(X293,bool)),X294),s(fun(X293,X295),X296),s(fun(X295,X293),X297))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X301)))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X301))),s(X293,X300))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f7051,plain,(
  ( ! [X294,X300,X302,X296,X298,X293,X295,X301,X297,X299] : (p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X300)))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X301))),s(X293,X300)))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X301)))) | s(bool,'F_0') = s(bool,X298) | s(bool,X302) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X293,fun(X293,bool)),X294),s(fun(X293,X295),X296),s(fun(X295,X293),X297))) | s(bool,'T_0') = s(bool,X302)) )),
  inference(subsumption_resolution,[],[f6125,f357])).
fof(f6125,plain,(
  ( ! [X294,X300,X302,X296,X298,X293,X295,X301,X297,X299] : (~p(s(bool,X298)) | p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X300)))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X301))),s(X293,X300)))) | ~p(s(bool,happ(s(fun(X293,bool),happ(s(fun(X293,fun(X293,bool)),X294),s(X293,X299))),s(X293,X301)))) | s(bool,'F_0') = s(bool,X298) | s(bool,X302) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X293,fun(X293,bool)),X294),s(fun(X293,X295),X296),s(fun(X295,X293),X297))) | s(bool,'T_0') = s(bool,X302)) )),
  inference(superposition,[],[f40,f401])).
fof(f6783,plain,(
  ( ! [X249,X250,X248,X247] : (s(bool,X247) != s(bool,X248) | s(bool,X248) != s(bool,X249) | s(bool,'T_0') = s(bool,X249) | s(bool,'F_0') = s(bool,X247) | s(bool,'T_0') = s(bool,X250)) ) | ($spl43 | $spl76)),
  inference(duplicate_literal_removal,[],[f6112])).
fof(f6112,plain,(
  ( ! [X249,X250,X248,X247] : (s(bool,X247) != s(bool,X248) | s(bool,X248) != s(bool,X249) | s(bool,'T_0') = s(bool,X249) | s(bool,'F_0') = s(bool,X247) | s(bool,'T_0') = s(bool,X250) | s(bool,'T_0') = s(bool,X250)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f401])).
fof(f5266,plain,(
  ( ! [X45,X44,X42] : (s(bool,'T_0') != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X45) | s(bool,X44) = s(bool,X45)) ) | $spl136),
  inference(cnf_transformation,[],[f5266_D])).
fof(f5266_D,plain,(
  ( ! [X45,X44,X42] : (s(bool,'T_0') != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X42) != s(bool,X45) | s(bool,X44) = s(bool,X45)) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f6784,plain,(
  ( ! [X243,X246,X244,X245] : (s(bool,X243) != s(bool,X244) | s(bool,'T_0') != s(bool,X245) | s(bool,X244) = s(bool,X245) | s(bool,'F_0') = s(bool,X243) | s(bool,'T_0') = s(bool,X246)) ) | $spl43),
  inference(duplicate_literal_removal,[],[f6111])).
fof(f6111,plain,(
  ( ! [X243,X246,X244,X245] : (s(bool,X243) != s(bool,X244) | s(bool,'T_0') != s(bool,X245) | s(bool,X244) = s(bool,X245) | s(bool,'F_0') = s(bool,X243) | s(bool,'T_0') = s(bool,X246) | s(bool,'T_0') = s(bool,X246)) ) | $spl43),
  inference(superposition,[],[f1619,f401])).
fof(f6785,plain,(
  ( ! [X241,X239,X242,X240] : (s(bool,X239) != s(bool,X240) | s(bool,X240) = s(bool,X241) | s(bool,'F_0') = s(bool,X241) | s(bool,'F_0') = s(bool,X239) | s(bool,'T_0') = s(bool,X242)) ) | $spl43),
  inference(duplicate_literal_removal,[],[f6110])).
fof(f6110,plain,(
  ( ! [X241,X239,X242,X240] : (s(bool,X239) != s(bool,X240) | s(bool,X240) = s(bool,X241) | s(bool,'F_0') = s(bool,X241) | s(bool,'F_0') = s(bool,X239) | s(bool,'T_0') = s(bool,X242) | s(bool,'T_0') = s(bool,X242)) ) | $spl43),
  inference(superposition,[],[f1169,f401])).
fof(f7044,plain,(
  ( ! [X212,X210,X213,X214] : (s(bool,'F_0') = s(bool,X212) | s(bool,'T_0') = s(bool,X214) | s(bool,X210) = s(bool,X214) | s(bool,X212) != s(bool,X213) | s(bool,X210) = s(bool,X213)) ) | $spl174),
  inference(cnf_transformation,[],[f7044_D])).
fof(f7044_D,plain,(
  ( ! [X212,X210,X213,X214] : (s(bool,'F_0') = s(bool,X212) | s(bool,'T_0') = s(bool,X214) | s(bool,X210) = s(bool,X214) | s(bool,X212) != s(bool,X213) | s(bool,X210) = s(bool,X213)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7043,plain,(
  ( ! [X212,X210,X213,X211,X214] : (s(bool,'F_0') = s(bool,X212) | s(bool,X210) = s(bool,X213) | s(bool,X212) != s(bool,X213) | s(bool,'F_0') = s(bool,X211) | s(bool,X210) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) ) | $spl112),
  inference(subsumption_resolution,[],[f6103,f357])).
fof(f6103,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,X211)) | s(bool,'F_0') = s(bool,X212) | s(bool,X210) = s(bool,X213) | s(bool,X212) != s(bool,X213) | s(bool,'F_0') = s(bool,X211) | s(bool,X210) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) ) | $spl112),
  inference(superposition,[],[f3711,f401])).
fof(f3711,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | s(bool,'F_0') = s(bool,X1) | s(bool,X0) = s(bool,X2) | s(bool,X0) != s(bool,X1)) ) | $spl112),
  inference(resolution,[],[f3596,f38])).
fof(f7042,plain,(
  ( ! [X194,X192,X195,X193] : (s(bool,X192) != s(bool,X194) | s(bool,'T_0') = s(bool,X194) | s(bool,'F_0') = s(bool,X193) | s(bool,X192) = s(bool,X195) | s(bool,'T_0') = s(bool,X195)) ) | $spl104),
  inference(subsumption_resolution,[],[f6099,f357])).
fof(f6099,plain,(
  ( ! [X194,X192,X195,X193] : (~p(s(bool,X193)) | s(bool,X192) != s(bool,X194) | s(bool,'T_0') = s(bool,X194) | s(bool,'F_0') = s(bool,X193) | s(bool,X192) = s(bool,X195) | s(bool,'T_0') = s(bool,X195)) ) | $spl104),
  inference(superposition,[],[f3326,f401])).
fof(f7041,plain,(
  ( ! [X191,X189,X187,X190] : (s(bool,X189) != s(bool,X190) | s(bool,'T_0') = s(bool,X191) | s(bool,X187) = s(bool,X191) | s(bool,X187) != s(bool,X189) | s(bool,'T_0') = s(bool,X190)) ) | $spl172),
  inference(cnf_transformation,[],[f7041_D])).
fof(f7041_D,plain,(
  ( ! [X191,X189,X187,X190] : (s(bool,X189) != s(bool,X190) | s(bool,'T_0') = s(bool,X191) | s(bool,X187) = s(bool,X191) | s(bool,X187) != s(bool,X189) | s(bool,'T_0') = s(bool,X190)) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f7040,plain,(
  ( ! [X191,X189,X187,X190,X188] : (s(bool,X189) != s(bool,X190) | s(bool,'T_0') = s(bool,X190) | s(bool,X187) != s(bool,X189) | s(bool,'F_0') = s(bool,X188) | s(bool,X187) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6098,f357])).
fof(f6098,plain,(
  ( ! [X191,X189,X187,X190,X188] : (~p(s(bool,X188)) | s(bool,X189) != s(bool,X190) | s(bool,'T_0') = s(bool,X190) | s(bool,X187) != s(bool,X189) | s(bool,'F_0') = s(bool,X188) | s(bool,X187) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3321,f401])).
fof(f3321,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3280,f43])).
fof(f3280,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,X18) != s(bool,X19) | s(bool,'T_0') = s(bool,X19) | ~p(s(bool,'T_0')) | ~p(s(bool,X17))) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f38])).
fof(f7038,plain,(
  ( ! [X167,X166,X164,X168] : (s(bool,X166) != s(bool,X167) | s(bool,'T_0') = s(bool,X168) | s(bool,X164) = s(bool,X168) | s(bool,'T_0') != s(bool,X166) | s(bool,X164) = s(bool,X167)) ) | $spl170),
  inference(cnf_transformation,[],[f7038_D])).
fof(f7038_D,plain,(
  ( ! [X167,X166,X164,X168] : (s(bool,X166) != s(bool,X167) | s(bool,'T_0') = s(bool,X168) | s(bool,X164) = s(bool,X168) | s(bool,'T_0') != s(bool,X166) | s(bool,X164) = s(bool,X167)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f7037,plain,(
  ( ! [X167,X165,X166,X164,X168] : (s(bool,X166) != s(bool,X167) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') != s(bool,X166) | s(bool,'F_0') = s(bool,X165) | s(bool,X164) = s(bool,X168) | s(bool,'T_0') = s(bool,X168)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6093,f357])).
fof(f6093,plain,(
  ( ! [X167,X165,X166,X164,X168] : (~p(s(bool,X165)) | s(bool,X166) != s(bool,X167) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') != s(bool,X166) | s(bool,'F_0') = s(bool,X165) | s(bool,X164) = s(bool,X168) | s(bool,'T_0') = s(bool,X168)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2565,f401])).
fof(f2565,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X12)) | s(bool,X10) != s(bool,X11) | s(bool,X11) = s(bool,X12) | s(bool,'T_0') != s(bool,X10)) ) | ($spl43 | $spl76)),
  inference(resolution,[],[f2494,f38])).
fof(f7036,plain,(
  ( ! [X158,X159,X157,X155] : (s(bool,'F_0') != s(bool,X157) | s(bool,'T_0') = s(bool,X159) | s(bool,X155) = s(bool,X159) | s(bool,X157) = s(bool,X158) | s(bool,X155) = s(bool,X158)) ) | $spl168),
  inference(cnf_transformation,[],[f7036_D])).
fof(f7036_D,plain,(
  ( ! [X158,X159,X157,X155] : (s(bool,'F_0') != s(bool,X157) | s(bool,'T_0') = s(bool,X159) | s(bool,X155) = s(bool,X159) | s(bool,X157) = s(bool,X158) | s(bool,X155) = s(bool,X158)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f4517,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) | $spl126),
  inference(cnf_transformation,[],[f4517_D])).
fof(f4517_D,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f7035,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X156)) | s(bool,'F_0') != s(bool,X157) | s(bool,X155) = s(bool,X158) | s(bool,X157) = s(bool,X158) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) ) | $spl76),
  inference(subsumption_resolution,[],[f6091,f2143])).
fof(f6091,plain,(
  ( ! [X158,X156,X159,X157,X155] : (~p(s(bool,X156)) | s(bool,'F_0') != s(bool,X157) | s(bool,X155) = s(bool,X158) | s(bool,X157) = s(bool,X158) | s(bool,'F_0') = s(bool,X156) | s(bool,X155) = s(bool,X159) | s(bool,'T_0') = s(bool,X159)) ) | $spl76),
  inference(superposition,[],[f2504,f401])).
fof(f2504,plain,(
  ( ! [X12,X10,X11] : (~p(s(bool,X12)) | s(bool,'F_0') != s(bool,X10) | s(bool,X11) = s(bool,X12) | s(bool,X10) = s(bool,X11)) ) | $spl76),
  inference(resolution,[],[f2404,f38])).
fof(f7034,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,X134) != s(bool,X136) | s(bool,'T_0') = s(bool,X136) | s(bool,'F_0') = s(bool,X135) | s(bool,X134) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f6086,f39])).
fof(f6086,plain,(
  ( ! [X134,X136,X135,X137] : (s(bool,'T_0') != s(bool,X135) | s(bool,X134) != s(bool,X136) | s(bool,'T_0') = s(bool,X136) | s(bool,'F_0') = s(bool,X135) | s(bool,X134) = s(bool,X137) | s(bool,'T_0') = s(bool,X137)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f401])).
fof(f7032,plain,(
  ( ! [X123,X121,X122,X119] : (s(bool,'T_0') != s(bool,X121) | s(bool,'T_0') = s(bool,X123) | s(bool,X119) = s(bool,X123) | s(bool,X119) != s(bool,X122) | s(bool,X121) = s(bool,X122)) ) | $spl166),
  inference(cnf_transformation,[],[f7032_D])).
fof(f7032_D,plain,(
  ( ! [X123,X121,X122,X119] : (s(bool,'T_0') != s(bool,X121) | s(bool,'T_0') = s(bool,X123) | s(bool,X119) = s(bool,X123) | s(bool,X119) != s(bool,X122) | s(bool,X121) = s(bool,X122)) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f7031,plain,(
  ( ! [X123,X121,X122,X120,X119] : (s(bool,'T_0') != s(bool,X121) | s(bool,X121) = s(bool,X122) | s(bool,X119) != s(bool,X122) | s(bool,'F_0') = s(bool,X120) | s(bool,X119) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) ) | $spl43),
  inference(subsumption_resolution,[],[f6082,f357])).
fof(f6082,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'T_0') != s(bool,X121) | s(bool,X121) = s(bool,X122) | s(bool,X119) != s(bool,X122) | s(bool,'F_0') = s(bool,X120) | s(bool,X119) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) ) | $spl43),
  inference(superposition,[],[f1811,f401])).
fof(f1811,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | s(bool,'T_0') != s(bool,X19) | s(bool,X18) = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl43),
  inference(subsumption_resolution,[],[f1774,f43])).
fof(f1774,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,'T_0') != s(bool,X19) | s(bool,X18) = s(bool,X19) | ~p(s(bool,'T_0')) | ~p(s(bool,X17))) ) | $spl43),
  inference(superposition,[],[f1619,f38])).
fof(f7029,plain,(
  ( ! [X118,X116,X114,X117,X115] : (s(bool,X116) = s(bool,X117) | s(bool,'F_0') = s(bool,X117) | s(bool,X114) != s(bool,X116) | s(bool,'F_0') = s(bool,X115) | s(bool,X114) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) ) | $spl43),
  inference(subsumption_resolution,[],[f6081,f357])).
fof(f6081,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,X115)) | s(bool,X116) = s(bool,X117) | s(bool,'F_0') = s(bool,X117) | s(bool,X114) != s(bool,X116) | s(bool,'F_0') = s(bool,X115) | s(bool,X114) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) ) | $spl43),
  inference(superposition,[],[f1733,f401])).
fof(f1733,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | s(bool,X18) = s(bool,X19) | s(bool,'F_0') = s(bool,X19) | s(bool,X17) != s(bool,X18)) ) | $spl43),
  inference(subsumption_resolution,[],[f1690,f43])).
fof(f1690,plain,(
  ( ! [X19,X17,X18] : (s(bool,X17) != s(bool,X18) | s(bool,X18) = s(bool,X19) | s(bool,'F_0') = s(bool,X19) | ~p(s(bool,'T_0')) | ~p(s(bool,X17))) ) | $spl43),
  inference(superposition,[],[f1169,f38])).
fof(f7028,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') != s(bool,X107) | s(bool,X105) = s(bool,X107) | s(bool,'F_0') = s(bool,X106) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) ) | $spl43),
  inference(subsumption_resolution,[],[f6079,f39])).
fof(f6079,plain,(
  ( ! [X107,X105,X108,X106] : (s(bool,'T_0') != s(bool,X106) | s(bool,'T_0') != s(bool,X107) | s(bool,X105) = s(bool,X107) | s(bool,'F_0') = s(bool,X106) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) ) | $spl43),
  inference(superposition,[],[f1619,f401])).
fof(f7027,plain,(
  ( ! [X103,X101,X102,X104] : (s(bool,X101) = s(bool,X103) | s(bool,'T_0') != s(bool,X103) | s(bool,'F_0') = s(bool,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) ) | $spl43),
  inference(subsumption_resolution,[],[f6078,f357])).
fof(f6078,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,X102)) | s(bool,X101) = s(bool,X103) | s(bool,'T_0') != s(bool,X103) | s(bool,'F_0') = s(bool,X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) ) | $spl43),
  inference(superposition,[],[f1223,f401])).
fof(f7022,plain,(
  ( ! [X61,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X61) = s(bool,X64) | s(bool,X63) = s(bool,X64) | s(bool,'F_0') = s(bool,X62) | s(bool,X61) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f6068,f357])).
fof(f6068,plain,(
  ( ! [X61,X64,X62,X65,X63] : (~p(s(bool,X62)) | s(bool,'T_0') = s(bool,X63) | s(bool,X61) = s(bool,X64) | s(bool,X63) = s(bool,X64) | s(bool,'F_0') = s(bool,X62) | s(bool,X61) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f267,f401])).
fof(f267,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,X4)) | s(bool,'T_0') = s(bool,X2) | s(bool,X3) = s(bool,X4) | s(bool,X2) = s(bool,X3)) )),
  inference(resolution,[],[f245,f38])).
fof(f1503,plain,(
  ( ! [X92,X88,X87,X93,X91,X89,X86] : (p(s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X89))),s(X86,X88)))) | s(bool,'F_0') = s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X88))),s(X86,X89))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X86,fun(X86,bool)),X87),s(fun(X86,X91),X92),s(fun(X91,X86),X93))))) ) | $spl62),
  inference(cnf_transformation,[],[f1503_D])).
fof(f1503_D,plain,(
  ( ! [X92,X88,X87,X93,X91,X89,X86] : (p(s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X89))),s(X86,X88)))) | s(bool,'F_0') = s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X88))),s(X86,X89))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X86,fun(X86,bool)),X87),s(fun(X86,X91),X92),s(fun(X91,X86),X93))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f7012,plain,(
  ( ! [X356,X358,X352,X354,X351,X357,X359,X353,X355] : (~p(s(bool,X355)) | p(s(bool,happ(s(fun(X351,bool),happ(s(fun(X351,fun(X351,bool)),X352),s(X351,X354))),s(X351,X353)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X351,fun(X351,bool)),X352),s(fun(X351,X356),X357),s(fun(X356,X351),X358)))) | s(bool,'F_0') = s(bool,happ(s(fun(X351,bool),happ(s(fun(X351,fun(X351,bool)),X352),s(X351,X353))),s(X351,X354))) | s(bool,X355) = s(bool,X359)) ) | $spl43),
  inference(subsumption_resolution,[],[f6038,f1223])).
fof(f6038,plain,(
  ( ! [X356,X358,X352,X354,X351,X357,X359,X353,X355] : (~p(s(bool,X355)) | p(s(bool,happ(s(fun(X351,bool),happ(s(fun(X351,fun(X351,bool)),X352),s(X351,X354))),s(X351,X353)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X351,fun(X351,bool)),X352),s(fun(X351,X356),X357),s(fun(X356,X351),X358)))) | s(bool,'F_0') = s(bool,happ(s(fun(X351,bool),happ(s(fun(X351,fun(X351,bool)),X352),s(X351,X353))),s(X351,X354))) | s(bool,X355) = s(bool,X359) | s(bool,'T_0') = s(bool,X359)) )),
  inference(superposition,[],[f41,f401])).
fof(f3252,plain,(
  ( ! [X111,X107,X105,X112,X110,X108,X106,X113] : (p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X108)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X107))))) ) | $spl98),
  inference(cnf_transformation,[],[f3252_D])).
fof(f3252_D,plain,(
  ( ! [X111,X107,X105,X112,X110,X108,X106,X113] : (p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X108)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X107))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f7011,plain,(
  ( ! [X333,X335,X331,X337,X339,X332,X334,X340,X336,X338] : (~p(s(bool,X335)) | p(s(bool,happ(s(fun(X331,bool),happ(s(fun(X331,fun(X331,bool)),X332),s(X331,X336))),s(X331,X334)))) | ~p(s(bool,happ(s(fun(X331,bool),happ(s(fun(X331,fun(X331,bool)),X332),s(X331,X336))),s(X331,X333)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X331,fun(X331,bool)),X332),s(fun(X331,X337),X338),s(fun(X337,X331),X339)))) | s(bool,'F_0') = s(bool,happ(s(fun(X331,bool),happ(s(fun(X331,fun(X331,bool)),X332),s(X331,X333))),s(X331,X334))) | s(bool,X335) = s(bool,X340)) ) | $spl43),
  inference(subsumption_resolution,[],[f6036,f1223])).
fof(f6036,plain,(
  ( ! [X333,X335,X331,X337,X339,X332,X334,X340,X336,X338] : (~p(s(bool,X335)) | p(s(bool,happ(s(fun(X331,bool),happ(s(fun(X331,fun(X331,bool)),X332),s(X331,X336))),s(X331,X334)))) | ~p(s(bool,happ(s(fun(X331,bool),happ(s(fun(X331,fun(X331,bool)),X332),s(X331,X336))),s(X331,X333)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X331,fun(X331,bool)),X332),s(fun(X331,X337),X338),s(fun(X337,X331),X339)))) | s(bool,'F_0') = s(bool,happ(s(fun(X331,bool),happ(s(fun(X331,fun(X331,bool)),X332),s(X331,X333))),s(X331,X334))) | s(bool,X335) = s(bool,X340) | s(bool,'T_0') = s(bool,X340)) )),
  inference(superposition,[],[f40,f401])).
fof(f3266,plain,(
  ( ! [X111,X107,X105,X112,X110,X108,X106,X113] : (p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X110)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X110))))) ) | $spl100),
  inference(cnf_transformation,[],[f3266_D])).
fof(f3266_D,plain,(
  ( ! [X111,X107,X105,X112,X110,X108,X106,X113] : (p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X110)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X110))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f7010,plain,(
  ( ! [X325,X327,X321,X323,X329,X324,X326,X322,X328,X330] : (~p(s(bool,X325)) | p(s(bool,happ(s(fun(X321,bool),happ(s(fun(X321,fun(X321,bool)),X322),s(X321,X323))),s(X321,X326)))) | ~p(s(bool,happ(s(fun(X321,bool),happ(s(fun(X321,fun(X321,bool)),X322),s(X321,X324))),s(X321,X326)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X321,fun(X321,bool)),X322),s(fun(X321,X327),X328),s(fun(X327,X321),X329)))) | s(bool,'F_0') = s(bool,happ(s(fun(X321,bool),happ(s(fun(X321,fun(X321,bool)),X322),s(X321,X323))),s(X321,X324))) | s(bool,X325) = s(bool,X330)) ) | $spl43),
  inference(subsumption_resolution,[],[f6035,f1223])).
fof(f6035,plain,(
  ( ! [X325,X327,X321,X323,X329,X324,X326,X322,X328,X330] : (~p(s(bool,X325)) | p(s(bool,happ(s(fun(X321,bool),happ(s(fun(X321,fun(X321,bool)),X322),s(X321,X323))),s(X321,X326)))) | ~p(s(bool,happ(s(fun(X321,bool),happ(s(fun(X321,fun(X321,bool)),X322),s(X321,X324))),s(X321,X326)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X321,fun(X321,bool)),X322),s(fun(X321,X327),X328),s(fun(X327,X321),X329)))) | s(bool,'F_0') = s(bool,happ(s(fun(X321,bool),happ(s(fun(X321,fun(X321,bool)),X322),s(X321,X323))),s(X321,X324))) | s(bool,X325) = s(bool,X330) | s(bool,'T_0') = s(bool,X330)) )),
  inference(superposition,[],[f40,f401])).
fof(f1510,plain,(
  ( ! [X94,X92,X90,X88,X95,X91,X89] : (p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X94))),s(X88,X95)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X88,fun(X88,bool)),X89),s(fun(X88,X90),X91),s(fun(X90,X88),X92))) | ~p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X95))),s(X88,X94))))) ) | $spl64),
  inference(cnf_transformation,[],[f1510_D])).
fof(f1510_D,plain,(
  ( ! [X94,X92,X90,X88,X95,X91,X89] : (p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X94))),s(X88,X95)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X88,fun(X88,bool)),X89),s(fun(X88,X90),X91),s(fun(X90,X88),X92))) | ~p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X95))),s(X88,X94))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f7005,plain,(
  ( ! [X292,X294,X300,X296,X298,X293,X295,X297,X299] : (~p(s(bool,X297)) | p(s(bool,happ(s(fun(X292,bool),happ(s(fun(X292,fun(X292,bool)),X293),s(X292,X298))),s(X292,X299)))) | ~p(s(bool,happ(s(fun(X292,bool),happ(s(fun(X292,fun(X292,bool)),X293),s(X292,X299))),s(X292,X298)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X292,fun(X292,bool)),X293),s(fun(X292,X294),X295),s(fun(X294,X292),X296))) | s(bool,X297) = s(bool,X300)) ) | $spl43),
  inference(subsumption_resolution,[],[f6024,f1223])).
fof(f6024,plain,(
  ( ! [X292,X294,X300,X296,X298,X293,X295,X297,X299] : (~p(s(bool,X297)) | p(s(bool,happ(s(fun(X292,bool),happ(s(fun(X292,fun(X292,bool)),X293),s(X292,X298))),s(X292,X299)))) | ~p(s(bool,happ(s(fun(X292,bool),happ(s(fun(X292,fun(X292,bool)),X293),s(X292,X299))),s(X292,X298)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X292,fun(X292,bool)),X293),s(fun(X292,X294),X295),s(fun(X294,X292),X296))) | s(bool,X297) = s(bool,X300) | s(bool,'T_0') = s(bool,X300)) )),
  inference(superposition,[],[f41,f401])).
fof(f3273,plain,(
  ( ! [X103,X101,X107,X102,X100,X108,X106,X104] : (p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X107)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X100,fun(X100,bool)),X101),s(fun(X100,X102),X103),s(fun(X102,X100),X104))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X108)))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X108))),s(X100,X107))))) ) | $spl102),
  inference(cnf_transformation,[],[f3273_D])).
fof(f3273_D,plain,(
  ( ! [X103,X101,X107,X102,X100,X108,X106,X104] : (p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X107)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X100,fun(X100,bool)),X101),s(fun(X100,X102),X103),s(fun(X102,X100),X104))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X108)))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X108))),s(X100,X107))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f7004,plain,(
  ( ! [X282,X288,X290,X285,X287,X283,X289,X291,X284,X286] : (~p(s(bool,X287)) | p(s(bool,happ(s(fun(X282,bool),happ(s(fun(X282,fun(X282,bool)),X283),s(X282,X288))),s(X282,X289)))) | ~p(s(bool,happ(s(fun(X282,bool),happ(s(fun(X282,fun(X282,bool)),X283),s(X282,X290))),s(X282,X289)))) | ~p(s(bool,happ(s(fun(X282,bool),happ(s(fun(X282,fun(X282,bool)),X283),s(X282,X288))),s(X282,X290)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X282,fun(X282,bool)),X283),s(fun(X282,X284),X285),s(fun(X284,X282),X286))) | s(bool,X287) = s(bool,X291)) ) | $spl43),
  inference(subsumption_resolution,[],[f6023,f1223])).
fof(f6023,plain,(
  ( ! [X282,X288,X290,X285,X287,X283,X289,X291,X284,X286] : (~p(s(bool,X287)) | p(s(bool,happ(s(fun(X282,bool),happ(s(fun(X282,fun(X282,bool)),X283),s(X282,X288))),s(X282,X289)))) | ~p(s(bool,happ(s(fun(X282,bool),happ(s(fun(X282,fun(X282,bool)),X283),s(X282,X290))),s(X282,X289)))) | ~p(s(bool,happ(s(fun(X282,bool),happ(s(fun(X282,fun(X282,bool)),X283),s(X282,X288))),s(X282,X290)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X282,fun(X282,bool)),X283),s(fun(X282,X284),X285),s(fun(X284,X282),X286))) | s(bool,X287) = s(bool,X291) | s(bool,'T_0') = s(bool,X291)) )),
  inference(superposition,[],[f40,f401])).
fof(f6998,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,X211)) | s(bool,'F_0') = s(bool,X212) | s(bool,X210) = s(bool,X213) | s(bool,X212) != s(bool,X213) | s(bool,'F_0') = s(bool,X210) | s(bool,X211) = s(bool,X214)) ) | ($spl43 | $spl112)),
  inference(subsumption_resolution,[],[f6001,f1223])).
fof(f6001,plain,(
  ( ! [X212,X210,X213,X211,X214] : (~p(s(bool,X211)) | s(bool,'F_0') = s(bool,X212) | s(bool,X210) = s(bool,X213) | s(bool,X212) != s(bool,X213) | s(bool,'F_0') = s(bool,X210) | s(bool,X211) = s(bool,X214) | s(bool,'T_0') = s(bool,X214)) ) | $spl112),
  inference(superposition,[],[f3711,f401])).
fof(f6995,plain,(
  ( ! [X191,X189,X187,X190,X188] : (~p(s(bool,X188)) | s(bool,X189) != s(bool,X190) | s(bool,'T_0') = s(bool,X190) | s(bool,X187) != s(bool,X189) | s(bool,'F_0') = s(bool,X187) | s(bool,X188) = s(bool,X191)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5996,f1223])).
fof(f5996,plain,(
  ( ! [X191,X189,X187,X190,X188] : (~p(s(bool,X188)) | s(bool,X189) != s(bool,X190) | s(bool,'T_0') = s(bool,X190) | s(bool,X187) != s(bool,X189) | s(bool,'F_0') = s(bool,X187) | s(bool,X188) = s(bool,X191) | s(bool,'T_0') = s(bool,X191)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3321,f401])).
fof(f6793,plain,(
  ( ! [X182,X180,X181,X179] : (s(bool,X180) != s(bool,X181) | s(bool,'T_0') = s(bool,X181) | s(bool,'F_0') = s(bool,X179) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl43 | $spl76)),
  inference(duplicate_literal_removal,[],[f5994])).
fof(f5994,plain,(
  ( ! [X182,X180,X181,X179] : (s(bool,X180) != s(bool,X181) | s(bool,'T_0') = s(bool,X181) | s(bool,'F_0') = s(bool,X179) | s(bool,'F_0') = s(bool,X179) | s(bool,X180) = s(bool,X182) | s(bool,'T_0') = s(bool,X182)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3318,f401])).
fof(f6993,plain,(
  ( ! [X167,X165,X166,X164,X168] : (~p(s(bool,X165)) | s(bool,X166) != s(bool,X167) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') != s(bool,X166) | s(bool,'F_0') = s(bool,X164) | s(bool,X165) = s(bool,X168)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5991,f1223])).
fof(f5991,plain,(
  ( ! [X167,X165,X166,X164,X168] : (~p(s(bool,X165)) | s(bool,X166) != s(bool,X167) | s(bool,X164) = s(bool,X167) | s(bool,'T_0') != s(bool,X166) | s(bool,'F_0') = s(bool,X164) | s(bool,X165) = s(bool,X168) | s(bool,'T_0') = s(bool,X168)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2565,f401])).
fof(f6983,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'T_0') != s(bool,X121) | s(bool,X121) = s(bool,X122) | s(bool,X119) != s(bool,X122) | s(bool,'F_0') = s(bool,X119) | s(bool,X120) = s(bool,X123)) ) | $spl43),
  inference(subsumption_resolution,[],[f5980,f1223])).
fof(f5980,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,X120)) | s(bool,'T_0') != s(bool,X121) | s(bool,X121) = s(bool,X122) | s(bool,X119) != s(bool,X122) | s(bool,'F_0') = s(bool,X119) | s(bool,X120) = s(bool,X123) | s(bool,'T_0') = s(bool,X123)) ) | $spl43),
  inference(superposition,[],[f1811,f401])).
fof(f6981,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,X115)) | s(bool,X116) = s(bool,X117) | s(bool,'F_0') = s(bool,X117) | s(bool,X114) != s(bool,X116) | s(bool,'F_0') = s(bool,X114) | s(bool,X115) = s(bool,X118)) ) | $spl43),
  inference(subsumption_resolution,[],[f5979,f1223])).
fof(f5979,plain,(
  ( ! [X118,X116,X114,X117,X115] : (~p(s(bool,X115)) | s(bool,X116) = s(bool,X117) | s(bool,'F_0') = s(bool,X117) | s(bool,X114) != s(bool,X116) | s(bool,'F_0') = s(bool,X114) | s(bool,X115) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) ) | $spl43),
  inference(superposition,[],[f1733,f401])).
fof(f6799,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) = s(bool,X65) | s(bool,X62) = s(bool,X66) | s(bool,'F_0') = s(bool,X66)) )),
  inference(duplicate_literal_removal,[],[f5942])).
fof(f5942,plain,(
  ( ! [X66,X64,X62,X65,X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,X64) = s(bool,X65) | s(bool,'F_0') = s(bool,X64) | s(bool,X62) = s(bool,X65) | s(bool,X62) = s(bool,X66) | s(bool,'F_0') = s(bool,X66) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f401,f401])).
fof(f6802,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,X68) = s(bool,X69) | s(bool,X70) = s(bool,X71) | s(bool,'F_0') = s(bool,X70) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X72)) )),
  inference(duplicate_literal_removal,[],[f5924])).
fof(f5924,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,X68) = s(bool,X69) | s(bool,X70) = s(bool,X71) | s(bool,'F_0') = s(bool,X70) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X68) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,X72)) )),
  inference(superposition,[],[f401,f401])).
fof(f6805,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X79) | s(bool,X80) = s(bool,X83) | s(bool,'F_0') = s(bool,X83)) )),
  inference(duplicate_literal_removal,[],[f5890])).
fof(f5890,plain,(
  ( ! [X80,X83,X81,X79,X82] : (s(bool,X80) = s(bool,X81) | s(bool,X81) = s(bool,X82) | s(bool,'F_0') = s(bool,X82) | s(bool,'T_0') = s(bool,X79) | s(bool,X80) = s(bool,X83) | s(bool,'F_0') = s(bool,X83) | s(bool,'T_0') = s(bool,X79)) )),
  inference(superposition,[],[f401,f401])).
fof(f6809,plain,(
  ( ! [X130,X131] : (s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X131)) ) | $spl8),
  inference(duplicate_literal_removal,[],[f5869])).
fof(f5869,plain,(
  ( ! [X130,X131] : (s(bool,'T_0') = s(bool,X130) | s(bool,X131) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X131) | s(bool,'T_0') = s(bool,X130)) ) | $spl8),
  inference(superposition,[],[f401,f165])).
fof(f6810,plain,(
  ( ! [X128,X129] : (s(bool,'T_0') = s(bool,X128) | s(bool,X129) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X129)) ) | $spl84),
  inference(duplicate_literal_removal,[],[f5868])).
fof(f5868,plain,(
  ( ! [X128,X129] : (s(bool,'T_0') = s(bool,X128) | s(bool,X129) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X129) | s(bool,'T_0') = s(bool,X128)) ) | $spl84),
  inference(superposition,[],[f401,f3925])).
fof(f6811,plain,(
  ( ! [X127,X126] : (s(bool,'T_0') = s(bool,X126) | s(bool,X127) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X127)) )),
  inference(duplicate_literal_removal,[],[f5867])).
fof(f5867,plain,(
  ( ! [X127,X126] : (s(bool,'T_0') = s(bool,X126) | s(bool,X127) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X127) | s(bool,'T_0') = s(bool,X126)) )),
  inference(superposition,[],[f401,f107])).
fof(f6812,plain,(
  ( ! [X125,X124] : (s(bool,'T_0') = s(bool,X124) | s(bool,X125) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X125)) ) | $spl2),
  inference(duplicate_literal_removal,[],[f5866])).
fof(f5866,plain,(
  ( ! [X125,X124] : (s(bool,'T_0') = s(bool,X124) | s(bool,X125) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'F_0') = s(bool,X125) | s(bool,'T_0') = s(bool,X124)) ) | $spl2),
  inference(superposition,[],[f401,f145])).
fof(f6813,plain,(
  ( ! [X121,X120] : (s(bool,'T_0') = s(bool,X120) | s(bool,X121) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X121)) ) | $spl24),
  inference(duplicate_literal_removal,[],[f5864])).
fof(f5864,plain,(
  ( ! [X121,X120] : (s(bool,'T_0') = s(bool,X120) | s(bool,X121) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X121) | s(bool,'T_0') = s(bool,X120)) ) | $spl24),
  inference(superposition,[],[f401,f235])).
fof(f6814,plain,(
  ( ! [X118,X119] : (s(bool,'T_0') = s(bool,X118) | s(bool,X119) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X119)) )),
  inference(duplicate_literal_removal,[],[f5863])).
fof(f5863,plain,(
  ( ! [X118,X119] : (s(bool,'T_0') = s(bool,X118) | s(bool,X119) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X119) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f401,f85])).
fof(f6815,plain,(
  ( ! [X116,X117] : (s(bool,'T_0') = s(bool,X116) | s(bool,X117) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X117)) )),
  inference(duplicate_literal_removal,[],[f5862])).
fof(f5862,plain,(
  ( ! [X116,X117] : (s(bool,'T_0') = s(bool,X116) | s(bool,X117) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X117) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f401,f62])).
fof(f6819,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,X100) = s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,'T_0') = s(bool,X101) | s(bool,X100) = s(bool,X103) | s(bool,X103) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(duplicate_literal_removal,[],[f5858])).
fof(f5858,plain,(
  ( ! [X103,X101,X102,X100,X104] : (s(bool,X100) = s(bool,X101) | s(bool,'F_0') = s(bool,X102) | s(bool,'F_0') = s(bool,X102) | s(bool,'T_0') = s(bool,X101) | s(bool,X100) = s(bool,X103) | s(bool,X103) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f401,f401])).
fof(f6820,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,X70) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X72) | s(bool,X69) = s(bool,X73)) )),
  inference(duplicate_literal_removal,[],[f5851])).
fof(f5851,plain,(
  ( ! [X70,X72,X71,X69,X73] : (s(bool,'T_0') = s(bool,X70) | s(bool,X69) = s(bool,X71) | s(bool,'F_0') = s(bool,X71) | s(bool,'T_0') = s(bool,X70) | s(bool,X72) = s(bool,X73) | s(bool,'F_0') = s(bool,X72) | s(bool,X69) = s(bool,X73)) )),
  inference(superposition,[],[f401,f401])).
fof(f6824,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,X52) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(duplicate_literal_removal,[],[f5823])).
fof(f5823,plain,(
  ( ! [X54,X52,X56,X55,X53] : (s(bool,'F_0') = s(bool,X53) | s(bool,X52) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'T_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X53) | s(bool,X52) = s(bool,X56) | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f401,f401])).
fof(f6827,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,X93) = s(bool,X94) | s(bool,X94) = s(bool,X95) | s(bool,X95) = s(bool,X96) | s(bool,'T_0') = s(bool,X96) | s(bool,'F_0') = s(bool,X97) | s(bool,'T_0') = s(bool,X93)) )),
  inference(duplicate_literal_removal,[],[f5809])).
fof(f5809,plain,(
  ( ! [X94,X97,X95,X93,X96] : (s(bool,X93) = s(bool,X94) | s(bool,X94) = s(bool,X95) | s(bool,X95) = s(bool,X96) | s(bool,'T_0') = s(bool,X96) | s(bool,'F_0') = s(bool,X97) | s(bool,'F_0') = s(bool,X97) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f401,f401])).
fof(f6830,plain,(
  ( ! [X132,X133] : (s(bool,'F_0') = s(bool,X132) | s(bool,X133) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X133)) ) | $spl87),
  inference(duplicate_literal_removal,[],[f5786])).
fof(f5786,plain,(
  ( ! [X132,X133] : (s(bool,'F_0') = s(bool,X132) | s(bool,'F_0') = s(bool,X132) | s(bool,X133) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X133)) ) | $spl87),
  inference(superposition,[],[f401,f3965])).
fof(f6831,plain,(
  ( ! [X123,X122] : (s(bool,'F_0') = s(bool,X122) | s(bool,X123) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X123)) ) | $spl22),
  inference(duplicate_literal_removal,[],[f5781])).
fof(f5781,plain,(
  ( ! [X123,X122] : (s(bool,'F_0') = s(bool,X122) | s(bool,'F_0') = s(bool,X122) | s(bool,X123) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X123)) ) | $spl22),
  inference(superposition,[],[f401,f234])).
fof(f6835,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,X84) = s(bool,X88)) )),
  inference(duplicate_literal_removal,[],[f5770])).
fof(f5770,plain,(
  ( ! [X88,X87,X85,X86,X84] : (s(bool,X84) = s(bool,X85) | s(bool,'F_0') = s(bool,X85) | s(bool,'T_0') = s(bool,X86) | s(bool,'T_0') = s(bool,X86) | s(bool,X87) = s(bool,X88) | s(bool,'F_0') = s(bool,X87) | s(bool,X84) = s(bool,X88)) )),
  inference(superposition,[],[f401,f401])).
fof(f6836,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,X54) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X54) = s(bool,X57) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(duplicate_literal_removal,[],[f5764])).
fof(f5764,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'F_0') = s(bool,X55) | s(bool,'F_0') = s(bool,X55) | s(bool,X54) = s(bool,X56) | s(bool,'T_0') = s(bool,X56) | s(bool,X54) = s(bool,X57) | s(bool,X57) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f401,f401])).
fof(f6840,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X59) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(duplicate_literal_removal,[],[f5728])).
fof(f5728,plain,(
  ( ! [X61,X59,X62,X60,X63] : (s(bool,X60) = s(bool,X61) | s(bool,'F_0') = s(bool,X59) | s(bool,X61) = s(bool,X62) | s(bool,'T_0') = s(bool,X62) | s(bool,'F_0') = s(bool,X59) | s(bool,X60) = s(bool,X63) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f401,f401])).
fof(f1048,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X47)) ) | $spl44),
  inference(cnf_transformation,[],[f1048_D])).
fof(f1048_D,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X47)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5712,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,X32) = s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5676,f2401])).
fof(f5676,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X31) | s(bool,X30) = s(bool,X32) | s(bool,X32) = s(bool,X33) | s(bool,'F_0') = s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl76),
  inference(superposition,[],[f2502,f266])).
fof(f5711,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X26)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5675,f2401])).
fof(f5675,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,X27) | s(bool,X26) = s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,'F_0') = s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl76),
  inference(superposition,[],[f2502,f266])).
fof(f5641,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | ($spl76 | $spl112)),
  inference(subsumption_resolution,[],[f5616,f210])).
fof(f5616,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X32) != s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | ($spl76 | $spl112)),
  inference(superposition,[],[f3800,f266])).
fof(f5640,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X26) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | ($spl76 | $spl112)),
  inference(subsumption_resolution,[],[f5615,f210])).
fof(f5615,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') = s(bool,X28) | s(bool,X28) != s(bool,X29) | s(bool,X26) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | ($spl76 | $spl112)),
  inference(superposition,[],[f3800,f266])).
fof(f5577,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl112),
  inference(subsumption_resolution,[],[f5554,f357])).
fof(f5554,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,'F_0') = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,X48) != s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl112),
  inference(superposition,[],[f3711,f400])).
fof(f5576,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'F_0') = s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl112),
  inference(subsumption_resolution,[],[f5553,f357])).
fof(f5553,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,'F_0') = s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,X44) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl112),
  inference(superposition,[],[f3711,f400])).
fof(f5516,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X30)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5490,f1190])).
fof(f5490,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') != s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3523,f266])).
fof(f5514,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') != s(bool,X29) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') = s(bool,X26)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5489,f1190])).
fof(f5489,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') != s(bool,X29) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3523,f266])).
fof(f5454,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl104),
  inference(subsumption_resolution,[],[f5429,f210])).
fof(f5429,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) != s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl104),
  inference(superposition,[],[f3393,f266])).
fof(f5453,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'T_0') = s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl104),
  inference(subsumption_resolution,[],[f5428,f210])).
fof(f5428,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,'T_0') = s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,X28) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl104),
  inference(superposition,[],[f3393,f266])).
fof(f5391,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X48) != s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X46) != s(bool,X48) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5367,f357])).
fof(f5367,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,X48) != s(bool,X49) | s(bool,'T_0') = s(bool,X49) | s(bool,X46) != s(bool,X48) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3321,f400])).
fof(f5389,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X44) != s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5366,f357])).
fof(f5366,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,X44) != s(bool,X45) | s(bool,'T_0') = s(bool,X45) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3321,f400])).
fof(f5329,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X48) != s(bool,X49) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') != s(bool,X48) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5305,f357])).
fof(f5305,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,X48) != s(bool,X49) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') != s(bool,X48) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2565,f400])).
fof(f5327,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X44) != s(bool,X45) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f5304,f357])).
fof(f5304,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,X44) != s(bool,X45) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2565,f400])).
fof(f5267,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'T_0') != s(bool,X48) | s(bool,X48) = s(bool,X49) | s(bool,X46) != s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl43),
  inference(subsumption_resolution,[],[f5243,f357])).
fof(f5243,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,'T_0') != s(bool,X48) | s(bool,X48) = s(bool,X49) | s(bool,X46) != s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl43),
  inference(superposition,[],[f1811,f400])).
fof(f5265,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') != s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,X42) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl43),
  inference(subsumption_resolution,[],[f5242,f357])).
fof(f5242,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,'T_0') != s(bool,X44) | s(bool,X44) = s(bool,X45) | s(bool,X42) != s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl43),
  inference(superposition,[],[f1811,f400])).
fof(f5204,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X30)) ) | $spl43),
  inference(subsumption_resolution,[],[f5178,f1190])).
fof(f5178,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'T_0') != s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl43),
  inference(superposition,[],[f1622,f266])).
fof(f5202,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,'T_0') != s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X26)) ) | $spl43),
  inference(subsumption_resolution,[],[f5177,f1190])).
fof(f5177,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,'T_0') != s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl43),
  inference(superposition,[],[f1622,f266])).
fof(f5115,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,X46) != s(bool,X48) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl43),
  inference(subsumption_resolution,[],[f5091,f357])).
fof(f5091,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X49) | s(bool,X46) != s(bool,X48) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl43),
  inference(superposition,[],[f1733,f400])).
fof(f5113,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X45) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl43),
  inference(subsumption_resolution,[],[f5090,f357])).
fof(f5090,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X45) | s(bool,X42) != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl43),
  inference(superposition,[],[f1733,f400])).
fof(f5041,plain,(
  $false | ($spl24 | $spl34 | $spl43)),
  inference(trivial_inequality_removal,[],[f5025])).
fof(f5025,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl34 | $spl43)),
  inference(superposition,[],[f4990,f62])).
fof(f4990,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2)) ) | ($spl24 | $spl34 | $spl43)),
  inference(subsumption_resolution,[],[f4989,f1190])).
fof(f4989,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,X2))) ) | ($spl24 | $spl34)),
  inference(forward_demodulation,[],[f244,f235])).
fof(f5040,plain,(
  $false | ($spl24 | $spl34 | $spl43)),
  inference(trivial_inequality_removal,[],[f5026])).
fof(f5026,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl34 | $spl43)),
  inference(superposition,[],[f4990,f85])).
fof(f5039,plain,(
  $false | ($spl24 | $spl34 | $spl43)),
  inference(trivial_inequality_removal,[],[f5027])).
fof(f5027,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl34 | $spl43)),
  inference(superposition,[],[f4990,f235])).
fof(f5038,plain,(
  $false | ($spl2 | $spl24 | $spl34 | $spl43)),
  inference(trivial_inequality_removal,[],[f5029])).
fof(f5029,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl2 | $spl24 | $spl34 | $spl43)),
  inference(superposition,[],[f4990,f145])).
fof(f5037,plain,(
  $false | ($spl24 | $spl34 | $spl43)),
  inference(trivial_inequality_removal,[],[f5030])).
fof(f5030,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl34 | $spl43)),
  inference(superposition,[],[f4990,f107])).
fof(f5036,plain,(
  $false | ($spl24 | $spl34 | $spl43 | $spl84)),
  inference(trivial_inequality_removal,[],[f5031])).
fof(f5031,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl24 | $spl34 | $spl43 | $spl84)),
  inference(superposition,[],[f4990,f3925])).
fof(f5035,plain,(
  $false | ($spl8 | $spl24 | $spl34 | $spl43)),
  inference(trivial_inequality_removal,[],[f5032])).
fof(f5032,plain,(
  s(bool,'T_0') != s(bool,'T_0') | ($spl8 | $spl24 | $spl34 | $spl43)),
  inference(superposition,[],[f4990,f165])).
fof(f5034,plain,(
  $false | ($spl24 | $spl34 | $spl43)),
  inference(equality_resolution,[],[f4990])).
fof(f4979,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl76),
  inference(subsumption_resolution,[],[f4953,f210])).
fof(f4953,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X30) != s(bool,X32) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl76),
  inference(superposition,[],[f2705,f266])).
fof(f4977,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl76),
  inference(subsumption_resolution,[],[f4952,f210])).
fof(f4952,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X29) | s(bool,X26) != s(bool,X28) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl76),
  inference(superposition,[],[f2705,f266])).
fof(f4919,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') != s(bool,X32) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl76),
  inference(subsumption_resolution,[],[f4893,f210])).
fof(f4893,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') != s(bool,X32) | s(bool,X32) = s(bool,X33) | s(bool,X30) != s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) ) | $spl76),
  inference(superposition,[],[f2648,f266])).
fof(f4917,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') != s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,X26) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl76),
  inference(subsumption_resolution,[],[f4892,f210])).
fof(f4892,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') != s(bool,X28) | s(bool,X28) = s(bool,X29) | s(bool,X26) != s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) ) | $spl76),
  inference(superposition,[],[f2648,f266])).
fof(f4859,plain,(
  ( ! [X47,X48,X46,X49] : (s(bool,'T_0') = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f4836,f357])).
fof(f4836,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,'T_0') = s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) )),
  inference(superposition,[],[f267,f400])).
fof(f4858,plain,(
  ( ! [X45,X43,X44,X42] : (s(bool,'T_0') = s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) )),
  inference(subsumption_resolution,[],[f4835,f357])).
fof(f4835,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,'T_0') = s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) )),
  inference(superposition,[],[f267,f400])).
fof(f4801,plain,(
  ( ! [X30,X33,X31,X32] : (s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f4776,f210])).
fof(f4776,plain,(
  ( ! [X30,X33,X31,X32] : (p(s(bool,X31)) | s(bool,'F_0') = s(bool,X32) | s(bool,X30) = s(bool,X33) | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f402,f266])).
fof(f4800,plain,(
  ( ! [X28,X26,X29,X27] : (s(bool,'F_0') = s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f4775,f210])).
fof(f4775,plain,(
  ( ! [X28,X26,X29,X27] : (p(s(bool,X27)) | s(bool,'F_0') = s(bool,X28) | s(bool,X26) = s(bool,X29) | s(bool,X28) = s(bool,X29) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f402,f266])).
fof(f3990,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0))))) ) | $spl120),
  inference(cnf_transformation,[],[f3990_D])).
fof(f3990_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f4726,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl87),
  inference(subsumption_resolution,[],[f4661,f42])).
fof(f4661,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl87),
  inference(superposition,[],[f40,f3965])).
fof(f3951,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7))))) ) | $spl116),
  inference(cnf_transformation,[],[f3951_D])).
fof(f3951_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4611,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl84),
  inference(subsumption_resolution,[],[f4546,f43])).
fof(f4546,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl84),
  inference(superposition,[],[f40,f3925])).
fof(f3950,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0))))) ) | $spl114),
  inference(cnf_transformation,[],[f3950_D])).
fof(f3950_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4610,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl84),
  inference(subsumption_resolution,[],[f4545,f43])).
fof(f4545,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl84),
  inference(superposition,[],[f40,f3925])).
fof(f4516,plain,(
  ( ! [X45,X44,X42] : (s(bool,'F_0') != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X44) = s(bool,X45) | s(bool,X42) = s(bool,X45)) ) | $spl124),
  inference(cnf_transformation,[],[f4516_D])).
fof(f4516_D,plain,(
  ( ! [X45,X44,X42] : (s(bool,'F_0') != s(bool,X44) | s(bool,'F_0') = s(bool,X42) | s(bool,X44) = s(bool,X45) | s(bool,X42) = s(bool,X45)) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f4518,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,'F_0') != s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X46)) ) | $spl76),
  inference(subsumption_resolution,[],[f4495,f2143])).
fof(f4495,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | s(bool,'F_0') != s(bool,X48) | s(bool,X46) = s(bool,X49) | s(bool,X48) = s(bool,X49) | s(bool,'F_0') = s(bool,X47) | s(bool,'F_0') = s(bool,X46)) ) | $spl76),
  inference(superposition,[],[f2504,f400])).
fof(f4515,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,'F_0') != s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X42)) ) | $spl76),
  inference(subsumption_resolution,[],[f4494,f2143])).
fof(f4494,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | s(bool,'F_0') != s(bool,X44) | s(bool,X42) = s(bool,X45) | s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X43)) ) | $spl76),
  inference(superposition,[],[f2504,f400])).
fof(f164,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK8))))) ) | $spl20),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK8))))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4435,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl8),
  inference(subsumption_resolution,[],[f4371,f43])).
fof(f4371,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl8),
  inference(superposition,[],[f40,f165])).
fof(f163,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X3))))) ) | $spl18),
  inference(cnf_transformation,[],[f163_D])).
fof(f163_D,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X3))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4434,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl8),
  inference(subsumption_resolution,[],[f4370,f43])).
fof(f4370,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl8),
  inference(superposition,[],[f40,f165])).
fof(f144,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK6))))) ) | $spl16),
  inference(cnf_transformation,[],[f144_D])).
fof(f144_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK6))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4325,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl2),
  inference(subsumption_resolution,[],[f4261,f43])).
fof(f4261,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl2),
  inference(superposition,[],[f40,f145])).
fof(f143,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X3))))) ) | $spl14),
  inference(cnf_transformation,[],[f143_D])).
fof(f143_D,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X3))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4324,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl2),
  inference(subsumption_resolution,[],[f4260,f43])).
fof(f4260,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl2),
  inference(superposition,[],[f40,f145])).
fof(f3471,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0))))) ) | $spl110),
  inference(cnf_transformation,[],[f3471_D])).
fof(f3471_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4191,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl22),
  inference(subsumption_resolution,[],[f4127,f42])).
fof(f4127,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl22),
  inference(superposition,[],[f40,f234])).
fof(f4073,plain,(
  $false | ($spl22 | $spl26 | $spl43)),
  inference(subsumption_resolution,[],[f3999,f696])).
fof(f3999,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl22 | $spl26)),
  inference(backward_demodulation,[],[f238,f234])).
fof(f3997,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | $spl122),
  inference(cnf_transformation,[],[f3997_D])).
fof(f3997_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f3998,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) ) | $spl87),
  inference(subsumption_resolution,[],[f3986,f357])).
fof(f3986,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) ) | $spl87),
  inference(superposition,[],[f2911,f400])).
fof(f3995,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X9)) ) | $spl87),
  inference(subsumption_resolution,[],[f3985,f357])).
fof(f3985,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6))) | s(bool,'F_0') = s(bool,X9)) ) | $spl87),
  inference(superposition,[],[f2911,f400])).
fof(f3960,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl87),
  inference(resolution,[],[f2911,f40])).
fof(f3956,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | $spl118),
  inference(cnf_transformation,[],[f3956_D])).
fof(f3956_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f3957,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) ) | $spl84),
  inference(subsumption_resolution,[],[f3942,f210])).
fof(f3942,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) ) | $spl84),
  inference(superposition,[],[f2909,f266])).
fof(f3954,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl84),
  inference(subsumption_resolution,[],[f3941,f210])).
fof(f3941,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl84),
  inference(superposition,[],[f2909,f266])).
fof(f3921,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl84),
  inference(resolution,[],[f2909,f40])).
fof(f3920,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl84),
  inference(resolution,[],[f2909,f40])).
fof(f2909,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) | $spl84),
  inference(cnf_transformation,[],[f2909_D])).
fof(f2909_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3893,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X12),X13),s(fun(X12,sK1),X14))))) ) | $spl24),
  inference(subsumption_resolution,[],[f3830,f43])).
fof(f3830,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X12),X13),s(fun(X12,sK1),X14))))) ) | $spl24),
  inference(superposition,[],[f41,f235])).
fof(f3224,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5))))) ) | $spl88),
  inference(cnf_transformation,[],[f3224_D])).
fof(f3224_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3892,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl24),
  inference(subsumption_resolution,[],[f3828,f43])).
fof(f3828,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl24),
  inference(superposition,[],[f40,f235])).
fof(f3226,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0))))) ) | $spl92),
  inference(cnf_transformation,[],[f3226_D])).
fof(f3226_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f3891,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl24),
  inference(subsumption_resolution,[],[f3827,f43])).
fof(f3827,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl24),
  inference(superposition,[],[f40,f235])).
fof(f2911,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) | $spl87),
  inference(cnf_transformation,[],[f2911_D])).
fof(f2911_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f3686,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X15),X16),s(fun(X15,sK1),X17))))) ) | $spl22),
  inference(subsumption_resolution,[],[f3626,f42])).
fof(f3626,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X15),X16),s(fun(X15,sK1),X17))))) ) | $spl22),
  inference(superposition,[],[f41,f234])).
fof(f3685,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl22),
  inference(subsumption_resolution,[],[f3624,f42])).
fof(f3624,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl22),
  inference(superposition,[],[f40,f234])).
fof(f3596,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) | $spl112),
  inference(cnf_transformation,[],[f3596_D])).
fof(f3596_D,plain,(
  ( ! [X2,X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,X2) | s(bool,'F_0') = s(bool,X2)) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f3597,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X4))) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3557,f210])).
fof(f3557,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) != s(bool,X5) | s(bool,'T_0') = s(bool,X3) | s(bool,'F_0') = s(bool,X5) | p(s(bool,X3)) | p(s(bool,X4))) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3318,f35])).
fof(f3595,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X1))) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3556,f210])).
fof(f3556,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) != s(bool,X2) | s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X2) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f3318,f35])).
fof(f3476,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl31),
  inference(subsumption_resolution,[],[f3467,f357])).
fof(f3467,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl31),
  inference(superposition,[],[f241,f400])).
fof(f3475,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) ) | $spl31),
  inference(subsumption_resolution,[],[f3466,f357])).
fof(f3466,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) ) | $spl31),
  inference(superposition,[],[f241,f400])).
fof(f3443,plain,(
  ( ! [X6,X4,X5] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X4),X5),s(fun(X4,sK1),X6))))) ) | $spl31),
  inference(resolution,[],[f241,f41])).
fof(f3442,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl31),
  inference(resolution,[],[f241,f40])).
fof(f3330,plain,(
  ( ! [X33,X31] : (s(bool,X31) != s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X33)) ) | $spl108),
  inference(cnf_transformation,[],[f3330_D])).
fof(f3330_D,plain,(
  ( ! [X33,X31] : (s(bool,X31) != s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') = s(bool,X33)) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f3435,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | $spl104),
  inference(subsumption_resolution,[],[f3418,f357])).
fof(f3418,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | $spl104),
  inference(superposition,[],[f3326,f400])).
fof(f3434,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | $spl104),
  inference(subsumption_resolution,[],[f3417,f357])).
fof(f3417,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | $spl104),
  inference(superposition,[],[f3326,f400])).
fof(f3374,plain,(
  $false | $spl106),
  inference(trivial_inequality_removal,[],[f3368])).
fof(f3368,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl106),
  inference(superposition,[],[f3327,f62])).
fof(f3373,plain,(
  $false | $spl106),
  inference(trivial_inequality_removal,[],[f3369])).
fof(f3369,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl106),
  inference(superposition,[],[f3327,f85])).
fof(f3372,plain,(
  $false | $spl106),
  inference(trivial_inequality_removal,[],[f3370])).
fof(f3370,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl106),
  inference(superposition,[],[f3327,f107])).
fof(f3371,plain,(
  $false | $spl106),
  inference(equality_resolution,[],[f3327])).
fof(f3331,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3301,f39])).
fof(f3301,plain,(
  ( ! [X35,X36,X34] : (s(bool,'T_0') != s(bool,X35) | s(bool,X34) != s(bool,X36) | s(bool,'T_0') = s(bool,X36) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f400])).
fof(f3329,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) != s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3300,f39])).
fof(f3300,plain,(
  ( ! [X33,X31,X32] : (s(bool,'T_0') != s(bool,X32) | s(bool,X31) != s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f400])).
fof(f3326,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) | $spl104),
  inference(cnf_transformation,[],[f3326_D])).
fof(f3326_D,plain,(
  ( ! [X6,X8] : (~p(s(bool,X6)) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X8)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f3327,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) | $spl106),
  inference(cnf_transformation,[],[f3327_D])).
fof(f3327_D,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,X7)) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f3328,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,X9) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X9))) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3292,f1190])).
fof(f3292,plain,(
  ( ! [X10,X11,X9] : (s(bool,'T_0') != s(bool,X10) | s(bool,X9) != s(bool,X11) | s(bool,'T_0') = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,X9))) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f38])).
fof(f3325,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6))) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f3291,f1190])).
fof(f3291,plain,(
  ( ! [X6,X8,X7] : (s(bool,'T_0') != s(bool,X7) | s(bool,X6) != s(bool,X8) | s(bool,'T_0') = s(bool,X8) | ~p(s(bool,X6)) | ~p(s(bool,X7))) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2449,f38])).
fof(f242,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | $spl30),
  inference(cnf_transformation,[],[f242_D])).
fof(f242_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f243,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,X0))) ) | $spl32),
  inference(cnf_transformation,[],[f243_D])).
fof(f243_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,X0))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2355,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f35])).
fof(f2356,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f33,f35])).
fof(f241,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | $spl31),
  inference(cnf_transformation,[],[f241_D])).
fof(f241_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f244,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,X2))) ) | $spl34),
  inference(cnf_transformation,[],[f244_D])).
fof(f244_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,X2))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2357,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f33,f38])).
fof(f2358,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,X3)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f38])).
fof(f3274,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X117,X115,X113] : (p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),X110),s(X109,X115))),s(X109,X116)))) | ~p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),X110),s(X109,X117))),s(X109,X116)))) | ~p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),X110),s(X109,X115))),s(X109,X117)))) | s(bool,'F_0') = s(bool,X114) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X109,fun(X109,bool)),X110),s(fun(X109,X111),X112),s(fun(X111,X109),X113)))) )),
  inference(subsumption_resolution,[],[f3214,f357])).
fof(f3214,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X117,X115,X113] : (~p(s(bool,X114)) | p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),X110),s(X109,X115))),s(X109,X116)))) | ~p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),X110),s(X109,X117))),s(X109,X116)))) | ~p(s(bool,happ(s(fun(X109,bool),happ(s(fun(X109,fun(X109,bool)),X110),s(X109,X115))),s(X109,X117)))) | s(bool,'F_0') = s(bool,X114) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X109,fun(X109,bool)),X110),s(fun(X109,X111),X112),s(fun(X111,X109),X113)))) )),
  inference(superposition,[],[f40,f400])).
fof(f3272,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X108,X106,X104] : (p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X107)))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X108))),s(X100,X107)))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X108)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X100,fun(X100,bool)),X101),s(fun(X100,X102),X103),s(fun(X102,X100),X104))) | s(bool,'F_0') = s(bool,X105)) )),
  inference(subsumption_resolution,[],[f3213,f357])).
fof(f3213,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X108,X106,X104] : (~p(s(bool,X105)) | p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X107)))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X108))),s(X100,X107)))) | ~p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X106))),s(X100,X108)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X100,fun(X100,bool)),X101),s(fun(X100,X102),X103),s(fun(X102,X100),X104))) | s(bool,'F_0') = s(bool,X105)) )),
  inference(superposition,[],[f40,f400])).
fof(f3267,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115] : (p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X116))),s(X114,X119)))) | ~p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X117))),s(X114,X119)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X114,fun(X114,bool)),X115),s(fun(X114,X120),X121),s(fun(X120,X114),X122)))) | s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X116))),s(X114,X117)))) )),
  inference(subsumption_resolution,[],[f3193,f357])).
fof(f3193,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115] : (~p(s(bool,X118)) | p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X116))),s(X114,X119)))) | ~p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X117))),s(X114,X119)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X114,fun(X114,bool)),X115),s(fun(X114,X120),X121),s(fun(X120,X114),X122)))) | s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X116))),s(X114,X117)))) )),
  inference(superposition,[],[f40,f400])).
fof(f3265,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X113] : (p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X110)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X110)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | s(bool,'F_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f3192,f357])).
fof(f3192,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X113] : (~p(s(bool,X109)) | p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X110)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X108))),s(X105,X110)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | s(bool,'F_0') = s(bool,X109)) )),
  inference(superposition,[],[f40,f400])).
fof(f105,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X3))))) ) | $spl10),
  inference(cnf_transformation,[],[f105_D])).
fof(f105_D,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X3))))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f3261,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X4)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X4)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) )),
  inference(subsumption_resolution,[],[f3180,f43])).
fof(f3180,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X4)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X4)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) )),
  inference(superposition,[],[f40,f107])).
fof(f83,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X3))))) ) | $spl4),
  inference(cnf_transformation,[],[f83_D])).
fof(f83_D,plain,(
  ( ! [X3] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X3)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X3))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f3260,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) )),
  inference(subsumption_resolution,[],[f3179,f43])).
fof(f3179,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) )),
  inference(superposition,[],[f40,f85])).
fof(f3253,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115] : (p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X119))),s(X114,X117)))) | ~p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X119))),s(X114,X116)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X114,fun(X114,bool)),X115),s(fun(X114,X120),X121),s(fun(X120,X114),X122)))) | s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X116))),s(X114,X117)))) )),
  inference(subsumption_resolution,[],[f3123,f357])).
fof(f3123,plain,(
  ( ! [X121,X118,X116,X114,X122,X120,X119,X117,X115] : (~p(s(bool,X118)) | p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X119))),s(X114,X117)))) | ~p(s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X119))),s(X114,X116)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X114,fun(X114,bool)),X115),s(fun(X114,X120),X121),s(fun(X120,X114),X122)))) | s(bool,'F_0') = s(bool,X118) | s(bool,'F_0') = s(bool,happ(s(fun(X114,bool),happ(s(fun(X114,fun(X114,bool)),X115),s(X114,X116))),s(X114,X117)))) )),
  inference(superposition,[],[f40,f400])).
fof(f3251,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X113] : (p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X108)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X107)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | s(bool,'F_0') = s(bool,X109)) )),
  inference(subsumption_resolution,[],[f3122,f357])).
fof(f3122,plain,(
  ( ! [X111,X109,X107,X105,X112,X110,X108,X106,X113] : (~p(s(bool,X109)) | p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X108)))) | ~p(s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X110))),s(X105,X107)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X105,fun(X105,bool)),X106),s(fun(X105,X111),X112),s(fun(X111,X105),X113)))) | s(bool,'F_0') = s(bool,happ(s(fun(X105,bool),happ(s(fun(X105,fun(X105,bool)),X106),s(X105,X107))),s(X105,X108))) | s(bool,'F_0') = s(bool,X109)) )),
  inference(superposition,[],[f40,f400])).
fof(f106,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK8))))) ) | $spl12),
  inference(cnf_transformation,[],[f106_D])).
fof(f106_D,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK7)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK8))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3247,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) )),
  inference(subsumption_resolution,[],[f3110,f43])).
fof(f3110,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) )),
  inference(superposition,[],[f40,f107])).
fof(f84,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK6))))) ) | $spl6),
  inference(cnf_transformation,[],[f84_D])).
fof(f84_D,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK5)))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK6))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3246,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) )),
  inference(subsumption_resolution,[],[f3109,f43])).
fof(f3109,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) )),
  inference(superposition,[],[f40,f85])).
fof(f3237,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X86,X84,X82] : (~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X83))),s(X78,X81)))) | ~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X83)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X78,fun(X78,bool)),X79),s(fun(X78,X84),X85),s(fun(X84,X78),X86)))) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X81)))) )),
  inference(subsumption_resolution,[],[f3049,f210])).
fof(f3049,plain,(
  ( ! [X80,X78,X85,X83,X81,X79,X86,X84,X82] : (p(s(bool,X82)) | ~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X83))),s(X78,X81)))) | ~p(s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X83)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X78,fun(X78,bool)),X79),s(fun(X78,X84),X85),s(fun(X84,X78),X86)))) | s(bool,'T_0') = s(bool,X82) | s(bool,'T_0') = s(bool,happ(s(fun(X78,bool),happ(s(fun(X78,fun(X78,bool)),X79),s(X78,X80))),s(X78,X81)))) )),
  inference(superposition,[],[f40,f266])).
fof(f3235,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X74))),s(X69,X72)))) | ~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X74)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X69,fun(X69,bool)),X70),s(fun(X69,X75),X76),s(fun(X75,X69),X77)))) | s(bool,'T_0') = s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X72))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f3048,f210])).
fof(f3048,plain,(
  ( ! [X70,X76,X74,X72,X71,X69,X77,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X74))),s(X69,X72)))) | ~p(s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X74)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X69,fun(X69,bool)),X70),s(fun(X69,X75),X76),s(fun(X75,X69),X77)))) | s(bool,'T_0') = s(bool,happ(s(fun(X69,bool),happ(s(fun(X69,fun(X69,bool)),X70),s(X69,X71))),s(X69,X72))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f40,f266])).
fof(f2967,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X12)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X12)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X13),X14),s(fun(X13,sK1),X15))))) ) | $spl8),
  inference(resolution,[],[f40,f104])).
fof(f2966,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl30),
  inference(resolution,[],[f40,f242])).
fof(f2965,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X4)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X4)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl2),
  inference(resolution,[],[f40,f82])).
fof(f2964,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X0)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X0)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl72),
  inference(resolution,[],[f40,f2140])).
fof(f2954,plain,(
  ( ! [X14,X12,X15,X13] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X12))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X12))),s(sK1,sK8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X13),X14),s(fun(X13,sK1),X15))))) ) | $spl8),
  inference(resolution,[],[f40,f104])).
fof(f2953,plain,(
  ( ! [X10,X8,X11,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X8))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X9),X10),s(fun(X9,sK1),X11))))) ) | $spl30),
  inference(resolution,[],[f40,f242])).
fof(f2952,plain,(
  ( ! [X6,X4,X7,X5] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X4))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X5),X6),s(fun(X5,sK1),X7))))) ) | $spl2),
  inference(resolution,[],[f40,f82])).
fof(f2951,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X0))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X1),X2),s(fun(X1,sK1),X3))))) ) | $spl72),
  inference(resolution,[],[f40,f2140])).
fof(f2941,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X102))),s(X96,X103)))) | ~p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X103))),s(X96,X102)))) | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X96,fun(X96,bool)),X97),s(fun(X96,X98),X99),s(fun(X98,X96),X100)))) )),
  inference(subsumption_resolution,[],[f2900,f357])).
fof(f2900,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,X101)) | p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X102))),s(X96,X103)))) | ~p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X103))),s(X96,X102)))) | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X96,fun(X96,bool)),X97),s(fun(X96,X98),X99),s(fun(X98,X96),X100)))) )),
  inference(superposition,[],[f41,f400])).
fof(f2940,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X94))),s(X88,X95)))) | ~p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X95))),s(X88,X94)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X88,fun(X88,bool)),X89),s(fun(X88,X90),X91),s(fun(X90,X88),X92))) | s(bool,'F_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f2899,f357])).
fof(f2899,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (~p(s(bool,X93)) | p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X94))),s(X88,X95)))) | ~p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X95))),s(X88,X94)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X88,fun(X88,bool)),X89),s(fun(X88,X90),X91),s(fun(X90,X88),X92))) | s(bool,'F_0') = s(bool,X93)) )),
  inference(superposition,[],[f41,f400])).
fof(f2935,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X106,X104] : (p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X103))),s(X100,X102)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X100,fun(X100,bool)),X101),s(fun(X100,X105),X106),s(fun(X105,X100),X107)))) | s(bool,'F_0') = s(bool,X104) | s(bool,'F_0') = s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)))) )),
  inference(subsumption_resolution,[],[f2879,f357])).
fof(f2879,plain,(
  ( ! [X103,X101,X107,X105,X102,X100,X106,X104] : (~p(s(bool,X104)) | p(s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X103))),s(X100,X102)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X100,fun(X100,bool)),X101),s(fun(X100,X105),X106),s(fun(X105,X100),X107)))) | s(bool,'F_0') = s(bool,X104) | s(bool,'F_0') = s(bool,happ(s(fun(X100,bool),happ(s(fun(X100,fun(X100,bool)),X101),s(X100,X102))),s(X100,X103)))) )),
  inference(superposition,[],[f41,f400])).
fof(f2934,plain,(
  ( ! [X94,X92,X99,X97,X95,X93,X98,X96] : (p(s(bool,happ(s(fun(X92,bool),happ(s(fun(X92,fun(X92,bool)),X93),s(X92,X95))),s(X92,X94)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X92,fun(X92,bool)),X93),s(fun(X92,X97),X98),s(fun(X97,X92),X99)))) | s(bool,'F_0') = s(bool,happ(s(fun(X92,bool),happ(s(fun(X92,fun(X92,bool)),X93),s(X92,X94))),s(X92,X95))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f2878,f357])).
fof(f2878,plain,(
  ( ! [X94,X92,X99,X97,X95,X93,X98,X96] : (~p(s(bool,X96)) | p(s(bool,happ(s(fun(X92,bool),happ(s(fun(X92,fun(X92,bool)),X93),s(X92,X95))),s(X92,X94)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X92,fun(X92,bool)),X93),s(fun(X92,X97),X98),s(fun(X97,X92),X99)))) | s(bool,'F_0') = s(bool,happ(s(fun(X92,bool),happ(s(fun(X92,fun(X92,bool)),X93),s(X92,X94))),s(X92,X95))) | s(bool,'F_0') = s(bool,X96)) )),
  inference(superposition,[],[f41,f400])).
fof(f2922,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (~p(s(bool,happ(s(fun(X68,bool),happ(s(fun(X68,fun(X68,bool)),X69),s(X68,X71))),s(X68,X70)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X68,fun(X68,bool)),X69),s(fun(X68,X73),X74),s(fun(X73,X68),X75)))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),happ(s(fun(X68,fun(X68,bool)),X69),s(X68,X70))),s(X68,X71)))) )),
  inference(subsumption_resolution,[],[f2805,f210])).
fof(f2805,plain,(
  ( ! [X70,X68,X74,X72,X71,X69,X75,X73] : (p(s(bool,X72)) | ~p(s(bool,happ(s(fun(X68,bool),happ(s(fun(X68,fun(X68,bool)),X69),s(X68,X71))),s(X68,X70)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X68,fun(X68,bool)),X69),s(fun(X68,X73),X74),s(fun(X73,X68),X75)))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,happ(s(fun(X68,bool),happ(s(fun(X68,fun(X68,bool)),X69),s(X68,X70))),s(X68,X71)))) )),
  inference(superposition,[],[f41,f266])).
fof(f2921,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (~p(s(bool,happ(s(fun(X60,bool),happ(s(fun(X60,fun(X60,bool)),X61),s(X60,X63))),s(X60,X62)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X60,fun(X60,bool)),X61),s(fun(X60,X65),X66),s(fun(X65,X60),X67)))) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),happ(s(fun(X60,fun(X60,bool)),X61),s(X60,X62))),s(X60,X63))) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f2804,f210])).
fof(f2804,plain,(
  ( ! [X61,X66,X64,X62,X60,X67,X65,X63] : (p(s(bool,X64)) | ~p(s(bool,happ(s(fun(X60,bool),happ(s(fun(X60,fun(X60,bool)),X61),s(X60,X63))),s(X60,X62)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X60,fun(X60,bool)),X61),s(fun(X60,X65),X66),s(fun(X65,X60),X67)))) | s(bool,'T_0') = s(bool,happ(s(fun(X60,bool),happ(s(fun(X60,fun(X60,bool)),X61),s(X60,X62))),s(X60,X63))) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f41,f266])).
fof(f2724,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X6),X7),s(fun(X6,sK1),X8))))) ) | $spl30),
  inference(resolution,[],[f41,f242])).
fof(f2722,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X0),X1),s(fun(X0,sK1),X2))))) ) | $spl72),
  inference(resolution,[],[f41,f2140])).
fof(f2315,plain,(
  ( ! [X21,X19] : (s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') != s(bool,X21)) ) | $spl78),
  inference(cnf_transformation,[],[f2315_D])).
fof(f2315_D,plain,(
  ( ! [X21,X19] : (s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'F_0') != s(bool,X21)) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2658,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X23) | s(bool,'F_0') != s(bool,X24) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X22)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f2630,f2401])).
fof(f2630,plain,(
  ( ! [X24,X23,X22] : (s(bool,'F_0') != s(bool,X23) | s(bool,'F_0') != s(bool,X24) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl76),
  inference(superposition,[],[f2500,f266])).
fof(f2656,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'F_0') != s(bool,X21) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X19)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f2629,f2401])).
fof(f2629,plain,(
  ( ! [X21,X19,X20] : (s(bool,'F_0') != s(bool,X20) | s(bool,'F_0') != s(bool,X21) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl76),
  inference(superposition,[],[f2500,f266])).
fof(f2598,plain,(
  ( ! [X21,X19] : (s(bool,'T_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,X19) != s(bool,X21)) ) | $spl80),
  inference(cnf_transformation,[],[f2598_D])).
fof(f2598_D,plain,(
  ( ! [X21,X19] : (s(bool,'T_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,X19) != s(bool,X21)) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2599,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X23)) | s(bool,'T_0') != s(bool,X24) | s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X22)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f2576,f1190])).
fof(f2576,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X23)) | s(bool,'T_0') != s(bool,X24) | s(bool,X22) != s(bool,X24) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2494,f266])).
fof(f2597,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19)) ) | ($spl43 | $spl76)),
  inference(subsumption_resolution,[],[f2575,f1190])).
fof(f2575,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | s(bool,'T_0') != s(bool,X21) | s(bool,X19) != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | ($spl43 | $spl76)),
  inference(superposition,[],[f2494,f266])).
fof(f235,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | $spl24),
  inference(cnf_transformation,[],[f235_D])).
fof(f235_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2560,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl22),
  inference(subsumption_resolution,[],[f2550,f39])).
fof(f2550,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl22),
  inference(superposition,[],[f2056,f266])).
fof(f2056,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | $spl22),
  inference(backward_demodulation,[],[f234,f33])).
fof(f2559,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(subsumption_resolution,[],[f2549,f39])).
fof(f2549,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(superposition,[],[f2056,f266])).
fof(f2534,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,X24) | s(bool,'F_0') != s(bool,X24) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl76),
  inference(subsumption_resolution,[],[f2515,f210])).
fof(f2515,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X23)) | s(bool,X22) = s(bool,X24) | s(bool,'F_0') != s(bool,X24) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl76),
  inference(superposition,[],[f2404,f266])).
fof(f2533,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) = s(bool,X21) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl76),
  inference(subsumption_resolution,[],[f2514,f210])).
fof(f2514,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl76),
  inference(superposition,[],[f2404,f266])).
fof(f2398,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl72),
  inference(subsumption_resolution,[],[f2388,f210])).
fof(f2388,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl72),
  inference(superposition,[],[f2140,f266])).
fof(f2397,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl72),
  inference(subsumption_resolution,[],[f2387,f210])).
fof(f2387,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl72),
  inference(superposition,[],[f2140,f266])).
fof(f2368,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X11) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f33,f400])).
fof(f2367,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f400])).
fof(f2366,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f33,f400])).
fof(f2365,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X8) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f33,f400])).
fof(f2364,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X7) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f33,f266])).
fof(f2363,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f266])).
fof(f2362,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f33,f266])).
fof(f2361,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X4) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f33,f266])).
fof(f2360,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f2359,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f2348,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl30),
  inference(subsumption_resolution,[],[f2338,f210])).
fof(f2338,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl30),
  inference(superposition,[],[f242,f266])).
fof(f2347,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | $spl30),
  inference(subsumption_resolution,[],[f2337,f210])).
fof(f2337,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) ) | $spl30),
  inference(superposition,[],[f242,f266])).
fof(f2316,plain,(
  ( ! [X24,X23,X22] : (s(bool,X22) = s(bool,X24) | s(bool,'F_0') != s(bool,X24) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl76),
  inference(subsumption_resolution,[],[f2295,f210])).
fof(f2295,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X23)) | s(bool,X22) = s(bool,X24) | s(bool,'F_0') != s(bool,X24) | s(bool,'T_0') = s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl76),
  inference(superposition,[],[f2183,f266])).
fof(f2183,plain,(
  ( ! [X4,X5] : (p(s(bool,X5)) | s(bool,X4) = s(bool,X5) | s(bool,'F_0') != s(bool,X4)) ) | $spl76),
  inference(resolution,[],[f2143,f35])).
fof(f2314,plain,(
  ( ! [X21,X19,X20] : (s(bool,X19) = s(bool,X21) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl76),
  inference(subsumption_resolution,[],[f2294,f210])).
fof(f2294,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X20)) | s(bool,X19) = s(bool,X21) | s(bool,'F_0') != s(bool,X21) | s(bool,'T_0') = s(bool,X19) | s(bool,'T_0') = s(bool,X20)) ) | $spl76),
  inference(superposition,[],[f2183,f266])).
fof(f2151,plain,(
  $false | $spl74),
  inference(resolution,[],[f2141,f42])).
fof(f2145,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl22),
  inference(subsumption_resolution,[],[f2129,f39])).
fof(f2129,plain,(
  ( ! [X6] : (s(bool,'F_0') != s(bool,X6) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) ) | $spl22),
  inference(superposition,[],[f2056,f266])).
fof(f2144,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(subsumption_resolution,[],[f2128,f39])).
fof(f2128,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl22),
  inference(superposition,[],[f2056,f266])).
fof(f2139,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | $spl73),
  inference(cnf_transformation,[],[f2139_D])).
fof(f2139_D,plain,(
  ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f2143,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') != s(bool,X2)) ) | $spl76),
  inference(cnf_transformation,[],[f2143_D])).
fof(f2143_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2124,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))))) ) | $spl22),
  inference(superposition,[],[f2056,f38])).
fof(f2123,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | ~p(s(bool,X2))) ) | $spl22),
  inference(superposition,[],[f2056,f38])).
fof(f2142,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | p(s(bool,X1))) ) | $spl22),
  inference(subsumption_resolution,[],[f2122,f357])).
fof(f2122,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7)))) | p(s(bool,X1))) ) | $spl22),
  inference(superposition,[],[f2056,f35])).
fof(f2138,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))))) ) | $spl22),
  inference(subsumption_resolution,[],[f2121,f357])).
fof(f2121,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))))) ) | $spl22),
  inference(superposition,[],[f2056,f35])).
fof(f1928,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl31),
  inference(subsumption_resolution,[],[f1920,f357])).
fof(f1920,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) ) | $spl31),
  inference(superposition,[],[f241,f400])).
fof(f1927,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) ) | $spl31),
  inference(subsumption_resolution,[],[f1919,f357])).
fof(f1919,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) ) | $spl31),
  inference(superposition,[],[f241,f400])).
fof(f1895,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | $spl70),
  inference(cnf_transformation,[],[f1895_D])).
fof(f1895_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1896,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7)))) ) | $spl8),
  inference(subsumption_resolution,[],[f1884,f210])).
fof(f1884,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7)))) ) | $spl8),
  inference(superposition,[],[f104,f266])).
fof(f1893,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(subsumption_resolution,[],[f1883,f210])).
fof(f1883,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X5)) ) | $spl8),
  inference(superposition,[],[f104,f266])).
fof(f1866,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X11) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f33,f400])).
fof(f1865,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f400])).
fof(f1864,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f33,f400])).
fof(f1863,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X8) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f33,f400])).
fof(f1862,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X7) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f33,f266])).
fof(f1861,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f266])).
fof(f1860,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f33,f266])).
fof(f1859,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X4) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f33,f266])).
fof(f1858,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f1857,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f1856,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,X3)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f38])).
fof(f1855,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f33,f38])).
fof(f1854,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f33,f35])).
fof(f1853,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f35])).
fof(f1845,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | $spl68),
  inference(cnf_transformation,[],[f1845_D])).
fof(f1845_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1846,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5)))) ) | $spl2),
  inference(subsumption_resolution,[],[f1834,f210])).
fof(f1834,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5)))) ) | $spl2),
  inference(superposition,[],[f82,f266])).
fof(f1843,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl2),
  inference(subsumption_resolution,[],[f1833,f210])).
fof(f1833,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5))) | s(bool,'T_0') = s(bool,X5)) ) | $spl2),
  inference(superposition,[],[f82,f266])).
fof(f1066,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | $spl48),
  inference(cnf_transformation,[],[f1066_D])).
fof(f1066_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f1766,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8)))) )),
  inference(subsumption_resolution,[],[f1756,f210])).
fof(f1756,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8)))) )),
  inference(superposition,[],[f32,f266])).
fof(f1765,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1755,f210])).
fof(f1755,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f32,f266])).
fof(f1063,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | $spl46),
  inference(cnf_transformation,[],[f1063_D])).
fof(f1063_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f1682,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6)))) )),
  inference(subsumption_resolution,[],[f1672,f210])).
fof(f1672,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6)))) )),
  inference(superposition,[],[f31,f266])).
fof(f1681,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f1671,f210])).
fof(f1671,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f31,f266])).
fof(f1652,plain,(
  ( ! [X33,X31] : (s(bool,X31) = s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') != s(bool,X33)) ) | $spl66),
  inference(cnf_transformation,[],[f1652_D])).
fof(f1652_D,plain,(
  ( ! [X33,X31] : (s(bool,X31) = s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'T_0') != s(bool,X33)) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f1653,plain,(
  ( ! [X35,X36,X34] : (s(bool,X34) = s(bool,X36) | s(bool,'T_0') != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | $spl43),
  inference(subsumption_resolution,[],[f1637,f357])).
fof(f1637,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,X35)) | s(bool,X34) = s(bool,X36) | s(bool,'T_0') != s(bool,X36) | s(bool,'F_0') = s(bool,X35) | s(bool,'F_0') = s(bool,X34)) ) | $spl43),
  inference(superposition,[],[f1223,f400])).
fof(f1651,plain,(
  ( ! [X33,X31,X32] : (s(bool,X31) = s(bool,X33) | s(bool,'T_0') != s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | $spl43),
  inference(subsumption_resolution,[],[f1636,f357])).
fof(f1636,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,X32)) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') != s(bool,X33) | s(bool,'F_0') = s(bool,X31) | s(bool,'F_0') = s(bool,X32)) ) | $spl43),
  inference(superposition,[],[f1223,f400])).
fof(f1609,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f1515,f42])).
fof(f1515,plain,(
  p(s(bool,'F_0')) | $spl60),
  inference(backward_demodulation,[],[f1502,f32])).
fof(f1608,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f1514,f1502])).
fof(f1514,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | $spl60),
  inference(backward_demodulation,[],[f1502,f33])).
fof(f1607,plain,(
  $false | $spl60),
  inference(subsumption_resolution,[],[f1513,f42])).
fof(f1513,plain,(
  p(s(bool,'F_0')) | $spl60),
  inference(backward_demodulation,[],[f1502,f31])).
fof(f1511,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X102))),s(X96,X103)))) | ~p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X103))),s(X96,X102)))) | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X96,fun(X96,bool)),X97),s(fun(X96,X98),X99),s(fun(X98,X96),X100)))) )),
  inference(subsumption_resolution,[],[f1480,f357])).
fof(f1480,plain,(
  ( ! [X103,X101,X99,X97,X102,X100,X98,X96] : (~p(s(bool,X101)) | p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X102))),s(X96,X103)))) | ~p(s(bool,happ(s(fun(X96,bool),happ(s(fun(X96,fun(X96,bool)),X97),s(X96,X103))),s(X96,X102)))) | s(bool,'F_0') = s(bool,X101) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X96,fun(X96,bool)),X97),s(fun(X96,X98),X99),s(fun(X98,X96),X100)))) )),
  inference(superposition,[],[f41,f400])).
fof(f1509,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X94))),s(X88,X95)))) | ~p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X95))),s(X88,X94)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X88,fun(X88,bool)),X89),s(fun(X88,X90),X91),s(fun(X90,X88),X92))) | s(bool,'F_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f1479,f357])).
fof(f1479,plain,(
  ( ! [X94,X92,X90,X88,X95,X93,X91,X89] : (~p(s(bool,X93)) | p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X94))),s(X88,X95)))) | ~p(s(bool,happ(s(fun(X88,bool),happ(s(fun(X88,fun(X88,bool)),X89),s(X88,X95))),s(X88,X94)))) | s(bool,'F_0') = s(bool,'const.quotient.QUOTIENT_3'(s(fun(X88,fun(X88,bool)),X89),s(fun(X88,X90),X91),s(fun(X90,X88),X92))) | s(bool,'F_0') = s(bool,X93)) )),
  inference(superposition,[],[f41,f400])).
fof(f1504,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (p(s(bool,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),X95),s(X94,X97))),s(X94,X96)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X99),X100),s(fun(X99,X94),X101)))) | s(bool,'F_0') = s(bool,X98) | s(bool,'F_0') = s(bool,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),X95),s(X94,X96))),s(X94,X97)))) )),
  inference(subsumption_resolution,[],[f1459,f357])).
fof(f1459,plain,(
  ( ! [X94,X101,X99,X97,X95,X100,X98,X96] : (~p(s(bool,X98)) | p(s(bool,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),X95),s(X94,X97))),s(X94,X96)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X94,fun(X94,bool)),X95),s(fun(X94,X99),X100),s(fun(X99,X94),X101)))) | s(bool,'F_0') = s(bool,X98) | s(bool,'F_0') = s(bool,happ(s(fun(X94,bool),happ(s(fun(X94,fun(X94,bool)),X95),s(X94,X96))),s(X94,X97)))) )),
  inference(superposition,[],[f41,f400])).
fof(f1501,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89,X86] : (p(s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X89))),s(X86,X88)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X86,fun(X86,bool)),X87),s(fun(X86,X91),X92),s(fun(X91,X86),X93)))) | s(bool,'F_0') = s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X88))),s(X86,X89))) | s(bool,'F_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f1458,f357])).
fof(f1458,plain,(
  ( ! [X92,X90,X88,X87,X93,X91,X89,X86] : (~p(s(bool,X90)) | p(s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X89))),s(X86,X88)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X86,fun(X86,bool)),X87),s(fun(X86,X91),X92),s(fun(X91,X86),X93)))) | s(bool,'F_0') = s(bool,happ(s(fun(X86,bool),happ(s(fun(X86,fun(X86,bool)),X87),s(X86,X88))),s(X86,X89))) | s(bool,'F_0') = s(bool,X90)) )),
  inference(superposition,[],[f41,f400])).
fof(f1493,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (~p(s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X65))),s(X62,X64)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X62,fun(X62,bool)),X63),s(fun(X62,X67),X68),s(fun(X67,X62),X69)))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X64))),s(X62,X65)))) )),
  inference(subsumption_resolution,[],[f1391,f210])).
fof(f1391,plain,(
  ( ! [X68,X66,X64,X62,X69,X67,X65,X63] : (p(s(bool,X66)) | ~p(s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X65))),s(X62,X64)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X62,fun(X62,bool)),X63),s(fun(X62,X67),X68),s(fun(X67,X62),X69)))) | s(bool,'T_0') = s(bool,X66) | s(bool,'T_0') = s(bool,happ(s(fun(X62,bool),happ(s(fun(X62,fun(X62,bool)),X63),s(X62,X64))),s(X62,X65)))) )),
  inference(superposition,[],[f41,f266])).
fof(f1491,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55] : (~p(s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X57))),s(X54,X56)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X54,fun(X54,bool)),X55),s(fun(X54,X59),X60),s(fun(X59,X54),X61)))) | s(bool,'T_0') = s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X56))),s(X54,X57))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1390,f210])).
fof(f1390,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55] : (p(s(bool,X58)) | ~p(s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X57))),s(X54,X56)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(X54,fun(X54,bool)),X55),s(fun(X54,X59),X60),s(fun(X59,X54),X61)))) | s(bool,'T_0') = s(bool,happ(s(fun(X54,bool),happ(s(fun(X54,fun(X54,bool)),X55),s(X54,X56))),s(X54,X57))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f41,f266])).
fof(f104,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7)))) | $spl8),
  inference(cnf_transformation,[],[f104_D])).
fof(f104_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7)))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1321,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X3),X4),s(fun(X3,sK1),X5))))) )),
  inference(resolution,[],[f41,f32])).
fof(f82,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5)))) | $spl2),
  inference(cnf_transformation,[],[f82_D])).
fof(f82_D,plain,(
  p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1320,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X0),X1),s(fun(X0,sK1),X2))))) )),
  inference(resolution,[],[f41,f31])).
fof(f1271,plain,(
  ( ! [X11] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X11) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X11)) )),
  inference(superposition,[],[f33,f400])).
fof(f1270,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,X10) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f400])).
fof(f1269,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X9)) )),
  inference(superposition,[],[f33,f400])).
fof(f1268,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X8) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'F_0') = s(bool,X8)) )),
  inference(superposition,[],[f33,f400])).
fof(f1267,plain,(
  ( ! [X7] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X7) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f33,f266])).
fof(f1266,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,X6) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) )),
  inference(superposition,[],[f33,f266])).
fof(f1265,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f33,f266])).
fof(f1264,plain,(
  ( ! [X4] : (s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,X4) = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f33,f266])).
fof(f1263,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f1262,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f1261,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,X3)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f38])).
fof(f1260,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | ~p(s(bool,X2))) )),
  inference(superposition,[],[f33,f38])).
fof(f1259,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | p(s(bool,X1))) )),
  inference(superposition,[],[f33,f35])).
fof(f1258,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,X0)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f35])).
fof(f1132,plain,(
  $false | $spl42),
  inference(subsumption_resolution,[],[f1089,f43])).
fof(f1089,plain,(
  ~p(s(bool,'T_0')) | $spl42),
  inference(backward_demodulation,[],[f697,f42])).
fof(f697,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl42),
  inference(cnf_transformation,[],[f697_D])).
fof(f697_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f965,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X47) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f400,f266])).
fof(f962,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X37) = s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f400,f266])).
fof(f948,plain,(
  ( ! [X45,X46,X44] : (s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f400,f266])).
fof(f947,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f400,f266])).
fof(f922,plain,(
  ( ! [X47,X48,X46] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X46) = s(bool,X47) | s(bool,'F_0') = s(bool,X47) | s(bool,X46) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f400,f266])).
fof(f919,plain,(
  ( ! [X39,X37,X38] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X37) = s(bool,X38) | s(bool,'F_0') = s(bool,X38) | s(bool,X37) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) )),
  inference(superposition,[],[f400,f266])).
fof(f905,plain,(
  ( ! [X45,X46,X44] : (s(bool,X44) = s(bool,X45) | s(bool,'F_0') = s(bool,X45) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X44) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f400,f266])).
fof(f904,plain,(
  ( ! [X43,X41,X42] : (s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X42) | s(bool,X41) = s(bool,X43) | s(bool,'T_0') = s(bool,X43)) )),
  inference(superposition,[],[f400,f266])).
fof(f1070,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8)))) )),
  inference(subsumption_resolution,[],[f887,f210])).
fof(f887,plain,(
  ( ! [X64] : (p(s(bool,X64)) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8)))) )),
  inference(superposition,[],[f32,f266])).
fof(f1069,plain,(
  ( ! [X63] : (s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6)))) )),
  inference(subsumption_resolution,[],[f886,f210])).
fof(f886,plain,(
  ( ! [X63] : (p(s(bool,X63)) | s(bool,'T_0') = s(bool,X63) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6)))) )),
  inference(superposition,[],[f31,f266])).
fof(f882,plain,(
  ( ! [X57,X58,X56] : (s(bool,X56) = s(bool,X57) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,'T_0') = s(bool,X56) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f400,f266])).
fof(f881,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,X53) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f400,f266])).
fof(f1064,plain,(
  ( ! [X64] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X64)) )),
  inference(subsumption_resolution,[],[f858,f210])).
fof(f858,plain,(
  ( ! [X64] : (p(s(bool,X64)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,sK8))) | s(bool,'T_0') = s(bool,X64)) )),
  inference(superposition,[],[f32,f266])).
fof(f1061,plain,(
  ( ! [X63] : (s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(subsumption_resolution,[],[f857,f210])).
fof(f857,plain,(
  ( ! [X63] : (p(s(bool,X63)) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK6))) | s(bool,'T_0') = s(bool,X63)) )),
  inference(superposition,[],[f31,f266])).
fof(f853,plain,(
  ( ! [X57,X58,X56] : (s(bool,X56) = s(bool,X57) | s(bool,X57) = s(bool,X58) | s(bool,'F_0') = s(bool,X58) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X56)) )),
  inference(superposition,[],[f400,f266])).
fof(f852,plain,(
  ( ! [X54,X55,X53] : (s(bool,X53) = s(bool,X54) | s(bool,X54) = s(bool,X55) | s(bool,'F_0') = s(bool,X55) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f400,f266])).
fof(f828,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f266,f400])).
fof(f825,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f266,f400])).
fof(f817,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f266,f400])).
fof(f816,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) = s(bool,X38) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f266,f400])).
fof(f800,plain,(
  ( ! [X33,X34,X32] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X32) = s(bool,X33) | s(bool,'T_0') = s(bool,X33) | s(bool,X32) = s(bool,X34) | s(bool,'F_0') = s(bool,X34)) )),
  inference(superposition,[],[f266,f400])).
fof(f797,plain,(
  ( ! [X24,X23,X25] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X23) = s(bool,X24) | s(bool,'T_0') = s(bool,X24) | s(bool,X23) = s(bool,X25) | s(bool,'F_0') = s(bool,X25)) )),
  inference(superposition,[],[f266,f400])).
fof(f789,plain,(
  ( ! [X41,X42,X40] : (s(bool,X40) = s(bool,X41) | s(bool,X41) = s(bool,X42) | s(bool,'T_0') = s(bool,X42) | s(bool,'F_0') = s(bool,X40) | s(bool,'T_0') = s(bool,'F_0')) )),
  inference(superposition,[],[f266,f400])).
fof(f788,plain,(
  ( ! [X39,X37,X38] : (s(bool,X37) = s(bool,X38) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'F_0') = s(bool,X37)) )),
  inference(superposition,[],[f266,f400])).
fof(f773,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,X50) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X49) = s(bool,X51) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f266,f400])).
fof(f772,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,'F_0') | s(bool,X46) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f266,f400])).
fof(f752,plain,(
  ( ! [X50,X51,X49] : (s(bool,X49) = s(bool,X50) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X50) | s(bool,X49) = s(bool,X51) | s(bool,'F_0') = s(bool,X51)) )),
  inference(superposition,[],[f266,f400])).
fof(f751,plain,(
  ( ! [X47,X48,X46] : (s(bool,X46) = s(bool,X47) | s(bool,'T_0') = s(bool,'F_0') | s(bool,'T_0') = s(bool,X47) | s(bool,X46) = s(bool,X48) | s(bool,'F_0') = s(bool,X48)) )),
  inference(superposition,[],[f266,f400])).
fof(f620,plain,(
  ( ! [X40] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f62,f400])).
fof(f585,plain,(
  ( ! [X40] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X40) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X40)) )),
  inference(superposition,[],[f62,f400])).
fof(f508,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X27) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f400,f62])).
fof(f487,plain,(
  ( ! [X27] : (s(bool,'T_0') = s(bool,'F_0') | s(bool,X27) = s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,sK0),sK3),s(fun(sK0,sK1),sK4))) | s(bool,'F_0') = s(bool,X27)) )),
  inference(superposition,[],[f400,f62])).
fof(f348,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f294,f288])).
fof(f294,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))) | $spl36),
  inference(backward_demodulation,[],[f288,f33])).
fof(f343,plain,(
  $false | $spl36),
  inference(subsumption_resolution,[],[f290,f43])).
fof(f290,plain,(
  ~p(s(bool,'T_0')) | $spl36),
  inference(backward_demodulation,[],[f288,f42])).
fof(f289,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) | $spl38),
  inference(cnf_transformation,[],[f289_D])).
fof(f289_D,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') = s(bool,X16)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f284,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,X17) | s(bool,'F_0') = s(bool,X16)) )),
  inference(duplicate_literal_removal,[],[f281])).
fof(f281,plain,(
  ( ! [X17,X16] : (p(s(bool,X16)) | s(bool,'T_0') = s(bool,X17) | s(bool,'T_0') = s(bool,X17) | s(bool,'F_0') = s(bool,X16)) )),
  inference(superposition,[],[f245,f39])).
fof(f232,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f33,f38])).
fof(f231,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f38])).
fof(f230,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8))))) )),
  inference(superposition,[],[f33,f35])).
fof(f229,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))) | p(s(bool,X0))) )),
  inference(superposition,[],[f33,f35])).
fof(f228,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'T_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f227,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,sK7))) | s(bool,'F_0') = s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK8)))),
  inference(superposition,[],[f33,f39])).
fof(f148,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK7)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK8)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) ) | $spl8),
  inference(resolution,[],[f104,f40])).
fof(f147,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X3)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X4),X5),s(fun(X4,sK1),X6))))) ) | $spl8),
  inference(resolution,[],[f104,f40])).
fof(f128,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK5)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK6)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) ) | $spl2),
  inference(resolution,[],[f82,f40])).
fof(f127,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X3)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X4),X5),s(fun(X4,sK1),X6))))) ) | $spl2),
  inference(resolution,[],[f82,f40])).
fof(f108,plain,(
  $false | $spl0),
  inference(resolution,[],[f80,f30])).
fof(f88,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK8)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) )),
  inference(resolution,[],[f32,f40])).
fof(f87,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK7))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,X3)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X4),X5),s(fun(X4,sK1),X6))))) )),
  inference(resolution,[],[f32,f40])).
fof(f86,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK8))),s(sK1,sK7)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X0),X1),s(fun(X0,sK1),X2))))) )),
  inference(resolution,[],[f32,f41])).
fof(f65,plain,(
  ( ! [X10,X8,X7,X9] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK6)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,X7))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X8),X9),s(fun(X8,sK1),X10))))) )),
  inference(resolution,[],[f31,f40])).
fof(f64,plain,(
  ( ! [X6,X4,X5,X3] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK5))),s(sK1,X3)))) | ~p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,X3)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X4),X5),s(fun(X4,sK1),X6))))) )),
  inference(resolution,[],[f31,f40])).
fof(f63,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(sK1,bool),happ(s(fun(sK1,fun(sK1,bool)),sK2),s(sK1,sK6))),s(sK1,sK5)))) | ~p(s(bool,'const.quotient.QUOTIENT_3'(s(fun(sK1,fun(sK1,bool)),sK2),s(fun(sK1,X0),X1),s(fun(X0,sK1),X2))))) )),
  inference(resolution,[],[f31,f41])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 130
% Passive clauses: 2356
% Generated clauses: 8950
% Final active clauses: 75
% Final passive clauses: 1774
% Input formulas: 8
% Initial clauses: 14
% 
% Selected by SInE selection: 8
% SInE iterations: 1
% 
% Duplicate literals: 492
% Trivial inequalities: 177
% Fw subsumption resolutions: 1503
% Fw demodulations: 325
% Bw demodulations: 32
% 
% Simple tautologies: 371
% Equational tautologies: 37
% Forward subsumptions: 3221
% 
% Binary resolution: 872
% Factoring: 1
% Forward superposition: 3345
% Backward superposition: 1755
% Self superposition: 96
% Equality factoring: 92
% Equality resolution: 15
% 
% Split clauses: 442
% Split components: 884
% SAT solver clauses: 467
% SAT solver unit clauses: 13
% SAT solver binary clauses: 166
% 
% Sat splits: 442
% Sat splitting refutations: 24
% 
% Memory used [KB]: 4093
% Time elapsed: 0.288 s
% ------------------------------
% ------------------------------
% Success in time 11.996 s
