Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.quotient
% (103613)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 965
% Passive clauses: 8719
% Generated clauses: 25302
% Final active clauses: 955
% Final passive clauses: 7419
% Input formulas: 52
% Initial clauses: 209
% 
% Duplicate literals: 649
% Trivial inequalities: 85
% Fw subsumption resolutions: 2852
% Fw demodulations: 1739
% Bw demodulations: 10
% 
% Simple tautologies: 295
% Equational tautologies: 45
% Forward subsumptions: 10215
% 
% Binary resolution: 10156
% Factoring: 1561
% Forward superposition: 4080
% Backward superposition: 3632
% Self superposition: 28
% Equality factoring: 28
% Equality resolution: 20
% 
% Split clauses: 627
% Split components: 1260
% SAT solver clauses: 1053
% SAT solver unit clauses: 17
% SAT solver binary clauses: 220
% 
% Sat splits: 627
% 
% TWLsolver clauses: 1053
% TWLsolver calls for satisfiability: 575
% 
% Memory used [KB]: 12409
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f12883,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f12882,f2855,f2874,f294,f508,f11707,f2872,f3052,f3054,f484,f526,f1200,f1202,f500,f502,f364,f535,f658,f664,f273,f495,f298,f510,f3058,f3078,f11698,f2558,f2575,f366,f539,f274,f499,f365,f540,f639,f643,f272,f497,f2847,f2852,f2850,f360,f528,f12872,f11715,f2848,f361,f527])).
fof(f527,plain,(
  ( ! [X4,X0,X1] : (sP0(X0,X1) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4))))) ) | $spl24),
  inference(cnf_transformation,[],[f527_D])).
fof(f527_D,plain,(
  ( ! [X4,X0,X1] : (sP0(X0,X1) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f361,plain,(
  ( ! [X4,X0,X1] : (sP0(X0,X1) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK15(X1,X0)))),s(X0,sK15(X1,X0))))) & ! [X5,X6] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))))))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK16(X1,X0))))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK16(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15,sK16,sK17])],[f213])).
fof(f213,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (? [X7] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X7))),s(X0,X7)))) & ! [X5,X6] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X5)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))),s(X0,X6)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X6))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))),s(X0,X6)))))))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | ? [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) | sP0(X0,X1)))),
  inference(rectify,[],[f212])).
fof(f212,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))))))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | ? [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) | sP0(X0,X1)))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))))))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) | ? [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f2848,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0))),s(sK13,X0))))) ) | $spl1098),
  inference(cnf_transformation,[],[f2848_D])).
fof(f2848_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0))),s(sK13,X0))))) ) <=> ~$spl1098),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1098])])).
fof(f11715,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0))),s(sK13,X0))))) ) | ($spl24 | $spl463 | $spl1101 | $spl1105)),
  inference(subsumption_resolution,[],[f11714,f2849])).
fof(f2849,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | $spl1101),
  inference(cnf_transformation,[],[f2849_D])).
fof(f2849_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) <=> ~$spl1101),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1101])])).
fof(f11714,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0))),s(sK13,X0))))) ) | ($spl24 | $spl463 | $spl1105)),
  inference(subsumption_resolution,[],[f2906,f2871])).
fof(f2871,plain,(
  s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) != s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | $spl1105),
  inference(cnf_transformation,[],[f2871_D])).
fof(f2871_D,plain,(
  s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) != s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) <=> ~$spl1105),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1105])])).
fof(f2906,plain,(
  ( ! [X0] : (s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) = s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0))),s(sK13,X0))))) ) | ($spl24 | $spl463)),
  inference(resolution,[],[f527,f1202])).
fof(f12872,plain,(
  $false | ($spl4 | $spl220 | $spl1022 | $spl1098)),
  inference(subsumption_resolution,[],[f12871,f499])).
fof(f12871,plain,(
  ~p(s(bool,'T_0')) | ($spl220 | $spl1022 | $spl1098)),
  inference(forward_demodulation,[],[f12813,f643])).
fof(f12813,plain,(
  ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) | ($spl1022 | $spl1098)),
  inference(resolution,[],[f2848,f2575])).
fof(f528,plain,(
  ( ! [X4,X0,X1] : (sP0(X0,X1) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4))))) ) | $spl26),
  inference(cnf_transformation,[],[f528_D])).
fof(f528_D,plain,(
  ( ! [X4,X0,X1] : (sP0(X0,X1) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f360,plain,(
  ( ! [X4,X0,X1] : (sP0(X0,X1) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f2850,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | $spl1100),
  inference(cnf_transformation,[],[f2850_D])).
fof(f2850_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) <=> ~$spl1100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1100])])).
fof(f2852,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | $spl1102),
  inference(cnf_transformation,[],[f2852_D])).
fof(f2852_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) <=> ~$spl1102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1102])])).
fof(f2847,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,X0))),s(sK13,X0))))) ) | ($spl26 | $spl463)),
  inference(resolution,[],[f528,f1202])).
fof(f497,plain,(
  p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) | $spl2),
  inference(cnf_transformation,[],[f497_D])).
fof(f497_D,plain,(
  p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f272,plain,(
  p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))))),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) & ~p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14])],[f128])).
fof(f128,plain,(
  ? [X0,X1] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) & ~p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ~! [X0,X1] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) => p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) => p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X6,fun(X6,bool)),X7)))) => p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X6,fun(X6,bool)),X7)))) => p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X6,fun(X6,bool)),X7)))))),
  file('thm.quotient.EQUIV_IMP_PARTIAL_EQUIV','thm.quotient.EQUIV_IMP_PARTIAL_EQUIV')).
fof(f643,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))) | $spl220),
  inference(cnf_transformation,[],[f643_D])).
fof(f643_D,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f639,plain,(
  s(bool,'T_0') = s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))) | ($spl2 | $spl10)),
  inference(resolution,[],[f508,f497])).
fof(f540,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl50),
  inference(cnf_transformation,[],[f540_D])).
fof(f540_D,plain,(
  ( ! [X4,X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f365,plain,(
  ( ! [X4,X0,X5,X1] : (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))) & (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))))) & (((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0)))),s(X0,sK19(X1,X0))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0)))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK19(X1,X0))))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0)))),s(X0,sK19(X1,X0))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK18(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK19(X1,X0)))))) | p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK18,sK19])],[f217])).
fof(f217,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X4,X5] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5)))) & (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5))))))) & (? [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) | p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f216])).
fof(f216,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) | ! [X2,X3] : ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))) & (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3))))))) & (? [X2,X3] : ((p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))))) | p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) <=> s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) <=> s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X6] : ! [X14] : (p(s(bool,'const.quotient.EQUIV_1'(s(fun(X6,fun(X6,bool)),X14)))) <=> ! [X11,X15] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X11))),s(X6,X15)))) <=> s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X11))) = s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X14),s(X6,X15)))))),
  file('thm.quotient.EQUIV_IMP_PARTIAL_EQUIV','thm.quotient.EQUIV_def')).
fof(f499,plain,(
  p(s(bool,'T_0')) | $spl4),
  inference(cnf_transformation,[],[f499_D])).
fof(f499_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f274,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.EQUIV_IMP_PARTIAL_EQUIV',aHL_TRUTH)).
fof(f539,plain,(
  ( ! [X4,X0,X5,X1] : (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl48),
  inference(cnf_transformation,[],[f539_D])).
fof(f539_D,plain,(
  ( ! [X4,X0,X5,X1] : (s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f366,plain,(
  ( ! [X4,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X5)))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X5))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f218])).
fof(f2575,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl1022),
  inference(cnf_transformation,[],[f2575_D])).
fof(f2575_D,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl1022),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1022])])).
fof(f2558,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl48),
  inference(equality_resolution,[],[f539])).
fof(f11698,plain,(
  $false | ($spl4 | $spl7 | $spl220 | $spl1022 | $spl1144)),
  inference(subsumption_resolution,[],[f11697,f499])).
fof(f11697,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl220 | $spl1022 | $spl1144)),
  inference(forward_demodulation,[],[f11696,f643])).
fof(f11696,plain,(
  ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) | ($spl7 | $spl1022 | $spl1144)),
  inference(subsumption_resolution,[],[f11652,f502])).
fof(f11652,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) | ($spl1022 | $spl1144)),
  inference(superposition,[],[f2575,f3078])).
fof(f3078,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13)))) | $spl1144),
  inference(cnf_transformation,[],[f3078_D])).
fof(f3078_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13)))) <=> ~$spl1144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1144])])).
fof(f3058,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13)))) | ($spl14 | $spl1141)),
  inference(resolution,[],[f3054,f510])).
fof(f510,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f510_D])).
fof(f510_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f298,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.quotient.EQUIV_IMP_PARTIAL_EQUIV','thm.bool.EQ_CLAUSES|split|2')).
fof(f495,plain,(
  ~p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) | $spl1),
  inference(cnf_transformation,[],[f495_D])).
fof(f495_D,plain,(
  ~p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f273,plain,(
  ~p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))))),
  inference(cnf_transformation,[],[f171])).
fof(f664,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))) | $spl226),
  inference(cnf_transformation,[],[f664_D])).
fof(f664_D,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f658,plain,(
  s(bool,'F_0') = s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14))) | ($spl1 | $spl14)),
  inference(resolution,[],[f510,f495])).
fof(f535,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~sP0(X0,X1)) ) | $spl40),
  inference(cnf_transformation,[],[f535_D])).
fof(f535_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~sP0(X0,X1)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f364,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) | sP0(X0,X1)) & (~sP0(X0,X1) | p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1] : (p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> sP0(X0,X1))),
  inference(definition_folding,[],[f104,f151])).
fof(f104,plain,(
  ! [X0,X1] : (p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4)))) & ! [X2,X3] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X3)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))),s(X0,X2)))) & p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3))),s(X0,X3)))) & s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X2))) = s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X3)))))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.quotient.PARTIAL_EQUIV_1'(s(fun(X6,fun(X6,bool)),X7)))) <=> (? [X11] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X11)))) & ! [X11,X15] : (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X15)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))),s(X6,X11)))) & p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X15))),s(X6,X15)))) & s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X11))) = s(fun(X6,bool),happ(s(fun(X6,fun(X6,bool)),X7),s(X6,X15)))))))),
  file('thm.quotient.EQUIV_IMP_PARTIAL_EQUIV','thm.quotient.PARTIAL_EQUIV_def')).
fof(f502,plain,(
  ~p(s(bool,'F_0')) | $spl7),
  inference(cnf_transformation,[],[f502_D])).
fof(f502_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f500,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f277,f274])).
fof(f277,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f32])).
fof(f32,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quotient.EQUIV_IMP_PARTIAL_EQUIV','thm.bool.NOT_CLAUSES|split|2')).
fof(f1202,plain,(
  ~sP0(sK13,sK14) | $spl463),
  inference(cnf_transformation,[],[f1202_D])).
fof(f1202_D,plain,(
  ~sP0(sK13,sK14) <=> ~$spl463),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl463])])).
fof(f1200,plain,(
  ~sP0(sK13,sK14) | ($spl7 | $spl40 | $spl226)),
  inference(subsumption_resolution,[],[f1188,f502])).
fof(f1188,plain,(
  p(s(bool,'F_0')) | ~sP0(sK13,sK14) | ($spl40 | $spl226)),
  inference(superposition,[],[f535,f664])).
fof(f526,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK16(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0)))))) ) | $spl22),
  inference(cnf_transformation,[],[f526_D])).
fof(f526_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK16(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0)))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f484,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | sP0(X0,X1) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK16(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0)))))) )),
  inference(condensation,[],[f362])).
fof(f362,plain,(
  ( ! [X4,X0,X1] : (sP0(X0,X1) | s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))) != s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK17(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK16(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,sK16(X1,X0)))),s(X0,sK17(X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X0,fun(X0,bool)),X1),s(X0,X4))),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f214])).
fof(f3054,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) | $spl1141),
  inference(cnf_transformation,[],[f3054_D])).
fof(f3054_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) <=> ~$spl1141),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1141])])).
fof(f3052,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) | ($spl4 | $spl22 | $spl463 | $spl1104 | $spl1106)),
  inference(subsumption_resolution,[],[f3051,f499])).
fof(f3051,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) | ($spl22 | $spl463 | $spl1104 | $spl1106)),
  inference(forward_demodulation,[],[f3050,f2874])).
fof(f3050,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) | ($spl22 | $spl463 | $spl1104)),
  inference(duplicate_literal_removal,[],[f3049])).
fof(f3049,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ($spl22 | $spl463 | $spl1104)),
  inference(forward_demodulation,[],[f3048,f2872])).
fof(f3048,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ($spl22 | $spl463 | $spl1104)),
  inference(subsumption_resolution,[],[f3047,f2872])).
fof(f3047,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) != s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK16(sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ($spl22 | $spl463)),
  inference(resolution,[],[f526,f1202])).
fof(f2872,plain,(
  s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) = s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | $spl1104),
  inference(cnf_transformation,[],[f2872_D])).
fof(f2872_D,plain,(
  s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) = s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) <=> ~$spl1104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1104])])).
fof(f11707,plain,(
  s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) = s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | ($spl4 | $spl50 | $spl220 | $spl1106)),
  inference(subsumption_resolution,[],[f11706,f499])).
fof(f11706,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) = s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | ($spl4 | $spl50 | $spl220 | $spl1106)),
  inference(forward_demodulation,[],[f11705,f643])).
fof(f11705,plain,(
  s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) = s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) | ($spl4 | $spl50 | $spl1106)),
  inference(subsumption_resolution,[],[f2976,f499])).
fof(f2976,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))) = s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK17(sK14,sK13)))) | ~p(s(bool,'const.quotient.EQUIV_1'(s(fun(sK13,fun(sK13,bool)),sK14)))) | ($spl50 | $spl1106)),
  inference(superposition,[],[f540,f2874])).
fof(f508,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f508_D])).
fof(f508_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f294,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.quotient.EQUIV_IMP_PARTIAL_EQUIV','thm.bool.EQ_CLAUSES|split|0')).
fof(f2874,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13)))) | $spl1106),
  inference(cnf_transformation,[],[f2874_D])).
fof(f2874_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13)))) <=> ~$spl1106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1106])])).
fof(f2855,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13)))) | ($spl10 | $spl1100)),
  inference(resolution,[],[f2850,f508])).
fof(f12882,plain,(
  $false | ($spl1101 | $spl1102 | $spl1104)),
  inference(subsumption_resolution,[],[f12881,f2849])).
fof(f12881,plain,(
  p(s(bool,happ(s(fun(sK13,bool),happ(s(fun(sK13,fun(sK13,bool)),sK14),s(sK13,sK16(sK14,sK13)))),s(sK13,sK17(sK14,sK13))))) | ($spl1102 | $spl1104)),
  inference(backward_demodulation,[],[f2872,f2852])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 448
% Passive clauses: 2510
% Generated clauses: 12576
% Final active clauses: 418
% Final passive clauses: 2046
% Input formulas: 52
% Initial clauses: 209
% 
% Duplicate literals: 251
% Trivial inequalities: 98
% Fw subsumption resolutions: 1070
% Fw demodulations: 926
% Bw demodulations: 18
% Condensations: 324
% Global subsumptions: 22
% 
% Simple tautologies: 298
% Equational tautologies: 30
% Forward subsumptions: 4117
% 
% Binary resolution: 3395
% Factoring: 186
% Forward superposition: 1722
% Backward superposition: 1711
% Self superposition: 30
% Equality factoring: 30
% Equality resolution: 13
% 
% Split clauses: 74
% Split components: 160
% SAT solver clauses: 3117
% SAT solver unit clauses: 229
% SAT solver binary clauses: 84
% SAT solver learnt clauses: 47
% SAT solver learnt literals: 13
% 
% Sat splits: 74
% Sat splitting refutations: 10
% 
% TWLsolver clauses: 3026
% TWLsolver calls for satisfiability: 1979
% 
% Memory used [KB]: 12537
% Time elapsed: 0.673 s
% ------------------------------
% ------------------------------
% Success in time 1.301 s
