Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (74515)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 258
% Passive clauses: 6286
% Generated clauses: 32458
% Final active clauses: 241
% Final passive clauses: 5730
% Input formulas: 69
% Initial clauses: 126
% 
% Duplicate literals: 645
% Trivial inequalities: 1326
% Fw subsumption resolutions: 4330
% Fw demodulations: 6677
% Bw demodulations: 17
% 
% Simple tautologies: 268
% Equational tautologies: 31
% Forward subsumptions: 13048
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 891
% Factoring: 1
% Forward superposition: 8244
% Backward superposition: 10006
% Self superposition: 99
% Equality factoring: 40
% Equality resolution: 18
% 
% Split clauses: 84
% Split components: 169
% SAT solver clauses: 127
% SAT solver unit clauses: 9
% SAT solver binary clauses: 57
% 
% Sat splits: 84
% 
% TWLsolver clauses: 127
% TWLsolver calls for satisfiability: 60
% 
% Memory used [KB]: 17654
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (74574)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 164
% Passive clauses: 2098
% Generated clauses: 17451
% Final active clauses: 161
% Final passive clauses: 1917
% Input formulas: 69
% Initial clauses: 126
% 
% Duplicate literals: 395
% Trivial inequalities: 313
% Fw subsumption resolutions: 1421
% Fw demodulations: 2447
% Bw demodulations: 19
% Condensations: 564
% Global subsumptions: 8
% 
% Simple tautologies: 285
% Equational tautologies: 28
% Forward subsumptions: 6937
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 544
% Factoring: 3
% Forward superposition: 4680
% Backward superposition: 4702
% Self superposition: 87
% Equality factoring: 34
% Equality resolution: 16
% 
% Split clauses: 7
% Split components: 14
% SAT solver clauses: 2905
% SAT solver unit clauses: 147
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 16
% 
% Sat splits: 7
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 2873
% TWLsolver calls for satisfiability: 485
% 
% Memory used [KB]: 22387
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f119420,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f119414])).
fof(f119414,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'('type.sum.sum'('type.list.list'(sK0),X0),X1),'const.sum.INL_1'(s('type.sum.sum'('type.list.list'(sK0),X0),'const.sum.INL_1'(s('type.list.list'(sK0),sK1))))) != s('type.sum.sum'('type.sum.sum'('type.list.list'(sK0),X0),X1),'const.sum.INL_1'(s('type.sum.sum'('type.list.list'(sK0),X0),'const.sum.INL_1'(s('type.list.list'(sK0),sK1)))))) )),
  inference(backward_demodulation,[],[f118887,f1467])).
fof(f1467,plain,(
  ( ! [X0,X1] : (s('type.sum.sum'('type.sum.sum'('type.list.list'(sK0),X0),X1),'const.sum.INL_1'(s('type.sum.sum'('type.list.list'(sK0),X0),'const.sum.INL_1'(s('type.list.list'(sK0),sK1))))) != s('type.sum.sum'('type.sum.sum'('type.list.list'(sK0),X0),X1),'const.sum.INL_1'(s('type.sum.sum'('type.list.list'(sK0),X0),'const.sum.INL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK0,sK0),'const.combin.I_0'),s('type.list.list'(sK0),sK1)))))))) )),
  inference(unit_resulting_resolution,[],[f1068,f349])).
fof(f349,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) )),
  inference(cnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1,X2,X3] : ((s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) != s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) | s(X1,X2) = s(X1,X3)) & (s(X1,X2) != s(X1,X3) | s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3)))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2,X3] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X2))) = s('type.sum.sum'(X1,X0),'const.sum.INL_1'(s(X1,X3))) <=> s(X1,X2) = s(X1,X3))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X8,X6] : ! [X27,X15] : (s('type.sum.sum'(X6,X8),'const.sum.INL_1'(s(X6,X15))) = s('type.sum.sum'(X6,X8),'const.sum.INL_1'(s(X6,X27))) <=> s(X6,X15) = s(X6,X27))),
  file('thm.list.MAP_ID|split|1','thm.sum.INR_INL_|49||49||split|0')).
fof(f1068,plain,(
  ( ! [X0] : (s('type.sum.sum'('type.list.list'(sK0),X0),'const.sum.INL_1'(s('type.list.list'(sK0),sK1))) != s('type.sum.sum'('type.list.list'(sK0),X0),'const.sum.INL_1'(s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK0,sK0),'const.combin.I_0'),s('type.list.list'(sK0),sK1)))))) )),
  inference(unit_resulting_resolution,[],[f251,f349])).
fof(f251,plain,(
  s('type.list.list'(sK0),sK1) != s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK0,sK0),'const.combin.I_0'),s('type.list.list'(sK0),sK1)))),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  s('type.list.list'(sK0),sK1) != s('type.list.list'(sK0),'const.list.MAP_2'(s(fun(sK0,sK0),'const.combin.I_0'),s('type.list.list'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f179])).
fof(f179,plain,(
  ? [X0,X1] : s('type.list.list'(X0),X1) != s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s('type.list.list'(X6),X7) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X6,X6),'const.combin.I_0'),s('type.list.list'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s('type.list.list'(X6),X7) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X6,X6),'const.combin.I_0'),s('type.list.list'(X6),X7)))),
  file('thm.list.MAP_ID|split|1','thm.list.MAP_ID|split|1')).
fof(f118887,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'const.combin.I_0'),s('type.list.list'(X0),X1)))) )),
  inference(backward_demodulation,[],[f118845,f303])).
fof(f303,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'_i_0_0'),s('type.list.list'(X0),X1)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'_i_0_0'),s('type.list.list'(X0),X1)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.MAP_2'(s(fun(X0,X0),'_i_0_0'),s('type.list.list'(X0),X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X7] : s('type.list.list'(X6),X7) = s('type.list.list'(X6),'const.list.MAP_2'(s(fun(X6,X6),'_i_0_0'),s('type.list.list'(X6),X7)))),
  file('thm.list.MAP_ID|split|1','thm.list.MAP_ID|split|0')).
fof(f118845,plain,(
  ( ! [X0] : (s(fun(X0,X0),'const.combin.I_0') = s(fun(X0,X0),'_i_0_0')) )),
  inference(unit_resulting_resolution,[],[f384,f10694])).
fof(f10694,plain,(
  ( ! [X0,X1] : (s(X0,sK14(X1,'_i_0_0',X0,X0)) != s(X0,happ(s(fun(X0,X0),X1),s(X0,sK14(X1,'_i_0_0',X0,X0)))) | s(fun(X0,X0),'_i_0_0') = s(fun(X0,X0),X1)) )),
  inference(superposition,[],[f348,f300])).
fof(f300,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,happ(s(fun(X0,X0),'_i_0_0'),s(X0,X1)))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X23] : ! [X25] : s(X23,X25) = s(X23,happ(s(fun(X23,X23),'_i_0_0'),s(X23,X25)))),
  file('thm.list.MAP_ID|split|1','_i0_ax')).
fof(f348,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK14(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (s(X0,happ(s(fun(X1,X0),X2),s(X1,sK14(X3,X2,X1,X0)))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,sK14(X3,X2,X1,X0)))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14])],[f245])).
fof(f245,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X5] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X5))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X5)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(rectify,[],[f244])).
fof(f244,plain,(
  ! [X0,X1,X2,X3] : ((s(fun(X1,X0),X2) != s(fun(X1,X0),X3) | ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4)))) & (? [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) != s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))) | s(fun(X1,X0),X2) = s(fun(X1,X0),X3)))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0,X1] : ! [X2,X3] : (s(fun(X1,X0),X2) = s(fun(X1,X0),X3) <=> ! [X4] : s(X0,happ(s(fun(X1,X0),X2),s(X1,X4))) = s(X0,happ(s(fun(X1,X0),X3),s(X1,X4))))),
  inference(rectify,[],[f59])).
fof(f59,axiom,(
  ! [X8,X6] : ! [X9,X39] : (s(fun(X6,X8),X9) = s(fun(X6,X8),X39) <=> ! [X15] : s(X8,happ(s(fun(X6,X8),X9),s(X6,X15))) = s(X8,happ(s(fun(X6,X8),X39),s(X6,X15))))),
  file('thm.list.MAP_ID|split|1','thm.bool.FUN_EQ_THM')).
fof(f384,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f302,f299])).
fof(f299,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6] : ! [X15] : s(X6,X15) = s(X6,'const.combin.I_1'(s(X6,X15)))),
  file('thm.list.MAP_ID|split|1','thm.combin.I_THM')).
fof(f302,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X14] : s(X6,'const.combin.I_1'(s(X6,X14))) = s(X6,happ(s(fun(X6,X6),'const.combin.I_0'),s(X6,X14)))),
  file('thm.list.MAP_ID|split|1','ARITY_0_1_const.combin.I')).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 485
% Passive clauses: 8963
% Generated clauses: 121701
% Final active clauses: 283
% Final passive clauses: 2917
% Input formulas: 69
% Initial clauses: 126
% Discarded non-redundant clauses: 2513
% 
% Duplicate literals: 1499
% Trivial inequalities: 127
% Fw subsumption resolutions: 6682
% Fw demodulations: 36476
% Bw demodulations: 3339
% Condensations: 1542
% Global subsumptions: 28
% 
% Simple tautologies: 411
% Equational tautologies: 503
% Forward subsumptions: 47590
% Fw demodulations to eq. taut.: 17107
% Bw demodulations to eq. taut.: 57
% 
% Unit resulting resolution: 20829
% Factoring: 1
% Forward superposition: 26244
% Backward superposition: 21221
% Self superposition: 136
% Equality factoring: 66
% Equality resolution: 34
% 
% Split clauses: 77
% Split components: 154
% SAT solver clauses: 9250
% SAT solver unit clauses: 532
% SAT solver binary clauses: 4506
% 
% Sat splits: 77
% Sat splitting refutations: 103
% 
% Memory used [KB]: 58463
% Time elapsed: 4.059 s
% ------------------------------
% ------------------------------
% Success in time 5.39 s
