Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.complex
% (9919)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 689
% Passive clauses: 8326
% Generated clauses: 22962
% Final active clauses: 688
% Final passive clauses: 7276
% Input formulas: 133
% Initial clauses: 326
% 
% Duplicate literals: 276
% Trivial inequalities: 30
% Fw subsumption resolutions: 707
% Fw demodulations: 1144
% Bw demodulations: 1
% 
% Simple tautologies: 234
% Equational tautologies: 55
% Forward subsumptions: 12111
% Fw demodulations to eq. taut.: 16
% 
% Binary resolution: 9523
% Factoring: 1518
% Forward superposition: 4952
% Backward superposition: 4316
% Self superposition: 74
% Equality factoring: 2
% Equality resolution: 20
% 
% Split clauses: 221
% Split components: 445
% SAT solver clauses: 433
% SAT solver unit clauses: 17
% SAT solver binary clauses: 84
% 
% Sat splits: 221
% 
% TWLsolver clauses: 433
% TWLsolver calls for satisfiability: 358
% 
% Memory used [KB]: 13816
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.complex
% (9952)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 426
% Passive clauses: 2593
% Generated clauses: 14848
% Final active clauses: 424
% Final passive clauses: 2150
% Input formulas: 133
% Initial clauses: 326
% 
% Duplicate literals: 125
% Trivial inequalities: 25
% Fw subsumption resolutions: 522
% Fw demodulations: 1408
% Bw demodulations: 17
% Condensations: 139
% Global subsumptions: 22
% 
% Simple tautologies: 216
% Equational tautologies: 24
% Forward subsumptions: 6581
% Fw demodulations to eq. taut.: 24
% 
% Binary resolution: 3344
% Factoring: 346
% Forward superposition: 3824
% Backward superposition: 2034
% Self superposition: 82
% Equality factoring: 2
% Equality resolution: 23
% 
% Split clauses: 12
% Split components: 25
% SAT solver clauses: 3498
% SAT solver unit clauses: 375
% SAT solver binary clauses: 114
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 12
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 3410
% TWLsolver calls for satisfiability: 2161
% 
% Memory used [KB]: 15991
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.complex
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.complex
% SZS output start Proof for thm.complex
fof(f78741,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f477,f828,f476,f830,f826,f475,f827,f829,f825,f857,f878,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f11617,f11517,f11619,f11521,f11621,f11524,f11623,f11527,f11625,f11530,f11627,f11533,f11629,f11536,f11631,f11539,f11633,f11542,f11635,f11545,f11637,f11548,f11639,f11551,f11641,f11554,f11643,f11557,f11645,f11560,f11647,f11563,f11649,f11566,f11651,f11569,f11653,f11572,f11655,f11575,f11657,f11578,f11659,f11518,f11581,f11801,f11802,f11803,f11804,f11805,f11806,f11815,f11819,f11821,f11823,f11825,f11827,f11829,f11831,f11833,f11835,f11837,f11839,f11841,f11843,f11845,f11847,f11849,f11851,f11853,f11855,f11857,f11859,f11861,f11863,f11865,f11866,f11867,f11868,f11869,f11870,f11871,f11872,f11873,f11874,f11875,f11876,f11877,f11879,f11883,f11887,f11891,f11894,f11898,f11902,f11905,f11909,f11913,f11916,f11920,f11924,f11927,f11930,f11933,f11936,f11939,f16214,f16217,f16216,f16589,f16591,f16590,f16599,f16601,f16603,f16607,f16609,f16613,f16615,f16617,f16619,f16626,f16628,f16635,f16637,f16641,f16645,f16659,f16678,f16713,f16715,f16767,f16776,f16796,f16804,f16840,f16852,f16867,f16879,f16881,f16888,f16894,f16901,f16907,f16909,f16911,f16913,f16915,f17595,f17596,f17597,f17598,f16932,f16934,f16938,f16940,f16942,f16944,f16948,f16950,f16969,f16971,f16977,f16983,f16992,f16999,f17001,f17013,f17021,f17023,f17033,f17039,f17046,f17052,f17054,f17065,f17067,f17074,f17076,f17088,f17090,f17097,f17099,f17111,f17113,f17120,f17122,f17129,f17135,f17142,f17148,f17159,f17170,f17180,f17194,f17197,f17205,f17207,f17209,f17215,f17217,f17223,f17225,f17231,f17233,f17239,f17241,f17253,f17262,f17271,f17279,f17281,f17304,f17306,f17308,f17310,f17323,f17325,f17338,f17340,f17351,f17353,f17364,f17366,f17368,f17377,f17385,f17393,f17395,f17403,f17405,f17413,f17422,f17432,f17441,f17454,f17456,f17464,f17466,f17468,f17470,f17476,f17478,f17480,f17482,f17484,f17486,f17488,f17490,f17495,f17502,f17507,f17510,f17515,f17518,f17531,f17536,f17541,f17546,f17548,f17550,f17555,f17557,f17562,f17568,f17572,f17574,f17576,f17599,f17600,f17601,f17602,f17603,f17604,f17605,f17606,f17607,f17608,f17609,f17610,f17611,f17612,f17613,f17614,f17615,f17616,f17617,f17618,f17619,f17620,f17621,f17622,f17623,f17624,f17625,f17626,f17627,f17628,f17629,f17630,f17631,f17632,f17633,f17634,f17635,f17636,f17637,f17638,f17639,f17640,f17641,f17642,f17643,f17644,f17645,f17646,f17647,f17648,f17649,f17650,f17651,f17652,f17653,f17654,f17655,f17656,f17657,f17658,f17659,f17660,f17661,f17662,f17663,f17664,f17665,f17666,f17667,f17668,f17669,f17670,f17671,f17672,f17673,f17674,f17675,f17676,f17677,f17678,f17679,f17680,f17681,f17682,f17683,f17684,f17685,f17686,f17687,f17688,f17689,f17690,f17691,f17692,f17693,f17694,f17695,f17697,f17698,f17699,f17700,f17701,f17702,f17703,f17704,f17705,f17706,f17707,f17708,f17709,f17710,f17711,f17712,f17713,f17714,f17715,f17716,f17717,f17718,f17719,f17720,f17721,f17722,f17723,f17724,f17725,f17726,f17727,f17728,f17729,f17730,f17731,f17732,f17733,f17734,f17735,f17736,f17737,f17738,f17739,f17740,f17741,f17742,f17743,f17744,f17745,f17746,f17747,f17748,f17749,f18298,f18300,f18299,f18866,f18868,f18867,f16215,f78688])).
fof(f78688,plain,(
  $false | ($spl0 | $spl3 | $spl5)),
  inference(subsumption_resolution,[],[f78687,f830])).
fof(f78687,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl0 | $spl3)),
  inference(forward_demodulation,[],[f78686,f512])).
fof(f512,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_RID')).
fof(f78686,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f77905,f828])).
fof(f77905,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl0),
  inference(superposition,[],[f77216,f529])).
fof(f529,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))))),
  inference(ennf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))))),
  inference(flattening,[],[f184])).
fof(f184,plain,(
  ! [X0] : (~s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : (~s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_RINV')).
fof(f77216,plain,(
  ( ! [X28] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK13),s('type.pair.prod'('type.realax.real','type.realax.real'),X28)))))) ) | $spl0),
  inference(forward_demodulation,[],[f76398,f507])).
fof(f507,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_LZERO')).
fof(f76398,plain,(
  ( ! [X28] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.realax.real','type.realax.real'),X28))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK13),s('type.pair.prod'('type.realax.real','type.realax.real'),X28)))))) ) | $spl0),
  inference(superposition,[],[f702,f825])).
fof(f702,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6,X7,X14] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.pair.prod'('type.realax.real','type.realax.real'),X14))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))),s('type.pair.prod'('type.realax.real','type.realax.real'),X14)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_ASSOC')).
fof(f16215,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | $spl52),
  inference(cnf_transformation,[],[f16215_D])).
fof(f16215_D,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f18867,plain,(
  ( ! [X10,X11,X9] : (sP1('const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9)))),'T_0') | s(bool,'T_0') = s(bool,X9)) ) | $spl60),
  inference(cnf_transformation,[],[f18867_D])).
fof(f18867_D,plain,(
  ( ! [X10,X11,X9] : (sP1('const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9)))),'T_0') | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f18868,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13)))),'T_0') | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13)) ) | $spl10),
  inference(subsumption_resolution,[],[f18850,f530])).
fof(f530,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f350])).
fof(f350,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.bool.EQ_CLAUSES|split|0')).
fof(f18850,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13)))),'T_0') | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | $spl10),
  inference(superposition,[],[f12330,f1030])).
fof(f1030,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f516,f516])).
fof(f516,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.complex.COMPLEX_ENTIRE',aHL_BOOL_CASES)).
fof(f12330,plain,(
  ( ! [X19,X20,X18] : (sP1('const.pair.SND_1'(s('type.pair.prod'(X18,bool),'const.pair.,_2'(s(X18,X19),s(bool,X20)))),'T_0') | p(s(bool,X20))) ) | $spl10),
  inference(superposition,[],[f11521,f688])).
fof(f688,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X8,X12] : ! [X9,X13] : s(X12,X13) = s(X12,'const.pair.SND_1'(s('type.pair.prod'(X8,X12),'const.pair.,_2'(s(X8,X9),s(X12,X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.pair.SND')).
fof(f18866,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9)))),'T_0') | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f18849,f530])).
fof(f18849,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9)))),'T_0') | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f12330,f1030])).
fof(f18299,plain,(
  ( ! [X10,X11,X9] : (sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11)))),'T_0') | s(bool,'T_0') = s(bool,X9)) ) | $spl58),
  inference(cnf_transformation,[],[f18299_D])).
fof(f18299_D,plain,(
  ( ! [X10,X11,X9] : (sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11)))),'T_0') | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f18300,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15)))),'T_0') | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13)) ) | $spl10),
  inference(subsumption_resolution,[],[f18213,f530])).
fof(f18213,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15)))),'T_0') | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | $spl10),
  inference(superposition,[],[f12329,f1030])).
fof(f12329,plain,(
  ( ! [X17,X15,X16] : (sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X16),s(X15,X17)))),'T_0') | p(s(bool,X16))) ) | $spl10),
  inference(superposition,[],[f11521,f687])).
fof(f687,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X12,X8] : ! [X9,X13] : s(X8,X9) = s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X12),'const.pair.,_2'(s(X8,X9),s(X12,X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.pair.FST')).
fof(f18298,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11)))),'T_0') | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(subsumption_resolution,[],[f18212,f530])).
fof(f18212,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11)))),'T_0') | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl10),
  inference(superposition,[],[f12329,f1030])).
fof(f17749,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16606,f478])).
fof(f478,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.complex.COMPLEX_ENTIRE',aHL_TRUTH)).
fof(f16606,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f911,f16215,f615])).
fof(f615,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f313])).
fof(f313,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X15,X16] : ((p(s(bool,X15)) => p(s(bool,X16))) => ((p(s(bool,X16)) => p(s(bool,X15))) => s(bool,X15) = s(bool,X16)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.bool.IMP_ANTISYM_AX')).
fof(f911,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f831,f530])).
fof(f831,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f482,f478])).
fof(f482,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f136])).
fof(f136,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f39])).
fof(f39,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.complex.COMPLEX_ENTIRE','thm.bool.NOT_CLAUSES|split|2')).
fof(f17748,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16612,f478])).
fof(f16612,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f911,f16215,f615])).
fof(f17747,plain,(
  $false | ($spl6 | $spl52)),
  inference(subsumption_resolution,[],[f16621,f16215])).
fof(f16621,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1))))))) ) | ($spl6 | $spl52)),
  inference(unit_resulting_resolution,[],[f16128,f16215,f669])).
fof(f669,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f422])).
fof(f422,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f421])).
fof(f421,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f420])).
fof(f420,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f329])).
fof(f329,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f16128,plain,(
  ( ! [X0,X1] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1)))))) ) | $spl6),
  inference(unit_resulting_resolution,[],[f831,f12313])).
fof(f12313,plain,(
  ( ! [X17,X15,X16] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X15),'const.pair.,_2'(s(bool,X16),s(X15,X17))))) | p(s(bool,X16))) ) | $spl6),
  inference(superposition,[],[f11517,f687])).
fof(f17746,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16622,f478])).
fof(f16622,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f929,f16215,f669])).
fof(f929,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f831,f478,f673])).
fof(f673,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f422])).
fof(f17745,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16623,f16215])).
fof(f16623,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1115,f16215,f669])).
fof(f1115,plain,(
  ( ! [X0] : (sP1('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f831,f1092])).
fof(f1092,plain,(
  ( ! [X12,X11] : (sP1(X12,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X11))) | p(s(bool,X12))) )),
  inference(subsumption_resolution,[],[f1079,f478])).
fof(f1079,plain,(
  ( ! [X12,X11] : (~p(s(bool,'T_0')) | p(s(bool,X12)) | sP1(X12,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real',X11)))) )),
  inference(superposition,[],[f673,f914])).
fof(f914,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f488,f531])).
fof(f531,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f350])).
fof(f488,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X9] : p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X9))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_LE_REFL')).
fof(f17744,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16629,f478])).
fof(f16629,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f924,f16215,f669])).
fof(f924,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f478,f670])).
fof(f670,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f422])).
fof(f17743,plain,(
  $false | ($spl6 | $spl52)),
  inference(subsumption_resolution,[],[f16630,f478])).
fof(f16630,plain,(
  ~p(s(bool,'T_0')) | ($spl6 | $spl52)),
  inference(unit_resulting_resolution,[],[f16128,f16215,f669])).
fof(f17742,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16633,f16215])).
fof(f16633,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1099,f16215,f669])).
fof(f1099,plain,(
  ( ! [X0] : (sP1('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f1091])).
fof(f1091,plain,(
  ( ! [X10,X9] : (sP1('const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X9)),X10) | p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f1078,f478])).
fof(f1078,plain,(
  ( ! [X10,X9] : (~p(s(bool,'T_0')) | p(s(bool,X10)) | sP1('const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X9)),X10)) )),
  inference(superposition,[],[f670,f914])).
fof(f17741,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16646,f478])).
fof(f16646,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1133,f831,f16215,f722])).
fof(f722,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f440])).
fof(f440,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f439])).
fof(f439,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f438])).
fof(f438,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f1133,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f911,f738])).
fof(f738,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f449])).
fof(f449,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f287,f334,f333,f332,f331])).
fof(f331,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f332,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f333,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f287,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X19,X17,X18] : ((p(s(bool,X18)) <=> s(bool,X17) = s(bool,X19)) <=> ((p(s(bool,X18)) | p(s(bool,X17)) | p(s(bool,X19))) & (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X17))) & (p(s(bool,X17)) | ~p(s(bool,X19)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X17)) | ~p(s(bool,X18)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.sat.sat_2|49|_prove_ltb')).
fof(f17740,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16647,f16215])).
fof(f16647,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f3601,f831,f16215,f722])).
fof(f3601,plain,(
  ( ! [X0] : (sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f905,f902,f886,f831,f725])).
fof(f725,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f440])).
fof(f886,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f831,f732])).
fof(f732,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f446])).
fof(f446,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f445])).
fof(f445,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f444])).
fof(f444,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f332])).
fof(f902,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f488,f727])).
fof(f727,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f443])).
fof(f443,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f442])).
fof(f442,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f441])).
fof(f441,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f333])).
fof(f905,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f488,f736])).
fof(f736,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f448])).
fof(f448,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f447])).
fof(f447,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f331])).
fof(f17739,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16649,f16215])).
fof(f16649,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1))))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1646,f831,f16215,f722])).
fof(f1646,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))),X1)) )),
  inference(unit_resulting_resolution,[],[f478,f687,f739])).
fof(f739,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f449])).
fof(f17738,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16650,f16215])).
fof(f16650,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0'))))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1783,f831,f16215,f722])).
fof(f1783,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))),X2)) )),
  inference(unit_resulting_resolution,[],[f478,f688,f739])).
fof(f17737,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16708,f478])).
fof(f16708,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f3531,f831,f16215,f722])).
fof(f3531,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f888,f883,f886,f831,f724])).
fof(f724,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f440])).
fof(f883,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f831,f728])).
fof(f728,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f443])).
fof(f888,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f735])).
fof(f735,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f448])).
fof(f17736,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16709,f478])).
fof(f16709,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1646,f831,f16215,f722])).
fof(f17735,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16710,f478])).
fof(f16710,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1783,f831,f16215,f722])).
fof(f17734,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16711,f16215])).
fof(f16711,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f3532,f831,f16215,f722])).
fof(f3532,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f904,f883,f886,f831,f724])).
fof(f904,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f488,f735])).
fof(f17733,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16716,f16908])).
fof(f16908,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f731])).
fof(f731,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f446])).
fof(f16716,plain,(
  ~sP3('T_0','F_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f882,f16215,f723])).
fof(f723,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f440])).
fof(f882,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f727])).
fof(f1155,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f478,f911,f740])).
fof(f740,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f449])).
fof(f889,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f478,f736])).
fof(f17732,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16726,f16908])).
fof(f16726,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s(X1,X2)))),'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f905,f1649,f902,f16215,f723])).
fof(f1649,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',X0,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))))) )),
  inference(unit_resulting_resolution,[],[f831,f687,f741])).
fof(f741,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f449])).
fof(f17731,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16727,f16908])).
fof(f16727,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))))),'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f905,f1786,f902,f16215,f723])).
fof(f1786,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',X0,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f831,f688,f741])).
fof(f17730,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16728,f16908])).
fof(f16728,plain,(
  ~sP3('T_0','F_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f883,f16215,f723])).
fof(f17729,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16730,f16908])).
fof(f16730,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s(X1,X2)))),'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f905,f1649,f884,f16215,f723])).
fof(f884,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f831,f729])).
fof(f729,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f443])).
fof(f17728,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16732,f16908])).
fof(f16732,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))))),'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f905,f1786,f884,f16215,f723])).
fof(f17727,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16747,f16880])).
fof(f16880,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f727])).
fof(f16747,plain,(
  ~sP4('T_0','F_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f885,f16215,f723])).
fof(f885,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f731])).
fof(f17726,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16766,f16914])).
fof(f16914,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f737])).
fof(f737,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f448])).
fof(f16766,plain,(
  ~sP2('F_0','T_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1155,f882,f885,f16215,f723])).
fof(f17725,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16775,f16914])).
fof(f16775,plain,(
  ~sP2('F_0','T_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1155,f883,f885,f16215,f723])).
fof(f17724,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16790,f16880])).
fof(f16790,plain,(
  ~sP4('T_0','F_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f886,f16215,f723])).
fof(f17723,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16793,f16968])).
fof(f16968,plain,(
  ( ! [X0] : (~sP5(X0,'T_0','F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f911,f16215,f740])).
fof(f16793,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f889,f883,f886,f16215,f723])).
fof(f17722,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16795,f16914])).
fof(f16795,plain,(
  ~sP2('F_0','T_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1155,f883,f886,f16215,f723])).
fof(f17721,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16800,f16968])).
fof(f16800,plain,(
  ( ! [X0] : (sP5(X0,'T_0','F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f889,f882,f886,f16215,f723])).
fof(f17720,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16801,f1155])).
fof(f16801,plain,(
  sP5('T_0','T_0','F_0') | $spl52),
  inference(unit_resulting_resolution,[],[f888,f882,f886,f16215,f723])).
fof(f17719,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16802,f16968])).
fof(f16802,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f904,f882,f886,f16215,f723])).
fof(f17718,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16803,f16914])).
fof(f16803,plain,(
  ~sP2('F_0','T_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1155,f882,f886,f16215,f723])).
fof(f17717,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16831,f16215])).
fof(f16831,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f884,f831,f16215,f726])).
fof(f726,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f443])).
fof(f17716,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16832,f16215])).
fof(f16832,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1))))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1951,f831,f16215,f726])).
fof(f1951,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))),X1)) )),
  inference(unit_resulting_resolution,[],[f1650,f721])).
fof(f721,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f440])).
fof(f1650,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))))) )),
  inference(unit_resulting_resolution,[],[f478,f687,f739])).
fof(f17715,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16833,f16215])).
fof(f16833,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0'))))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f2176,f831,f16215,f726])).
fof(f2176,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))),X2)) )),
  inference(unit_resulting_resolution,[],[f1787,f721])).
fof(f1787,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',X0,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f478,f688,f739])).
fof(f17714,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16851,f16880])).
fof(f16851,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f16215,f726])).
fof(f17713,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16853,f16215])).
fof(f16853,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f883,f831,f16215,f726])).
fof(f17712,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16855,f478])).
fof(f16855,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1951,f831,f16215,f726])).
fof(f17711,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16856,f478])).
fof(f16856,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f2176,f831,f16215,f726])).
fof(f17710,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16866,f16880])).
fof(f16866,plain,(
  ( ! [X0] : (~sP4('T_0',X0,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f478,f16215,f726])).
fof(f17709,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16878,f16880])).
fof(f16878,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f16215,f726])).
fof(f17708,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16883,f16215])).
fof(f16883,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f887,f831,f16215,f730])).
fof(f730,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f446])).
fof(f887,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f831,f733])).
fof(f733,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f446])).
fof(f17707,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16893,f16908])).
fof(f16893,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f16215,f730])).
fof(f17706,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16895,f16215])).
fof(f16895,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f886,f831,f16215,f730])).
fof(f17705,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16900,f16908])).
fof(f16900,plain,(
  ( ! [X0] : (~sP3('T_0',X0,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f478,f16215,f730])).
fof(f17704,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16906,f16908])).
fof(f16906,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f16215,f730])).
fof(f17703,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16919,f687])).
fof(f16919,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1649,f16215,f739])).
fof(f17702,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16920,f688])).
fof(f16920,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1786,f16215,f739])).
fof(f17701,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16921,f914])).
fof(f16921,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1164,f16215,f739])).
fof(f1164,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f831,f914,f741])).
fof(f17700,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16924,f687])).
fof(f16924,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1645,f16215,f739])).
fof(f1645,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,X1),s(X0,X2)))),X1)) )),
  inference(unit_resulting_resolution,[],[f831,f687,f741])).
fof(f17699,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16926,f688])).
fof(f16926,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1782,f16215,f739])).
fof(f1782,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,X2)))),X2)) )),
  inference(unit_resulting_resolution,[],[f831,f688,f741])).
fof(f17698,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16928,f914])).
fof(f16928,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1162,f16215,f739])).
fof(f1162,plain,(
  ( ! [X0] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f831,f914,f741])).
fof(f17697,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17696,f914])).
fof(f17696,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl52),
  inference(forward_demodulation,[],[f16929,f914])).
fof(f16929,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))) != s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f8976,f16215,f739])).
fof(f8976,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f8891,f720])).
fof(f720,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f440])).
fof(f8891,plain,(
  ( ! [X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f5427])).
fof(f5427,plain,(
  ( ! [X6,X8,X7] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f5423,f478])).
fof(f5423,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP3('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7)) )),
  inference(superposition,[],[f1346,f914])).
fof(f1346,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP3(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8)) )),
  inference(subsumption_resolution,[],[f1344,f478])).
fof(f1344,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP3(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8)) )),
  inference(superposition,[],[f730,f914])).
fof(f17695,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16951,f911])).
fof(f16951,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1132,f16215,f740])).
fof(f1132,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f831,f911,f738])).
fof(f17694,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16952,f17442])).
fof(f17442,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f812])).
fof(f812,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | s(bool,X0) != s(bool,X1) | p(s(bool,X1))) )),
  inference(condensation,[],[f811])).
fof(f811,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1) | s(bool,X0) != s(bool,X2)) )),
  inference(condensation,[],[f800])).
fof(f800,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f474])).
fof(f474,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f473])).
fof(f473,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f326])).
fof(f326,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f325])).
fof(f325,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f294])).
fof(f294,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f132])).
fof(f132,axiom,(
  ! [X9,X25,X13,X26] : ((s(bool,X9) = s(bool,X25) & (p(s(bool,X25)) => s(bool,X13) = s(bool,X26))) => ((p(s(bool,X9)) => p(s(bool,X13))) <=> (p(s(bool,X25)) => p(s(bool,X26)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.bool.IMP_CONG')).
fof(f16952,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f3518,f16215,f740])).
fof(f3518,plain,(
  ( ! [X0] : (sP5('F_0','F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f906,f884,f887,f831,f724])).
fof(f906,plain,(
  ( ! [X2,X0,X1] : (sP2('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f488,f737])).
fof(f17693,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16953,f911])).
fof(f16953,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1133,f16215,f740])).
fof(f17692,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16954,f17442])).
fof(f16954,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f3601,f16215,f740])).
fof(f17691,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16976,f17121])).
fof(f17121,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f16215,f16215,f755])).
fof(f755,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f456])).
fof(f456,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f289,f337,f336])).
fof(f336,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f337,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f289,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f101])).
fof(f101,axiom,(
  ! [X19,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) & p(s(bool,X19)))) <=> ((p(s(bool,X18)) | ~p(s(bool,X17)) | ~p(s(bool,X19))) & (p(s(bool,X17)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X18)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.sat.sat_23_prove_ltb')).
fof(f16976,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f743])).
fof(f743,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f451])).
fof(f451,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f450])).
fof(f450,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f337])).
fof(f17690,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16982,f17121])).
fof(f16982,plain,(
  ( ! [X0,X1] : (~sP7(X0,'F_0',X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f744])).
fof(f744,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f452])).
fof(f17689,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16985,f16215])).
fof(f16985,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f893,f942,f16215,f745])).
fof(f745,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f452])).
fof(f942,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f478,f831,f744])).
fof(f893,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f831,f752])).
fof(f752,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f454])).
fof(f454,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f453])).
fof(f453,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f336])).
fof(f17688,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16986,f16215])).
fof(f16986,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f891,f942,f16215,f745])).
fof(f891,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f750])).
fof(f750,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f454])).
fof(f17687,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16991,f17053])).
fof(f17053,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f750])).
fof(f16991,plain,(
  ~sP6('F_0','T_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f942,f478,f16215,f745])).
fof(f17686,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f16998,f17053])).
fof(f16998,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f745])).
fof(f941,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f478,f831,f743])).
fof(f17685,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17000,f17053])).
fof(f17000,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f16215,f16215,f745])).
fof(f17684,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17005,f16215])).
fof(f17005,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f892,f941,f16215,f745])).
fof(f892,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f831,f751])).
fof(f751,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f454])).
fof(f17683,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17006,f16215])).
fof(f17006,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f891,f941,f16215,f745])).
fof(f17682,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17012,f17053])).
fof(f17012,plain,(
  ~sP6('T_0','F_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f941,f478,f16215,f745])).
fof(f17681,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17020,f17053])).
fof(f17020,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','T_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f745])).
fof(f17680,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17022,f17053])).
fof(f17022,plain,(
  ( ! [X0] : (~sP6('F_0',X0,'T_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f16215,f16215,f745])).
fof(f17679,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17028,f16215])).
fof(f17028,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f893,f831,f16215,f749])).
fof(f749,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f454])).
fof(f17678,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17038,f17053])).
fof(f17038,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f16215,f749])).
fof(f17677,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17040,f16215])).
fof(f17040,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f892,f831,f16215,f749])).
fof(f17676,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17045,f17053])).
fof(f17045,plain,(
  ( ! [X0] : (~sP6('T_0',X0,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f478,f16215,f749])).
fof(f17675,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17051,f17053])).
fof(f17051,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f16215,f749])).
fof(f17674,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17056,f16215])).
fof(f17056,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f478,f16215,f755])).
fof(f17673,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17058,f16215])).
fof(f17058,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f16215,f16215,f755])).
fof(f17672,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17059,f478])).
fof(f17059,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1384,f16215,f16215,f755])).
fof(f1384,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1375,f742])).
fof(f742,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f452])).
fof(f1375,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f478,f478,f749])).
fof(f17671,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17060,f16215])).
fof(f17060,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f5449,f16215,f16215,f755])).
fof(f5449,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f5428,f742])).
fof(f5428,plain,(
  ( ! [X0] : (~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f478,f831,f1380])).
fof(f1380,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP6(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8)) )),
  inference(subsumption_resolution,[],[f1378,f478])).
fof(f1378,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP6(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8)) )),
  inference(superposition,[],[f749,f914])).
fof(f17670,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17061,f16215])).
fof(f17061,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f9782,f16215,f16215,f755])).
fof(f9782,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f9376,f742])).
fof(f9376,plain,(
  ( ! [X0,X1] : (~sP6('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f5438])).
fof(f5438,plain,(
  ( ! [X6,X8,X7] : (~sP6('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f5434,f478])).
fof(f5434,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP6('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7)) )),
  inference(superposition,[],[f1380,f914])).
fof(f17669,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17062,f478])).
fof(f17062,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f478,f16215,f755])).
fof(f17668,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17068,f478])).
fof(f17068,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f755])).
fof(f17667,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17069,f478])).
fof(f17069,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f16215,f16215,f755])).
fof(f17666,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17077,f16215])).
fof(f17077,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f478,f16215,f755])).
fof(f17665,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17079,f16215])).
fof(f17079,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f755])).
fof(f17664,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17081,f478])).
fof(f17081,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1384,f16215,f16215,f755])).
fof(f17663,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17082,f478])).
fof(f17082,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f5449,f16215,f16215,f755])).
fof(f17662,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17083,f16215])).
fof(f17083,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f9782,f16215,f16215,f755])).
fof(f17661,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17084,f478])).
fof(f17084,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f478,f16215,f755])).
fof(f17660,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17091,f478])).
fof(f17091,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f755])).
fof(f17659,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17092,f478])).
fof(f17092,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f16215,f16215,f755])).
fof(f17658,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17101,f478])).
fof(f17101,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1384,f478,f16215,f755])).
fof(f17657,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17102,f16215])).
fof(f17102,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f755])).
fof(f17656,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17104,f478])).
fof(f17104,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1384,f16215,f16215,f755])).
fof(f17655,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17105,f478])).
fof(f17105,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f5449,f16215,f16215,f755])).
fof(f17654,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17106,f16215])).
fof(f17106,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f9782,f16215,f16215,f755])).
fof(f17653,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17108,f478])).
fof(f17108,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1384,f478,f16215,f755])).
fof(f17652,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17109,f16215])).
fof(f17109,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f5449,f478,f16215,f755])).
fof(f17651,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17115,f16215])).
fof(f17115,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f942,f16215,f16215,f755])).
fof(f17650,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17116,f478])).
fof(f17116,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1384,f16215,f16215,f755])).
fof(f17649,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17117,f16215])).
fof(f17117,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f5449,f16215,f16215,f755])).
fof(f17648,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17118,f16215])).
fof(f17118,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f9782,f16215,f16215,f755])).
fof(f17647,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17124,f16215])).
fof(f17124,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f951,f831,f16215,f758])).
fof(f758,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f456])).
fof(f951,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f831,f831,f754])).
fof(f754,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f456])).
fof(f17646,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17136,f16215])).
fof(f17136,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f945,f831,f16215,f758])).
fof(f945,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f831,f753])).
fof(f753,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f456])).
fof(f17645,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17158,f17483])).
fof(f17483,plain,(
  ( ! [X0,X1] : (sP9('F_0',X0,X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f896,f16215,f814])).
fof(f814,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f762])).
fof(f762,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f458])).
fof(f458,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f457])).
fof(f457,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f340])).
fof(f340,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f896,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f831,f769])).
fof(f769,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f462])).
fof(f462,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f461])).
fof(f461,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f460])).
fof(f460,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f17158,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f759])).
fof(f759,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f459])).
fof(f17644,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17169,f17483])).
fof(f17169,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f760])).
fof(f760,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f459])).
fof(f17643,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17179,f17208])).
fof(f17208,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f768])).
fof(f768,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f462])).
fof(f17179,plain,(
  ~sP8('F_0','T_0','F_0') | $spl52),
  inference(unit_resulting_resolution,[],[f956,f831,f16215,f763])).
fof(f763,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f459])).
fof(f956,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f478,f831,f760])).
fof(f17642,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17190,f17208])).
fof(f17190,plain,(
  ~sP8('F_0','F_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f952,f831,f16215,f764])).
fof(f764,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f459])).
fof(f952,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f831,f759])).
fof(f17641,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17204,f17208])).
fof(f17204,plain,(
  ( ! [X0] : (~sP8(X0,'F_0','F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f831,f16215,f766])).
fof(f766,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f462])).
fof(f17640,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17212,f478])).
fof(f17212,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1441,f16215,f771])).
fof(f771,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f463])).
fof(f463,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f341])).
fof(f341,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f291,f340,f339])).
fof(f291,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X19,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) | p(s(bool,X19)))) <=> ((p(s(bool,X18)) | ~p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X19))) & (p(s(bool,X17)) | p(s(bool,X19)) | ~p(s(bool,X18)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.sat.sat_25_prove_ltb')).
fof(f1441,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1428,f761])).
fof(f761,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f459])).
fof(f1428,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f831,f478,f766])).
fof(f17639,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17213,f16215])).
fof(f17213,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f5599,f16215,f771])).
fof(f5599,plain,(
  ( ! [X0] : (~sP9('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f5524,f761])).
fof(f5524,plain,(
  ( ! [X0] : (~sP8('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f831,f1434])).
fof(f1434,plain,(
  ( ! [X6,X8,X7] : (~sP8('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X7,X8) | p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f1431,f478])).
fof(f1431,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | p(s(bool,X8)) | ~sP8('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X7,X8)) )),
  inference(superposition,[],[f766,f914])).
fof(f17638,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17218,f478])).
fof(f17218,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f952,f16215,f771])).
fof(f17637,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17219,f16215])).
fof(f17219,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f956,f16215,f771])).
fof(f17636,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17228,f478])).
fof(f17228,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1441,f16215,f772])).
fof(f772,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f464])).
fof(f17635,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17229,f16215])).
fof(f17229,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f5599,f16215,f772])).
fof(f17634,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17234,f16215])).
fof(f17234,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f952,f16215,f772])).
fof(f17633,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17235,f478])).
fof(f17235,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f956,f16215,f772])).
fof(f17632,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17261,f17469])).
fof(f17469,plain,(
  ( ! [X0,X1] : (sP11('F_0',X0,X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f899,f16215,f813])).
fof(f813,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f779])).
fof(f779,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((~sP11(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP10(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP10(X0,X1,X2) | sP11(X0,X1,X2)))),
  inference(rectify,[],[f466])).
fof(f466,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1) | sP11(X2,X0,X1)))),
  inference(flattening,[],[f465])).
fof(f465,plain,(
  ! [X2,X0,X1] : ((~sP11(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP10(X2,X0,X1)) | sP11(X2,X0,X1)))),
  inference(nnf_transformation,[],[f343])).
fof(f343,plain,(
  ! [X2,X0,X1] : (sP11(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP10(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f899,plain,(
  ( ! [X0,X1] : (sP10('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f831,f786])).
fof(f786,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f469])).
fof(f469,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X2,X0,X1] : ((~sP10(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP10(X2,X0,X1)))),
  inference(nnf_transformation,[],[f342])).
fof(f342,plain,(
  ! [X2,X0,X1] : (sP10(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f17261,plain,(
  ( ! [X0,X1] : (~sP11('F_0',X0,X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215,f777])).
fof(f777,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f467])).
fof(f17631,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17264,f16215])).
fof(f17264,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f899,f970,f16215,f780])).
fof(f780,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f467])).
fof(f970,plain,(
  ( ! [X0] : (~sP11('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f478,f831,f777])).
fof(f17630,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17265,f16215])).
fof(f17265,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f898,f970,f16215,f780])).
fof(f898,plain,(
  ( ! [X0,X1] : (sP10(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f478,f785])).
fof(f785,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f470])).
fof(f17629,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17270,f17367])).
fof(f17367,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f785])).
fof(f17270,plain,(
  ~sP10('F_0','T_0','T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f970,f478,f16215,f780])).
fof(f17628,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17278,f17367])).
fof(f17278,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f969,f16215,f16215,f780])).
fof(f969,plain,(
  ( ! [X0] : (~sP11('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f831,f776])).
fof(f776,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f467])).
fof(f17627,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17280,f17367])).
fof(f17280,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f970,f16215,f16215,f780])).
fof(f17626,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17297,f17367])).
fof(f17297,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f969,f16215,f16215,f780])).
fof(f17625,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17298,f17367])).
fof(f17298,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f970,f16215,f16215,f780])).
fof(f17624,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17317,f17367])).
fof(f17317,plain,(
  ~sP10('F_0','F_0','F_0') | $spl52),
  inference(unit_resulting_resolution,[],[f969,f831,f16215,f781])).
fof(f781,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f467])).
fof(f17623,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17332,f17367])).
fof(f17332,plain,(
  ~sP10('F_0','F_0','F_0') | $spl52),
  inference(unit_resulting_resolution,[],[f969,f831,f16215,f782])).
fof(f782,plain,(
  ( ! [X2,X0,X1] : (sP11(X0,X1,X2) | ~sP10(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f467])).
fof(f17622,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17349,f16215])).
fof(f17349,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f899,f831,f16215,f783])).
fof(f783,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f470])).
fof(f17621,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17352,f17367])).
fof(f17352,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f831,f16215,f783])).
fof(f17620,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17361,f16215])).
fof(f17361,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f897,f831,f16215,f783])).
fof(f897,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f784])).
fof(f784,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f470])).
fof(f17619,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17363,f17367])).
fof(f17363,plain,(
  ( ! [X0] : (~sP10('T_0','F_0',X0)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f478,f831,f16215,f783])).
fof(f17618,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17365,f17367])).
fof(f17365,plain,(
  ( ! [X0,X1] : (~sP10(X0,'F_0',X1)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f831,f16215,f783])).
fof(f17617,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17388,f478])).
fof(f17388,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1488,f16215,f789])).
fof(f789,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f472])).
fof(f472,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f471])).
fof(f471,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP11(X2,X0,X1)) & (~sP11(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f344])).
fof(f344,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP11(X2,X0,X1))),
  inference(definition_folding,[],[f324,f343,f342])).
fof(f324,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X19,X17,X18] : ((p(s(bool,X18)) <=> (p(s(bool,X17)) => p(s(bool,X19)))) <=> ((p(s(bool,X18)) | p(s(bool,X17))) & (p(s(bool,X18)) | ~p(s(bool,X19))) & (~p(s(bool,X17)) | p(s(bool,X19)) | ~p(s(bool,X18)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.sat.sat_27_prove_ltb')).
fof(f1488,plain,(
  ~sP11('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1476,f778])).
fof(f778,plain,(
  ( ! [X2,X0,X1] : (~sP11(X0,X1,X2) | sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f467])).
fof(f1476,plain,(
  ~sP10('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f478,f831,f478,f783])).
fof(f17616,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17389,f478])).
fof(f17389,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f6007,f16215,f789])).
fof(f6007,plain,(
  ( ! [X0] : (~sP11('T_0','F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f5730,f778])).
fof(f5730,plain,(
  ( ! [X0] : (~sP10('T_0','F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f478,f831,f1481])).
fof(f1481,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X8)) | p(s(bool,X7)) | ~sP10(X8,X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(subsumption_resolution,[],[f1479,f478])).
fof(f1479,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~p(s(bool,X8)) | ~sP10(X8,X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f783,f914])).
fof(f17615,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17390,f16215])).
fof(f17390,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f6231,f16215,f789])).
fof(f6231,plain,(
  ( ! [X0] : (~sP11('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f831,f1505])).
fof(f1505,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8,X7)) )),
  inference(subsumption_resolution,[],[f1503,f478])).
fof(f1503,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8,X7)) )),
  inference(superposition,[],[f790,f914])).
fof(f790,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f472])).
fof(f17614,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17391,f16215])).
fof(f17391,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f10256,f16215,f789])).
fof(f10256,plain,(
  ( ! [X0,X1] : (~sP11('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f10096,f778])).
fof(f10096,plain,(
  ( ! [X0,X1] : (~sP10('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f831,f5740])).
fof(f5740,plain,(
  ( ! [X6,X8,X7] : (~sP10('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X7,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8))) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f5736,f478])).
fof(f5736,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP10('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X7,'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)))) )),
  inference(superposition,[],[f1481,f914])).
fof(f17613,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17396,f16215])).
fof(f17396,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f969,f16215,f789])).
fof(f17612,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17397,f478])).
fof(f17397,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f970,f16215,f789])).
fof(f17611,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17406,f478])).
fof(f17406,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f972,f831,f16215,f790])).
fof(f972,plain,(
  ( ! [X0] : (sP11('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f478,f788])).
fof(f788,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP11(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f472])).
fof(f17610,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17407,f16215])).
fof(f17407,plain,(
  ( ! [X0] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f1225,f831,f16215,f790])).
fof(f1225,plain,(
  ( ! [X0,X1] : (sP11('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f1095])).
fof(f1095,plain,(
  ( ! [X21,X19,X20] : (sP11('const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real',X19)),X21,X20) | p(s(bool,X20))) )),
  inference(subsumption_resolution,[],[f1082,f478])).
fof(f1082,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | p(s(bool,X20)) | sP11('const.real.real_lte_2'(s('type.realax.real',X19),s('type.realax.real',X19)),X21,X20)) )),
  inference(superposition,[],[f788,f914])).
fof(f17609,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17425,f478])).
fof(f17425,plain,(
  ~p(s(bool,'T_0')) | $spl52),
  inference(unit_resulting_resolution,[],[f1468,f831,f16215,f790])).
fof(f1468,plain,(
  sP11('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f899,f478,f831,f782])).
fof(f17608,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17457,f897])).
fof(f17457,plain,(
  ( ! [X0] : (~sP10('F_0',X0,'F_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f969,f16215,f813])).
fof(f17607,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17458,f898])).
fof(f17458,plain,(
  ( ! [X0] : (~sP10('F_0','T_0',X0)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f970,f16215,f813])).
fof(f17606,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17471,f894])).
fof(f894,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f767])).
fof(f767,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f462])).
fof(f17471,plain,(
  ( ! [X0] : (~sP8('F_0',X0,'T_0')) ) | $spl52),
  inference(unit_resulting_resolution,[],[f952,f16215,f814])).
fof(f17605,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17472,f895])).
fof(f895,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f478,f768])).
fof(f17472,plain,(
  ( ! [X0] : (~sP8('F_0','T_0',X0)) ) | $spl52),
  inference(unit_resulting_resolution,[],[f956,f16215,f814])).
fof(f17604,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17565,f16215])).
fof(f17565,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl52),
  inference(unit_resulting_resolution,[],[f486,f16215,f646])).
fof(f646,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(flattening,[],[f410])).
fof(f410,plain,(
  ! [X0,X1] : (((~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) | s('type.realax.real',X0) = s('type.realax.real',X1)) & (s('type.realax.real',X0) != s('type.realax.real',X1) | (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))))))),
  inference(nnf_transformation,[],[f250])).
fof(f250,plain,(
  ! [X0,X1] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) <=> s('type.realax.real',X0) = s('type.realax.real',X1))),
  inference(rectify,[],[f128])).
fof(f128,axiom,(
  ! [X9,X13] : ((p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real',X13)))) & p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X9))))) <=> s('type.realax.real',X9) = s('type.realax.real',X13))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_LE_ANTISYM')).
fof(f486,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f108])).
fof(f108,axiom,(
  ~s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_|49||48|')).
fof(f17603,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17566,f16215])).
fof(f17566,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl52),
  inference(unit_resulting_resolution,[],[f486,f16215,f646])).
fof(f17602,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17569,f16215])).
fof(f17569,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl52),
  inference(unit_resulting_resolution,[],[f486,f16215,f646])).
fof(f17601,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17570,f16215])).
fof(f17570,plain,(
  ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl52),
  inference(unit_resulting_resolution,[],[f486,f16215,f646])).
fof(f17600,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17573,f16215])).
fof(f17573,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f616])).
fof(f616,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f393])).
fof(f393,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) & (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) | p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))))),
  inference(nnf_transformation,[],[f237])).
fof(f237,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(flattening,[],[f236])).
fof(f236,plain,(
  ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f116])).
fof(f116,axiom,(
  ! [X13,X9] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X9),s('type.realax.real',X13)))) <=> ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X9)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.real_lt')).
fof(f17599,plain,(
  $false | $spl52),
  inference(subsumption_resolution,[],[f17575,f16215])).
fof(f17575,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f616])).
fof(f17576,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f16215,f616])).
fof(f17574,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f16215,f616])).
fof(f17572,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f486,f16215,f16215,f646])).
fof(f17568,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f486,f16215,f16215,f646])).
fof(f17562,plain,(
  $false | ($spl34 | $spl52)),
  inference(unit_resulting_resolution,[],[f972,f831,f16215,f11557])).
fof(f17557,plain,(
  $false | ($spl32 | $spl52)),
  inference(unit_resulting_resolution,[],[f1488,f16215,f11554])).
fof(f17555,plain,(
  $false | ($spl28 | $spl52)),
  inference(unit_resulting_resolution,[],[f899,f831,f16215,f11548])).
fof(f17550,plain,(
  $false | ($spl24 | $spl52)),
  inference(unit_resulting_resolution,[],[f1441,f16215,f11542])).
fof(f17548,plain,(
  $false | ($spl22 | $spl52)),
  inference(unit_resulting_resolution,[],[f1441,f16215,f11539])).
fof(f17546,plain,(
  $false | ($spl18 | $spl52)),
  inference(unit_resulting_resolution,[],[f951,f831,f16215,f11533])).
fof(f17541,plain,(
  $false | ($spl16 | $spl52)),
  inference(unit_resulting_resolution,[],[f893,f831,f16215,f11530])).
fof(f17536,plain,(
  $false | ($spl14 | $spl52)),
  inference(unit_resulting_resolution,[],[f887,f831,f16215,f11527])).
fof(f17531,plain,(
  $false | ($spl12 | $spl52)),
  inference(unit_resulting_resolution,[],[f884,f831,f16215,f11524])).
fof(f17518,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f1225,f16215,f1505])).
fof(f17515,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f899,f16215,f1481])).
fof(f17510,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f951,f16215,f1410])).
fof(f1410,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP7(X8,X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(subsumption_resolution,[],[f1408,f478])).
fof(f1408,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP7(X8,X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f758,f914])).
fof(f17507,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f893,f16215,f1380])).
fof(f17502,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f887,f16215,f1346])).
fof(f17495,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f884,f16215,f1332])).
fof(f1332,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8)) )),
  inference(subsumption_resolution,[],[f1330,f478])).
fof(f1330,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,X7)) | p(s(bool,X8)) | ~sP4(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8)) )),
  inference(superposition,[],[f726,f914])).
fof(f17490,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f6231,f16215,f1096])).
fof(f1096,plain,(
  ( ! [X24,X23,X22] : (sP11('const.real.real_lte_2'(s('type.realax.real',X22),s('type.realax.real',X22)),X23,X24) | ~p(s(bool,X23))) )),
  inference(subsumption_resolution,[],[f1083,f478])).
fof(f1083,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP11('const.real.real_lte_2'(s('type.realax.real',X22),s('type.realax.real',X22)),X23,X24)) )),
  inference(superposition,[],[f789,f914])).
fof(f17488,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f5599,f16215,f1094])).
fof(f1094,plain,(
  ( ! [X17,X18,X16] : (sP9('const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X16)),X17,X18) | ~p(s(bool,X17))) )),
  inference(subsumption_resolution,[],[f1081,f478])).
fof(f1081,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | sP9('const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X16)),X17,X18)) )),
  inference(superposition,[],[f772,f914])).
fof(f17486,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f5599,f16215,f1093])).
fof(f1093,plain,(
  ( ! [X14,X15,X13] : (sP9('const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)),X15,X14) | ~p(s(bool,X14))) )),
  inference(subsumption_resolution,[],[f1080,f478])).
fof(f1080,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | sP9('const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)),X15,X14)) )),
  inference(superposition,[],[f771,f914])).
fof(f17484,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f896,f16215,f814])).
fof(f17482,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f909,f16215,f814])).
fof(f909,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f488,f768])).
fof(f17480,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f956,f895,f16215,f814])).
fof(f17478,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f956,f908,f16215,f814])).
fof(f908,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f488,f767])).
fof(f17476,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f894,f16215,f814])).
fof(f17470,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f899,f16215,f813])).
fof(f17468,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f910,f16215,f813])).
fof(f910,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f488,f785])).
fof(f17466,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f970,f898,f16215,f813])).
fof(f17464,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f897,f16215,f813])).
fof(f17456,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f688,f16215,f812])).
fof(f17454,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f687,f16215,f812])).
fof(f17441,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f972,f831,f16215,f16215,f790])).
fof(f17432,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1468,f831,f478,f16215,f790])).
fof(f17422,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f972,f831,f16215,f16215,f790])).
fof(f17413,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f972,f831,f478,f16215,f790])).
fof(f17405,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f16215,f16215,f789])).
fof(f17403,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f970,f478,f16215,f789])).
fof(f17395,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f16215,f16215,f789])).
fof(f17393,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1488,f478,f16215,f789])).
fof(f17385,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f831,f16215,f788])).
fof(f17377,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f831,f831,f16215,f787])).
fof(f787,plain,(
  ( ! [X2,X0,X1] : (sP11(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f472])).
fof(f17368,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1476,f16215,f785])).
fof(f17366,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f897,f16215,f831,f16215,f783])).
fof(f17364,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f897,f478,f831,f16215,f783])).
fof(f17353,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f897,f16215,f831,f16215,f783])).
fof(f17351,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f899,f478,f831,f16215,f783])).
fof(f17340,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f899,f831,f16215,f782])).
fof(f17338,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f897,f831,f16215,f782])).
fof(f17325,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f899,f831,f16215,f781])).
fof(f17323,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f897,f831,f16215,f781])).
fof(f17310,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f899,f16215,f16215,f780])).
fof(f17308,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f910,f16215,f16215,f780])).
fof(f17306,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f970,f898,f16215,f16215,f780])).
fof(f17304,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f969,f897,f16215,f16215,f780])).
fof(f17281,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f899,f970,f16215,f16215,f780])).
fof(f17279,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f899,f969,f16215,f16215,f780])).
fof(f17271,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f899,f970,f478,f16215,f780])).
fof(f17262,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1468,f831,f16215,f777])).
fof(f17253,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1419,f831,f831,f16215,f773])).
fof(f773,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f464])).
fof(f1419,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f896,f831,f831,f765])).
fof(f765,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f459])).
fof(f17241,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f16215,f16215,f772])).
fof(f17239,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f956,f478,f16215,f772])).
fof(f17233,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f16215,f16215,f772])).
fof(f17231,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1441,f478,f16215,f772])).
fof(f17225,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f16215,f16215,f771])).
fof(f17223,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f478,f16215,f771])).
fof(f17217,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f16215,f16215,f771])).
fof(f17215,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1441,f478,f16215,f771])).
fof(f17209,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1428,f16215,f768])).
fof(f17207,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1428,f16215,f767])).
fof(f17205,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f896,f831,f831,f16215,f766])).
fof(f17197,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f896,f831,f16215,f764])).
fof(f17194,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f952,f894,f831,f16215,f764])).
fof(f17180,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f896,f956,f831,f16215,f763])).
fof(f17170,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1419,f831,f16215,f760])).
fof(f17159,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1419,f831,f16215,f759])).
fof(f17148,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f945,f831,f16215,f16215,f758])).
fof(f17142,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f945,f831,f478,f16215,f758])).
fof(f17135,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f945,f831,f16215,f16215,f758])).
fof(f17129,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f951,f831,f478,f16215,f758])).
fof(f17122,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f16215,f755])).
fof(f17120,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f942,f478,f16215,f16215,f755])).
fof(f17113,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f478,f16215,f755])).
fof(f17111,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1384,f478,f478,f16215,f755])).
fof(f17099,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f16215,f755])).
fof(f17097,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f478,f16215,f16215,f755])).
fof(f17090,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f478,f16215,f755])).
fof(f17088,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f478,f478,f16215,f755])).
fof(f17076,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f16215,f16215,f16215,f755])).
fof(f17074,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f941,f478,f16215,f16215,f755])).
fof(f17067,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f942,f16215,f478,f16215,f755])).
fof(f17065,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f942,f478,f478,f16215,f755])).
fof(f17054,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1375,f16215,f750])).
fof(f17052,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f892,f831,f16215,f16215,f749])).
fof(f17046,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f892,f831,f478,f16215,f749])).
fof(f17039,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f892,f831,f16215,f16215,f749])).
fof(f17033,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f893,f831,f478,f16215,f749])).
fof(f17023,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f893,f942,f16215,f16215,f745])).
fof(f17021,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f892,f941,f16215,f16215,f745])).
fof(f17013,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f892,f941,f478,f16215,f745])).
fof(f17001,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f893,f942,f16215,f16215,f745])).
fof(f16999,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f892,f941,f16215,f16215,f745])).
fof(f16992,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f893,f942,f478,f16215,f745])).
fof(f16983,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f951,f831,f16215,f744])).
fof(f16977,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f945,f831,f16215,f743])).
fof(f16971,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1132,f911,f16215,f740])).
fof(f16969,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1133,f911,f16215,f740])).
fof(f16950,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1786,f688,f16215,f739])).
fof(f16948,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1649,f687,f16215,f739])).
fof(f16944,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f8976,f600,f16215,f739])).
fof(f600,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f121])).
fof(f121,axiom,(
  ! [X9,X13] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X9))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_LE_NEG')).
fof(f16942,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1164,f914,f16215,f739])).
fof(f16940,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1782,f688,f16215,f739])).
fof(f16938,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1645,f687,f16215,f739])).
fof(f16934,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f8976,f600,f16215,f739])).
fof(f16932,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1162,f914,f16215,f739])).
fof(f17598,plain,(
  $false | $spl52),
  inference(trivial_inequality_removal,[],[f16927])).
fof(f16927,plain,(
  ( ! [X0,X1] : (s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),X1))) != s(bool,'const.pair.SND_1'(s('type.pair.prod'(X0,bool),X1)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f9491,f16215,f739])).
fof(f9491,plain,(
  ( ! [X21,X22] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X21,bool),X22)),'const.pair.SND_1'(s('type.pair.prod'(X21,bool),X22)))) )),
  inference(superposition,[],[f1786,f686])).
fof(f686,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))))),
  inference(flattening,[],[f262])).
fof(f262,plain,(
  ! [X0,X1] : ! [X2] : s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))),s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X8,X12] : ! [X9] : s('type.pair.prod'(X8,X12),X9) = s('type.pair.prod'(X8,X12),'const.pair.,_2'(s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X12),X9))),s(X12,'const.pair.SND_1'(s('type.pair.prod'(X8,X12),X9)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.pair.PAIR')).
fof(f17597,plain,(
  $false | $spl52),
  inference(trivial_inequality_removal,[],[f16925])).
fof(f16925,plain,(
  ( ! [X0,X1] : (s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1))) != s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X0),X1)))) ) | $spl52),
  inference(unit_resulting_resolution,[],[f9497,f16215,f739])).
fof(f9497,plain,(
  ( ! [X33,X34] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X33),X34)),'const.pair.FST_1'(s('type.pair.prod'(bool,X33),X34)))) )),
  inference(superposition,[],[f1649,f686])).
fof(f17596,plain,(
  $false | $spl52),
  inference(trivial_inequality_removal,[],[f16923])).
fof(f16923,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1364,f16215,f739])).
fof(f1364,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1351,f719])).
fof(f719,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f440])).
fof(f1351,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f831,f831,f734])).
fof(f734,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f448])).
fof(f17595,plain,(
  $false | $spl52),
  inference(trivial_inequality_removal,[],[f16922])).
fof(f16922,plain,(
  s(bool,'T_0') != s(bool,'T_0') | $spl52),
  inference(unit_resulting_resolution,[],[f1360,f16215,f739])).
fof(f1360,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1341,f720])).
fof(f1341,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f478,f478,f730])).
fof(f16915,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1351,f16215,f737])).
fof(f16913,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1351,f16215,f736])).
fof(f16911,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1351,f16215,f735])).
fof(f16909,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1341,f16215,f731])).
fof(f16907,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f886,f831,f16215,f16215,f730])).
fof(f16901,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f886,f831,f478,f16215,f730])).
fof(f16894,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f886,f831,f16215,f16215,f730])).
fof(f16888,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f887,f831,f478,f16215,f730])).
fof(f16881,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1327,f16215,f727])).
fof(f1327,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f478,f478,f726])).
fof(f16879,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f883,f831,f16215,f16215,f726])).
fof(f16867,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1951,f831,f478,f16215,f726])).
fof(f16852,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f883,f831,f16215,f16215,f726])).
fof(f16840,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f884,f831,f478,f16215,f726])).
fof(f16804,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f882,f886,f16215,f723])).
fof(f16796,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f883,f886,f16215,f723])).
fof(f16776,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f883,f885,f16215,f723])).
fof(f16767,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f889,f1155,f882,f885,f16215,f723])).
fof(f16715,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1133,f16215,f831,f16215,f722])).
fof(f16713,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f3531,f478,f831,f16215,f722])).
fof(f16678,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1133,f831,f16215,f16215,f722])).
fof(f16659,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f1133,f831,f478,f16215,f722])).
fof(f16645,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f920,f831,f16215,f673])).
fof(f920,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f831,f668])).
fof(f668,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f422])).
fof(f16641,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f920,f831,f16215,f670])).
fof(f16637,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f924,f16215,f16215,f669])).
fof(f16635,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f924,f478,f16215,f669])).
fof(f16628,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f924,f16215,f16215,f669])).
fof(f16626,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f929,f478,f16215,f669])).
fof(f16619,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f936,f16215,f655])).
fof(f655,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f415])).
fof(f415,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f414])).
fof(f414,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f327])).
fof(f327,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f936,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f831,f831,f835])).
fof(f835,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f653,f831])).
fof(f653,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f415])).
fof(f16617,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f936,f16215,f654])).
fof(f654,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f415])).
fof(f16615,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f911,f16215,f615])).
fof(f16613,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f478,f911,f16215,f615])).
fof(f16609,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f16215,f911,f16215,f615])).
fof(f16607,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f478,f911,f16215,f615])).
fof(f16603,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f687,f16215,f534])).
fof(f534,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f352])).
fof(f352,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.bool.EQ_CLAUSES|split|2')).
fof(f16601,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f911,f16215,f531])).
fof(f16599,plain,(
  $false | $spl52),
  inference(unit_resulting_resolution,[],[f831,f16215])).
fof(f16590,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl56),
  inference(cnf_transformation,[],[f16590_D])).
fof(f16590_D,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9))))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f16591,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('T_0','const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13))))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13)) ) | $spl6),
  inference(subsumption_resolution,[],[f16573,f530])).
fof(f16573,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('T_0','const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X13))))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f12314,f1030])).
fof(f12314,plain,(
  ( ! [X19,X20,X18] : (sP1('T_0','const.pair.SND_1'(s('type.pair.prod'(X18,bool),'const.pair.,_2'(s(X18,X19),s(bool,X20))))) | p(s(bool,X20))) ) | $spl6),
  inference(superposition,[],[f11517,f688])).
fof(f16589,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('T_0','const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9))))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f16572,f530])).
fof(f16572,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('T_0','const.pair.SND_1'(s('type.pair.prod'(X10,bool),'const.pair.,_2'(s(X10,X11),s(bool,X9))))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f12314,f1030])).
fof(f16216,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11))))) | s(bool,'T_0') = s(bool,X9)) ) | $spl54),
  inference(cnf_transformation,[],[f16216_D])).
fof(f16216_D,plain,(
  ( ! [X10,X11,X9] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11))))) | s(bool,'T_0') = s(bool,X9)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f16217,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15))))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13)) ) | $spl6),
  inference(subsumption_resolution,[],[f16132,f530])).
fof(f16132,plain,(
  ( ! [X14,X12,X15,X13] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X14),'const.pair.,_2'(s(bool,X13),s(X14,X15))))) | p(s(bool,X12)) | s(bool,'T_0') = s(bool,X13) | s(bool,'T_0') = s(bool,X12)) ) | $spl6),
  inference(superposition,[],[f12313,f1030])).
fof(f16214,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11))))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(subsumption_resolution,[],[f16131,f530])).
fof(f16131,plain,(
  ( ! [X10,X8,X11,X9] : (sP1('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X10),'const.pair.,_2'(s(bool,X9),s(X10,X11))))) | p(s(bool,X8)) | s(bool,'T_0') = s(bool,X8) | s(bool,'T_0') = s(bool,X9)) ) | $spl6),
  inference(superposition,[],[f12313,f1030])).
fof(f11939,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11938,f478])).
fof(f11938,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11937,f11518])).
fof(f11937,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X27))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11790,f478])).
fof(f11790,plain,(
  ( ! [X28,X27] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X27))))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f5911])).
fof(f5911,plain,(
  ( ! [X28,X27] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X27)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X27),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X28))))) )),
  inference(superposition,[],[f616,f506])).
fof(f506,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f97])).
fof(f97,axiom,(
  ! [X9] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X9))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X9),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.POW_2')).
fof(f11936,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11935,f478])).
fof(f11935,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11934,f11518])).
fof(f11934,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X23))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11784,f478])).
fof(f11784,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X23))))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3070])).
fof(f3070,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X23)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X23),s('type.realax.real',X24))),s('type.realax.real',X25))))) )),
  inference(superposition,[],[f616,f595])).
fof(f595,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X9,X13] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X9)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.realax.REAL_ADD_SYM')).
fof(f11933,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11932,f478])).
fof(f11932,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11931,f11518])).
fof(f11931,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X30),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11781,f478])).
fof(f11781,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X30),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f5799])).
fof(f5799,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X31),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X30),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30))),s('type.realax.real',X31))))) )),
  inference(superposition,[],[f616,f506])).
fof(f11930,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11929,f478])).
fof(f11929,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11928,f11518])).
fof(f11928,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11776,f478])).
fof(f11776,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11))))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3065])).
fof(f3065,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X12))),s('type.realax.real',X13))))) )),
  inference(superposition,[],[f616,f596])).
fof(f596,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0)))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X9,X13] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X9)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.realax.REAL_MUL_SYM')).
fof(f11927,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11926,f478])).
fof(f11926,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11925,f11518])).
fof(f11925,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X29))),s('type.realax.real',X30))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11768,f478])).
fof(f11768,plain,(
  ( ! [X30,X29] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X29))),s('type.realax.real',X30))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f5912])).
fof(f5912,plain,(
  ( ! [X30,X29] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X29),s('type.realax.real',X29))),s('type.realax.real',X30)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X29),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f616,f506])).
fof(f11924,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886,f11890,f11897,f11901,f11908,f11912,f11919,f11923])).
fof(f11923,plain,(
  ( ! [X72,X71] : (s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11922,f478])).
fof(f11922,plain,(
  ( ! [X72,X71] : (~p(s(bool,'T_0')) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(forward_demodulation,[],[f11921,f11518])).
fof(f11921,plain,(
  ( ! [X72,X71] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X71))),s('type.realax.real',X72)))) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11765,f478])).
fof(f11765,plain,(
  ( ! [X72,X71] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X71))),s('type.realax.real',X72)))) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9088])).
fof(f11919,plain,(
  ( ! [X30,X31] : (s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11918,f478])).
fof(f11918,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(forward_demodulation,[],[f11917,f11518])).
fof(f11917,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30))),s('type.realax.real',X31)))) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11764,f478])).
fof(f11764,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30))),s('type.realax.real',X31)))) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9110])).
fof(f11912,plain,(
  ( ! [X61,X62,X63] : (s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11911,f478])).
fof(f11911,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(forward_demodulation,[],[f11910,f11518])).
fof(f11910,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X62),s('type.realax.real',X61))),s('type.realax.real',X63)))) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11759,f478])).
fof(f11759,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X62),s('type.realax.real',X61))),s('type.realax.real',X63)))) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9084])).
fof(f11908,plain,(
  ( ! [X61,X62,X63] : (s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11907,f478])).
fof(f11907,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(forward_demodulation,[],[f11906,f11518])).
fof(f11906,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62))),s('type.realax.real',X63)))) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11758,f478])).
fof(f11758,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62))),s('type.realax.real',X63)))) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9123])).
fof(f11901,plain,(
  ( ! [X30,X31] : (s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11900,f478])).
fof(f11900,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(forward_demodulation,[],[f11899,f11518])).
fof(f11899,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X30),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X31)))) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11754,f478])).
fof(f11754,plain,(
  ( ! [X30,X31] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X30),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X31)))) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9071])).
fof(f11897,plain,(
  ( ! [X72,X71] : (s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11896,f478])).
fof(f11896,plain,(
  ( ! [X72,X71] : (~p(s(bool,'T_0')) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(forward_demodulation,[],[f11895,f11518])).
fof(f11895,plain,(
  ( ! [X72,X71] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X72)))) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11753,f478])).
fof(f11753,plain,(
  ( ! [X72,X71] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X72)))) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9127])).
fof(f11890,plain,(
  ( ! [X33,X34,X32] : (s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11889,f478])).
fof(f11889,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(forward_demodulation,[],[f11888,f11518])).
fof(f11888,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X32))),s('type.realax.real',X34)))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11748,f478])).
fof(f11748,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X32))),s('type.realax.real',X34)))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9072])).
fof(f11886,plain,(
  ( ! [X33,X34,X32] : (s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11885,f478])).
fof(f11885,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(forward_demodulation,[],[f11884,f11518])).
fof(f11884,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33))),s('type.realax.real',X34)))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11747,f478])).
fof(f11747,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33))),s('type.realax.real',X34)))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9111])).
fof(f11882,plain,(
  ( ! [X2,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11881,f478])).
fof(f11881,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))) ) | $spl8),
  inference(forward_demodulation,[],[f11880,f11518])).
fof(f11880,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11742,f478])).
fof(f11742,plain,(
  ( ! [X2,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f9290])).
fof(f11878,plain,(
  ( ! [X12,X10,X11,X9] : (p(s(bool,X11)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X12)),X11,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11741,f478])).
fof(f11741,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X12)),X11,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6237])).
fof(f11740,plain,(
  ( ! [X180,X178,X181,X179] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X180,bool),'const.pair.,_2'(s(X180,X181),s(bool,'T_0')))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X178))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X179)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6189])).
fof(f11739,plain,(
  ( ! [X177,X175,X176,X174] : (sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X174))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X175)))),'const.pair.SND_1'(s('type.pair.prod'(X176,bool),'const.pair.,_2'(s(X176,X177),s(bool,'T_0')))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6188])).
fof(f11738,plain,(
  ( ! [X173,X171,X172,X170] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X172),'const.pair.,_2'(s(bool,'T_0'),s(X172,X173)))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X170))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X171)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6187])).
fof(f11737,plain,(
  ( ! [X167,X169,X166,X168] : (sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X166))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X167)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X168),'const.pair.,_2'(s(bool,'T_0'),s(X168,X169)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6186])).
fof(f11736,plain,(
  ( ! [X154,X152,X155,X153] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X152))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X153)))),'const.pair.SND_1'(s('type.pair.prod'(X154,bool),'const.pair.,_2'(s(X154,X155),s(bool,'T_0')))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6183])).
fof(f11735,plain,(
  ( ! [X151,X149,X150,X148] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X148))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X149)))),'const.pair.SND_1'(s('type.pair.prod'(X150,bool),'const.pair.,_2'(s(X150,X151),s(bool,'T_0')))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6182])).
fof(f11734,plain,(
  ( ! [X146,X144,X147,X145] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X146,bool),'const.pair.,_2'(s(X146,X147),s(bool,'T_0')))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X144))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X145)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6181])).
fof(f11733,plain,(
  ( ! [X142,X140,X143,X141] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X142,bool),'const.pair.,_2'(s(X142,X143),s(bool,'T_0')))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X140))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X141)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6180])).
fof(f11732,plain,(
  ( ! [X127,X128,X126,X129] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X126))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X127)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X128),'const.pair.,_2'(s(bool,'T_0'),s(X128,X129)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6177])).
fof(f11731,plain,(
  ( ! [X125,X123,X124,X122] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X122))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X123)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,'T_0'),s(X124,X125)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6176])).
fof(f11730,plain,(
  ( ! [X121,X118,X120,X119] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X120),'const.pair.,_2'(s(bool,'T_0'),s(X120,X121)))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X118))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X119)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6175])).
fof(f11729,plain,(
  ( ! [X116,X114,X117,X115] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X116),'const.pair.,_2'(s(bool,'T_0'),s(X116,X117)))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X114))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X115)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6174])).
fof(f11864,plain,(
  ( ! [X111,X112,X110,X113] : (p(s(bool,X112)) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X110))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X111)))),X112,'const.real.real_lte_2'(s('type.realax.real',X113),s('type.realax.real',X113)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11728,f478])).
fof(f11728,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'T_0')) | p(s(bool,X112)) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X110))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X111)))),X112,'const.real.real_lte_2'(s('type.realax.real',X113),s('type.realax.real',X113)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6173])).
fof(f11862,plain,(
  ( ! [X109,X107,X108,X106] : (p(s(bool,X108)) | ~sP7(X108,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X106))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X107)))),'const.real.real_lte_2'(s('type.realax.real',X109),s('type.realax.real',X109)))) ) | $spl8),
  inference(subsumption_resolution,[],[f11727,f478])).
fof(f11727,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,X108)) | ~sP7(X108,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X106))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X107)))),'const.real.real_lte_2'(s('type.realax.real',X109),s('type.realax.real',X109)))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6172])).
fof(f11860,plain,(
  ( ! [X103,X105,X102,X104] : (p(s(bool,X104)) | ~sP6('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X102))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X103)))),'const.real.real_lte_2'(s('type.realax.real',X105),s('type.realax.real',X105)),X104)) ) | $spl8),
  inference(subsumption_resolution,[],[f11726,f478])).
fof(f11726,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,X104)) | ~sP6('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X102))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X103)))),'const.real.real_lte_2'(s('type.realax.real',X105),s('type.realax.real',X105)),X104)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6171])).
fof(f11858,plain,(
  ( ! [X101,X99,X100,X98] : (p(s(bool,X100)) | ~sP3('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X98))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X99)))),'const.real.real_lte_2'(s('type.realax.real',X101),s('type.realax.real',X101)),X100)) ) | $spl8),
  inference(subsumption_resolution,[],[f11725,f478])).
fof(f11725,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X100)) | ~sP3('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X98))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X99)))),'const.real.real_lte_2'(s('type.realax.real',X101),s('type.realax.real',X101)),X100)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6170])).
fof(f11856,plain,(
  ( ! [X94,X97,X95,X96] : (p(s(bool,X96)) | ~sP4('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X94))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X95)))),'const.real.real_lte_2'(s('type.realax.real',X97),s('type.realax.real',X97)),X96)) ) | $spl8),
  inference(subsumption_resolution,[],[f11724,f478])).
fof(f11724,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'T_0')) | p(s(bool,X96)) | ~sP4('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X94))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X95)))),'const.real.real_lte_2'(s('type.realax.real',X97),s('type.realax.real',X97)),X96)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6169])).
fof(f11854,plain,(
  ( ! [X90,X88,X87,X89] : (~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X87))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X88)))),X89,X90) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X87))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X88)))),X89,X90)) ) | $spl8),
  inference(subsumption_resolution,[],[f11723,f478])).
fof(f11723,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'T_0')) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X87))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X88)))),X89,X90) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X87))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X88)))),X89,X90)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6167])).
fof(f11852,plain,(
  ( ! [X85,X83,X86,X84] : (~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X83))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X84)))),X85,X86) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X83))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X84)))),X85,X86)) ) | $spl8),
  inference(subsumption_resolution,[],[f11722,f478])).
fof(f11722,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X83))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X84)))),X85,X86) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X83))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X84)))),X85,X86)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6166])).
fof(f11850,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,X78)) | p(s(bool,X79)) | ~sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X76))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X77)))),X79,X78)) ) | $spl8),
  inference(subsumption_resolution,[],[f11721,f478])).
fof(f11721,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'T_0')) | ~p(s(bool,X78)) | p(s(bool,X79)) | ~sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X76))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X77)))),X79,X78)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6164])).
fof(f11848,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,X74)) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X72))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X73)))),X74,X75)) ) | $spl8),
  inference(subsumption_resolution,[],[f11720,f478])).
fof(f11720,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X74)) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X72))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X73)))),X74,X75)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6163])).
fof(f11846,plain,(
  ( ! [X70,X68,X71,X69] : (p(s(bool,X70)) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X68))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X69)))),X71,X70)) ) | $spl8),
  inference(subsumption_resolution,[],[f11719,f478])).
fof(f11719,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X70)) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X68))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X69)))),X71,X70)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6162])).
fof(f11844,plain,(
  ( ! [X66,X64,X67,X65] : (p(s(bool,X66)) | ~p(s(bool,X67)) | ~sP10(X67,X66,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X64))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X65)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11718,f478])).
fof(f11718,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'T_0')) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~sP10(X67,X66,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X64))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X65)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6161])).
fof(f11842,plain,(
  ( ! [X61,X62,X60,X63] : (p(s(bool,X62)) | p(s(bool,X63)) | ~sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X60))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X61)))),X63,X62)) ) | $spl8),
  inference(subsumption_resolution,[],[f11717,f478])).
fof(f11717,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'T_0')) | p(s(bool,X62)) | p(s(bool,X63)) | ~sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X60))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X61)))),X63,X62)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6160])).
fof(f11840,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X58)) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X56))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X57)))),X58,X59)) ) | $spl8),
  inference(subsumption_resolution,[],[f11716,f478])).
fof(f11716,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'T_0')) | ~p(s(bool,X58)) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X56))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X57)))),X58,X59)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6159])).
fof(f11838,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X54)) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X52))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X53)))),X55,X54)) ) | $spl8),
  inference(subsumption_resolution,[],[f11715,f478])).
fof(f11715,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X54)) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X52))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X53)))),X55,X54)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6158])).
fof(f11836,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,X50)) | p(s(bool,X51)) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X48))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X49)))),X50,X51)) ) | $spl8),
  inference(subsumption_resolution,[],[f11714,f478])).
fof(f11714,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,X50)) | p(s(bool,X51)) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X48))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X49)))),X50,X51)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6157])).
fof(f11834,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X46)) | p(s(bool,X47)) | ~sP7(X47,X46,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X44))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X45)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11713,f478])).
fof(f11713,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP7(X47,X46,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X44))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X45)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6156])).
fof(f11832,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,X42)) | p(s(bool,X43)) | ~sP6(X42,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X40))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41)))),X43)) ) | $spl8),
  inference(subsumption_resolution,[],[f11712,f478])).
fof(f11712,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP6(X42,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X40))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41)))),X43)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6155])).
fof(f11830,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,X38)) | p(s(bool,X39)) | ~sP3(X38,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X36))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X37)))),X39)) ) | $spl8),
  inference(subsumption_resolution,[],[f11711,f478])).
fof(f11711,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP3(X38,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X36))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X37)))),X39)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6154])).
fof(f11828,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,X34)) | p(s(bool,X35)) | ~sP4(X34,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X32))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33)))),X35)) ) | $spl8),
  inference(subsumption_resolution,[],[f11710,f478])).
fof(f11710,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP4(X34,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X32))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33)))),X35)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6153])).
fof(f11826,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X31)) | sP1(X31,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11709,f478])).
fof(f11709,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'T_0')) | p(s(bool,X31)) | sP1(X31,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6152])).
fof(f11824,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X28)) | sP1('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X27)))),X28)) ) | $spl8),
  inference(subsumption_resolution,[],[f11708,f478])).
fof(f11708,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'T_0')) | p(s(bool,X28)) | sP1('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X27)))),X28)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6151])).
fof(f11822,plain,(
  ( ! [X24,X23,X25] : (sP0('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X23))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24)))),X25)) ) | $spl8),
  inference(subsumption_resolution,[],[f11707,f478])).
fof(f11707,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | sP0('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X23))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24)))),X25)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6150])).
fof(f11820,plain,(
  ( ! [X21,X22,X20] : (sP0(X22,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X21)))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11706,f478])).
fof(f11706,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'T_0')) | sP0(X22,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X21)))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f6149])).
fof(f11818,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f11817,f478])).
fof(f11817,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl8),
  inference(forward_demodulation,[],[f11816,f11518])).
fof(f11816,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl8),
  inference(subsumption_resolution,[],[f11703,f478])).
fof(f11703,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s('type.realax.real',X0) = s('type.realax.real',X1)) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f646])).
fof(f11680,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0')))),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3599])).
fof(f11679,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1)))),'F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3598])).
fof(f11678,plain,(
  ( ! [X0,X1] : (~sP2('const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0')))),'F_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3521])).
fof(f11677,plain,(
  ( ! [X0,X1] : (~sP2('const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1)))),'F_0','F_0')) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3520])).
fof(f11800,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f11660,f11518])).
fof(f11660,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f534])).
fof(f11056,plain,(
  ( ! [X350,X352,X351] : (~p(s(bool,happ(s(fun(bool,bool),X350),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X350),s(bool,X351)))) | s(bool,X352) = s(bool,sK15(X350,bool)) | s(bool,'T_0') = s(bool,X352)) )),
  inference(superposition,[],[f590,f1030])).
fof(f11055,plain,(
  ( ! [X349,X347,X348] : (~p(s(bool,happ(s(fun(bool,bool),X347),s(bool,'T_0')))) | p(s(bool,happ(s(fun(bool,bool),X347),s(bool,X348)))) | s(bool,X349) = s(bool,sK14(X347,bool)) | s(bool,'T_0') = s(bool,X349)) )),
  inference(superposition,[],[f589,f1030])).
fof(f11054,plain,(
  ( ! [X344,X346,X345] : (s('type.pair.prod'(X344,bool),X345) = s('type.pair.prod'(X344,bool),'const.pair.,_2'(s(X344,'const.pair.FST_1'(s('type.pair.prod'(X344,bool),X345))),s(bool,'T_0'))) | s(bool,X346) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X344,bool),X345))) | s(bool,'T_0') = s(bool,X346)) )),
  inference(superposition,[],[f686,f1030])).
fof(f11052,plain,(
  ( ! [X337,X339,X338] : (s('type.pair.prod'(bool,X337),X338) = s('type.pair.prod'(bool,X337),'const.pair.,_2'(s(bool,'T_0'),s(X337,'const.pair.SND_1'(s('type.pair.prod'(bool,X337),X338))))) | s(bool,X339) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X337),X338))) | s(bool,'T_0') = s(bool,X339)) )),
  inference(superposition,[],[f686,f1030])).
fof(f11611,plain,(
  ( ! [X251,X252,X250] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X251),s('type.realax.real',X250)))) | s(bool,X252) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X250),s('type.realax.real',X251))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(subsumption_resolution,[],[f11022,f478])).
fof(f11022,plain,(
  ( ! [X251,X252,X250] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X251),s('type.realax.real',X250)))) | s(bool,X252) = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X250),s('type.realax.real',X251))) | s(bool,'T_0') = s(bool,X252)) )),
  inference(superposition,[],[f616,f1030])).
fof(f11610,plain,(
  ( ! [X243,X241,X242] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X242),s('type.realax.real',X241)))) | s('type.realax.real',X241) = s('type.realax.real',X242) | s(bool,X243) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X241),s('type.realax.real',X242))) | s(bool,'T_0') = s(bool,X243)) )),
  inference(subsumption_resolution,[],[f11019,f478])).
fof(f11019,plain,(
  ( ! [X243,X241,X242] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X242),s('type.realax.real',X241)))) | s('type.realax.real',X241) = s('type.realax.real',X242) | s(bool,X243) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X241),s('type.realax.real',X242))) | s(bool,'T_0') = s(bool,X243)) )),
  inference(superposition,[],[f646,f1030])).
fof(f11609,plain,(
  ( ! [X239,X237,X238,X236] : (p(s(bool,happ(s(fun(X236,bool),X237),s(X236,X238)))) | s(bool,X239) = s(bool,happ(s(fun(X236,bool),X237),s(X236,sK15(X237,X236)))) | s(bool,'T_0') = s(bool,X239)) )),
  inference(subsumption_resolution,[],[f11018,f478])).
fof(f11018,plain,(
  ( ! [X239,X237,X238,X236] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X236,bool),X237),s(X236,X238)))) | s(bool,X239) = s(bool,happ(s(fun(X236,bool),X237),s(X236,sK15(X237,X236)))) | s(bool,'T_0') = s(bool,X239)) )),
  inference(superposition,[],[f590,f1030])).
fof(f11608,plain,(
  ( ! [X235,X233,X234,X232] : (p(s(bool,happ(s(fun(X232,bool),X233),s(X232,X234)))) | s(bool,X235) = s(bool,happ(s(fun(X232,bool),X233),s(X232,sK14(X233,X232)))) | s(bool,'T_0') = s(bool,X235)) )),
  inference(subsumption_resolution,[],[f11017,f478])).
fof(f11017,plain,(
  ( ! [X235,X233,X234,X232] : (~p(s(bool,'T_0')) | p(s(bool,happ(s(fun(X232,bool),X233),s(X232,X234)))) | s(bool,X235) = s(bool,happ(s(fun(X232,bool),X233),s(X232,sK14(X233,X232)))) | s(bool,'T_0') = s(bool,X235)) )),
  inference(superposition,[],[f589,f1030])).
fof(f10999,plain,(
  ( ! [X182,X184,X183,X181] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X182,bool),'const.pair.,_2'(s(X182,X183),s(bool,'T_0')))),X181) | s(bool,X181) = s(bool,X184) | s(bool,'T_0') = s(bool,X184)) )),
  inference(superposition,[],[f2176,f1030])).
fof(f10998,plain,(
  ( ! [X177,X180,X178,X179] : (sP4('T_0',X177,'const.pair.SND_1'(s('type.pair.prod'(X178,bool),'const.pair.,_2'(s(X178,X179),s(bool,'T_0'))))) | s(bool,X177) = s(bool,X180) | s(bool,'T_0') = s(bool,X180)) )),
  inference(superposition,[],[f2058,f1030])).
fof(f10997,plain,(
  ( ! [X175,X173,X176,X174] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X174),'const.pair.,_2'(s(bool,'T_0'),s(X174,X175)))),X173) | s(bool,X173) = s(bool,X176) | s(bool,'T_0') = s(bool,X176)) )),
  inference(superposition,[],[f1951,f1030])).
fof(f10996,plain,(
  ( ! [X171,X169,X172,X170] : (sP4('T_0',X169,'const.pair.FST_1'(s('type.pair.prod'(bool,X170),'const.pair.,_2'(s(bool,'T_0'),s(X170,X171))))) | s(bool,X169) = s(bool,X172) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f1908,f1030])).
fof(f10993,plain,(
  ( ! [X158,X156,X157,X155] : (sP5('T_0',X155,'const.pair.SND_1'(s('type.pair.prod'(X156,bool),'const.pair.,_2'(s(X156,X157),s(bool,'T_0'))))) | s(bool,X155) = s(bool,X158) | s(bool,'T_0') = s(bool,X158)) )),
  inference(superposition,[],[f1787,f1030])).
fof(f10992,plain,(
  ( ! [X154,X152,X151,X153] : (~sP5('F_0',X151,'const.pair.SND_1'(s('type.pair.prod'(X152,bool),'const.pair.,_2'(s(X152,X153),s(bool,'T_0'))))) | s(bool,X151) = s(bool,X154) | s(bool,'T_0') = s(bool,X154)) )),
  inference(superposition,[],[f1786,f1030])).
fof(f10991,plain,(
  ( ! [X149,X147,X150,X148] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X148,bool),'const.pair.,_2'(s(X148,X149),s(bool,'T_0')))),X147) | s(bool,X147) = s(bool,X150) | s(bool,'T_0') = s(bool,X150)) )),
  inference(superposition,[],[f1783,f1030])).
fof(f10990,plain,(
  ( ! [X146,X144,X145,X143] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X144,bool),'const.pair.,_2'(s(X144,X145),s(bool,'T_0')))),X143) | s(bool,X143) = s(bool,X146) | s(bool,'T_0') = s(bool,X146)) )),
  inference(superposition,[],[f1782,f1030])).
fof(f10987,plain,(
  ( ! [X132,X130,X131,X129] : (sP5('T_0',X129,'const.pair.FST_1'(s('type.pair.prod'(bool,X130),'const.pair.,_2'(s(bool,'T_0'),s(X130,X131))))) | s(bool,X129) = s(bool,X132) | s(bool,'T_0') = s(bool,X132)) )),
  inference(superposition,[],[f1650,f1030])).
fof(f10986,plain,(
  ( ! [X127,X125,X128,X126] : (~sP5('F_0',X125,'const.pair.FST_1'(s('type.pair.prod'(bool,X126),'const.pair.,_2'(s(bool,'T_0'),s(X126,X127))))) | s(bool,X125) = s(bool,X128) | s(bool,'T_0') = s(bool,X128)) )),
  inference(superposition,[],[f1649,f1030])).
fof(f10985,plain,(
  ( ! [X123,X121,X124,X122] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X122),'const.pair.,_2'(s(bool,'T_0'),s(X122,X123)))),X121) | s(bool,X121) = s(bool,X124) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f1646,f1030])).
fof(f10984,plain,(
  ( ! [X118,X120,X119,X117] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X118),'const.pair.,_2'(s(bool,'T_0'),s(X118,X119)))),X117) | s(bool,X117) = s(bool,X120) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f1645,f1030])).
fof(f11607,plain,(
  ( ! [X116,X114,X115,X113] : (p(s(bool,X114)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X115),s('type.realax.real',X115)),X114,X113) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(subsumption_resolution,[],[f10983,f478])).
fof(f10983,plain,(
  ( ! [X116,X114,X115,X113] : (~p(s(bool,'T_0')) | p(s(bool,X114)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X115),s('type.realax.real',X115)),X114,X113) | s(bool,X113) = s(bool,X116) | s(bool,'T_0') = s(bool,X116)) )),
  inference(superposition,[],[f1505,f1030])).
fof(f11606,plain,(
  ( ! [X111,X109,X112,X110] : (p(s(bool,X110)) | ~sP10(X109,X110,'const.real.real_lte_2'(s('type.realax.real',X111),s('type.realax.real',X111))) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(subsumption_resolution,[],[f10982,f478])).
fof(f10982,plain,(
  ( ! [X111,X109,X112,X110] : (~p(s(bool,'T_0')) | p(s(bool,X110)) | ~sP10(X109,X110,'const.real.real_lte_2'(s('type.realax.real',X111),s('type.realax.real',X111))) | s(bool,X109) = s(bool,X112) | s(bool,'T_0') = s(bool,X112)) )),
  inference(superposition,[],[f1481,f1030])).
fof(f11605,plain,(
  ( ! [X107,X105,X108,X106] : (p(s(bool,X106)) | ~sP7(X106,X105,'const.real.real_lte_2'(s('type.realax.real',X107),s('type.realax.real',X107))) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(subsumption_resolution,[],[f10981,f478])).
fof(f10981,plain,(
  ( ! [X107,X105,X108,X106] : (~p(s(bool,'T_0')) | p(s(bool,X106)) | ~sP7(X106,X105,'const.real.real_lte_2'(s('type.realax.real',X107),s('type.realax.real',X107))) | s(bool,X105) = s(bool,X108) | s(bool,'T_0') = s(bool,X108)) )),
  inference(superposition,[],[f1410,f1030])).
fof(f11604,plain,(
  ( ! [X103,X101,X102,X104] : (p(s(bool,X102)) | ~sP6(X101,'const.real.real_lte_2'(s('type.realax.real',X103),s('type.realax.real',X103)),X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(subsumption_resolution,[],[f10980,f478])).
fof(f10980,plain,(
  ( ! [X103,X101,X102,X104] : (~p(s(bool,'T_0')) | p(s(bool,X102)) | ~sP6(X101,'const.real.real_lte_2'(s('type.realax.real',X103),s('type.realax.real',X103)),X102) | s(bool,X101) = s(bool,X104) | s(bool,'T_0') = s(bool,X104)) )),
  inference(superposition,[],[f1380,f1030])).
fof(f11603,plain,(
  ( ! [X99,X97,X100,X98] : (p(s(bool,X98)) | ~sP3(X97,'const.real.real_lte_2'(s('type.realax.real',X99),s('type.realax.real',X99)),X98) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(subsumption_resolution,[],[f10979,f478])).
fof(f10979,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | p(s(bool,X98)) | ~sP3(X97,'const.real.real_lte_2'(s('type.realax.real',X99),s('type.realax.real',X99)),X98) | s(bool,X97) = s(bool,X100) | s(bool,'T_0') = s(bool,X100)) )),
  inference(superposition,[],[f1346,f1030])).
fof(f11602,plain,(
  ( ! [X94,X95,X93,X96] : (p(s(bool,X94)) | ~sP4(X93,'const.real.real_lte_2'(s('type.realax.real',X95),s('type.realax.real',X95)),X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(subsumption_resolution,[],[f10978,f478])).
fof(f10978,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'T_0')) | p(s(bool,X94)) | ~sP4(X93,'const.real.real_lte_2'(s('type.realax.real',X95),s('type.realax.real',X95)),X94) | s(bool,X93) = s(bool,X96) | s(bool,'T_0') = s(bool,X96)) )),
  inference(superposition,[],[f1332,f1030])).
fof(f11601,plain,(
  ( ! [X88,X87,X89,X86] : (~sP8(X86,X87,X88) | sP9(X86,X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f10976,f478])).
fof(f10976,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,'T_0')) | ~sP8(X86,X87,X88) | sP9(X86,X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f814,f1030])).
fof(f11600,plain,(
  ( ! [X85,X83,X84,X82] : (~sP10(X82,X83,X84) | sP11(X82,X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f10975,f478])).
fof(f10975,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,'T_0')) | ~sP10(X82,X83,X84) | sP11(X82,X83,X84) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f813,f1030])).
fof(f11599,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f10973,f478])).
fof(f10973,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | p(s(bool,X77)) | ~sP11(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f790,f1030])).
fof(f11598,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP11(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f10972,f478])).
fof(f10972,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP11(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f789,f1030])).
fof(f11597,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | sP11(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f10971,f478])).
fof(f10971,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | sP11(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f788,f1030])).
fof(f11596,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP10(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f10970,f478])).
fof(f10970,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP10(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f783,f1030])).
fof(f11595,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f10969,f478])).
fof(f10969,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X60)) | p(s(bool,X61)) | ~sP9(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f773,f1030])).
fof(f11594,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f10968,f478])).
fof(f10968,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f772,f1030])).
fof(f11593,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | sP9(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f10967,f478])).
fof(f10967,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | sP9(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f771,f1030])).
fof(f11592,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | ~sP8(X47,X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f10966,f478])).
fof(f10966,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | p(s(bool,X49)) | ~sP8(X47,X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f766,f1030])).
fof(f11591,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP7(X45,X44,X43) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f10965,f478])).
fof(f10965,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP7(X45,X44,X43) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f758,f1030])).
fof(f11590,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP6(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f10964,f478])).
fof(f10964,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP6(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f749,f1030])).
fof(f11589,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP3(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f10963,f478])).
fof(f10963,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP3(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f730,f1030])).
fof(f11588,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP4(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f10962,f478])).
fof(f10962,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP4(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f726,f1030])).
fof(f11587,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f10961,f478])).
fof(f10961,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f673,f1030])).
fof(f11586,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f10960,f478])).
fof(f10960,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f670,f1030])).
fof(f11585,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f10959,f478])).
fof(f10959,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f655,f1030])).
fof(f11584,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f10958,f478])).
fof(f10958,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f654,f1030])).
fof(f11583,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f10957,f478])).
fof(f10957,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f615,f1030])).
fof(f11582,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f10953,f478])).
fof(f10953,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f534,f1030])).
fof(f10920,plain,(
  ( ! [X218,X220,X219,X217] : (sP4('T_0','T_0','const.pair.SND_1'(s('type.pair.prod'(X218,bool),'const.pair.,_2'(s(X218,X219),s(bool,X217))))) | s(bool,X217) = s(bool,X220) | s(bool,'T_0') = s(bool,X220)) )),
  inference(superposition,[],[f2058,f1030])).
fof(f10918,plain,(
  ( ! [X212,X210,X211,X209] : (sP4('T_0','T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X210),'const.pair.,_2'(s(bool,X209),s(X210,X211))))) | s(bool,X209) = s(bool,X212) | s(bool,'T_0') = s(bool,X212)) )),
  inference(superposition,[],[f1908,f1030])).
fof(f10915,plain,(
  ( ! [X198,X196,X197,X195] : (sP5('T_0','T_0','const.pair.SND_1'(s('type.pair.prod'(X196,bool),'const.pair.,_2'(s(X196,X197),s(bool,X195))))) | s(bool,X195) = s(bool,X198) | s(bool,'T_0') = s(bool,X198)) )),
  inference(superposition,[],[f1787,f1030])).
fof(f10914,plain,(
  ( ! [X191,X194,X192,X193] : (~sP5('F_0','T_0','const.pair.SND_1'(s('type.pair.prod'(X192,bool),'const.pair.,_2'(s(X192,X193),s(bool,X191))))) | s(bool,X191) = s(bool,X194) | s(bool,'T_0') = s(bool,X194)) )),
  inference(superposition,[],[f1786,f1030])).
fof(f10913,plain,(
  ( ! [X189,X187,X190,X188] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X188,bool),'const.pair.,_2'(s(X188,X189),s(bool,X187)))),'T_0') | s(bool,X187) = s(bool,X190) | s(bool,'T_0') = s(bool,X190)) )),
  inference(superposition,[],[f1783,f1030])).
fof(f10912,plain,(
  ( ! [X185,X186,X184,X183] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X184,bool),'const.pair.,_2'(s(X184,X185),s(bool,X183)))),'T_0') | s(bool,X183) = s(bool,X186) | s(bool,'T_0') = s(bool,X186)) )),
  inference(superposition,[],[f1782,f1030])).
fof(f10909,plain,(
  ( ! [X171,X169,X172,X170] : (sP5('T_0','T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X170),'const.pair.,_2'(s(bool,X169),s(X170,X171))))) | s(bool,X169) = s(bool,X172) | s(bool,'T_0') = s(bool,X172)) )),
  inference(superposition,[],[f1650,f1030])).
fof(f10908,plain,(
  ( ! [X167,X165,X166,X168] : (~sP5('F_0','T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X166),'const.pair.,_2'(s(bool,X165),s(X166,X167))))) | s(bool,X165) = s(bool,X168) | s(bool,'T_0') = s(bool,X168)) )),
  inference(superposition,[],[f1649,f1030])).
fof(f10907,plain,(
  ( ! [X163,X161,X164,X162] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X162),'const.pair.,_2'(s(bool,X161),s(X162,X163)))),'T_0') | s(bool,X161) = s(bool,X164) | s(bool,'T_0') = s(bool,X164)) )),
  inference(superposition,[],[f1646,f1030])).
fof(f10906,plain,(
  ( ! [X158,X159,X157,X160] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X158),'const.pair.,_2'(s(bool,X157),s(X158,X159)))),'T_0') | s(bool,X157) = s(bool,X160) | s(bool,'T_0') = s(bool,X160)) )),
  inference(superposition,[],[f1645,f1030])).
fof(f10872,plain,(
  ( ! [X39,X41,X38,X40] : (~sP3('T_0','const.pair.SND_1'(s('type.pair.prod'(X39,bool),'const.pair.,_2'(s(X39,X40),s(bool,X38)))),'F_0') | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f3585,f1030])).
fof(f10871,plain,(
  ( ! [X37,X35,X36,X34] : (~sP3('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,X34),s(X35,X36)))),'F_0') | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f3584,f1030])).
fof(f10870,plain,(
  ( ! [X30,X33,X31,X32] : (~sP3('const.pair.SND_1'(s('type.pair.prod'(X31,bool),'const.pair.,_2'(s(X31,X32),s(bool,X30)))),'T_0','F_0') | s(bool,X30) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f3460,f1030])).
fof(f10869,plain,(
  ( ! [X28,X26,X29,X27] : (~sP3('const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X26),s(X27,X28)))),'T_0','F_0') | s(bool,X26) = s(bool,X29) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f3459,f1030])).
fof(f10868,plain,(
  ( ! [X24,X23,X21,X25,X22] : (sP0('const.pair.SND_1'(s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X23),s(bool,X21)))),X24) | s(bool,X21) = s(bool,X25) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f3333,f1030])).
fof(f10867,plain,(
  ( ! [X19,X17,X20,X18,X16] : (sP0(X17,'const.pair.SND_1'(s('type.pair.prod'(X18,bool),'const.pair.,_2'(s(X18,X19),s(bool,X16))))) | s(bool,X16) = s(bool,X20) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f3225,f1030])).
fof(f10866,plain,(
  ( ! [X14,X12,X15,X13,X11] : (sP0('const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X11),s(X12,X13)))),X14) | s(bool,X11) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f3007,f1030])).
fof(f10865,plain,(
  ( ! [X6,X10,X8,X7,X9] : (sP0(X7,'const.pair.FST_1'(s('type.pair.prod'(bool,X8),'const.pair.,_2'(s(bool,X6),s(X8,X9))))) | s(bool,X6) = s(bool,X10) | s(bool,'T_0') = s(bool,X10)) )),
  inference(superposition,[],[f2952,f1030])).
fof(f10864,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f530,f1030])).
fof(f10754,plain,(
  ( ! [X352,X354,X353] : (~p(s(bool,happ(s(fun(bool,bool),X352),s(bool,X353)))) | p(s(bool,happ(s(fun(bool,bool),X352),s(bool,X354)))) | s(bool,'T_0') = s(bool,X353) | s(bool,'T_0') = s(bool,sK15(X352,bool))) )),
  inference(superposition,[],[f590,f1030])).
fof(f10753,plain,(
  ( ! [X350,X349,X351] : (~p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X350)))) | p(s(bool,happ(s(fun(bool,bool),X349),s(bool,X351)))) | s(bool,'T_0') = s(bool,X350) | s(bool,'T_0') = s(bool,sK14(X349,bool))) )),
  inference(superposition,[],[f589,f1030])).
fof(f10752,plain,(
  ( ! [X346,X347,X348] : (s('type.pair.prod'(X346,bool),X347) = s('type.pair.prod'(X346,bool),'const.pair.,_2'(s(X346,'const.pair.FST_1'(s('type.pair.prod'(X346,bool),X347))),s(bool,X348))) | s(bool,'T_0') = s(bool,X348) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X346,bool),X347)))) )),
  inference(superposition,[],[f686,f1030])).
fof(f10750,plain,(
  ( ! [X341,X339,X340] : (s('type.pair.prod'(bool,X339),X340) = s('type.pair.prod'(bool,X339),'const.pair.,_2'(s(bool,X341),s(X339,'const.pair.SND_1'(s('type.pair.prod'(bool,X339),X340))))) | s(bool,'T_0') = s(bool,X341) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X339),X340)))) )),
  inference(superposition,[],[f686,f1030])).
fof(f10696,plain,(
  ( ! [X182,X184,X183,X181] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X183,bool),'const.pair.,_2'(s(X183,X184),s(bool,X182)))),X181) | s(bool,'T_0') = s(bool,X182) | s(bool,'T_0') = s(bool,X181)) )),
  inference(superposition,[],[f2176,f1030])).
fof(f10695,plain,(
  ( ! [X177,X180,X178,X179] : (sP4('T_0',X177,'const.pair.SND_1'(s('type.pair.prod'(X179,bool),'const.pair.,_2'(s(X179,X180),s(bool,X178))))) | s(bool,'T_0') = s(bool,X178) | s(bool,'T_0') = s(bool,X177)) )),
  inference(superposition,[],[f2058,f1030])).
fof(f10694,plain,(
  ( ! [X175,X173,X176,X174] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X175),'const.pair.,_2'(s(bool,X174),s(X175,X176)))),X173) | s(bool,'T_0') = s(bool,X174) | s(bool,'T_0') = s(bool,X173)) )),
  inference(superposition,[],[f1951,f1030])).
fof(f10693,plain,(
  ( ! [X171,X169,X172,X170] : (sP4('T_0',X169,'const.pair.FST_1'(s('type.pair.prod'(bool,X171),'const.pair.,_2'(s(bool,X170),s(X171,X172))))) | s(bool,'T_0') = s(bool,X170) | s(bool,'T_0') = s(bool,X169)) )),
  inference(superposition,[],[f1908,f1030])).
fof(f10690,plain,(
  ( ! [X158,X156,X157,X155] : (sP5('T_0',X155,'const.pair.SND_1'(s('type.pair.prod'(X157,bool),'const.pair.,_2'(s(X157,X158),s(bool,X156))))) | s(bool,'T_0') = s(bool,X156) | s(bool,'T_0') = s(bool,X155)) )),
  inference(superposition,[],[f1787,f1030])).
fof(f10689,plain,(
  ( ! [X154,X152,X151,X153] : (~sP5('F_0',X151,'const.pair.SND_1'(s('type.pair.prod'(X153,bool),'const.pair.,_2'(s(X153,X154),s(bool,X152))))) | s(bool,'T_0') = s(bool,X152) | s(bool,'T_0') = s(bool,X151)) )),
  inference(superposition,[],[f1786,f1030])).
fof(f10688,plain,(
  ( ! [X149,X147,X150,X148] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X149,bool),'const.pair.,_2'(s(X149,X150),s(bool,X148)))),X147) | s(bool,'T_0') = s(bool,X148) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f1783,f1030])).
fof(f10687,plain,(
  ( ! [X146,X144,X145,X143] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X145,bool),'const.pair.,_2'(s(X145,X146),s(bool,X144)))),X143) | s(bool,'T_0') = s(bool,X144) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f1782,f1030])).
fof(f10684,plain,(
  ( ! [X132,X130,X131,X129] : (sP5('T_0',X129,'const.pair.FST_1'(s('type.pair.prod'(bool,X131),'const.pair.,_2'(s(bool,X130),s(X131,X132))))) | s(bool,'T_0') = s(bool,X130) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f1650,f1030])).
fof(f10683,plain,(
  ( ! [X127,X125,X128,X126] : (~sP5('F_0',X125,'const.pair.FST_1'(s('type.pair.prod'(bool,X127),'const.pair.,_2'(s(bool,X126),s(X127,X128))))) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f1649,f1030])).
fof(f10682,plain,(
  ( ! [X123,X121,X124,X122] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X123),'const.pair.,_2'(s(bool,X122),s(X123,X124)))),X121) | s(bool,'T_0') = s(bool,X122) | s(bool,'T_0') = s(bool,X121)) )),
  inference(superposition,[],[f1646,f1030])).
fof(f10681,plain,(
  ( ! [X118,X120,X119,X117] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X119),'const.pair.,_2'(s(bool,X118),s(X119,X120)))),X117) | s(bool,'T_0') = s(bool,X118) | s(bool,'T_0') = s(bool,X117)) )),
  inference(superposition,[],[f1645,f1030])).
fof(f10633,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,'T_0') = s(bool,X24) | s(bool,X25) = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X22,bool),'const.pair.,_2'(s(X22,X23),s(bool,X24))))) | s(bool,'T_0') = s(bool,X25)) )),
  inference(superposition,[],[f1030,f688])).
fof(f10632,plain,(
  ( ! [X21,X19,X20,X18] : (s(bool,'T_0') = s(bool,X19) | s(bool,X21) = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X18),'const.pair.,_2'(s(bool,X19),s(X18,X20))))) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f1030,f687])).
fof(f10631,plain,(
  ( ! [X17,X15,X16] : (s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real',X15))) | s(bool,X17) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X16))))) | s(bool,'T_0') = s(bool,X17)) )),
  inference(superposition,[],[f1030,f600])).
fof(f10621,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1030,f1030])).
fof(f10608,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X29,bool),'const.pair.,_2'(s(X29,X30),s(bool,X31)))))) )),
  inference(superposition,[],[f1030,f688])).
fof(f10607,plain,(
  ( ! [X28,X26,X27,X25] : (s(bool,X26) = s(bool,X28) | s(bool,'T_0') = s(bool,X28) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X25),'const.pair.,_2'(s(bool,X26),s(X25,X27)))))) )),
  inference(superposition,[],[f1030,f687])).
fof(f10606,plain,(
  ( ! [X24,X23,X22] : (s(bool,X24) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X23),s('type.realax.real',X22))) | s(bool,'T_0') = s(bool,X24) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X22))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X23)))))) )),
  inference(superposition,[],[f1030,f600])).
fof(f10199,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f509,f505])).
fof(f9999,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X8))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X8) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X7)))) )),
  inference(forward_demodulation,[],[f9972,f688])).
fof(f9972,plain,(
  ( ! [X8,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X8) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X7))) | s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X8))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f691,f503])).
fof(f9998,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X6) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X5)))) )),
  inference(forward_demodulation,[],[f9971,f4723])).
fof(f9971,plain,(
  ( ! [X6,X5] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X5))) | s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f691,f499])).
fof(f9684,plain,(
  ( ! [X8,X7] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X8))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X8) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X7)))) )),
  inference(forward_demodulation,[],[f9657,f1814])).
fof(f9657,plain,(
  ( ! [X8,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X8) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X7))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X8))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f619,f503])).
fof(f9683,plain,(
  ( ! [X6,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X6) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X5)))) )),
  inference(forward_demodulation,[],[f9656,f4718])).
fof(f9656,plain,(
  ( ! [X6,X5] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X5))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f619,f499])).
fof(f9496,plain,(
  ( ! [X31,X32] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X31),X32)),'const.pair.FST_1'(s('type.pair.prod'(bool,X31),X32)))) )),
  inference(superposition,[],[f1650,f686])).
fof(f9494,plain,(
  ( ! [X28,X27] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28)),'const.pair.FST_1'(s('type.pair.prod'(bool,X27),X28)))) )),
  inference(superposition,[],[f1951,f686])).
fof(f9490,plain,(
  ( ! [X19,X20] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X19,bool),X20)),'const.pair.SND_1'(s('type.pair.prod'(X19,bool),X20)))) )),
  inference(superposition,[],[f1787,f686])).
fof(f9488,plain,(
  ( ! [X15,X16] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X15,bool),X16)),'const.pair.SND_1'(s('type.pair.prod'(X15,bool),X16)))) )),
  inference(superposition,[],[f2176,f686])).
fof(f9475,plain,(
  ( ! [X4,X5] : (s('type.pair.prod'(X4,bool),X5) = s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,'const.pair.FST_1'(s('type.pair.prod'(X4,bool),X5))),s(bool,'F_0'))) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),X5)))) )),
  inference(superposition,[],[f686,f516])).
fof(f9532,plain,(
  ( ! [X8] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X8))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X8))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f9531,f6554])).
fof(f9531,plain,(
  ( ! [X8] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X8))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X8))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X8)))))))) )),
  inference(forward_demodulation,[],[f9462,f494])).
fof(f9462,plain,(
  ( ! [X8] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X8))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X8))),s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X8)))))))) )),
  inference(superposition,[],[f686,f4780])).
fof(f9528,plain,(
  ( ! [X6] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))))) )),
  inference(forward_demodulation,[],[f9460,f494])).
fof(f9460,plain,(
  ( ! [X6] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))))) )),
  inference(superposition,[],[f686,f493])).
fof(f9459,plain,(
  ( ! [X4,X5] : (s('type.pair.prod'(bool,X4),X5) = s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,'const.pair.SND_1'(s('type.pair.prod'(bool,X4),X5))))) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),X5)))) )),
  inference(superposition,[],[f686,f516])).
fof(f9290,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X2)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))) )),
  inference(forward_demodulation,[],[f9136,f600])).
fof(f9136,plain,(
  ( ! [X2,X1] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X1)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))) )),
  inference(superposition,[],[f646,f600])).
fof(f9127,plain,(
  ( ! [X72,X71] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X72),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X71)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X72)))) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f646,f506])).
fof(f9123,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X63),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X62),s('type.realax.real',X61)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62))),s('type.realax.real',X63)))) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) )),
  inference(superposition,[],[f646,f595])).
fof(f9111,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X34),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X32)))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33))),s('type.realax.real',X34)))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) )),
  inference(superposition,[],[f646,f596])).
fof(f9110,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X31),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X30),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30))),s('type.realax.real',X31)))) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) )),
  inference(superposition,[],[f646,f506])).
fof(f9088,plain,(
  ( ! [X72,X71] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X71),s('type.realax.real',X71))),s('type.realax.real',X72)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X72),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.realax.real',X72) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X71),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f646,f506])).
fof(f9084,plain,(
  ( ! [X61,X62,X63] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X62),s('type.realax.real',X61))),s('type.realax.real',X63)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X63),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))))) | s('type.realax.real',X63) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X61),s('type.realax.real',X62)))) )),
  inference(superposition,[],[f646,f595])).
fof(f9072,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X32))),s('type.realax.real',X34)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X34),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))))) | s('type.realax.real',X34) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X33)))) )),
  inference(superposition,[],[f646,f596])).
fof(f9071,plain,(
  ( ! [X30,X31] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X30),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X31)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X31),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))))) | s('type.realax.real',X31) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X30),s('type.realax.real',X30)))) )),
  inference(superposition,[],[f646,f506])).
fof(f8684,plain,(
  ( ! [X0,X1] : (~sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f8557,f721])).
fof(f8557,plain,(
  ( ! [X0,X1] : (~sP4('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f831,f3390])).
fof(f8332,plain,(
  ( ! [X103,X102,X104] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X102),s('type.realax.real',X102))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X102),s('type.realax.real',X102)))),X103,X104) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X102),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X102),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X103,X104)) )),
  inference(superposition,[],[f1098,f506])).
fof(f8328,plain,(
  ( ! [X90,X88,X91,X89] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X89),s('type.realax.real',X88))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X89),s('type.realax.real',X88)))),X90,X91) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X88),s('type.realax.real',X89))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X88),s('type.realax.real',X89)))),X90,X91)) )),
  inference(superposition,[],[f1098,f595])).
fof(f8316,plain,(
  ( ! [X47,X50,X48,X49] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X48),s('type.realax.real',X47))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X48),s('type.realax.real',X47)))),X49,X50) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X47),s('type.realax.real',X48))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X47),s('type.realax.real',X48)))),X49,X50)) )),
  inference(superposition,[],[f1098,f596])).
fof(f8315,plain,(
  ( ! [X45,X46,X44] : (sP9('const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X45,X46) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X44),s('type.realax.real',X44))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X44),s('type.realax.real',X44)))),X45,X46)) )),
  inference(superposition,[],[f1098,f506])).
fof(f8297,plain,(
  ( ! [X2,X0,X1] : (sP9('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f908,f1098])).
fof(f8299,plain,(
  ( ! [X2,X0,X1] : (sP9('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f909,f1098])).
fof(f8288,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))))) | s('type.num.num',X7) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X4,'type.num.num'),'const.pair.,_2'(s(X4,X5),s('type.num.num',X6)))))) )),
  inference(superposition,[],[f633,f688])).
fof(f8287,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X3))))) | s('type.num.num',X3) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),'const.pair.,_2'(s('type.num.num',X1),s(X0,X2)))))) )),
  inference(superposition,[],[f633,f687])).
fof(f8156,plain,(
  ( ! [X103,X102,X104] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X102),s('type.realax.real',X102))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X102),s('type.realax.real',X102)))),X103,X104) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X102),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X102),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X103,X104)) )),
  inference(superposition,[],[f1097,f506])).
fof(f8152,plain,(
  ( ! [X90,X88,X91,X89] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X89),s('type.realax.real',X88))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X89),s('type.realax.real',X88)))),X90,X91) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X88),s('type.realax.real',X89))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X88),s('type.realax.real',X89)))),X90,X91)) )),
  inference(superposition,[],[f1097,f595])).
fof(f8140,plain,(
  ( ! [X47,X50,X48,X49] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X48),s('type.realax.real',X47))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X48),s('type.realax.real',X47)))),X49,X50) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X47),s('type.realax.real',X48))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X47),s('type.realax.real',X48)))),X49,X50)) )),
  inference(superposition,[],[f1097,f596])).
fof(f8139,plain,(
  ( ! [X45,X46,X44] : (sP11('const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),X45,X46) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X44),s('type.realax.real',X44))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X44),s('type.realax.real',X44)))),X45,X46)) )),
  inference(superposition,[],[f1097,f506])).
fof(f8123,plain,(
  ( ! [X2,X0,X1] : (sP11('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f910,f1097])).
fof(f8117,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK15(X0,bool))) )),
  inference(superposition,[],[f590,f516])).
fof(f8120,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X8),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK15('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5)))))) )),
  inference(forward_demodulation,[],[f8116,f688])).
fof(f8116,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK15('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | p(s(bool,happ(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))),s(X5,X9))))) )),
  inference(superposition,[],[f590,f688])).
fof(f8119,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0)))))) )),
  inference(forward_demodulation,[],[f8115,f687])).
fof(f8115,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15('const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3))))),s(X0,X4))))) )),
  inference(superposition,[],[f590,f687])).
fof(f7744,plain,(
  ( ! [X0,X1] : (~p(s(bool,happ(s(fun(bool,bool),X0),s(bool,'F_0')))) | p(s(bool,happ(s(fun(bool,bool),X0),s(bool,X1)))) | s(bool,'T_0') = s(bool,sK14(X0,bool))) )),
  inference(superposition,[],[f589,f516])).
fof(f7747,plain,(
  ( ! [X6,X8,X7,X5,X9] : (p(s(bool,happ(s(fun(X5,bool),X8),s(X5,X9)))) | ~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK14('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5)))))) )),
  inference(forward_demodulation,[],[f7743,f688])).
fof(f7743,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,happ(s(fun(X5,bool),X8),s(X5,sK14('const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8)))),X5))))) | p(s(bool,happ(s(fun(X5,bool),'const.pair.SND_1'(s('type.pair.prod'(X6,fun(X5,bool)),'const.pair.,_2'(s(X6,X7),s(fun(X5,bool),X8))))),s(X5,X9))))) )),
  inference(superposition,[],[f589,f688])).
fof(f7746,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK14('const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0)))))) )),
  inference(forward_demodulation,[],[f7742,f687])).
fof(f7742,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK14('const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3)))),X0))))) | p(s(bool,happ(s(fun(X0,bool),'const.pair.FST_1'(s('type.pair.prod'(fun(X0,bool),X1),'const.pair.,_2'(s(fun(X0,bool),X2),s(X1,X3))))),s(X0,X4))))) )),
  inference(superposition,[],[f589,f687])).
fof(f7310,plain,(
  ( ! [X17,X15,X16] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X15))),s('type.realax.real',X16))),s('type.realax.real',X17))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X15),s('type.realax.real',X16))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X17)))))) )),
  inference(superposition,[],[f832,f833])).
fof(f7309,plain,(
  ( ! [X14,X12,X13] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X12),s('type.realax.real',X13))),s('type.realax.real',X14))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12))),s('type.realax.real',X13)))))) )),
  inference(superposition,[],[f600,f833])).
fof(f7308,plain,(
  ( ! [X10,X11,X9] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X10))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X10))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X11)))))) )),
  inference(superposition,[],[f600,f833])).
fof(f7307,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X6),s('type.realax.real',X7))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X6))),s('type.realax.real',X7)))))) )),
  inference(superposition,[],[f598,f833])).
fof(f7306,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))),s('type.realax.real',X5))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X5)))))) )),
  inference(superposition,[],[f497,f833])).
fof(f7305,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X3))),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3)))))) )),
  inference(superposition,[],[f496,f833])).
fof(f7304,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f489,f833])).
fof(f7300,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))),s('type.realax.real',X5))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real',X4)))))) )),
  inference(superposition,[],[f833,f595])).
fof(f7440,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X3),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) )),
  inference(forward_demodulation,[],[f7439,f4343])).
fof(f7439,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real',X3))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3)))))) )),
  inference(forward_demodulation,[],[f7299,f4340])).
fof(f7299,plain,(
  ( ! [X2,X3] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))) )),
  inference(superposition,[],[f833,f598])).
fof(f7435,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))))))))) )),
  inference(forward_demodulation,[],[f7226,f6105])).
fof(f7226,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))))))) )),
  inference(unit_resulting_resolution,[],[f833,f648])).
fof(f6903,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))))) )),
  inference(backward_demodulation,[],[f6902,f514])).
fof(f6902,plain,(
  ( ! [X10,X9] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X10))))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X10)))) )),
  inference(forward_demodulation,[],[f6763,f6344])).
fof(f6763,plain,(
  ( ! [X10,X9] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X10))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10)))))))) )),
  inference(superposition,[],[f832,f597])).
fof(f6759,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))))) )),
  inference(superposition,[],[f832,f596])).
fof(f6900,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real',X2))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f6758,f489])).
fof(f6758,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f832,f506])).
fof(f6723,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f832,f489])).
fof(f6554,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f4718,f499])).
fof(f6350,plain,(
  ( ! [X14,X13] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X14))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X14)))))))) )),
  inference(superposition,[],[f489,f601])).
fof(f6349,plain,(
  ( ! [X12,X11] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X12))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X12)))))))) )),
  inference(superposition,[],[f496,f601])).
fof(f6348,plain,(
  ( ! [X10,X9] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) )),
  inference(superposition,[],[f497,f601])).
fof(f6347,plain,(
  ( ! [X6,X8,X7] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real',X7))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X6),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7)))))))) )),
  inference(superposition,[],[f598,f601])).
fof(f6346,plain,(
  ( ! [X4,X5,X3] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X5),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real',X4))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X5)))))) )),
  inference(superposition,[],[f600,f601])).
fof(f6345,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))))) )),
  inference(superposition,[],[f600,f601])).
fof(f6344,plain,(
  ( ! [X10,X9] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X10))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X10)))))))) )),
  inference(superposition,[],[f601,f597])).
fof(f6343,plain,(
  ( ! [X8] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X8))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X8),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) )),
  inference(superposition,[],[f601,f502])).
fof(f6340,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X3),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X4),s('type.realax.real',X3)))))) )),
  inference(superposition,[],[f601,f596])).
fof(f6339,plain,(
  ( ! [X2] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X2),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X2),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f601,f506])).
fof(f6243,plain,(
  ( ! [X0] : (~sP10('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f831,f6231,f782])).
fof(f6239,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,X20)) | p(s(bool,X21)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X22),s('type.realax.real',X22)),X21,'const.pair.SND_1'(s('type.pair.prod'(X18,bool),'const.pair.,_2'(s(X18,X19),s(bool,X20)))))) )),
  inference(superposition,[],[f1505,f688])).
fof(f6238,plain,(
  ( ! [X14,X17,X15,X13,X16] : (~p(s(bool,X14)) | p(s(bool,X16)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X17),s('type.realax.real',X17)),X16,'const.pair.FST_1'(s('type.pair.prod'(bool,X13),'const.pair.,_2'(s(bool,X14),s(X13,X15)))))) )),
  inference(superposition,[],[f1505,f687])).
fof(f6237,plain,(
  ( ! [X12,X10,X11,X9] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X10),s('type.realax.real',X9)))) | p(s(bool,X11)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X12),s('type.realax.real',X12)),X11,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X10)))))) )),
  inference(superposition,[],[f1505,f600])).
fof(f6240,plain,(
  ( ! [X6,X8,X7] : (~sP11('const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6))) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f6236,f478])).
fof(f6236,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f1505,f914])).
fof(f6189,plain,(
  ( ! [X180,X178,X181,X179] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X180,bool),'const.pair.,_2'(s(X180,X181),s(bool,'const.real.real_lte_2'(s('type.realax.real',X179),s('type.realax.real',X178)))))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X178))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X179)))))) )),
  inference(superposition,[],[f2176,f600])).
fof(f6188,plain,(
  ( ! [X177,X175,X176,X174] : (sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X174))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X175)))),'const.pair.SND_1'(s('type.pair.prod'(X176,bool),'const.pair.,_2'(s(X176,X177),s(bool,'const.real.real_lte_2'(s('type.realax.real',X175),s('type.realax.real',X174)))))))) )),
  inference(superposition,[],[f2058,f600])).
fof(f6187,plain,(
  ( ! [X173,X171,X172,X170] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X172),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X171),s('type.realax.real',X170))),s(X172,X173)))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X170))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X171)))))) )),
  inference(superposition,[],[f1951,f600])).
fof(f6186,plain,(
  ( ! [X167,X169,X166,X168] : (sP4('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X166))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X167)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X168),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X167),s('type.realax.real',X166))),s(X168,X169)))))) )),
  inference(superposition,[],[f1908,f600])).
fof(f6183,plain,(
  ( ! [X154,X152,X155,X153] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X152))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X153)))),'const.pair.SND_1'(s('type.pair.prod'(X154,bool),'const.pair.,_2'(s(X154,X155),s(bool,'const.real.real_lte_2'(s('type.realax.real',X153),s('type.realax.real',X152)))))))) )),
  inference(superposition,[],[f1787,f600])).
fof(f6182,plain,(
  ( ! [X151,X149,X150,X148] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X148))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X149)))),'const.pair.SND_1'(s('type.pair.prod'(X150,bool),'const.pair.,_2'(s(X150,X151),s(bool,'const.real.real_lte_2'(s('type.realax.real',X149),s('type.realax.real',X148)))))))) )),
  inference(superposition,[],[f1786,f600])).
fof(f6181,plain,(
  ( ! [X146,X144,X147,X145] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X146,bool),'const.pair.,_2'(s(X146,X147),s(bool,'const.real.real_lte_2'(s('type.realax.real',X145),s('type.realax.real',X144)))))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X144))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X145)))))) )),
  inference(superposition,[],[f1783,f600])).
fof(f6180,plain,(
  ( ! [X142,X140,X143,X141] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X142,bool),'const.pair.,_2'(s(X142,X143),s(bool,'const.real.real_lte_2'(s('type.realax.real',X141),s('type.realax.real',X140)))))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X140))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X141)))))) )),
  inference(superposition,[],[f1782,f600])).
fof(f6177,plain,(
  ( ! [X127,X128,X126,X129] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X126))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X127)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X128),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X127),s('type.realax.real',X126))),s(X128,X129)))))) )),
  inference(superposition,[],[f1650,f600])).
fof(f6176,plain,(
  ( ! [X125,X123,X124,X122] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X122))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X123)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X124),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X123),s('type.realax.real',X122))),s(X124,X125)))))) )),
  inference(superposition,[],[f1649,f600])).
fof(f6175,plain,(
  ( ! [X121,X118,X120,X119] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X120),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X119),s('type.realax.real',X118))),s(X120,X121)))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X118))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X119)))))) )),
  inference(superposition,[],[f1646,f600])).
fof(f6174,plain,(
  ( ! [X116,X114,X117,X115] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X116),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X115),s('type.realax.real',X114))),s(X116,X117)))),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X114))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X115)))))) )),
  inference(superposition,[],[f1645,f600])).
fof(f6173,plain,(
  ( ! [X111,X112,X110,X113] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X111),s('type.realax.real',X110)))) | p(s(bool,X112)) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X110))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X111)))),X112,'const.real.real_lte_2'(s('type.realax.real',X113),s('type.realax.real',X113)))) )),
  inference(superposition,[],[f1481,f600])).
fof(f6172,plain,(
  ( ! [X109,X107,X108,X106] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X107),s('type.realax.real',X106)))) | p(s(bool,X108)) | ~sP7(X108,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X106))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X107)))),'const.real.real_lte_2'(s('type.realax.real',X109),s('type.realax.real',X109)))) )),
  inference(superposition,[],[f1410,f600])).
fof(f6171,plain,(
  ( ! [X103,X105,X102,X104] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X103),s('type.realax.real',X102)))) | p(s(bool,X104)) | ~sP6('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X102))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X103)))),'const.real.real_lte_2'(s('type.realax.real',X105),s('type.realax.real',X105)),X104)) )),
  inference(superposition,[],[f1380,f600])).
fof(f6170,plain,(
  ( ! [X101,X99,X100,X98] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X99),s('type.realax.real',X98)))) | p(s(bool,X100)) | ~sP3('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X98))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X99)))),'const.real.real_lte_2'(s('type.realax.real',X101),s('type.realax.real',X101)),X100)) )),
  inference(superposition,[],[f1346,f600])).
fof(f6169,plain,(
  ( ! [X94,X97,X95,X96] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X95),s('type.realax.real',X94)))) | p(s(bool,X96)) | ~sP4('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X94))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X95)))),'const.real.real_lte_2'(s('type.realax.real',X97),s('type.realax.real',X97)),X96)) )),
  inference(superposition,[],[f1332,f600])).
fof(f6167,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X88),s('type.realax.real',X87)))) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X87))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X88)))),X89,X90) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X87))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X88)))),X89,X90)) )),
  inference(superposition,[],[f814,f600])).
fof(f6166,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X84),s('type.realax.real',X83)))) | ~sP10('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X83))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X84)))),X85,X86) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X83))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X84)))),X85,X86)) )),
  inference(superposition,[],[f813,f600])).
fof(f6164,plain,(
  ( ! [X78,X76,X79,X77] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X77),s('type.realax.real',X76)))) | ~p(s(bool,X78)) | p(s(bool,X79)) | ~sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X76))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X77)))),X79,X78)) )),
  inference(superposition,[],[f790,f600])).
fof(f6163,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X73),s('type.realax.real',X72)))) | ~p(s(bool,X74)) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X72))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X73)))),X74,X75)) )),
  inference(superposition,[],[f789,f600])).
fof(f6162,plain,(
  ( ! [X70,X68,X71,X69] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X69),s('type.realax.real',X68)))) | p(s(bool,X70)) | sP11('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X68))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X69)))),X71,X70)) )),
  inference(superposition,[],[f788,f600])).
fof(f6161,plain,(
  ( ! [X66,X64,X67,X65] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X65),s('type.realax.real',X64)))) | p(s(bool,X66)) | ~p(s(bool,X67)) | ~sP10(X67,X66,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X64))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X65)))))) )),
  inference(superposition,[],[f783,f600])).
fof(f6160,plain,(
  ( ! [X61,X62,X60,X63] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X61),s('type.realax.real',X60)))) | p(s(bool,X62)) | p(s(bool,X63)) | ~sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X60))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X61)))),X63,X62)) )),
  inference(superposition,[],[f773,f600])).
fof(f6159,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X57),s('type.realax.real',X56)))) | ~p(s(bool,X58)) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X56))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X57)))),X58,X59)) )),
  inference(superposition,[],[f772,f600])).
fof(f6158,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X53),s('type.realax.real',X52)))) | ~p(s(bool,X54)) | sP9('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X52))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X53)))),X55,X54)) )),
  inference(superposition,[],[f771,f600])).
fof(f6157,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X49),s('type.realax.real',X48)))) | p(s(bool,X50)) | p(s(bool,X51)) | ~sP8('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X48))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X49)))),X50,X51)) )),
  inference(superposition,[],[f766,f600])).
fof(f6156,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X45),s('type.realax.real',X44)))) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP7(X47,X46,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X44))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X45)))))) )),
  inference(superposition,[],[f758,f600])).
fof(f6155,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X41),s('type.realax.real',X40)))) | ~p(s(bool,X42)) | p(s(bool,X43)) | ~sP6(X42,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X40))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X41)))),X43)) )),
  inference(superposition,[],[f749,f600])).
fof(f6154,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X37),s('type.realax.real',X36)))) | ~p(s(bool,X38)) | p(s(bool,X39)) | ~sP3(X38,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X36))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X37)))),X39)) )),
  inference(superposition,[],[f730,f600])).
fof(f6153,plain,(
  ( ! [X35,X33,X34,X32] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X33),s('type.realax.real',X32)))) | ~p(s(bool,X34)) | p(s(bool,X35)) | ~sP4(X34,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X32))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X33)))),X35)) )),
  inference(superposition,[],[f726,f600])).
fof(f6152,plain,(
  ( ! [X30,X31,X29] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X30),s('type.realax.real',X29)))) | p(s(bool,X31)) | sP1(X31,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X29))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X30)))))) )),
  inference(superposition,[],[f673,f600])).
fof(f6151,plain,(
  ( ! [X28,X26,X27] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X27),s('type.realax.real',X26)))) | p(s(bool,X28)) | sP1('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X26))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X27)))),X28)) )),
  inference(superposition,[],[f670,f600])).
fof(f6150,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X24),s('type.realax.real',X23)))) | sP0('const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X23))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X24)))),X25)) )),
  inference(superposition,[],[f655,f600])).
fof(f6149,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X21),s('type.realax.real',X20)))) | sP0(X22,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X20))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X21)))))) )),
  inference(superposition,[],[f654,f600])).
fof(f6139,plain,(
  ( ! [X4,X3] : (s(bool,'F_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real',X4),s('type.realax.real',X3))) | s(bool,'T_0') = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4)))))) )),
  inference(superposition,[],[f600,f516])).
fof(f6137,plain,(
  ( ! [X2] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X2))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f600,f2113])).
fof(f6136,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f600,f489])).
fof(f6106,plain,(
  ( ! [X2] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X2)))))) )),
  inference(superposition,[],[f600,f2113])).
fof(f6105,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f600,f489])).
fof(f6072,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f831,f600,f741])).
fof(f6073,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f478,f600,f739])).
fof(f6074,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(unit_resulting_resolution,[],[f831,f600,f741])).
fof(f6075,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1)),'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(unit_resulting_resolution,[],[f478,f600,f739])).
fof(f5919,plain,(
  ( ! [X43,X44] : (s('type.realax.real',X44) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X43),s('type.realax.real',X43))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X43),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X44))))) )),
  inference(superposition,[],[f648,f506])).
fof(f5918,plain,(
  ( ! [X41,X42] : (s('type.realax.real',X42) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X41),s('type.realax.real',X41))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X42),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X41),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f648,f506])).
fof(f5917,plain,(
  ( ! [X39,X40] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X39),s('type.realax.real',X39))) | s('type.realax.real',X40) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X40),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X39),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f624,f506])).
fof(f5916,plain,(
  ( ! [X37,X38] : (s('type.realax.real',X38) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X38),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X37),s('type.realax.real',X37))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real',X37),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f623,f506])).
fof(f5914,plain,(
  ( ! [X33,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X33),s('type.realax.real',X33))),s('type.realax.real',X34)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X34),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X33),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(superposition,[],[f617,f506])).
fof(f5913,plain,(
  ( ! [X31,X32] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X32),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X31),s('type.realax.real',X31)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X31),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X32))))) )),
  inference(superposition,[],[f617,f506])).
fof(f5807,plain,(
  ( ! [X47,X46] : (s('type.realax.real',X47) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X46),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X46),s('type.realax.real',X46))),s('type.realax.real',X47))))) )),
  inference(superposition,[],[f648,f506])).
fof(f5806,plain,(
  ( ! [X45,X44] : (s('type.realax.real',X45) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X44),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X45),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X44),s('type.realax.real',X44))))))) )),
  inference(superposition,[],[f648,f506])).
fof(f5805,plain,(
  ( ! [X43,X42] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.real.pow_2'(s('type.realax.real',X42),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real',X43) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X43),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X42),s('type.realax.real',X42)))))) )),
  inference(superposition,[],[f624,f506])).
fof(f5804,plain,(
  ( ! [X41,X40] : (s('type.realax.real',X41) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X41),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X40),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X40),s('type.realax.real',X40)))) )),
  inference(superposition,[],[f623,f506])).
fof(f5802,plain,(
  ( ! [X37,X36] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X36),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X37)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X37),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X36),s('type.realax.real',X36))))))) )),
  inference(superposition,[],[f617,f506])).
fof(f5801,plain,(
  ( ! [X35,X34] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X35),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X34),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X34),s('type.realax.real',X34))),s('type.realax.real',X35))))) )),
  inference(superposition,[],[f617,f506])).
fof(f5800,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X32),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X33)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X33),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X32),s('type.realax.real',X32))))))) )),
  inference(superposition,[],[f616,f506])).
fof(f5747,plain,(
  ( ! [X3] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.realax.real',X3))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X3))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f506,f597])).
fof(f5743,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f506,f501])).
fof(f5741,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f506,f648])).
fof(f5742,plain,(
  ( ! [X0] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.real.pow_2'(s('type.realax.real',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) )),
  inference(unit_resulting_resolution,[],[f506,f648])).
fof(f5738,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | ~sP10('const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),X17,'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X18)))) )),
  inference(superposition,[],[f1481,f688])).
fof(f5737,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X10)) | p(s(bool,X12)) | ~sP10('const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))),X12,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1481,f687])).
fof(f5522,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | ~sP7(X17,'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X18)))) )),
  inference(superposition,[],[f1410,f688])).
fof(f5521,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X10)) | p(s(bool,X12)) | ~sP7(X12,'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))),'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)))) )),
  inference(superposition,[],[f1410,f687])).
fof(f5523,plain,(
  ( ! [X6,X8,X7] : (~sP7(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8))) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f5520,f478])).
fof(f5520,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP7(X7,'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)))) )),
  inference(superposition,[],[f1410,f914])).
fof(f5514,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X4,'type.num.num'),'const.pair.,_2'(s(X4,X5),s('type.num.num',X6)))))) )),
  inference(superposition,[],[f628,f688])).
fof(f5513,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X3))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),'const.pair.,_2'(s('type.num.num',X1),s(X0,X2)))))) )),
  inference(superposition,[],[f628,f687])).
fof(f5509,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(unit_resulting_resolution,[],[f3396,f628])).
fof(f5442,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X4,'type.num.num'),'const.pair.,_2'(s(X4,X5),s('type.num.num',X6)))))) )),
  inference(superposition,[],[f627,f688])).
fof(f5441,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X3))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),'const.pair.,_2'(s('type.num.num',X1),s(X0,X2)))))) )),
  inference(superposition,[],[f627,f687])).
fof(f5439,plain,(
  ( ! [X0] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(unit_resulting_resolution,[],[f3396,f627])).
fof(f5436,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | ~sP6('const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X18)),X17)) )),
  inference(superposition,[],[f1380,f688])).
fof(f5435,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X10)) | p(s(bool,X12)) | ~sP6('const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))),'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)),X12)) )),
  inference(superposition,[],[f1380,f687])).
fof(f5425,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | ~sP3('const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X18)),X17)) )),
  inference(superposition,[],[f1346,f688])).
fof(f5424,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X10)) | p(s(bool,X12)) | ~sP3('const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))),'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)),X12)) )),
  inference(superposition,[],[f1346,f687])).
fof(f5353,plain,(
  ( ! [X45,X46,X44] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X45),s('type.realax.real',X44))) | s('type.realax.real',X46) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X46),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X44),s('type.realax.real',X45)))))) )),
  inference(superposition,[],[f624,f595])).
fof(f5345,plain,(
  ( ! [X26,X24,X25] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real',X24))) | s('type.realax.real',X26) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X26),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X24),s('type.realax.real',X25)))))) )),
  inference(superposition,[],[f624,f596])).
fof(f5190,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f4723,f499])).
fof(f5099,plain,(
  ( ! [X9] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))) | s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f623,f497])).
fof(f5096,plain,(
  ( ! [X4,X5] : (s('type.realax.real',X4) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X5),s('type.realax.real',X4))) | s('type.realax.real',X5) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f623,f595])).
fof(f5095,plain,(
  ( ! [X2,X3] : (s('type.realax.real',X2) != s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X2),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X3)))) )),
  inference(superposition,[],[f623,f598])).
fof(f5094,plain,(
  ( ! [X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))) )),
  inference(superposition,[],[f623,f496])).
fof(f5092,plain,(
  ( ! [X61,X59,X62,X60] : (s('type.realax.real',X62) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X62),s('type.realax.real',X61))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'(X59,'type.realax.real'),'const.pair.,_2'(s(X59,X60),s('type.realax.real',X61)))))) )),
  inference(superposition,[],[f623,f688])).
fof(f5091,plain,(
  ( ! [X57,X58,X56,X55] : (s('type.realax.real',X58) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X58),s('type.realax.real',X56))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real',X55),'const.pair.,_2'(s('type.realax.real',X56),s(X55,X57)))))) )),
  inference(superposition,[],[f623,f687])).
fof(f5089,plain,(
  ( ! [X52,X53] : (s('type.realax.real',X53) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X53),s('type.realax.real',X52))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X52)))))) )),
  inference(superposition,[],[f623,f489])).
fof(f5087,plain,(
  ( ! [X48,X49] : (s('type.realax.real',X49) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X49),s('type.realax.real',X48))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X48)))) )),
  inference(superposition,[],[f623,f491])).
fof(f5085,plain,(
  ( ! [X43,X44,X42] : (s('type.realax.real',X44) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X44),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X43),s('type.realax.real',X42))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X42),s('type.realax.real',X43)))) )),
  inference(superposition,[],[f623,f595])).
fof(f5082,plain,(
  ( ! [X35,X36] : (s('type.realax.real',X36) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X36),s('type.realax.real',X35))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X35),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f623,f490])).
fof(f5080,plain,(
  ( ! [X31,X32] : (s('type.realax.real',X32) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X32),s('type.realax.real',X31))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X31),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f623,f492])).
fof(f5077,plain,(
  ( ! [X24,X23,X22] : (s('type.realax.real',X24) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X23),s('type.realax.real',X22))))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X22),s('type.realax.real',X23)))) )),
  inference(superposition,[],[f623,f596])).
fof(f5076,plain,(
  ( ! [X21,X20] : (s('type.realax.real',X21) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X21),s('type.realax.real',X20))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X20),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f623,f502])).
fof(f5074,plain,(
  ( ! [X17,X16] : (s('type.realax.real',X17) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X17),s('type.realax.real',X16))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X16)))) )),
  inference(superposition,[],[f623,f501])).
fof(f5071,plain,(
  ( ! [X10,X11,X9] : (s('type.realax.real',X11) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X11),s('type.realax.real',X10))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real',X10)))))) )),
  inference(superposition,[],[f623,f1814])).
fof(f5070,plain,(
  ( ! [X8,X7] : (s('type.realax.real',X8) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X8),s('type.realax.real',X7))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X7)))))) )),
  inference(superposition,[],[f623,f4724])).
fof(f5068,plain,(
  ( ! [X4,X3] : (s('type.realax.real',X4) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X4),s('type.realax.real',X3))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X3)))))) )),
  inference(superposition,[],[f623,f4719])).
fof(f5067,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real',X2) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X2),s('type.realax.real',X0))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f623,f1675])).
fof(f5037,plain,(
  ( ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(unit_resulting_resolution,[],[f486,f623])).
fof(f4927,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f4719,f499])).
fof(f4780,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f4724,f499])).
fof(f4724,plain,(
  ( ! [X6] : (s('type.realax.real',X6) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X6)))))) )),
  inference(superposition,[],[f687,f503])).
fof(f4723,plain,(
  ( ! [X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X5)))))) )),
  inference(superposition,[],[f688,f503])).
fof(f4719,plain,(
  ( ! [X1] : (s('type.realax.real',X1) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X1)))))) )),
  inference(superposition,[],[f1675,f503])).
fof(f4718,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f1814,f503])).
fof(f4559,plain,(
  ( ! [X2,X0,X1] : (~sP2('F_0','const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3599,f516])).
fof(f4532,plain,(
  ( ! [X8] : (~sP2('F_0','const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X8),s(bool,X8)))),'F_0') | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f3599,f516])).
fof(f4494,plain,(
  ( ! [X2,X0,X1] : (~sP2('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3598,f516])).
fof(f4345,plain,(
  ( ! [X7] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X7))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X7)))) )),
  inference(superposition,[],[f598,f491])).
fof(f4343,plain,(
  ( ! [X4,X3] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X3),s('type.realax.real',X4))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X4))),s('type.realax.real',X3)))) )),
  inference(superposition,[],[f598,f595])).
fof(f4342,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f598,f496])).
fof(f4340,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(superposition,[],[f598,f489])).
fof(f4256,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3521,f516])).
fof(f4230,plain,(
  ( ! [X8] : (~sP2('const.pair.SND_1'(s('type.pair.prod'(bool,bool),'const.pair.,_2'(s(bool,X8),s(bool,X8)))),'F_0','F_0') | s(bool,'T_0') = s(bool,X8)) )),
  inference(superposition,[],[f3521,f516])).
fof(f4121,plain,(
  ( ! [X4,X5] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X4),s('type.realax.real',X5))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X5))),s('type.realax.real',X4)))) )),
  inference(superposition,[],[f597,f596])).
fof(f4120,plain,(
  ( ! [X1] : (s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f597,f501])).
fof(f4119,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(superposition,[],[f597,f495])).
fof(f4043,plain,(
  ( ! [X2,X0,X1] : (~sP2('const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))),'F_0','F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f3520,f516])).
fof(f3611,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X2)))),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'F_0')) )),
  inference(forward_demodulation,[],[f3576,f914])).
fof(f3576,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1))),s(X0,X2)))),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f906,f1645,f884,f831,f725])).
fof(f3610,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0')))),'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),'F_0')) )),
  inference(forward_demodulation,[],[f3578,f914])).
fof(f3578,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))))),'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f906,f1782,f884,f831,f725])).
fof(f3582,plain,(
  ( ! [X0] : (~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f889,f1164,f882,f831,f725])).
fof(f3584,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f889,f1649,f882,f831,f725])).
fof(f3585,plain,(
  ( ! [X0,X1] : (~sP3('T_0','const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f889,f1786,f882,f831,f725])).
fof(f3609,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2)))),'F_0')) )),
  inference(forward_demodulation,[],[f3588,f914])).
fof(f3588,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s(X1,X2)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f905,f1649,f902,f831,f725])).
fof(f3608,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'T_0')))),'F_0')) )),
  inference(forward_demodulation,[],[f3589,f914])).
fof(f3589,plain,(
  ( ! [X2,X0,X1] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f905,f1786,f902,f831,f725])).
fof(f3598,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1645,f883,f886,f831,f725])).
fof(f3599,plain,(
  ( ! [X0,X1] : (~sP2('F_0','const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0')))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1782,f883,f886,f831,f725])).
fof(f3520,plain,(
  ( ! [X0,X1] : (~sP2('const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'F_0'),s(X0,X1)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1649,f884,f887,f831,f724])).
fof(f3521,plain,(
  ( ! [X0,X1] : (~sP2('const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'F_0')))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1786,f884,f887,f831,f724])).
fof(f3459,plain,(
  ( ! [X0,X1] : (~sP3('const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f890,f1645,f884,f478,f723])).
fof(f3460,plain,(
  ( ! [X0,X1] : (~sP3('const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0')))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f890,f1782,f884,f478,f723])).
fof(f3462,plain,(
  ( ! [X0] : (~sP3('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f905,f1162,f884,f478,f723])).
fof(f3468,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f890,f882,f885,f478,f723])).
fof(f3474,plain,(
  ( ! [X0] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f904,f882,f903,f478,f723])).
fof(f3477,plain,(
  ( ! [X0,X1] : (sP5('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f904,f902,f903,f478,f723])).
fof(f3409,plain,(
  ( ! [X0] : (~sP5('T_0','F_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f3374,f721])).
fof(f3396,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(unit_resulting_resolution,[],[f486,f622])).
fof(f3384,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,X16)) | p(s(bool,X17)) | ~sP4('const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.real.real_lte_2'(s('type.realax.real',X18),s('type.realax.real',X18)),X17)) )),
  inference(superposition,[],[f1332,f688])).
fof(f3383,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~p(s(bool,X10)) | p(s(bool,X12)) | ~sP4('const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))),'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real',X13)),X12)) )),
  inference(superposition,[],[f1332,f687])).
fof(f3390,plain,(
  ( ! [X6,X8,X7] : (~sP4('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f3382,f478])).
fof(f3382,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | ~sP4('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.real.real_lte_2'(s('type.realax.real',X8),s('type.realax.real',X8)),X7)) )),
  inference(superposition,[],[f1332,f914])).
fof(f3374,plain,(
  ( ! [X0] : (~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f478,f831,f1332])).
fof(f3344,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X6))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X7))) | s('type.num.num',X7) = s('type.num.num','const.pair.SND_1'(s('type.pair.prod'(X4,'type.num.num'),'const.pair.,_2'(s(X4,X5),s('type.num.num',X6)))))) )),
  inference(superposition,[],[f621,f688])).
fof(f3343,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X3))) | s('type.num.num',X3) = s('type.num.num','const.pair.FST_1'(s('type.pair.prod'('type.num.num',X0),'const.pair.,_2'(s('type.num.num',X1),s(X0,X2)))))) )),
  inference(superposition,[],[f621,f687])).
fof(f3333,plain,(
  ( ! [X2,X0,X1] : (sP0('const.pair.SND_1'(s('type.pair.prod'(X0,bool),'const.pair.,_2'(s(X0,X1),s(bool,'T_0')))),X2)) )),
  inference(unit_resulting_resolution,[],[f478,f1827])).
fof(f3306,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(subsumption_resolution,[],[f3273,f831])).
fof(f3273,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | s(bool,'T_0') = s(bool,'const.realax.real_lt_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(superposition,[],[f617,f516])).
fof(f3265,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X28),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X27),s('type.realax.real',X26)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X26),s('type.realax.real',X27))),s('type.realax.real',X28))))) )),
  inference(superposition,[],[f617,f595])).
fof(f3260,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X14)))))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X15))),s('type.realax.real',X16))))) )),
  inference(superposition,[],[f617,f596])).
fof(f3246,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X27),s('type.realax.real',X26))),s('type.realax.real',X28)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X26),s('type.realax.real',X27))))))) )),
  inference(superposition,[],[f617,f595])).
fof(f3241,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X15),s('type.realax.real',X14))),s('type.realax.real',X16)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X16),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X14),s('type.realax.real',X15))))))) )),
  inference(superposition,[],[f617,f596])).
fof(f3225,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,'T_0')))))) )),
  inference(unit_resulting_resolution,[],[f478,f1826])).
fof(f3052,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X23))),s('type.realax.real',X25)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X23),s('type.realax.real',X24))))))) )),
  inference(superposition,[],[f616,f595])).
fof(f3047,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11))),s('type.realax.real',X13)))) | ~p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X13),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X11),s('type.realax.real',X12))))))) )),
  inference(superposition,[],[f616,f596])).
fof(f3007,plain,(
  ( ! [X2,X0,X1] : (sP0('const.pair.FST_1'(s('type.pair.prod'(bool,X0),'const.pair.,_2'(s(bool,'T_0'),s(X0,X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f478,f1685])).
fof(f2952,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,'T_0'),s(X1,X2)))))) )),
  inference(unit_resulting_resolution,[],[f478,f1684])).
fof(f2837,plain,(
  ( ! [X28,X29,X27] : (s('type.realax.real',X29) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X28),s('type.realax.real',X27))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X27),s('type.realax.real',X28))),s('type.realax.real',X29))))) )),
  inference(superposition,[],[f648,f596])).
fof(f2836,plain,(
  ( ! [X26,X24,X25] : (s('type.realax.real',X26) != s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X25),s('type.realax.real',X24))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X26),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X24),s('type.realax.real',X25))))))) )),
  inference(superposition,[],[f648,f596])).
fof(f2783,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f596,f648])).
fof(f2561,plain,(
  ( ! [X24,X23,X25] : (s('type.realax.real',X25) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X23))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X23),s('type.realax.real',X24))),s('type.realax.real',X25))))) )),
  inference(superposition,[],[f648,f595])).
fof(f2560,plain,(
  ( ! [X21,X22,X20] : (s('type.realax.real',X22) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X21),s('type.realax.real',X20))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X22),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X20),s('type.realax.real',X21))))))) )),
  inference(superposition,[],[f648,f595])).
fof(f2512,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X0))))))) )),
  inference(unit_resulting_resolution,[],[f595,f648])).
fof(f2437,plain,(
  ( ! [X4,X5,X3] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,'F_0')))),X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f2176,f516])).
fof(f2436,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f2176,f516])).
fof(f2341,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.pair.SND_1'(s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X18),s(bool,X16)))))) )),
  inference(superposition,[],[f2058,f688])).
fof(f2339,plain,(
  ( ! [X6,X8,X7] : (sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.pair.SND_1'(s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X8),s(bool,'T_0')))))) )),
  inference(superposition,[],[f2058,f914])).
fof(f2338,plain,(
  ( ! [X4,X5,X3] : (sP4('T_0',X3,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f2058,f516])).
fof(f2304,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X18)))),'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))))) )),
  inference(superposition,[],[f1951,f688])).
fof(f2301,plain,(
  ( ! [X4,X5,X3] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5)))),X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1951,f516])).
fof(f2300,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1951,f516])).
fof(f2204,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP4('T_0','const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X18)))))) )),
  inference(superposition,[],[f1908,f688])).
fof(f2203,plain,(
  ( ! [X12,X10,X13,X11,X9] : (sP4('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X10),s(X12,X13)))))) )),
  inference(superposition,[],[f1908,f687])).
fof(f2202,plain,(
  ( ! [X6,X8,X7] : (sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.pair.FST_1'(s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,'T_0'),s(X7,X8)))))) )),
  inference(superposition,[],[f1908,f914])).
fof(f2201,plain,(
  ( ! [X4,X5,X3] : (sP4('T_0',X3,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1908,f516])).
fof(f2195,plain,(
  ( ! [X6,X8,X7] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.pair.SND_1'(s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X8),s(bool,'T_0')))))) )),
  inference(superposition,[],[f1787,f914])).
fof(f2194,plain,(
  ( ! [X4,X5,X3] : (sP5('T_0',X3,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1787,f516])).
fof(f2193,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','F_0','const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1787,f516])).
fof(f2113,plain,(
  s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f496,f491])).
fof(f2097,plain,(
  ( ! [X6,X8,X7] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.pair.SND_1'(s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X8),s(bool,'T_0')))))) )),
  inference(superposition,[],[f1786,f914])).
fof(f2096,plain,(
  ( ! [X4,X5,X3] : (~sP5('F_0',X3,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,'F_0'))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1786,f516])).
fof(f2095,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','F_0','const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1786,f516])).
fof(f2078,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X18),s(bool,X16)))),'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))))) )),
  inference(superposition,[],[f1783,f688])).
fof(f2076,plain,(
  ( ! [X6,X8,X7] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X8),s(bool,'T_0')))),'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f1783,f914])).
fof(f2075,plain,(
  ( ! [X4,X5,X3] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,'F_0')))),X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1783,f516])).
fof(f2074,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1783,f516])).
fof(f2058,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0,'const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))))) )),
  inference(unit_resulting_resolution,[],[f1783,f721])).
fof(f1990,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X17,bool),'const.pair.,_2'(s(X17,X18),s(bool,X16)))),'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))))) )),
  inference(superposition,[],[f1782,f688])).
fof(f1988,plain,(
  ( ! [X6,X8,X7] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X7,bool),'const.pair.,_2'(s(X7,X8),s(bool,'T_0')))),'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f1782,f914])).
fof(f1987,plain,(
  ( ! [X4,X5,X3] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,'F_0')))),X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1782,f516])).
fof(f1986,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X1,bool),'const.pair.,_2'(s(X1,X2),s(bool,X0)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1782,f516])).
fof(f1958,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP5('T_0','const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X18)))))) )),
  inference(superposition,[],[f1650,f688])).
fof(f1956,plain,(
  ( ! [X6,X8,X7] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.pair.FST_1'(s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,'T_0'),s(X7,X8)))))) )),
  inference(superposition,[],[f1650,f914])).
fof(f1955,plain,(
  ( ! [X4,X5,X3] : (sP5('T_0',X3,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1650,f516])).
fof(f1954,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1650,f516])).
fof(f1935,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~sP5('F_0','const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X18)))))) )),
  inference(superposition,[],[f1649,f688])).
fof(f1933,plain,(
  ( ! [X6,X8,X7] : (~sP5('F_0','const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),'const.pair.FST_1'(s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,'T_0'),s(X7,X8)))))) )),
  inference(superposition,[],[f1649,f914])).
fof(f1932,plain,(
  ( ! [X4,X5,X3] : (~sP5('F_0',X3,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1649,f516])).
fof(f1931,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1649,f516])).
fof(f1915,plain,(
  ( ! [X14,X17,X15,X18,X16] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X17),'const.pair.,_2'(s(bool,X16),s(X17,X18)))),'const.pair.SND_1'(s('type.pair.prod'(X14,bool),'const.pair.,_2'(s(X14,X15),s(bool,X16)))))) )),
  inference(superposition,[],[f1646,f688])).
fof(f1914,plain,(
  ( ! [X12,X10,X13,X11,X9] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X10),s(X12,X13)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))))) )),
  inference(superposition,[],[f1646,f687])).
fof(f1913,plain,(
  ( ! [X6,X8,X7] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,'T_0'),s(X7,X8)))),'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f1646,f914])).
fof(f1912,plain,(
  ( ! [X4,X5,X3] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5)))),X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1646,f516])).
fof(f1911,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1646,f516])).
fof(f1908,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',X0,'const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))))) )),
  inference(unit_resulting_resolution,[],[f1646,f721])).
fof(f1845,plain,(
  ( ! [X134,X138,X136,X135,X137] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X137),'const.pair.,_2'(s(bool,X136),s(X137,X138)))),'const.pair.SND_1'(s('type.pair.prod'(X134,bool),'const.pair.,_2'(s(X134,X135),s(bool,X136)))))) )),
  inference(superposition,[],[f1645,f688])).
fof(f1844,plain,(
  ( ! [X132,X130,X133,X131,X129] : (~p(s(bool,X131)) | ~sP8('const.pair.SND_1'(s('type.pair.prod'(X129,bool),'const.pair.,_2'(s(X129,X130),s(bool,X131)))),X132,X133) | sP9('const.pair.SND_1'(s('type.pair.prod'(X129,bool),'const.pair.,_2'(s(X129,X130),s(bool,X131)))),X132,X133)) )),
  inference(superposition,[],[f814,f688])).
fof(f1843,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,X126)) | ~sP10('const.pair.SND_1'(s('type.pair.prod'(X124,bool),'const.pair.,_2'(s(X124,X125),s(bool,X126)))),X127,X128) | sP11('const.pair.SND_1'(s('type.pair.prod'(X124,bool),'const.pair.,_2'(s(X124,X125),s(bool,X126)))),X127,X128)) )),
  inference(superposition,[],[f813,f688])).
fof(f1841,plain,(
  ( ! [X118,X116,X119,X117,X115] : (~p(s(bool,X117)) | ~p(s(bool,X118)) | p(s(bool,X119)) | ~sP11('const.pair.SND_1'(s('type.pair.prod'(X115,bool),'const.pair.,_2'(s(X115,X116),s(bool,X117)))),X119,X118)) )),
  inference(superposition,[],[f790,f688])).
fof(f1840,plain,(
  ( ! [X111,X114,X112,X110,X113] : (~p(s(bool,X112)) | ~p(s(bool,X113)) | sP11('const.pair.SND_1'(s('type.pair.prod'(X110,bool),'const.pair.,_2'(s(X110,X111),s(bool,X112)))),X113,X114)) )),
  inference(superposition,[],[f789,f688])).
fof(f1839,plain,(
  ( ! [X109,X107,X105,X108,X106] : (~p(s(bool,X107)) | p(s(bool,X108)) | sP11('const.pair.SND_1'(s('type.pair.prod'(X105,bool),'const.pair.,_2'(s(X105,X106),s(bool,X107)))),X109,X108)) )),
  inference(superposition,[],[f788,f688])).
fof(f1838,plain,(
  ( ! [X103,X101,X102,X100,X104] : (~p(s(bool,X102)) | p(s(bool,X103)) | ~p(s(bool,X104)) | ~sP10(X104,X103,'const.pair.SND_1'(s('type.pair.prod'(X100,bool),'const.pair.,_2'(s(X100,X101),s(bool,X102)))))) )),
  inference(superposition,[],[f783,f688])).
fof(f1837,plain,(
  ( ! [X99,X97,X95,X98,X96] : (~p(s(bool,X97)) | p(s(bool,X98)) | p(s(bool,X99)) | ~sP9('const.pair.SND_1'(s('type.pair.prod'(X95,bool),'const.pair.,_2'(s(X95,X96),s(bool,X97)))),X99,X98)) )),
  inference(superposition,[],[f773,f688])).
fof(f1836,plain,(
  ( ! [X94,X92,X90,X93,X91] : (~p(s(bool,X92)) | ~p(s(bool,X93)) | sP9('const.pair.SND_1'(s('type.pair.prod'(X90,bool),'const.pair.,_2'(s(X90,X91),s(bool,X92)))),X93,X94)) )),
  inference(superposition,[],[f772,f688])).
fof(f1835,plain,(
  ( ! [X88,X87,X85,X89,X86] : (~p(s(bool,X87)) | ~p(s(bool,X88)) | sP9('const.pair.SND_1'(s('type.pair.prod'(X85,bool),'const.pair.,_2'(s(X85,X86),s(bool,X87)))),X89,X88)) )),
  inference(superposition,[],[f771,f688])).
fof(f1834,plain,(
  ( ! [X80,X83,X81,X84,X82] : (~p(s(bool,X82)) | p(s(bool,X83)) | p(s(bool,X84)) | ~sP8('const.pair.SND_1'(s('type.pair.prod'(X80,bool),'const.pair.,_2'(s(X80,X81),s(bool,X82)))),X83,X84)) )),
  inference(superposition,[],[f766,f688])).
fof(f1833,plain,(
  ( ! [X78,X76,X79,X77,X75] : (~p(s(bool,X77)) | ~p(s(bool,X78)) | p(s(bool,X79)) | ~sP7(X79,X78,'const.pair.SND_1'(s('type.pair.prod'(X75,bool),'const.pair.,_2'(s(X75,X76),s(bool,X77)))))) )),
  inference(superposition,[],[f758,f688])).
fof(f1832,plain,(
  ( ! [X70,X74,X72,X71,X73] : (~p(s(bool,X72)) | ~p(s(bool,X73)) | p(s(bool,X74)) | ~sP6(X73,'const.pair.SND_1'(s('type.pair.prod'(X70,bool),'const.pair.,_2'(s(X70,X71),s(bool,X72)))),X74)) )),
  inference(superposition,[],[f749,f688])).
fof(f1831,plain,(
  ( ! [X68,X66,X69,X67,X65] : (~p(s(bool,X67)) | ~p(s(bool,X68)) | p(s(bool,X69)) | ~sP3(X68,'const.pair.SND_1'(s('type.pair.prod'(X65,bool),'const.pair.,_2'(s(X65,X66),s(bool,X67)))),X69)) )),
  inference(superposition,[],[f730,f688])).
fof(f1830,plain,(
  ( ! [X61,X64,X62,X60,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | p(s(bool,X64)) | ~sP4(X63,'const.pair.SND_1'(s('type.pair.prod'(X60,bool),'const.pair.,_2'(s(X60,X61),s(bool,X62)))),X64)) )),
  inference(superposition,[],[f726,f688])).
fof(f1829,plain,(
  ( ! [X59,X57,X58,X56] : (~p(s(bool,X58)) | p(s(bool,X59)) | sP1(X59,'const.pair.SND_1'(s('type.pair.prod'(X56,bool),'const.pair.,_2'(s(X56,X57),s(bool,X58)))))) )),
  inference(superposition,[],[f673,f688])).
fof(f1828,plain,(
  ( ! [X54,X52,X55,X53] : (~p(s(bool,X54)) | p(s(bool,X55)) | sP1('const.pair.SND_1'(s('type.pair.prod'(X52,bool),'const.pair.,_2'(s(X52,X53),s(bool,X54)))),X55)) )),
  inference(superposition,[],[f670,f688])).
fof(f1827,plain,(
  ( ! [X50,X48,X51,X49] : (~p(s(bool,X50)) | sP0('const.pair.SND_1'(s('type.pair.prod'(X48,bool),'const.pair.,_2'(s(X48,X49),s(bool,X50)))),X51)) )),
  inference(superposition,[],[f655,f688])).
fof(f1826,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,X46)) | sP0(X47,'const.pair.SND_1'(s('type.pair.prod'(X44,bool),'const.pair.,_2'(s(X44,X45),s(bool,X46)))))) )),
  inference(superposition,[],[f654,f688])).
fof(f1814,plain,(
  ( ! [X8,X7] : (s('type.realax.real',X8) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) )),
  inference(superposition,[],[f688,f494])).
fof(f1813,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X6) | s(bool,'T_0') = s(bool,'const.pair.SND_1'(s('type.pair.prod'(X4,bool),'const.pair.,_2'(s(X4,X5),s(bool,X6)))))) )),
  inference(superposition,[],[f688,f516])).
fof(f1767,plain,(
  ( ! [X12,X10,X13,X11,X9] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X12),'const.pair.,_2'(s(bool,X10),s(X12,X13)))),'const.pair.FST_1'(s('type.pair.prod'(bool,X9),'const.pair.,_2'(s(bool,X10),s(X9,X11)))))) )),
  inference(superposition,[],[f1645,f687])).
fof(f1766,plain,(
  ( ! [X6,X8,X7] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X7),'const.pair.,_2'(s(bool,'T_0'),s(X7,X8)))),'const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)))) )),
  inference(superposition,[],[f1645,f914])).
fof(f1765,plain,(
  ( ! [X4,X5,X3] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,'F_0'),s(X4,X5)))),X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f1645,f516])).
fof(f1764,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.pair.FST_1'(s('type.pair.prod'(bool,X1),'const.pair.,_2'(s(bool,X0),s(X1,X2)))),'F_0') | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f1645,f516])).
fof(f1702,plain,(
  ( ! [X116,X114,X112,X115,X113] : (~p(s(bool,X113)) | ~sP8('const.pair.FST_1'(s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,X113),s(X112,X114)))),X115,X116) | sP9('const.pair.FST_1'(s('type.pair.prod'(bool,X112),'const.pair.,_2'(s(bool,X113),s(X112,X114)))),X115,X116)) )),
  inference(superposition,[],[f814,f687])).
fof(f1701,plain,(
  ( ! [X111,X109,X107,X110,X108] : (~p(s(bool,X108)) | ~sP10('const.pair.FST_1'(s('type.pair.prod'(bool,X107),'const.pair.,_2'(s(bool,X108),s(X107,X109)))),X110,X111) | sP11('const.pair.FST_1'(s('type.pair.prod'(bool,X107),'const.pair.,_2'(s(bool,X108),s(X107,X109)))),X110,X111)) )),
  inference(superposition,[],[f813,f687])).
fof(f1699,plain,(
  ( ! [X101,X99,X102,X100,X98] : (~p(s(bool,X99)) | ~p(s(bool,X101)) | p(s(bool,X102)) | ~sP11('const.pair.FST_1'(s('type.pair.prod'(bool,X98),'const.pair.,_2'(s(bool,X99),s(X98,X100)))),X102,X101)) )),
  inference(superposition,[],[f790,f687])).
fof(f1698,plain,(
  ( ! [X94,X97,X95,X93,X96] : (~p(s(bool,X94)) | ~p(s(bool,X96)) | sP11('const.pair.FST_1'(s('type.pair.prod'(bool,X93),'const.pair.,_2'(s(bool,X94),s(X93,X95)))),X96,X97)) )),
  inference(superposition,[],[f789,f687])).
fof(f1697,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,X89)) | p(s(bool,X91)) | sP11('const.pair.FST_1'(s('type.pair.prod'(bool,X88),'const.pair.,_2'(s(bool,X89),s(X88,X90)))),X92,X91)) )),
  inference(superposition,[],[f788,f687])).
fof(f1696,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X86)) | ~p(s(bool,X87)) | ~sP10(X87,X86,'const.pair.FST_1'(s('type.pair.prod'(bool,X83),'const.pair.,_2'(s(bool,X84),s(X83,X85)))))) )),
  inference(superposition,[],[f783,f687])).
fof(f1695,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X79)) | p(s(bool,X81)) | p(s(bool,X82)) | ~sP9('const.pair.FST_1'(s('type.pair.prod'(bool,X78),'const.pair.,_2'(s(bool,X79),s(X78,X80)))),X82,X81)) )),
  inference(superposition,[],[f773,f687])).
fof(f1694,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,X74)) | ~p(s(bool,X76)) | sP9('const.pair.FST_1'(s('type.pair.prod'(bool,X73),'const.pair.,_2'(s(bool,X74),s(X73,X75)))),X76,X77)) )),
  inference(superposition,[],[f772,f687])).
fof(f1693,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,X69)) | ~p(s(bool,X71)) | sP9('const.pair.FST_1'(s('type.pair.prod'(bool,X68),'const.pair.,_2'(s(bool,X69),s(X68,X70)))),X72,X71)) )),
  inference(superposition,[],[f771,f687])).
fof(f1692,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X64)) | p(s(bool,X66)) | p(s(bool,X67)) | ~sP8('const.pair.FST_1'(s('type.pair.prod'(bool,X63),'const.pair.,_2'(s(bool,X64),s(X63,X65)))),X66,X67)) )),
  inference(superposition,[],[f766,f687])).
fof(f1691,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X59)) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP7(X62,X61,'const.pair.FST_1'(s('type.pair.prod'(bool,X58),'const.pair.,_2'(s(bool,X59),s(X58,X60)))))) )),
  inference(superposition,[],[f758,f687])).
fof(f1690,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X54)) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP6(X56,'const.pair.FST_1'(s('type.pair.prod'(bool,X53),'const.pair.,_2'(s(bool,X54),s(X53,X55)))),X57)) )),
  inference(superposition,[],[f749,f687])).
fof(f1689,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,X49)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP3(X51,'const.pair.FST_1'(s('type.pair.prod'(bool,X48),'const.pair.,_2'(s(bool,X49),s(X48,X50)))),X52)) )),
  inference(superposition,[],[f730,f687])).
fof(f1688,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,X44)) | ~p(s(bool,X46)) | p(s(bool,X47)) | ~sP4(X46,'const.pair.FST_1'(s('type.pair.prod'(bool,X43),'const.pair.,_2'(s(bool,X44),s(X43,X45)))),X47)) )),
  inference(superposition,[],[f726,f687])).
fof(f1687,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X42)) | sP1(X42,'const.pair.FST_1'(s('type.pair.prod'(bool,X39),'const.pair.,_2'(s(bool,X40),s(X39,X41)))))) )),
  inference(superposition,[],[f673,f687])).
fof(f1686,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X38)) | sP1('const.pair.FST_1'(s('type.pair.prod'(bool,X35),'const.pair.,_2'(s(bool,X36),s(X35,X37)))),X38)) )),
  inference(superposition,[],[f670,f687])).
fof(f1685,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | sP0('const.pair.FST_1'(s('type.pair.prod'(bool,X31),'const.pair.,_2'(s(bool,X32),s(X31,X33)))),X34)) )),
  inference(superposition,[],[f655,f687])).
fof(f1684,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,X28)) | sP0(X30,'const.pair.FST_1'(s('type.pair.prod'(bool,X27),'const.pair.,_2'(s(bool,X28),s(X27,X29)))))) )),
  inference(superposition,[],[f654,f687])).
fof(f1675,plain,(
  ( ! [X8,X7] : (s('type.realax.real',X7) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X7),s('type.realax.real',X8)))))) )),
  inference(superposition,[],[f687,f493])).
fof(f1674,plain,(
  ( ! [X6,X4,X5] : (s(bool,'F_0') = s(bool,X5) | s(bool,'T_0') = s(bool,'const.pair.FST_1'(s('type.pair.prod'(bool,X4),'const.pair.,_2'(s(bool,X5),s(X4,X6)))))) )),
  inference(superposition,[],[f687,f516])).
fof(f1460,plain,(
  ( ! [X6,X8,X7] : (~sP9('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8,X7) | p(s(bool,X8)) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f1457,f478])).
fof(f1457,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | p(s(bool,X8)) | ~sP9('const.real.real_lte_2'(s('type.realax.real',X6),s('type.realax.real',X6)),X8,X7)) )),
  inference(superposition,[],[f773,f914])).
fof(f1367,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f891,f478,f478,f745])).
fof(f1368,plain,(
  ( ! [X0] : (sP7('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f907,f478,f478,f745])).
fof(f1269,plain,(
  ( ! [X0,X1] : (sP11('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f478,f1096])).
fof(f1187,plain,(
  ( ! [X0,X1] : (sP9('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f478,f1094])).
fof(f1153,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f478,f911,f740])).
fof(f1140,plain,(
  ( ! [X0] : (sP5('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f914,f739])).
fof(f1141,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)))) )),
  inference(unit_resulting_resolution,[],[f478,f914,f739])).
fof(f1130,plain,(
  ( ! [X0,X1] : (sP9('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f1093])).
fof(f1111,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(subsumption_resolution,[],[f1106,f831])).
fof(f1106,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,X3)) | s(bool,X2) = s(bool,X3) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f612,f516])).
fof(f1098,plain,(
  ( ! [X30,X28,X29] : (sP9('const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real',X28)),X29,X30) | ~sP8('const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real',X28)),X29,X30)) )),
  inference(subsumption_resolution,[],[f1085,f478])).
fof(f1085,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | ~sP8('const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real',X28)),X29,X30) | sP9('const.real.real_lte_2'(s('type.realax.real',X28),s('type.realax.real',X28)),X29,X30)) )),
  inference(superposition,[],[f814,f914])).
fof(f1097,plain,(
  ( ! [X26,X27,X25] : (sP11('const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',X25)),X26,X27) | ~sP10('const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',X25)),X26,X27)) )),
  inference(subsumption_resolution,[],[f1084,f478])).
fof(f1084,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | ~sP10('const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',X25)),X26,X27) | sP11('const.real.real_lte_2'(s('type.realax.real',X25),s('type.realax.real',X25)),X26,X27)) )),
  inference(superposition,[],[f813,f914])).
fof(f1000,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f908,f478,f814])).
fof(f1002,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f909,f478,f814])).
fof(f991,plain,(
  ( ! [X0,X1] : (sP11('T_0','const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f910,f478,f813])).
fof(f979,plain,(
  ( ! [X0] : (sP11('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f478,f478,f789])).
fof(f962,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f478,f478,f772])).
fof(f959,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f478,f771])).
fof(f921,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f478,f478,f669])).
fof(f900,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X1)))) )),
  inference(unit_resulting_resolution,[],[f488,f654])).
fof(f901,plain,(
  ( ! [X0,X1] : (sP0('const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X0)),X1)) )),
  inference(unit_resulting_resolution,[],[f488,f655])).
fof(f903,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f488,f731])).
fof(f907,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(unit_resulting_resolution,[],[f488,f750])).
fof(f890,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f478,f737])).
fof(f881,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f478,f655])).
fof(f880,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f478,f654])).
fof(f487,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(flattening,[],[f36])).
fof(f36,axiom,(
  ~s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_|49||48|')).
fof(f504,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_RID')).
fof(f505,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_LID')).
fof(f508,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_RZERO')).
fof(f509,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_RINV')).
fof(f510,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_LINV')).
fof(f528,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f305])).
fof(f305,plain,(
  ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),
  inference(ennf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),
  inference(flattening,[],[f182])).
fof(f182,plain,(
  ! [X0] : (~s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : (~s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.pair.prod'('type.realax.real','type.realax.real'),X6))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_LINV')).
fof(f538,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X0] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) & (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X6] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_|48|_THM')).
fof(f539,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f354])).
fof(f635,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(nnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_LID_UNIQ')).
fof(f636,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f404])).
fof(f637,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f405])).
fof(f405,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(nnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_RID_UNIQ')).
fof(f638,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f405])).
fof(f639,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f406])).
fof(f406,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(nnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_LNEG_UNIQ')).
fof(f640,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f406])).
fof(f641,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(nnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_RNEG_UNIQ')).
fof(f642,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f770,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f464])).
fof(f815,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f748])).
fof(f748,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f452])).
fof(f705,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),X2)) )),
  inference(cnf_transformation,[],[f433])).
fof(f433,plain,(
  ! [X0,X1,X2] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),X2)) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),X2) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))))),
  inference(nnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),X2))),
  inference(rectify,[],[f112])).
fof(f112,axiom,(
  ! [X6,X7,X14] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X14))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),X14))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_EQ_LADD')).
fof(f706,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),X2) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))) )),
  inference(cnf_transformation,[],[f433])).
fof(f704,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))))) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X6,X7,X14] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))),s('type.pair.prod'('type.realax.real','type.realax.real'),X14))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X14))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.pair.prod'('type.realax.real','type.realax.real'),X14)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_RDISTRIB')).
fof(f703,plain,(
  ( ! [X2,X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X2))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X2)))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6,X7,X14] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.pair.prod'('type.realax.real','type.realax.real'),X14))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X14)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_LDISTRIB')).
fof(f701,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f111])).
fof(f111,axiom,(
  ! [X9,X13,X6] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X13),s('type.realax.real',X6))))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X13))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X6)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_SUB_LDISTRIB')).
fof(f700,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2)))))),
  inference(rectify,[],[f100])).
fof(f100,axiom,(
  ! [X9,X13,X6] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X6))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X13))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X6)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.realax.REAL_LDISTRIB')).
fof(f699,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))),
  inference(rectify,[],[f106])).
fof(f106,axiom,(
  ! [X9,X13,X6] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X13))),s('type.realax.real',X6))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X6)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_RDISTRIB')).
fof(f698,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))) )),
  inference(cnf_transformation,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X2))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2)))))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X9,X13,X6] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X9),s('type.realax.real',X13))),s('type.realax.real',X6))) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X6))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X6)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_SUB_RDISTRIB')).
fof(f697,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X9,X13,X6] : s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X13),s('type.realax.real',X6))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X13))),s('type.realax.real',X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.realax.REAL_ADD_ASSOC')).
fof(f696,plain,(
  ( ! [X2,X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))) )),
  inference(cnf_transformation,[],[f274])).
fof(f274,plain,(
  ! [X0,X1,X2] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X2))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real',X2)))),
  inference(rectify,[],[f107])).
fof(f107,axiom,(
  ! [X9,X13,X6] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X6))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X13))),s('type.realax.real',X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.realax.REAL_MUL_ASSOC')).
fof(f693,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X0,X2) = s(X0,X4)) )),
  inference(cnf_transformation,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(flattening,[],[f431])).
fof(f431,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5))) & ((s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5)) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))),
  inference(nnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) <=> (s(X0,X2) = s(X0,X4) & s(X1,X3) = s(X1,X5)))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X8,X12] : ! [X9,X13,X22,X23] : (s('type.pair.prod'(X8,X12),'const.pair.,_2'(s(X8,X9),s(X12,X13))) = s('type.pair.prod'(X8,X12),'const.pair.,_2'(s(X8,X22),s(X12,X23))) <=> (s(X8,X9) = s(X8,X22) & s(X12,X13) = s(X12,X23)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.pair.CLOSED_PAIR_EQ')).
fof(f694,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) != s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))) | s(X1,X3) = s(X1,X5)) )),
  inference(cnf_transformation,[],[f432])).
fof(f695,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,X2) != s(X0,X4) | s(X1,X3) != s(X1,X5) | s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3))) = s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f432])).
fof(f690,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f430])).
fof(f430,plain,(
  ! [X0,X1,X2,X3] : ((s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))) & (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) != s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) | s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3)))),
  inference(flattening,[],[f429])).
fof(f429,plain,(
  ! [X0,X1,X2,X3] : ((s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))))) & ((s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) != s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) | s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))) | s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3)))),
  inference(nnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2,X3] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3) <=> (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1] : ! [X2,X3] : (s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3) <=> (s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) = s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3))) & s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))))),
  inference(rectify,[],[f103])).
fof(f103,axiom,(
  ! [X8,X12] : ! [X18,X17] : (s('type.pair.prod'(X8,X12),X17) = s('type.pair.prod'(X8,X12),X18) <=> (s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X12),X17))) = s(X8,'const.pair.FST_1'(s('type.pair.prod'(X8,X12),X18))) & s(X12,'const.pair.SND_1'(s('type.pair.prod'(X8,X12),X17))) = s(X12,'const.pair.SND_1'(s('type.pair.prod'(X8,X12),X18)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.pair.PAIR_FST_SND_EQ')).
fof(f691,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.pair.prod'(X0,X1),X2) != s('type.pair.prod'(X0,X1),X3) | s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f430])).
fof(f692,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X2))) != s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X3))) | s('type.pair.prod'(X0,X1),X2) = s('type.pair.prod'(X0,X1),X3) | s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X2))) != s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X3)))) )),
  inference(cnf_transformation,[],[f430])).
fof(f689,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f428])).
fof(f428,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK16(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK16(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f321])).
fof(f321,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f268])).
fof(f268,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.complex.COMPLEX_ENTIRE',aHL_EXT)).
fof(f648,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real',X1) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f643,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f409])).
fof(f409,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1))))) | (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1)))))))),
  inference(flattening,[],[f408])).
fof(f408,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1))))) | (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & ((s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1)))))))),
  inference(nnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1))))) <=> (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f102])).
fof(f102,axiom,(
  ! [X9,X13] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X9))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X13),s('type.realax.real',X13))))) <=> (s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.realax.real',X13) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_SUMSQ')).
fof(f644,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1))))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f409])).
fof(f645,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X0))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X1),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f409])).
fof(f633,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f403])).
fof(f403,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))))),
  inference(nnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X10,X24] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X10))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X24))))) <=> s('type.num.num',X10) = s('type.num.num',X24))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.eq_ints|split|3')).
fof(f634,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f403])).
fof(f627,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f400])).
fof(f400,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))))),
  inference(flattening,[],[f399])).
fof(f399,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1))) & ((s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1)) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))))),
  inference(nnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X0) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)))),
  inference(rectify,[],[f131])).
fof(f131,axiom,(
  ! [X10,X24] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X10))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X24))))) <=> (s('type.num.num','const.num.|48|_0') = s('type.num.num',X10) & s('type.num.num','const.num.|48|_0') = s('type.num.num',X24)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.eq_ints|split|2')).
fof(f628,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f400])).
fof(f629,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f400])).
fof(f626,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f398])).
fof(f398,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))))),
  inference(nnf_transformation,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),
  inference(rectify,[],[f134])).
fof(f134,axiom,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) <=> s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_NEG_EQ')).
fof(f623,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X0) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f397])).
fof(f397,plain,(
  ! [X0,X1] : ((s('type.realax.real',X0) != s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) | s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(nnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0,X1] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))) <=> s('type.realax.real',X1) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f105])).
fof(f105,axiom,(
  ! [X9,X13] : (s('type.realax.real',X9) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X13))) <=> s('type.realax.real',X13) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_ADD_RID_UNIQ')).
fof(f624,plain,(
  ( ! [X0,X1] : (s('type.realax.real',X1) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))) )),
  inference(cnf_transformation,[],[f397])).
fof(f621,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f396])).
fof(f396,plain,(
  ! [X0,X1] : ((s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f239])).
fof(f239,plain,(
  ! [X0,X1] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f120])).
fof(f120,axiom,(
  ! [X10,X24] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X10))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X24))) <=> s('type.num.num',X10) = s('type.num.num',X24))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.eq_ints|split|0')).
fof(f622,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f396])).
fof(f618,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f395])).
fof(f395,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) & (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))),
  inference(flattening,[],[f394])).
fof(f394,plain,(
  ! [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) & ((s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))),
  inference(nnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) <=> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),X7) <=> (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) & s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_RE_IM_EQ')).
fof(f619,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f395])).
fof(f620,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),X1) | s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) != s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(cnf_transformation,[],[f395])).
fof(f617,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0)))) | p(s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) )),
  inference(cnf_transformation,[],[f393])).
fof(f612,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f314])).
fof(f611,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))))) )),
  inference(cnf_transformation,[],[f234])).
fof(f234,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))),s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.complex_mul')).
fof(f610,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))) )),
  inference(cnf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.complex_add')).
fof(f834,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))) )),
  inference(backward_demodulation,[],[f609,f607])).
fof(f609,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))) )),
  inference(cnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(rectify,[],[f91])).
fof(f91,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7))))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_NEG_RMUL')).
fof(f608,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))) )),
  inference(cnf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_NEG_MUL2')).
fof(f607,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))) )),
  inference(cnf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1)))))),
  inference(rectify,[],[f89])).
fof(f89,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_NEG_LMUL')).
fof(f606,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_add_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ADD_COMM')).
fof(f605,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f228])).
fof(f228,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f115])).
fof(f115,axiom,(
  ! [X9,X13] : s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X13))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_LE_LNEG')).
fof(f604,plain,(
  ( ! [X0,X1] : (s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f227])).
fof(f227,plain,(
  ! [X0,X1] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f117])).
fof(f117,axiom,(
  ! [X9,X13] : s(bool,'const.real.real_lte_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))))) = s(bool,'const.real.real_lte_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X13))),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_LE_RNEG')).
fof(f603,plain,(
  ( ! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X1),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X7),s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_COMM')).
fof(f833,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(forward_demodulation,[],[f602,f598])).
fof(f602,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X9,X13] : s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real',X13))))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_NEG_ADD')).
fof(f832,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1)))) )),
  inference(backward_demodulation,[],[f601,f599])).
fof(f601,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f119])).
fof(f119,axiom,(
  ! [X9,X13] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13))))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_MUL_RNEG')).
fof(f599,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f222])).
fof(f222,plain,(
  ! [X0,X1] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real',X1)))))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X9,X13] : s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real',X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_MUL_LNEG')).
fof(f598,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X9,X13] : s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X9),s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.real_sub')).
fof(f597,plain,(
  ( ! [X0,X1] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))) )),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  ! [X0,X1] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X1))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X1)))))),
  inference(rectify,[],[f126])).
fof(f126,axiom,(
  ! [X9,X13] : s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X13))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.inv_1'(s('type.realax.real',X13)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.real_div')).
fof(f589,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f389])).
fof(f389,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK14(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15])],[f388])).
fof(f388,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f387])).
fof(f387,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f312])).
fof(f312,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X8] : ! [X20] : (~! [X9] : p(s(bool,happ(s(fun(X8,bool),X20),s(X8,X9)))) <=> ? [X9] : ~p(s(bool,happ(s(fun(X8,bool),X20),s(X8,X9)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.bool.NOT_FORALL_THM')).
fof(f590,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK15(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f389])).
fof(f535,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f352])).
fof(f527,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(ennf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0] : (s('type.realax.real',X0) != s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0] : (~s('type.realax.real',X0) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X0),s('type.realax.real',X0))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X9] : (~s('type.realax.real',X9) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.realax.real','const.real.|47|_2'(s('type.realax.real',X9),s('type.realax.real',X9))) = s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_DIV_REFL')).
fof(f514,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_inv_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),s('type.realax.real','const.real.|47|_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))))),s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real','const.real.pow_2'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.complex_inv')).
fof(f513,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))))))),
  inference(rectify,[],[f104])).
fof(f104,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))))),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.complex_neg')).
fof(f511,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_MUL_LID')).
fof(f503,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real',X9))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.pair.,_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.complex_of_real')).
fof(f502,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f99])).
fof(f99,axiom,(
  ! [X9] : s('type.realax.real',X9) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_MUL_RID')).
fof(f501,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X0)))),
  inference(rectify,[],[f129])).
fof(f129,axiom,(
  ! [X9] : s('type.realax.real',X9) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X9)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_MUL_LID')).
fof(f500,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X6] : s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_neg_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_NEGNEG')).
fof(f499,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X0))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10] : s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num',X10))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num',X10)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.complex_of_num')).
fof(f498,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X9] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_MUL_RZERO')).
fof(f497,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0))),s('type.realax.real',X0)))),
  inference(rectify,[],[f130])).
fof(f130,axiom,(
  ! [X9] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9))),s('type.realax.real',X9)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_ADD_LINV')).
fof(f496,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f98])).
fof(f98,axiom,(
  ! [X9] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_ADD_RINV')).
fof(f495,plain,(
  ( ! [X0] : (s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X9] : s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_MUL_LZERO')).
fof(f494,plain,(
  ( ! [X0] : (s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0] : s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : s('type.realax.real','const.complex.IM_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.pair.SND_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.IM')).
fof(f493,plain,(
  ( ! [X0] : (s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X6] : s('type.realax.real','const.complex.RE_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6))) = s('type.realax.real','const.pair.FST_1'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.RE')).
fof(f492,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X9] : s('type.realax.real',X9) = s('type.realax.real','const.real.real_sub_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_SUB_RZERO')).
fof(f491,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X0)))),
  inference(rectify,[],[f92])).
fof(f92,axiom,(
  ! [X9] : s('type.realax.real',X9) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.realax.real',X9)))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_ADD_LID')).
fof(f490,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X0),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X9] : s('type.realax.real',X9) = s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X9),s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_ADD_RID')).
fof(f489,plain,(
  ( ! [X0] : (s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))) )),
  inference(cnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0] : s('type.realax.real',X0) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X0)))))),
  inference(rectify,[],[f124])).
fof(f124,axiom,(
  ! [X9] : s('type.realax.real',X9) = s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real','const.realax.real_neg_1'(s('type.realax.real',X9)))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.real.REAL_NEGNEG')).
fof(f480,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(cnf_transformation,[],[f46])).
fof(f46,axiom,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_real_1'(s('type.realax.real','const.real.real_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_|48|')).
fof(f11920,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886,f11890,f11897,f11901,f11908,f11912,f11919])).
fof(f11916,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11915,f478])).
fof(f11915,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11914,f11518])).
fof(f11914,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X23))),s('type.realax.real',X25))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11762,f478])).
fof(f11762,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_add_2'(s('type.realax.real',X24),s('type.realax.real',X23))),s('type.realax.real',X25))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3052])).
fof(f11913,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886,f11890,f11897,f11901,f11908,f11912])).
fof(f11909,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886,f11890,f11897,f11901,f11908])).
fof(f11905,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11904,f478])).
fof(f11904,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11903,f11518])).
fof(f11903,plain,(
  ( ! [X33,X32] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X32),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X33))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11757,f478])).
fof(f11757,plain,(
  ( ! [X33,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.real.pow_2'(s('type.realax.real',X32),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.realax.real',X33))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f5800])).
fof(f11902,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886,f11890,f11897,f11901])).
fof(f11898,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886,f11890,f11897])).
fof(f11894,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11893,f478])).
fof(f11893,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11892,f11518])).
fof(f11892,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11))),s('type.realax.real',X13))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11751,f478])).
fof(f11751,plain,(
  ( ! [X12,X13,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real','const.realax.real_mul_2'(s('type.realax.real',X12),s('type.realax.real',X11))),s('type.realax.real',X13))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f3047])).
fof(f11891,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886,f11890])).
fof(f11887,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882,f11886])).
fof(f11883,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878,f11882])).
fof(f11879,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740,f11878])).
fof(f11877,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739,f11740])).
fof(f11876,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738,f11739])).
fof(f11875,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737,f11738])).
fof(f11874,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736,f11737])).
fof(f11873,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735,f11736])).
fof(f11872,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734,f11735])).
fof(f11871,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733,f11734])).
fof(f11870,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732,f11733])).
fof(f11869,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731,f11732])).
fof(f11868,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730,f11731])).
fof(f11867,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729,f11730])).
fof(f11866,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864,f11729])).
fof(f11865,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862,f11864])).
fof(f11863,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860,f11862])).
fof(f11861,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858,f11860])).
fof(f11859,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856,f11858])).
fof(f11857,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854,f11856])).
fof(f11855,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852,f11854])).
fof(f11853,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850,f11852])).
fof(f11851,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848,f11850])).
fof(f11849,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846,f11848])).
fof(f11847,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844,f11846])).
fof(f11845,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842,f11844])).
fof(f11843,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840,f11842])).
fof(f11841,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838,f11840])).
fof(f11839,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836,f11838])).
fof(f11837,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834,f11836])).
fof(f11835,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832,f11834])).
fof(f11833,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830,f11832])).
fof(f11831,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828,f11830])).
fof(f11829,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826,f11828])).
fof(f11827,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824,f11826])).
fof(f11825,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822,f11824])).
fof(f11823,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820,f11822])).
fof(f11821,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818,f11820])).
fof(f11819,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680,f11818])).
fof(f11815,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11814,f478])).
fof(f11814,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(forward_demodulation,[],[f11813,f11518])).
fof(f11813,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | $spl8),
  inference(subsumption_resolution,[],[f11701,f478])).
fof(f11701,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.realax.real_lt_2'(s('type.realax.real',X1),s('type.realax.real',X0))))) ) | $spl8),
  inference(backward_demodulation,[],[f11518,f616])).
fof(f11806,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679,f11680])).
fof(f11805,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678,f11679])).
fof(f11804,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677,f11678])).
fof(f11803,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800,f11677])).
fof(f11802,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f11662,f478])).
fof(f11662,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f11518,f831])).
fof(f11801,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f477,f476,f475,f478,f480,f831,f486,f488,f489,f490,f491,f492,f493,f494,f495,f496,f497,f498,f499,f500,f501,f502,f503,f506,f511,f512,f513,f514,f516,f527,f531,f530,f535,f534,f590,f589,f595,f596,f597,f598,f599,f600,f601,f832,f833,f603,f604,f605,f606,f607,f608,f609,f834,f610,f611,f615,f612,f617,f616,f620,f619,f618,f622,f621,f624,f623,f626,f629,f628,f627,f634,f633,f645,f644,f643,f648,f646,f655,f654,f835,f673,f670,f669,f668,f686,f687,f688,f689,f692,f691,f690,f695,f694,f693,f696,f697,f698,f699,f700,f701,f702,f703,f704,f706,f705,f725,f724,f723,f722,f721,f720,f719,f729,f728,f727,f726,f733,f732,f731,f730,f737,f736,f735,f734,f741,f740,f739,f738,f815,f745,f744,f743,f742,f752,f751,f750,f749,f758,f755,f754,f753,f765,f764,f763,f814,f761,f760,f759,f769,f768,f767,f766,f773,f772,f771,f770,f782,f781,f780,f813,f778,f777,f776,f786,f785,f784,f783,f790,f789,f788,f787,f812,f642,f641,f640,f639,f638,f637,f636,f635,f539,f538,f529,f528,f510,f509,f508,f507,f505,f504,f487,f880,f881,f882,f883,f884,f885,f886,f887,f888,f889,f890,f891,f892,f893,f894,f895,f896,f897,f898,f899,f910,f909,f908,f907,f906,f905,f904,f903,f902,f901,f900,f911,f914,f920,f921,f924,f929,f936,f941,f942,f945,f951,f952,f956,f959,f962,f969,f970,f972,f979,f991,f1002,f1000,f1030,f1091,f1092,f1093,f1094,f1095,f1096,f1097,f1098,f1099,f1111,f1115,f1130,f1133,f1132,f1141,f1140,f1155,f1153,f1164,f1162,f1187,f1225,f1269,f1327,f1332,f1341,f1346,f1351,f1360,f1364,f1368,f1367,f1375,f1380,f1384,f1410,f1419,f1428,f1434,f1441,f1460,f1468,f1476,f1481,f1488,f1505,f1650,f1649,f1646,f1645,f1674,f1675,f1684,f1685,f1686,f1687,f1688,f1689,f1690,f1691,f1692,f1693,f1694,f1695,f1696,f1697,f1698,f1699,f1701,f1702,f1764,f1765,f1766,f1767,f1787,f1786,f1783,f1782,f1813,f1814,f1826,f1827,f1828,f1829,f1830,f1831,f1832,f1833,f1834,f1835,f1836,f1837,f1838,f1839,f1840,f1841,f1843,f1844,f1845,f1908,f1911,f1912,f1913,f1914,f1915,f1931,f1932,f1933,f1935,f1951,f1954,f1955,f1956,f1958,f1986,f1987,f1988,f1990,f2058,f2074,f2075,f2076,f2078,f2095,f2096,f2097,f2113,f2176,f2193,f2194,f2195,f2201,f2202,f2203,f2204,f2300,f2301,f2304,f2338,f2339,f2341,f2436,f2437,f2512,f2560,f2561,f2783,f2836,f2837,f2952,f3007,f3047,f3052,f3065,f3070,f3225,f3241,f3246,f3260,f3265,f3306,f3333,f3343,f3344,f3374,f3390,f3383,f3384,f3396,f3409,f3477,f3474,f3468,f3462,f3460,f3459,f3532,f3531,f3521,f3520,f3518,f3601,f3599,f3598,f3608,f3609,f3585,f3584,f3582,f3610,f3611,f4043,f4119,f4120,f4121,f4230,f4256,f4340,f4342,f4343,f4345,f4494,f4532,f4559,f4718,f4719,f4723,f4724,f4780,f4927,f5037,f5067,f5068,f5070,f5071,f5074,f5076,f5077,f5080,f5082,f5085,f5087,f5089,f5091,f5092,f5094,f5095,f5096,f5099,f5190,f5345,f5353,f5427,f5424,f5425,f5428,f5438,f5435,f5436,f5439,f5441,f5442,f5449,f5509,f5513,f5514,f5523,f5521,f5522,f5524,f5599,f5730,f5740,f5737,f5738,f5742,f5741,f5743,f5747,f5799,f5800,f5801,f5802,f5804,f5805,f5806,f5807,f5911,f5912,f5913,f5914,f5916,f5917,f5918,f5919,f6007,f6075,f6074,f6073,f6072,f6105,f6106,f6136,f6137,f6139,f6149,f6150,f6151,f6152,f6153,f6154,f6155,f6156,f6157,f6158,f6159,f6160,f6161,f6162,f6163,f6164,f6166,f6167,f6169,f6170,f6171,f6172,f6173,f6174,f6175,f6176,f6177,f6180,f6181,f6182,f6183,f6186,f6187,f6188,f6189,f6231,f6240,f6237,f6238,f6239,f6243,f6339,f6340,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6554,f6723,f6900,f6759,f6902,f6903,f7435,f7440,f7300,f7304,f7305,f7306,f7307,f7308,f7309,f7310,f7746,f7747,f7744,f8119,f8120,f8117,f8123,f8139,f8140,f8152,f8156,f8287,f8288,f8299,f8297,f8315,f8316,f8328,f8332,f8557,f8684,f8891,f8976,f9071,f9072,f9084,f9088,f9110,f9111,f9123,f9127,f9290,f9376,f9459,f9528,f9532,f9475,f9488,f9490,f9491,f9494,f9496,f9497,f9683,f9684,f9782,f9998,f9999,f10096,f10199,f10256,f10606,f10607,f10608,f10621,f10631,f10632,f10633,f10681,f10682,f10683,f10684,f10687,f10688,f10689,f10690,f10693,f10694,f10695,f10696,f10750,f10752,f10753,f10754,f10864,f10865,f10866,f10867,f10868,f10869,f10870,f10871,f10872,f11516,f11520,f11523,f11526,f11529,f11532,f11535,f11538,f11541,f11544,f11547,f11550,f11553,f11556,f11559,f11562,f11565,f11568,f11571,f11574,f11577,f11580,f10906,f10907,f10908,f10909,f10912,f10913,f10914,f10915,f10918,f10920,f11582,f11583,f11584,f11585,f11586,f11587,f11588,f11589,f11590,f11591,f11592,f11593,f11594,f11595,f11596,f11597,f11598,f11599,f11600,f11601,f11602,f11603,f11604,f11605,f11606,f11607,f10984,f10985,f10986,f10987,f10990,f10991,f10992,f10993,f10996,f10997,f10998,f10999,f11608,f11609,f11610,f11611,f11052,f11054,f11055,f11056,f11617,f11619,f11621,f11623,f11625,f11627,f11629,f11631,f11633,f11635,f11637,f11639,f11641,f11643,f11645,f11647,f11649,f11651,f11653,f11655,f11657,f11659,f11518,f11800])).
fof(f11581,plain,(
  ( ! [X2,X1] : (~sP11('const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1,'T_0') | p(s(bool,X1))) ) | $spl50),
  inference(cnf_transformation,[],[f11581_D])).
fof(f11581_D,plain,(
  ( ! [X2,X1] : (~sP11('const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1,'T_0') | p(s(bool,X1))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f11518,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f11518_D])).
fof(f11518_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f11659,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1,'T_0')) )),
  inference(condensation,[],[f11658])).
fof(f11658,plain,(
  ( ! [X156,X154,X155,X153] : (p(s(bool,X154)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X155),s('type.realax.real',X155)),X154,'T_0') | s(bool,X153) = s(bool,X156) | s(bool,'T_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f11099,f1111])).
fof(f11099,plain,(
  ( ! [X156,X154,X155,X153] : (~p(s(bool,X153)) | p(s(bool,X154)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X155),s('type.realax.real',X155)),X154,'T_0') | s(bool,X153) = s(bool,X156) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f1505,f1030])).
fof(f11578,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2))) | p(s(bool,X1))) ) | $spl48),
  inference(cnf_transformation,[],[f11578_D])).
fof(f11578_D,plain,(
  ( ! [X2,X1] : (~sP10('T_0',X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2))) | p(s(bool,X1))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f11657,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP10('T_0',X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(condensation,[],[f11656])).
fof(f11656,plain,(
  ( ! [X152,X151,X149,X150] : (p(s(bool,X150)) | ~sP10('T_0',X150,'const.real.real_lte_2'(s('type.realax.real',X151),s('type.realax.real',X151))) | s(bool,X149) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(subsumption_resolution,[],[f11098,f1111])).
fof(f11098,plain,(
  ( ! [X152,X151,X149,X150] : (~p(s(bool,X149)) | p(s(bool,X150)) | ~sP10('T_0',X150,'const.real.real_lte_2'(s('type.realax.real',X151),s('type.realax.real',X151))) | s(bool,X149) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f1481,f1030])).
fof(f11575,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2))) | p(s(bool,X1))) ) | $spl46),
  inference(cnf_transformation,[],[f11575_D])).
fof(f11575_D,plain,(
  ( ! [X2,X1] : (~sP7(X1,'T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2))) | p(s(bool,X1))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f11655,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7(X1,'T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(condensation,[],[f11654])).
fof(f11654,plain,(
  ( ! [X146,X147,X145,X148] : (p(s(bool,X146)) | ~sP7(X146,'T_0','const.real.real_lte_2'(s('type.realax.real',X147),s('type.realax.real',X147))) | s(bool,X145) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(subsumption_resolution,[],[f11097,f1111])).
fof(f11097,plain,(
  ( ! [X146,X147,X145,X148] : (~p(s(bool,X145)) | p(s(bool,X146)) | ~sP7(X146,'T_0','const.real.real_lte_2'(s('type.realax.real',X147),s('type.realax.real',X147))) | s(bool,X145) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f1410,f1030])).
fof(f11572,plain,(
  ( ! [X2,X1] : (~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1) | p(s(bool,X1))) ) | $spl44),
  inference(cnf_transformation,[],[f11572_D])).
fof(f11572_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1) | p(s(bool,X1))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f11653,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1)) )),
  inference(condensation,[],[f11652])).
fof(f11652,plain,(
  ( ! [X144,X142,X143,X141] : (p(s(bool,X142)) | ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X143),s('type.realax.real',X143)),X142) | s(bool,X141) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f11096,f1111])).
fof(f11096,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,X141)) | p(s(bool,X142)) | ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X143),s('type.realax.real',X143)),X142) | s(bool,X141) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f1380,f1030])).
fof(f11569,plain,(
  ( ! [X2,X1] : (~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1) | p(s(bool,X1))) ) | $spl42),
  inference(cnf_transformation,[],[f11569_D])).
fof(f11569_D,plain,(
  ( ! [X2,X1] : (~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1) | p(s(bool,X1))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f11651,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1)) )),
  inference(condensation,[],[f11650])).
fof(f11650,plain,(
  ( ! [X140,X138,X139,X137] : (p(s(bool,X138)) | ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X139),s('type.realax.real',X139)),X138) | s(bool,X137) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) )),
  inference(subsumption_resolution,[],[f11095,f1111])).
fof(f11095,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,X137)) | p(s(bool,X138)) | ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X139),s('type.realax.real',X139)),X138) | s(bool,X137) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f1346,f1030])).
fof(f11566,plain,(
  ( ! [X2,X1] : (~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1) | p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f11566_D])).
fof(f11566_D,plain,(
  ( ! [X2,X1] : (~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1) | p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f11649,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1)) )),
  inference(condensation,[],[f11648])).
fof(f11648,plain,(
  ( ! [X134,X136,X135,X133] : (p(s(bool,X134)) | ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X135),s('type.realax.real',X135)),X134) | s(bool,X133) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f11094,f1111])).
fof(f11094,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,X133)) | p(s(bool,X134)) | ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X135),s('type.realax.real',X135)),X134) | s(bool,X133) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f1332,f1030])).
fof(f11563,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f11563_D])).
fof(f11563_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f11647,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f11646])).
fof(f11646,plain,(
  ( ! [X127,X128,X126,X129] : (~sP8('T_0',X127,X128) | sP9('T_0',X127,X128) | s(bool,X126) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f11092,f1111])).
fof(f11092,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X126)) | ~sP8('T_0',X127,X128) | sP9('T_0',X127,X128) | s(bool,X126) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f814,f1030])).
fof(f11560,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) | $spl36),
  inference(cnf_transformation,[],[f11560_D])).
fof(f11560_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~sP10('T_0',X1,X2)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f11645,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f11644])).
fof(f11644,plain,(
  ( ! [X125,X123,X124,X122] : (~sP10('T_0',X123,X124) | sP11('T_0',X123,X124) | s(bool,X122) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(subsumption_resolution,[],[f11091,f1111])).
fof(f11091,plain,(
  ( ! [X125,X123,X124,X122] : (~p(s(bool,X122)) | ~sP10('T_0',X123,X124) | sP11('T_0',X123,X124) | s(bool,X122) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f813,f1030])).
fof(f11557,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f11557_D])).
fof(f11557_D,plain,(
  ( ! [X2,X1] : (~sP11('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f11643,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f11642])).
fof(f11642,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,X116)) | p(s(bool,X117)) | ~sP11('T_0',X117,X116) | s(bool,X115) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f11089,f1111])).
fof(f11089,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,X115)) | ~p(s(bool,X116)) | p(s(bool,X117)) | ~sP11('T_0',X117,X116) | s(bool,X115) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f790,f1030])).
fof(f11554,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f11554_D])).
fof(f11554_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f11641,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f11640])).
fof(f11640,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,X112)) | sP11('T_0',X112,X113) | s(bool,X111) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f11088,f1111])).
fof(f11088,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,X111)) | ~p(s(bool,X112)) | sP11('T_0',X112,X113) | s(bool,X111) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f789,f1030])).
fof(f11551,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f11551_D])).
fof(f11551_D,plain,(
  ( ! [X2,X1] : (sP11('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f11639,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f11638])).
fof(f11638,plain,(
  ( ! [X109,X107,X110,X108] : (p(s(bool,X108)) | sP11('T_0',X109,X108) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f11087,f1111])).
fof(f11087,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | sP11('T_0',X109,X108) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f788,f1030])).
fof(f11548,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f11548_D])).
fof(f11548_D,plain,(
  ( ! [X2,X1] : (~sP10(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f11637,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f11636])).
fof(f11636,plain,(
  ( ! [X103,X105,X106,X104] : (p(s(bool,X104)) | ~p(s(bool,X105)) | ~sP10(X105,X104,'T_0') | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f11086,f1111])).
fof(f11086,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~p(s(bool,X105)) | ~sP10(X105,X104,'T_0') | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f783,f1030])).
fof(f11545,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f11545_D])).
fof(f11545_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f11635,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f11634])).
fof(f11634,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,X100)) | p(s(bool,X101)) | ~sP9('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f11085,f1111])).
fof(f11085,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X99)) | p(s(bool,X100)) | p(s(bool,X101)) | ~sP9('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f773,f1030])).
fof(f11542,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f11542_D])).
fof(f11542_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f11633,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f11632])).
fof(f11632,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f11084,f1111])).
fof(f11084,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~p(s(bool,X96)) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f772,f1030])).
fof(f11539,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl22),
  inference(cnf_transformation,[],[f11539_D])).
fof(f11539_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f11631,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f11630])).
fof(f11630,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | sP9('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f11083,f1111])).
fof(f11083,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~p(s(bool,X92)) | sP9('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f771,f1030])).
fof(f11536,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl20),
  inference(cnf_transformation,[],[f11536_D])).
fof(f11536_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f11629,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f11628])).
fof(f11628,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | p(s(bool,X89)) | ~sP8('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f11082,f1111])).
fof(f11082,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | p(s(bool,X88)) | p(s(bool,X89)) | ~sP8('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f766,f1030])).
fof(f11533,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl18),
  inference(cnf_transformation,[],[f11533_D])).
fof(f11533_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f11627,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f11626])).
fof(f11626,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP7(X85,X84,'T_0') | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f11081,f1111])).
fof(f11081,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP7(X85,X84,'T_0') | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f758,f1030])).
fof(f11530,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f11530_D])).
fof(f11530_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f11625,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f11624])).
fof(f11624,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP6(X80,'T_0',X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f11080,f1111])).
fof(f11080,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP6(X80,'T_0',X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f749,f1030])).
fof(f11527,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f11527_D])).
fof(f11527_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f11623,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f11622])).
fof(f11622,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~sP3(X76,'T_0',X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f11079,f1111])).
fof(f11079,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | ~sP3(X76,'T_0',X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f730,f1030])).
fof(f11524,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl12),
  inference(cnf_transformation,[],[f11524_D])).
fof(f11524_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f11621,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f11620])).
fof(f11620,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~sP4(X72,'T_0',X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f11078,f1111])).
fof(f11078,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | p(s(bool,X73)) | ~sP4(X72,'T_0',X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f726,f1030])).
fof(f11521,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f11521_D])).
fof(f11521_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f11619,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f11618])).
fof(f11618,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X69)) | sP1(X69,'T_0') | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f11077,f1111])).
fof(f11077,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | sP1(X69,'T_0') | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f673,f1030])).
fof(f11517,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl6),
  inference(cnf_transformation,[],[f11517_D])).
fof(f11517_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f11617,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f11616])).
fof(f11616,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | sP1('T_0',X66) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f11076,f1111])).
fof(f11076,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,X65)) | p(s(bool,X66)) | sP1('T_0',X66) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f670,f1030])).
fof(f11580,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1,'T_0')) )),
  inference(condensation,[],[f11579])).
fof(f11579,plain,(
  ( ! [X156,X154,X155,X153] : (p(s(bool,X154)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X155),s('type.realax.real',X155)),X154,'T_0') | s(bool,X153) = s(bool,X156) | s(bool,'T_0') = s(bool,X156)) )),
  inference(subsumption_resolution,[],[f10905,f1111])).
fof(f10905,plain,(
  ( ! [X156,X154,X155,X153] : (~p(s(bool,X153)) | p(s(bool,X154)) | ~sP11('const.real.real_lte_2'(s('type.realax.real',X155),s('type.realax.real',X155)),X154,'T_0') | s(bool,X153) = s(bool,X156) | s(bool,'T_0') = s(bool,X156)) )),
  inference(superposition,[],[f1505,f1030])).
fof(f11577,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP10('T_0',X1,'const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(condensation,[],[f11576])).
fof(f11576,plain,(
  ( ! [X152,X151,X149,X150] : (p(s(bool,X150)) | ~sP10('T_0',X150,'const.real.real_lte_2'(s('type.realax.real',X151),s('type.realax.real',X151))) | s(bool,X149) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(subsumption_resolution,[],[f10904,f1111])).
fof(f10904,plain,(
  ( ! [X152,X151,X149,X150] : (~p(s(bool,X149)) | p(s(bool,X150)) | ~sP10('T_0',X150,'const.real.real_lte_2'(s('type.realax.real',X151),s('type.realax.real',X151))) | s(bool,X149) = s(bool,X152) | s(bool,'T_0') = s(bool,X152)) )),
  inference(superposition,[],[f1481,f1030])).
fof(f11574,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP7(X1,'T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)))) )),
  inference(condensation,[],[f11573])).
fof(f11573,plain,(
  ( ! [X146,X147,X145,X148] : (p(s(bool,X146)) | ~sP7(X146,'T_0','const.real.real_lte_2'(s('type.realax.real',X147),s('type.realax.real',X147))) | s(bool,X145) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(subsumption_resolution,[],[f10903,f1111])).
fof(f10903,plain,(
  ( ! [X146,X147,X145,X148] : (~p(s(bool,X145)) | p(s(bool,X146)) | ~sP7(X146,'T_0','const.real.real_lte_2'(s('type.realax.real',X147),s('type.realax.real',X147))) | s(bool,X145) = s(bool,X148) | s(bool,'T_0') = s(bool,X148)) )),
  inference(superposition,[],[f1410,f1030])).
fof(f11571,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1)) )),
  inference(condensation,[],[f11570])).
fof(f11570,plain,(
  ( ! [X144,X142,X143,X141] : (p(s(bool,X142)) | ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X143),s('type.realax.real',X143)),X142) | s(bool,X141) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(subsumption_resolution,[],[f10902,f1111])).
fof(f10902,plain,(
  ( ! [X144,X142,X143,X141] : (~p(s(bool,X141)) | p(s(bool,X142)) | ~sP6('T_0','const.real.real_lte_2'(s('type.realax.real',X143),s('type.realax.real',X143)),X142) | s(bool,X141) = s(bool,X144) | s(bool,'T_0') = s(bool,X144)) )),
  inference(superposition,[],[f1380,f1030])).
fof(f11568,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1)) )),
  inference(condensation,[],[f11567])).
fof(f11567,plain,(
  ( ! [X140,X138,X139,X137] : (p(s(bool,X138)) | ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X139),s('type.realax.real',X139)),X138) | s(bool,X137) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) )),
  inference(subsumption_resolution,[],[f10901,f1111])).
fof(f10901,plain,(
  ( ! [X140,X138,X139,X137] : (~p(s(bool,X137)) | p(s(bool,X138)) | ~sP3('T_0','const.real.real_lte_2'(s('type.realax.real',X139),s('type.realax.real',X139)),X138) | s(bool,X137) = s(bool,X140) | s(bool,'T_0') = s(bool,X140)) )),
  inference(superposition,[],[f1346,f1030])).
fof(f11565,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X2),s('type.realax.real',X2)),X1)) )),
  inference(condensation,[],[f11564])).
fof(f11564,plain,(
  ( ! [X134,X136,X135,X133] : (p(s(bool,X134)) | ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X135),s('type.realax.real',X135)),X134) | s(bool,X133) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(subsumption_resolution,[],[f10900,f1111])).
fof(f10900,plain,(
  ( ! [X134,X136,X135,X133] : (~p(s(bool,X133)) | p(s(bool,X134)) | ~sP4('T_0','const.real.real_lte_2'(s('type.realax.real',X135),s('type.realax.real',X135)),X134) | s(bool,X133) = s(bool,X136) | s(bool,'T_0') = s(bool,X136)) )),
  inference(superposition,[],[f1332,f1030])).
fof(f11562,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f11561])).
fof(f11561,plain,(
  ( ! [X127,X128,X126,X129] : (~sP8('T_0',X127,X128) | sP9('T_0',X127,X128) | s(bool,X126) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(subsumption_resolution,[],[f10898,f1111])).
fof(f10898,plain,(
  ( ! [X127,X128,X126,X129] : (~p(s(bool,X126)) | ~sP8('T_0',X127,X128) | sP9('T_0',X127,X128) | s(bool,X126) = s(bool,X129) | s(bool,'T_0') = s(bool,X129)) )),
  inference(superposition,[],[f814,f1030])).
fof(f11559,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP10('T_0',X1,X2) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f11558])).
fof(f11558,plain,(
  ( ! [X125,X123,X124,X122] : (~sP10('T_0',X123,X124) | sP11('T_0',X123,X124) | s(bool,X122) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(subsumption_resolution,[],[f10897,f1111])).
fof(f10897,plain,(
  ( ! [X125,X123,X124,X122] : (~p(s(bool,X122)) | ~sP10('T_0',X123,X124) | sP11('T_0',X123,X124) | s(bool,X122) = s(bool,X125) | s(bool,'T_0') = s(bool,X125)) )),
  inference(superposition,[],[f813,f1030])).
fof(f11556,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f11555])).
fof(f11555,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,X116)) | p(s(bool,X117)) | ~sP11('T_0',X117,X116) | s(bool,X115) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(subsumption_resolution,[],[f10895,f1111])).
fof(f10895,plain,(
  ( ! [X118,X116,X117,X115] : (~p(s(bool,X115)) | ~p(s(bool,X116)) | p(s(bool,X117)) | ~sP11('T_0',X117,X116) | s(bool,X115) = s(bool,X118) | s(bool,'T_0') = s(bool,X118)) )),
  inference(superposition,[],[f790,f1030])).
fof(f11553,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP11('T_0',X1,X2)) )),
  inference(condensation,[],[f11552])).
fof(f11552,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,X112)) | sP11('T_0',X112,X113) | s(bool,X111) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(subsumption_resolution,[],[f10894,f1111])).
fof(f10894,plain,(
  ( ! [X111,X114,X112,X113] : (~p(s(bool,X111)) | ~p(s(bool,X112)) | sP11('T_0',X112,X113) | s(bool,X111) = s(bool,X114) | s(bool,'T_0') = s(bool,X114)) )),
  inference(superposition,[],[f789,f1030])).
fof(f11550,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP11('T_0',X2,X1)) )),
  inference(condensation,[],[f11549])).
fof(f11549,plain,(
  ( ! [X109,X107,X110,X108] : (p(s(bool,X108)) | sP11('T_0',X109,X108) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(subsumption_resolution,[],[f10893,f1111])).
fof(f10893,plain,(
  ( ! [X109,X107,X110,X108] : (~p(s(bool,X107)) | p(s(bool,X108)) | sP11('T_0',X109,X108) | s(bool,X107) = s(bool,X110) | s(bool,'T_0') = s(bool,X110)) )),
  inference(superposition,[],[f788,f1030])).
fof(f11547,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP10(X2,X1,'T_0')) )),
  inference(condensation,[],[f11546])).
fof(f11546,plain,(
  ( ! [X103,X105,X106,X104] : (p(s(bool,X104)) | ~p(s(bool,X105)) | ~sP10(X105,X104,'T_0') | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(subsumption_resolution,[],[f10892,f1111])).
fof(f10892,plain,(
  ( ! [X103,X105,X106,X104] : (~p(s(bool,X103)) | p(s(bool,X104)) | ~p(s(bool,X105)) | ~sP10(X105,X104,'T_0') | s(bool,X103) = s(bool,X106) | s(bool,'T_0') = s(bool,X106)) )),
  inference(superposition,[],[f783,f1030])).
fof(f11544,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f11543])).
fof(f11543,plain,(
  ( ! [X101,X99,X102,X100] : (p(s(bool,X100)) | p(s(bool,X101)) | ~sP9('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(subsumption_resolution,[],[f10891,f1111])).
fof(f10891,plain,(
  ( ! [X101,X99,X102,X100] : (~p(s(bool,X99)) | p(s(bool,X100)) | p(s(bool,X101)) | ~sP9('T_0',X101,X100) | s(bool,X99) = s(bool,X102) | s(bool,'T_0') = s(bool,X102)) )),
  inference(superposition,[],[f773,f1030])).
fof(f11541,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f11540])).
fof(f11540,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X96)) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(subsumption_resolution,[],[f10890,f1111])).
fof(f10890,plain,(
  ( ! [X97,X95,X98,X96] : (~p(s(bool,X95)) | ~p(s(bool,X96)) | sP9('T_0',X96,X97) | s(bool,X95) = s(bool,X98) | s(bool,'T_0') = s(bool,X98)) )),
  inference(superposition,[],[f772,f1030])).
fof(f11538,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f11537])).
fof(f11537,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X92)) | sP9('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(subsumption_resolution,[],[f10889,f1111])).
fof(f10889,plain,(
  ( ! [X94,X92,X93,X91] : (~p(s(bool,X91)) | ~p(s(bool,X92)) | sP9('T_0',X93,X92) | s(bool,X91) = s(bool,X94) | s(bool,'T_0') = s(bool,X94)) )),
  inference(superposition,[],[f771,f1030])).
fof(f11535,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f11534])).
fof(f11534,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,X88)) | p(s(bool,X89)) | ~sP8('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(subsumption_resolution,[],[f10888,f1111])).
fof(f10888,plain,(
  ( ! [X90,X88,X87,X89] : (~p(s(bool,X87)) | p(s(bool,X88)) | p(s(bool,X89)) | ~sP8('T_0',X88,X89) | s(bool,X87) = s(bool,X90) | s(bool,'T_0') = s(bool,X90)) )),
  inference(superposition,[],[f766,f1030])).
fof(f11532,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f11531])).
fof(f11531,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X84)) | p(s(bool,X85)) | ~sP7(X85,X84,'T_0') | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f10887,f1111])).
fof(f10887,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,X83)) | ~p(s(bool,X84)) | p(s(bool,X85)) | ~sP7(X85,X84,'T_0') | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f758,f1030])).
fof(f11529,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f11528])).
fof(f11528,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X80)) | p(s(bool,X81)) | ~sP6(X80,'T_0',X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f10886,f1111])).
fof(f10886,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,X79)) | ~p(s(bool,X80)) | p(s(bool,X81)) | ~sP6(X80,'T_0',X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f749,f1030])).
fof(f11526,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f11525])).
fof(f11525,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~sP3(X76,'T_0',X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f10885,f1111])).
fof(f10885,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | ~sP3(X76,'T_0',X77) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f730,f1030])).
fof(f11523,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f11522])).
fof(f11522,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | p(s(bool,X73)) | ~sP4(X72,'T_0',X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f10884,f1111])).
fof(f10884,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X71)) | ~p(s(bool,X72)) | p(s(bool,X73)) | ~sP4(X72,'T_0',X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f726,f1030])).
fof(f11520,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f11519])).
fof(f11519,plain,(
  ( ! [X70,X68,X69] : (p(s(bool,X69)) | sP1(X69,'T_0') | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f10883,f1111])).
fof(f10883,plain,(
  ( ! [X70,X68,X69] : (~p(s(bool,X68)) | p(s(bool,X69)) | sP1(X69,'T_0') | s(bool,X68) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f673,f1030])).
fof(f11516,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f11515])).
fof(f11515,plain,(
  ( ! [X66,X67,X65] : (p(s(bool,X66)) | sP1('T_0',X66) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(subsumption_resolution,[],[f10882,f1111])).
fof(f10882,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,X65)) | p(s(bool,X66)) | sP1('T_0',X66) | s(bool,X65) = s(bool,X67) | s(bool,'T_0') = s(bool,X67)) )),
  inference(superposition,[],[f670,f1030])).
fof(f878,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f877,f861])).
fof(f861,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),X0)))) ) | $spl4),
  inference(backward_demodulation,[],[f829,f507])).
fof(f877,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13))) | ($spl1 | $spl4)),
  inference(backward_demodulation,[],[f829,f826])).
fof(f857,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f856,f827])).
fof(f856,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f826,f841])).
fof(f841,plain,(
  ( ! [X0] : (s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13)))) ) | $spl2),
  inference(backward_demodulation,[],[f827,f508])).
fof(f825,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13))) | $spl0),
  inference(cnf_transformation,[],[f825_D])).
fof(f825_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f829,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl4),
  inference(cnf_transformation,[],[f829_D])).
fof(f829_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f827,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl2),
  inference(cnf_transformation,[],[f827_D])).
fof(f827_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f475,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13)))),
  inference(cnf_transformation,[],[f347])).
fof(f347,plain,(
  (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13))) | (s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13])],[f346])).
fof(f346,plain,(
  ? [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(flattening,[],[f345])).
fof(f345,plain,(
  ? [X0,X1] : ((s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))))) & (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) | (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) & s('type.pair.prod'('type.realax.real','type.realax.real'),X1) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))))))),
  inference(nnf_transformation,[],[f295])).
fof(f295,plain,(
  ? [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) <~> (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X0),s('type.pair.prod'('type.realax.real','type.realax.real'),X1))) <=> (s('type.pair.prod'('type.realax.real','type.realax.real'),X0) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X1) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) <=> (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),X6),s('type.pair.prod'('type.realax.real','type.realax.real'),X7))) <=> (s('type.pair.prod'('type.realax.real','type.realax.real'),X6) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),X7) = s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.complex.COMPLEX_ENTIRE','thm.complex.COMPLEX_ENTIRE')).
fof(f826,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13))) | $spl1),
  inference(cnf_transformation,[],[f826_D])).
fof(f826_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f830,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(cnf_transformation,[],[f830_D])).
fof(f830_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f476,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK12) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13)))),
  inference(cnf_transformation,[],[f347])).
fof(f828,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl3),
  inference(cnf_transformation,[],[f828_D])).
fof(f828_D,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f477,plain,(
  s('type.pair.prod'('type.realax.real','type.realax.real'),sK13) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.pair.prod'('type.realax.real','type.realax.real'),'const.complex.complex_mul_2'(s('type.pair.prod'('type.realax.real','type.realax.real'),sK12),s('type.pair.prod'('type.realax.real','type.realax.real'),sK13)))),
  inference(cnf_transformation,[],[f347])).
% SZS output end Proof for thm.complex
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 477
% Passive clauses: 2144
% Generated clauses: 78386
% Final active clauses: 461
% Final passive clauses: 1382
% Input formulas: 133
% Initial clauses: 326
% 
% Duplicate literals: 83
% Trivial inequalities: 59
% Fw subsumption resolutions: 377
% Fw demodulations: 32970
% Bw demodulations: 180
% Condensations: 427
% Global subsumptions: 53
% 
% Simple tautologies: 102
% Equational tautologies: 82
% Forward subsumptions: 25126
% Fw demodulations to eq. taut.: 16538
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 23455
% Factoring: 1
% Forward superposition: 11879
% Backward superposition: 8189
% Self superposition: 124
% Equality factoring: 26
% Equality resolution: 27
% 
% Split clauses: 55
% Split components: 111
% SAT solver clauses: 2679
% SAT solver unit clauses: 1064
% SAT solver binary clauses: 1004
% 
% Sat splits: 55
% Sat splitting refutations: 370
% 
% Memory used [KB]: 40297
% Time elapsed: 2.436 s
% ------------------------------
% ------------------------------
% Success in time 3.77 s
