Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.integer
% (66143)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1095
% Passive clauses: 6070
% Generated clauses: 34621
% Final active clauses: 1065
% Final passive clauses: 4632
% Input formulas: 96
% Initial clauses: 224
% 
% Duplicate literals: 328
% Trivial inequalities: 50
% Fw subsumption resolutions: 2572
% Fw demodulations: 2661
% Bw demodulations: 11
% 
% Simple tautologies: 218
% Equational tautologies: 36
% Forward subsumptions: 22326
% Fw demodulations to eq. taut.: 24
% 
% Binary resolution: 6737
% Factoring: 1516
% Forward superposition: 10146
% Backward superposition: 10070
% Self superposition: 67
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 338
% Split components: 676
% SAT solver clauses: 533
% SAT solver unit clauses: 11
% SAT solver binary clauses: 93
% 
% Sat splits: 338
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 533
% TWLsolver calls for satisfiability: 395
% 
% Memory used [KB]: 16886
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (66161)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 548
% Passive clauses: 3494
% Generated clauses: 23184
% Final active clauses: 449
% Final passive clauses: 1724
% Input formulas: 96
% Initial clauses: 224
% 
% Duplicate literals: 226
% Trivial inequalities: 139
% Fw subsumption resolutions: 1634
% Fw demodulations: 2731
% Bw demodulations: 1191
% Condensations: 219
% Global subsumptions: 15
% 
% Simple tautologies: 289
% Equational tautologies: 21
% Forward subsumptions: 10968
% Fw demodulations to eq. taut.: 26
% Bw demodulations to eq. taut.: 114
% 
% Binary resolution: 2215
% Factoring: 113
% Forward superposition: 6145
% Backward superposition: 4229
% Self superposition: 72
% Equality factoring: 30
% Equality resolution: 15
% 
% Split clauses: 98
% Split components: 200
% SAT solver clauses: 3401
% SAT solver unit clauses: 272
% SAT solver binary clauses: 73
% SAT solver learnt clauses: 30
% SAT solver learnt literals: 12
% 
% Sat splits: 98
% Sat splitting refutations: 29
% 
% TWLsolver clauses: 3341
% TWLsolver calls for satisfiability: 1579
% 
% Memory used [KB]: 16247
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f37861,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f1220,f1224,f1222,f1869,f1871,f1875,f1873,f1876,f1866,f1881,f1879,f10444,f10445,f10465,f10463,f10454,f10474,f10455,f10477,f10456,f10478,f10457,f10479,f10458,f10480,f10459,f10481,f10460,f10476,f10482,f10486,f10590,f10591,f10611,f10609,f10600,f10620,f10601,f10623,f10602,f10624,f10603,f10625,f10604,f10626,f10605,f10627,f10606,f10622,f10628,f10632,f14151,f14155,f14158,f14161,f14164,f14180,f14152,f14182,f14156,f14184,f14159,f14186,f14162,f14188,f14153,f14165,f14530,f14531,f14534,f14550,f14552,f14554,f14556,f14558,f14560,f14570,f14596,f14598,f14861,f14957,f14958,f14959,f14960,f25812,f26169,f25896,f26187,f26185,f26191,f26193,f26171,f26713,f26716,f26714,f33995,f34106,f34104,f34197,f34289,f34287,f35073,f35077,f35075,f37709,f37712,f37713,f37714,f37718,f37736,f37737,f37739,f37846])).
fof(f37846,plain,(
  $false | ($spl0 | $spl2)),
  inference(subsumption_resolution,[],[f37845,f318])).
fof(f318,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.integer.LT_LADD',aHL_TRUTH)).
fof(f37845,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2)),
  inference(forward_demodulation,[],[f37844,f37789])).
fof(f37789,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11))) | $spl2),
  inference(subsumption_resolution,[],[f37762,f545])).
fof(f545,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f324,f318])).
fof(f324,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f99])).
fof(f99,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f33])).
fof(f33,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.integer.LT_LADD','thm.bool.NOT_CLAUSES|split|1')).
fof(f37762,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11))) | $spl2),
  inference(superposition,[],[f3248,f1223])).
fof(f1223,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) | $spl2),
  inference(cnf_transformation,[],[f1223_D])).
fof(f1223_D,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3248,plain,(
  ( ! [X4,X5] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X4)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X5)))) )),
  inference(subsumption_resolution,[],[f3207,f545])).
fof(f3207,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X5),s('type.num.num',X4)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X5)))) )),
  inference(superposition,[],[f458,f360])).
fof(f360,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.integer.LT_LADD',aHL_BOOL_CASES)).
fof(f458,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1] : ((p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))))))),
  inference(nnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X9,X10] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10)))) <=> p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.arithmetic.NOT_LESS')).
fof(f37844,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11)))) | $spl0),
  inference(forward_demodulation,[],[f37843,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))))) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X9,X10,X11] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X11)))))),
  file('thm.integer.LT_LADD','thm.arithmetic.ADD_MONO_LESS_EQ')).
fof(f37843,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl0),
  inference(subsumption_resolution,[],[f37810,f318])).
fof(f37810,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl0),
  inference(superposition,[],[f459,f1222])).
fof(f459,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f280])).
fof(f37739,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37738,f37724])).
fof(f37724,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl3),
  inference(subsumption_resolution,[],[f37723,f371])).
fof(f371,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f244,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f140])).
fof(f140,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X13] : (s(bool,'F_0') = s(bool,X13) <=> ~p(s(bool,X13)))),
  file('thm.integer.LT_LADD','thm.bool.EQ_CLAUSES|split|2')).
fof(f37723,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,X1))) ) | $spl3),
  inference(forward_demodulation,[],[f37722,f3335])).
fof(f3335,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11))) | $spl3),
  inference(unit_resulting_resolution,[],[f545,f3328,f428])).
fof(f428,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f266])).
fof(f266,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1))))) & (((~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X0)) | p(s(bool,X1)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) & p(s(bool,X1))) | (~p(s(bool,X0)) & ~p(s(bool,X1)))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X16,X17] : (s(bool,X16) = s(bool,X17) <=> ((p(s(bool,X16)) & p(s(bool,X17))) | (~p(s(bool,X16)) & ~p(s(bool,X17)))))),
  file('thm.integer.LT_LADD','thm.bool.EQ_EXPAND')).
fof(f3328,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11)))) | $spl3),
  inference(subsumption_resolution,[],[f3292,f318])).
fof(f3292,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11)))) | $spl3),
  inference(superposition,[],[f459,f1225])).
fof(f1225,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) | $spl3),
  inference(unit_resulting_resolution,[],[f1224,f360])).
fof(f37722,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11))) | p(s(bool,X1))) ) | $spl3),
  inference(forward_demodulation,[],[f37721,f470])).
fof(f37721,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f37681,f318])).
fof(f37681,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,X1)) | s(bool,X1) != s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) ) | $spl3),
  inference(superposition,[],[f425,f37654])).
fof(f37654,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) | $spl3),
  inference(subsumption_resolution,[],[f37561,f545])).
fof(f37561,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))))) | $spl3),
  inference(superposition,[],[f3248,f1240])).
fof(f1240,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f1237,f360])).
fof(f1237,plain,(
  s(bool,'T_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))))) | $spl3),
  inference(backward_demodulation,[],[f1225,f317])).
fof(f317,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))))),
  inference(cnf_transformation,[],[f239])).
fof(f239,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f207])).
fof(f207,plain,(
  ? [X0,X1,X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) != s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f98])).
fof(f98,plain,(
  ~! [X0,X1,X2] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X7),s('type.num.num',X8))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X8)))))),
  file('thm.integer.LT_LADD','thm.integer.LT_LADD')).
fof(f425,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f37738,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl3),
  inference(subsumption_resolution,[],[f37731,f37724])).
fof(f37731,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))),s('type.num.num',X17))))) ) | $spl3),
  inference(backward_demodulation,[],[f37726,f5023])).
fof(f5023,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))))))) )),
  inference(forward_demodulation,[],[f4984,f407])).
fof(f407,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X9,X10] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.arithmetic.LESS_EQ')).
fof(f4984,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))))),s('type.num.num',X17))))) )),
  inference(superposition,[],[f431,f416])).
fof(f416,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X10,X9] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.arithmetic.ADD_CLAUSES|split|3')).
fof(f431,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) & (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))))),
  inference(nnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f87])).
fof(f87,axiom,(
  ! [X10,X9] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X9)))))) <=> ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_lte|split|5')).
fof(f37726,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f37725,f37724])).
fof(f37725,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) ) | $spl3),
  inference(subsumption_resolution,[],[f37682,f318])).
fof(f37682,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | ~p(s(bool,X2)) | s(bool,X2) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) ) | $spl3),
  inference(superposition,[],[f427,f37654])).
fof(f427,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f267])).
fof(f37737,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37733,f37724])).
fof(f37733,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl3),
  inference(condensation,[],[f37730])).
fof(f37730,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) ) | $spl3),
  inference(backward_demodulation,[],[f37726,f4958])).
fof(f4958,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f431,f406])).
fof(f406,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X9,X10] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X9)))),
  file('thm.integer.LT_LADD','thm.arithmetic.ADD_SYM')).
fof(f37736,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37735,f37724])).
fof(f37735,plain,(
  ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl3),
  inference(subsumption_resolution,[],[f37728,f37724])).
fof(f37728,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))))))))) ) | $spl3),
  inference(backward_demodulation,[],[f37726,f35976])).
fof(f35976,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26))))),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))))))))) )),
  inference(forward_demodulation,[],[f35975,f6323])).
fof(f6323,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f422,f406])).
fof(f422,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X10,X9] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X9)))))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|2')).
fof(f35975,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))))))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f35930,f35053])).
fof(f35053,plain,(
  ( ! [X8,X9] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))))) )),
  inference(backward_demodulation,[],[f35007,f6334])).
fof(f6334,plain,(
  ( ! [X8,X9] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))))) )),
  inference(superposition,[],[f357,f422])).
fof(f357,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X10] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X10))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|21')).
fof(f35007,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f34794,f407])).
fof(f34794,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f34783,f4508])).
fof(f4508,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f4507,f357])).
fof(f4507,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f4464,f407])).
fof(f4464,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f414,f322])).
fof(f322,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(cnf_transformation,[],[f59])).
fof(f59,axiom,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|13')).
fof(f414,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|28')).
fof(f34783,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f34782,f322])).
fof(f34782,plain,(
  s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(subsumption_resolution,[],[f34756,f32905])).
fof(f32905,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f32904,f322])).
fof(f32904,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f32349,f548])).
fof(f548,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f354,f322])).
fof(f354,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : s('type.num.num',X9) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  file('thm.integer.LT_LADD','thm.arithmetic.MULT_RIGHT_|49|')).
fof(f32349,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(superposition,[],[f5878,f322])).
fof(f5878,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))))) )),
  inference(forward_demodulation,[],[f5739,f547])).
fof(f547,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f350,f322])).
fof(f350,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : s('type.num.num',X9) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X9)))),
  file('thm.integer.LT_LADD','thm.arithmetic.MULT_LEFT_|49|')).
fof(f5739,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f420,f322])).
fof(f420,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f51])).
fof(f51,axiom,(
  ! [X10,X9] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X10),s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|5')).
fof(f34756,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f2917,f322])).
fof(f2917,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num',X0) = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f429,f322])).
fof(f429,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1] : ((s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X0) = s('type.num.num',X1)) & (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(nnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) <=> s('type.num.num',X0) = s('type.num.num',X1))),
  inference(rectify,[],[f65])).
fof(f65,axiom,(
  ! [X10,X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) <=> s('type.num.num',X9) = s('type.num.num',X10))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|19')).
fof(f35930,plain,(
  ( ! [X26,X27] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))))))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(superposition,[],[f35246,f422])).
fof(f35246,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f35219,f359])).
fof(f359,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X10] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|27')).
fof(f35219,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))))) )),
  inference(superposition,[],[f459,f35009])).
fof(f35009,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(backward_demodulation,[],[f35007,f357])).
fof(f37718,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37717,f545])).
fof(f37717,plain,(
  p(s(bool,'F_0')) | $spl3),
  inference(forward_demodulation,[],[f37716,f3335])).
fof(f37716,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11)))) | $spl3),
  inference(forward_demodulation,[],[f37702,f470])).
fof(f37702,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl3),
  inference(trivial_inequality_removal,[],[f37677])).
fof(f37677,plain,(
  s(bool,'T_0') != s(bool,'T_0') | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl3),
  inference(superposition,[],[f366,f37654])).
fof(f366,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X13] : (s(bool,'T_0') = s(bool,X13) <=> p(s(bool,X13)))),
  file('thm.integer.LT_LADD','thm.bool.EQ_CLAUSES|split|0')).
fof(f37714,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37673,f3332])).
fof(f3332,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11))) | $spl3),
  inference(unit_resulting_resolution,[],[f318,f3328,f425])).
fof(f37673,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11))) | $spl3),
  inference(superposition,[],[f470,f37654])).
fof(f37713,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37668,f3332])).
fof(f37668,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11))) | $spl3),
  inference(superposition,[],[f37654,f470])).
fof(f37712,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37711,f545])).
fof(f37711,plain,(
  p(s(bool,'F_0')) | $spl3),
  inference(forward_demodulation,[],[f37710,f3335])).
fof(f37710,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11)))) | $spl3),
  inference(forward_demodulation,[],[f37662,f470])).
fof(f37662,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl3),
  inference(unit_resulting_resolution,[],[f37654,f366])).
fof(f37709,plain,(
  $false | $spl3),
  inference(subsumption_resolution,[],[f37708,f545])).
fof(f37708,plain,(
  p(s(bool,'F_0')) | $spl3),
  inference(forward_demodulation,[],[f37707,f3335])).
fof(f37707,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',sK12),s('type.num.num',sK11)))) | $spl3),
  inference(forward_demodulation,[],[f37667,f470])).
fof(f37667,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11)))))) | $spl3),
  inference(unit_resulting_resolution,[],[f318,f37654,f425])).
fof(f35075,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl89),
  inference(cnf_transformation,[],[f35075_D])).
fof(f35075_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f35077,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl90),
  inference(cnf_transformation,[],[f35077_D])).
fof(f35077_D,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f35073,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f35072,f34783])).
fof(f35072,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(forward_demodulation,[],[f34758,f34785])).
fof(f34785,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(backward_demodulation,[],[f34783,f2134])).
fof(f2134,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f349,f322])).
fof(f349,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X0))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f62])).
fof(f62,axiom,(
  ! [X10] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num',X10))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|16')).
fof(f34758,plain,(
  s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f2917,f2134])).
fof(f34287,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl84),
  inference(cnf_transformation,[],[f34287_D])).
fof(f34287_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f34289,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl86),
  inference(cnf_transformation,[],[f34289_D])).
fof(f34289_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f34197,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f1376,f2134])).
fof(f1376,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f346,f360])).
fof(f346,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f79])).
fof(f79,axiom,(
  ! [X10] : s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X10))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|33')).
fof(f34104,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl80),
  inference(cnf_transformation,[],[f34104_D])).
fof(f34104_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f34106,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl82),
  inference(cnf_transformation,[],[f34106_D])).
fof(f34106_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f33995,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f1298,f2134])).
fof(f1298,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f345,f360])).
fof(f345,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X0))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f78])).
fof(f78,axiom,(
  ! [X10] : s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X10))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|32')).
fof(f26714,plain,(
  ( ! [X7] : (sP0(X7,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) ) | $spl76),
  inference(cnf_transformation,[],[f26714_D])).
fof(f26714_D,plain,(
  ( ! [X7] : (sP0(X7,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f26716,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl79),
  inference(cnf_transformation,[],[f26716_D])).
fof(f26716_D,plain,(
  ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f26713,plain,(
  ( ! [X7] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) | sP0(X7,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f26672,f25810])).
fof(f25810,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f357,f2134])).
fof(f26672,plain,(
  ( ! [X7] : (sP0(X7,'const.arithmetic.>_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) )),
  inference(superposition,[],[f3767,f2134])).
fof(f3767,plain,(
  ( ! [X14,X13] : (sP0(X14,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X13))))) )),
  inference(superposition,[],[f434,f356])).
fof(f356,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X10] : s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X10))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|24')).
fof(f434,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f224])).
fof(f224,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f26171,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl69),
  inference(cnf_transformation,[],[f26171_D])).
fof(f26171_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f26193,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) | $spl74),
  inference(cnf_transformation,[],[f26193_D])).
fof(f26193_D,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f26191,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f26190,f25811])).
fof(f25811,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f359,f2134])).
fof(f26190,plain,(
  p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(forward_demodulation,[],[f25899,f4298])).
fof(f4298,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f412,f1935])).
fof(f1935,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f1922,f548])).
fof(f1922,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))),
  inference(unit_resulting_resolution,[],[f548,f375])).
fof(f375,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f246])).
fof(f246,plain,(
  ! [X0] : ((s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) & (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num',X0) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0))),
  inference(rectify,[],[f64])).
fof(f64,axiom,(
  ! [X10] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))) <=> s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X10))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|18')).
fof(f412,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f77])).
fof(f77,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|31')).
fof(f25899,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))),
  inference(superposition,[],[f379,f2134])).
fof(f379,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0] : ((~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) & (s('type.num.num','const.num.|48|_0') != s('type.num.num',X0) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X0))),
  inference(rectify,[],[f76])).
fof(f76,axiom,(
  ! [X10] : (p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) <=> s('type.num.num','const.num.|48|_0') = s('type.num.num',X10))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|30')).
fof(f26185,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl70),
  inference(cnf_transformation,[],[f26185_D])).
fof(f26185_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f26187,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl73),
  inference(cnf_transformation,[],[f26187_D])).
fof(f26187_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f25896,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f374,f2134])).
fof(f374,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f246])).
fof(f26169,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl66),
  inference(cnf_transformation,[],[f26169_D])).
fof(f26169_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f25812,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f374,f2134])).
fof(f14960,plain,(
  $false | ($spl7 | $spl56)),
  inference(subsumption_resolution,[],[f14427,f318])).
fof(f14427,plain,(
  ~p(s(bool,'T_0')) | ($spl7 | $spl56)),
  inference(backward_demodulation,[],[f14153,f7923])).
fof(f7923,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1873,f378])).
fof(f378,plain,(
  ( ! [X0] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X0)) )),
  inference(cnf_transformation,[],[f248])).
fof(f14959,plain,(
  $false | ($spl13 | $spl56)),
  inference(subsumption_resolution,[],[f14426,f318])).
fof(f14426,plain,(
  ~p(s(bool,'T_0')) | ($spl13 | $spl56)),
  inference(backward_demodulation,[],[f14153,f2690])).
fof(f2690,plain,(
  ~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1881,f378])).
fof(f14958,plain,(
  $false | ($spl7 | $spl56)),
  inference(subsumption_resolution,[],[f14844,f14604])).
fof(f14604,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))) ) | $spl56),
  inference(backward_demodulation,[],[f14603,f5878])).
fof(f14603,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51)))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14416,f318])).
fof(f14416,plain,(
  ( ! [X50,X51] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51)))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f5886])).
fof(f5886,plain,(
  ( ! [X50,X51] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f5885,f420])).
fof(f5885,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(forward_demodulation,[],[f5795,f4298])).
fof(f5795,plain,(
  ( ! [X50,X51] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X50),s('type.num.num',X51)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51)))))) )),
  inference(superposition,[],[f378,f420])).
fof(f14844,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl7 | $spl56)),
  inference(backward_demodulation,[],[f14604,f7925])).
fof(f7925,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl7),
  inference(unit_resulting_resolution,[],[f1873,f429])).
fof(f14957,plain,(
  $false | ($spl13 | $spl56)),
  inference(subsumption_resolution,[],[f14839,f14604])).
fof(f14839,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl13 | $spl56)),
  inference(backward_demodulation,[],[f14604,f2897])).
fof(f2897,plain,(
  s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl13),
  inference(unit_resulting_resolution,[],[f1881,f429])).
fof(f14861,plain,(
  $false | ($spl13 | $spl56)),
  inference(subsumption_resolution,[],[f14860,f1881])).
fof(f14860,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl56),
  inference(forward_demodulation,[],[f14618,f14604])).
fof(f14618,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) ) | $spl56),
  inference(backward_demodulation,[],[f14604,f353])).
fof(f353,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))),
  inference(rectify,[],[f60])).
fof(f60,axiom,(
  ! [X10] : s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|14')).
fof(f14598,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14597,f14538])).
fof(f14538,plain,(
  ( ! [X32] : (p(s(bool,X32))) ) | $spl56),
  inference(global_subsumption,[],[f317,f318,f320,f321,f322,f545,f328,f329,f330,f331,f332,f333,f334,f335,f336,f338,f339,f340,f343,f345,f346,f347,f348,f349,f547,f351,f352,f353,f548,f549,f356,f357,f550,f359,f360,f367,f366,f371,f370,f375,f374,f379,f378,f406,f407,f408,f409,f410,f411,f412,f413,f414,f415,f416,f551,f417,f418,f419,f420,f421,f422,f428,f427,f425,f430,f429,f432,f431,f435,f434,f552,f453,f450,f449,f448,f459,f458,f468,f469,f470,f471,f490,f489,f488,f487,f486,f485,f484,f494,f493,f492,f491,f498,f497,f496,f495,f502,f501,f500,f499,f506,f505,f504,f503,f542,f510,f509,f508,f507,f517,f516,f515,f514,f523,f520,f519,f518,f530,f529,f528,f541,f526,f525,f524,f534,f533,f532,f531,f538,f537,f536,f535,f554,f553,f556,f555,f558,f557,f560,f559,f562,f561,f564,f563,f566,f565,f568,f567,f570,f569,f572,f571,f574,f573,f576,f575,f578,f577,f580,f579,f582,f581,f584,f583,f586,f585,f588,f587,f591,f608,f621,f622,f625,f628,f631,f634,f637,f646,f651,f652,f655,f658,f661,f668,f669,f671,f674,f681,f684,f688,f699,f697,f702,f712,f710,f736,f786,f787,f788,f789,f790,f791,f792,f883,f884,f885,f886,f887,f888,f889,f914,f915,f916,f917,f918,f919,f920,f942,f1035,f1077,f1116,f1129,f1149,f1148,f1147,f1146,f1169,f1168,f1167,f1166,f1165,f1164,f1163,f1162,f1161,f1160,f1159,f1158,f1157,f1156,f1182,f1181,f1178,f1177,f1205,f1204,f1203,f1202,f1201,f1200,f1199,f1198,f1197,f1196,f1195,f1194,f1193,f1192,f1219,f1218,f1217,f1216,f1220,f1288,f1287,f1286,f1285,f1298,f1308,f1309,f1310,f1311,f1312,f1313,f1314,f1366,f1365,f1364,f1363,f1376,f1386,f1387,f1388,f1389,f1390,f1391,f1392,f1410,f1409,f1407,f1406,f1452,f1451,f1449,f1448,f1501,f1518,f1527,f1534,f1546,f1551,f1550,f1560,f1569,f1619,f1618,f1629,f1638,f1751,f1766,f1823,f1855,f1856,f1859,f1860,f1869,f1862,f1875,f1866,f1935,f2009,f2025,f2134,f2328,f2435,f2436,f2437,f2438,f2591,f2592,f2593,f2594,f2760,f2761,f2762,f2786,f2850,f2898,f2899,f2901,f2943,f2903,f2904,f2905,f2906,f2944,f2909,f2910,f2911,f2946,f2913,f2914,f2915,f2917,f2918,f3018,f3032,f3129,f3128,f3127,f3126,f3158,f3159,f3160,f3161,f3162,f3163,f3164,f3166,f3211,f3185,f3246,f3247,f3248,f3254,f3273,f3595,f3593,f3592,f3589,f3588,f3586,f3576,f3575,f3574,f3573,f3651,f3649,f3645,f3644,f3635,f3632,f3631,f3629,f3725,f3723,f3712,f3704,f3701,f3700,f3687,f3686,f3685,f3684,f3734,f3733,f3732,f3731,f3754,f3755,f3774,f3757,f3767,f3768,f3769,f3770,f3771,f3772,f3773,f3794,f3793,f3792,f3791,f3814,f3835,f3815,f3817,f3828,f3829,f3830,f3831,f3832,f3833,f3834,f3859,f3858,f3857,f3856,f3900,f3903,f3882,f3893,f3894,f3895,f3896,f3897,f3898,f3899,f3933,f3932,f3931,f3930,f3972,f3983,f3984,f3985,f3986,f3987,f3988,f3989,f4024,f4023,f4022,f4021,f4063,f4064,f4074,f4075,f4076,f4077,f4078,f4079,f4080,f4104,f4103,f4102,f4101,f4124,f4125,f4146,f4147,f4149,f4160,f4161,f4162,f4163,f4164,f4165,f4166,f4190,f4189,f4188,f4187,f4210,f4211,f4232,f4233,f4235,f4245,f4246,f4247,f4248,f4249,f4250,f4251,f4276,f4275,f4274,f4273,f4296,f4297,f4298,f4318,f4338,f4321,f4331,f4332,f4333,f4334,f4335,f4336,f4337,f4363,f4362,f4361,f4360,f4402,f4403,f4413,f4414,f4415,f4416,f4417,f4418,f4419,f4444,f4443,f4442,f4441,f4508,f4509,f4510,f4486,f4487,f4489,f4500,f4501,f4502,f4503,f4504,f4505,f4506,f4536,f4535,f4534,f4533,f4575,f4576,f4586,f4587,f4588,f4589,f4590,f4591,f4592,f4750,f4651,f4655,f4762,f4829,f4943,f4958,f4979,f5023,f5054,f5075,f5119,f5134,f5166,f5167,f5295,f5191,f5192,f5296,f5216,f5299,f5240,f5483,f5349,f5374,f5377,f5402,f5428,f5552,f5553,f5554,f5579,f5580,f5687,f5604,f5690,f5632,f5878,f5740,f5879,f5880,f5768,f5881,f5882,f5886,f5961,f6084,f5967,f5968,f5994,f6026,f6108,f6175,f6134,f6234,f6258,f6165,f6166,f6167,f6168,f6322,f6323,f6466,f6326,f6467,f6471,f6330,f6331,f6332,f6333,f6334,f6335,f6336,f6337,f6338,f6339,f6340,f6341,f6342,f6343,f6344,f6345,f6346,f6347,f6348,f6349,f6350,f6351,f6475,f6552,f6554,f6715,f6559,f6589,f6591,f6719,f6743,f6742,f6741,f6740,f6865,f6773,f6775,f6866,f6867,f6868,f6780,f6837,f6838,f6840,f6845,f6846,f6847,f6848,f6858,f6859,f6860,f6861,f6862,f6863,f6864,f6931,f6932,f7018,f7055,f7093,f7134,f7137,f7138,f7139,f7140,f7183,f7186,f7187,f7188,f7189,f7227,f7258,f7259,f7260,f7261,f7274,f7305,f7306,f7307,f7308,f7317,f7320,f7321,f7322,f7323,f7353,f7358,f7363,f7366,f7367,f7368,f7369,f7398,f7409,f7440,f7441,f7442,f7443,f7456,f7487,f7488,f7489,f7490,f7502,f7587,f7624,f7712,f7813,f7858,f8210,f8209,f8208,f8207,f8206,f8205,f8204,f8203,f8198,f8197,f8192,f8191,f8190,f8185,f8176,f8171,f8167,f8166,f8161,f8160,f8159,f8158,f8157,f8153,f8151,f8145,f8144,f8138,f8137,f8130,f8127,f8126,f8120,f8117,f8116,f8115,f8112,f8111,f8108,f8107,f8106,f8105,f8104,f8103,f8102,f8101,f8100,f8099,f8098,f8097,f8096,f8095,f8094,f8093,f8092,f8090,f8089,f8087,f8086,f8085,f8084,f8083,f8082,f8080,f8079,f8078,f8077,f8076,f8075,f8074,f8073,f8072,f8071,f8070,f8069,f8068,f8067,f8066,f8065,f8063,f8062,f8060,f8059,f8058,f8057,f8056,f8055,f8053,f8052,f8051,f8050,f8044,f8043,f8036,f8035,f8029,f8028,f8021,f8020,f8018,f8016,f8014,f8012,f8011,f8008,f8007,f8006,f8005,f8003,f8002,f8001,f8000,f7999,f8602,f8601,f8600,f8599,f8598,f8597,f8596,f8595,f8590,f8589,f8584,f8583,f8582,f8577,f8568,f8563,f8559,f8558,f8553,f8552,f8551,f8550,f8549,f8545,f8543,f8537,f8536,f8530,f8529,f8522,f8519,f8518,f8512,f8509,f8508,f8507,f8504,f8503,f8500,f8499,f8498,f8497,f8496,f8495,f8494,f8493,f8492,f8491,f8490,f8489,f8488,f8487,f8486,f8485,f8484,f8482,f8481,f8479,f8478,f8477,f8476,f8475,f8474,f8472,f8471,f8470,f8469,f8468,f8467,f8466,f8465,f8464,f8463,f8462,f8461,f8460,f8459,f8458,f8457,f8455,f8454,f8452,f8451,f8450,f8449,f8448,f8447,f8445,f8444,f8443,f8442,f8436,f8435,f8428,f8427,f8421,f8420,f8413,f8412,f8410,f8408,f8406,f8404,f8403,f8400,f8399,f8398,f8397,f8395,f8394,f8393,f8392,f8391,f8670,f8671,f8751,f8750,f8748,f8747,f8930,f8929,f8956,f8955,f8961,f9000,f9034,f9071,f9168,f9265,f9359,f9395,f9397,f9407,f9413,f9389,f9487,f9523,f9525,f9535,f9541,f9517,f9621,f9657,f9659,f9669,f9672,f9676,f9651,f9721,f9757,f9759,f9769,f9772,f9776,f9751,f9881,f10090,f10089,f10087,f10086,f10161,f10160,f10159,f10158,f10194,f10195,f10196,f10197,f10198,f10199,f10200,f10240,f10239,f10238,f10237,f10269,f10270,f10271,f10272,f10273,f10274,f10275,f10382,f10383,f10390,f10391,f10392,f10398,f10402,f10444,f10445,f10454,f10455,f10456,f10457,f10458,f10459,f10460,f1297,f10590,f10591,f10600,f10601,f10602,f10603,f10604,f10605,f10606,f1375,f10824,f11027,f11029,f10832,f11032,f11036,f11041,f11044,f11050,f11053,f11132,f11133,f11134,f11135,f11245,f11248,f11249,f11250,f11251,f11382,f11383,f11384,f11385,f11465,f11498,f11499,f11500,f11501,f11573,f11574,f11575,f11576,f11681,f11684,f11685,f11686,f11687,f11759,f11758,f11756,f11755,f11782,f11815,f11816,f11817,f11818,f11819,f11831,f11864,f11865,f11866,f11867,f11869,f11874,f11879,f11880,f11881,f11882,f11925,f11928,f11929,f11930,f11931,f12032,f12033,f12034,f12035,f12108,f12141,f12142,f12143,f12144,f12157,f12290,f12293,f12394,f12497,f12530,f12643,f12640,f12635,f12632,f12709,f12711,f12712,f12713,f12714,f12716,f12717,f12718,f12748,f12749,f12758,f12759,f12760,f12761,f12762,f12763,f12764,f12808,f12806,f12809,f12873,f12879,f12877,f12880,f12940,f12944,f12950,f12948,f12951,f13015,f13018,f13054,f13061,f13071,f13077,f13051,f13096,f13099,f13135,f13142,f13152,f13158,f13132,f13173,f13178,f13211,f13247,f13254,f13264,f13267,f13271,f13244,f13317,f13350,f13386,f13393,f13403,f13406,f13410,f13383,f13420,f13523,f13525,f13529,f13530,f13532,f13536,f13537,f13538,f13539,f13542,f13543,f13546,f13550,f13554,f13598,f13611,f13612,f13613,f13614,f13615,f13619,f13620,f13621,f13622,f13625,f13626,f13627,f13631,f13633,f13736,f13737,f13807,f14151,f14155,f14158,f14161,f14164,f14166,f14167,f14168,f14169,f14170,f14171,f14172,f14173,f14174,f14175,f14176,f14177,f14178,f13881,f13882,f14180,f14182,f14184,f14186,f14188,f14153,f14535,f14536,f14537])).
fof(f14537,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X32)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.num.num',X31)),X32)) ) | $spl56),
  inference(subsumption_resolution,[],[f14234,f318])).
fof(f14234,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.num.num',X31)),X32)) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f3985])).
fof(f14536,plain,(
  ( ! [X28,X29,X27] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num',X28)),X29)) ) | $spl56),
  inference(subsumption_resolution,[],[f14233,f318])).
fof(f14233,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num',X28)),X29)) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f3984])).
fof(f14535,plain,(
  ( ! [X26,X24,X25] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num',X25)))) ) | $spl56),
  inference(subsumption_resolution,[],[f14232,f318])).
fof(f14232,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'T_0')) | sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num',X25)))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f3983])).
fof(f13882,plain,(
  ( ! [X146,X144,X147,X145] : (s(bool,'T_0') != s(bool,f(s(X144,sK13(X145,X146,bool,X144)))) | s(fun(X144,bool),X145) = s(fun(X144,bool),X146) | s(bool,X147) = s(bool,g(s(X144,sK13(X145,X146,bool,X144)))) | s(bool,'T_0') = s(bool,X147)) )),
  inference(superposition,[],[f468,f942])).
fof(f13881,plain,(
  ( ! [X142,X140,X143,X141] : (s(X142,f(s(bool,'T_0'))) != s(X142,g(s(bool,'T_0'))) | s(fun(bool,X142),X140) = s(fun(bool,X142),X141) | s(bool,X143) = s(bool,sK13(X140,X141,X142,bool)) | s(bool,'T_0') = s(bool,X143)) )),
  inference(superposition,[],[f468,f942])).
fof(f14178,plain,(
  ( ! [X123,X124] : (s('type.num.num','const.num.|48|_0') = s('type.num.num',X123) | s(bool,X124) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X123))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(subsumption_resolution,[],[f13872,f318])).
fof(f13872,plain,(
  ( ! [X123,X124] : (~p(s(bool,'T_0')) | s('type.num.num','const.num.|48|_0') = s('type.num.num',X123) | s(bool,X124) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X123))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f378,f942])).
fof(f14177,plain,(
  ( ! [X103,X101,X102] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X102),s('type.num.num',X101)))) | s(bool,X103) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X101))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,X103)) )),
  inference(subsumption_resolution,[],[f13863,f318])).
fof(f13863,plain,(
  ( ! [X103,X101,X102] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X102),s('type.num.num',X101)))) | s(bool,X103) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X101))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X102))))) | s(bool,'T_0') = s(bool,X103)) )),
  inference(superposition,[],[f431,f942])).
fof(f14176,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X73),s('type.num.num',X74)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X72),s('type.num.num',X74)))) | s(bool,X75) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X72),s('type.num.num',X73))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(subsumption_resolution,[],[f13853,f318])).
fof(f13853,plain,(
  ( ! [X74,X72,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X73),s('type.num.num',X74)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X72),s('type.num.num',X74)))) | s(bool,X75) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X72),s('type.num.num',X73))) | s(bool,'T_0') = s(bool,X75)) )),
  inference(superposition,[],[f471,f942])).
fof(f14175,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X59),s('type.num.num',X58)))) | s(bool,X60) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X58),s('type.num.num',X59))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f13847,f318])).
fof(f13847,plain,(
  ( ! [X59,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X59),s('type.num.num',X58)))) | s(bool,X60) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X58),s('type.num.num',X59))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f459,f942])).
fof(f14174,plain,(
  ( ! [X45,X43,X44,X42] : (~sP8(X42,X43,X44) | sP9(X42,X43,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f13839,f318])).
fof(f13839,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'T_0')) | ~sP8(X42,X43,X44) | sP9(X42,X43,X44) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f541,f942])).
fof(f14173,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | sP9(X38,X39,X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f13838,f318])).
fof(f13838,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X39)) | sP9(X38,X39,X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f537,f942])).
fof(f14172,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | sP9(X34,X36,X35) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f13837,f318])).
fof(f13837,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'T_0')) | ~p(s(bool,X35)) | sP9(X34,X36,X35) | s(bool,X34) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f536,f942])).
fof(f14171,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X32)) | sP1(X32,X31) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f13836,f318])).
fof(f13836,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | p(s(bool,X32)) | sP1(X32,X31) | s(bool,X31) = s(bool,X33) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f453,f942])).
fof(f14170,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X28,X29) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f13835,f318])).
fof(f13835,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X28,X29) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f450,f942])).
fof(f14169,plain,(
  ( ! [X26,X27,X25] : (sP0(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f13834,f318])).
fof(f13834,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | sP0(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f435,f942])).
fof(f14168,plain,(
  ( ! [X24,X23,X22] : (sP0(X23,X22) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f13833,f318])).
fof(f13833,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X23,X22) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f434,f942])).
fof(f14167,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(subsumption_resolution,[],[f13830,f318])).
fof(f13830,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X14)) | s(bool,X13) = s(bool,X14) | s(bool,X13) = s(bool,X15) | s(bool,'T_0') = s(bool,X15)) )),
  inference(superposition,[],[f427,f942])).
fof(f14166,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f13827,f318])).
fof(f13827,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f370,f942])).
fof(f13807,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f366,f942])).
fof(f13737,plain,(
  ( ! [X158,X161,X159,X160] : (s(bool,X161) != s(bool,f(s(X158,sK13(X159,X160,bool,X158)))) | s(fun(X158,bool),X159) = s(fun(X158,bool),X160) | s(bool,'T_0') = s(bool,X161) | s(bool,'T_0') = s(bool,g(s(X158,sK13(X159,X160,bool,X158))))) )),
  inference(superposition,[],[f468,f942])).
fof(f13736,plain,(
  ( ! [X156,X154,X157,X155] : (s(X156,f(s(bool,X157))) != s(X156,g(s(bool,X157))) | s(fun(bool,X156),X154) = s(fun(bool,X156),X155) | s(bool,'T_0') = s(bool,X157) | s(bool,'T_0') = s(bool,sK13(X154,X155,X156,bool))) )),
  inference(superposition,[],[f468,f942])).
fof(f13633,plain,(
  ( ! [X72,X73] : (s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X72))) | s(bool,X73) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X72))))) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f942,f346])).
fof(f13631,plain,(
  ( ! [X70,X69] : (s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X69))) | s(bool,X70) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X69))))) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f942,f345])).
fof(f13627,plain,(
  ( ! [X64,X65,X63] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X64),s('type.num.num',X63))) | s(bool,X65) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X63))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X64))))) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f942,f412])).
fof(f13626,plain,(
  ( ! [X61,X59,X60] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X60),s('type.num.num',X59))) | s(bool,X61) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X59))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X60))))) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f942,f411])).
fof(f13625,plain,(
  ( ! [X57,X58] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X57))) | s(bool,X58) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X57))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f942,f356])).
fof(f13622,plain,(
  ( ! [X52,X53,X51] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num',X52))) | s(bool,X53) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X51))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X52))))) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f942,f415])).
fof(f13621,plain,(
  ( ! [X50,X48,X49] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X48),s('type.num.num',X49))) | s(bool,X50) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X49))))) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f942,f413])).
fof(f13620,plain,(
  ( ! [X47,X45,X46] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X45),s('type.num.num',X46))) | s(bool,X47) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))))) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f942,f414])).
fof(f13619,plain,(
  ( ! [X43,X44,X42] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X42),s('type.num.num',X43))) | s(bool,X44) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X43))))) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f942,f408])).
fof(f13615,plain,(
  ( ! [X35,X36] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X35),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,X36) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X36)) )),
  inference(superposition,[],[f942,f359])).
fof(f13614,plain,(
  ( ! [X30,X31,X29,X32] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31))) | s(bool,X32) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X31))))) | s(bool,'T_0') = s(bool,X32)) )),
  inference(superposition,[],[f942,f470])).
fof(f13613,plain,(
  ( ! [X28,X26,X27] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X27))) | s(bool,X28) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X26))),s('type.num.num',X27))) | s(bool,'T_0') = s(bool,X28)) )),
  inference(superposition,[],[f942,f407])).
fof(f13612,plain,(
  ( ! [X23,X21,X22] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X21),s('type.num.num',X22))) | s(bool,X23) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f942,f409])).
fof(f13611,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X19))) | s(bool,X20) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))) | s(bool,'T_0') = s(bool,X20)) )),
  inference(superposition,[],[f942,f357])).
fof(f13598,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f942,f942])).
fof(f13554,plain,(
  ( ! [X88,X87] : (s(bool,X88) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X87))) | s(bool,'T_0') = s(bool,X88) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X87)))))) )),
  inference(superposition,[],[f942,f346])).
fof(f13550,plain,(
  ( ! [X83,X82] : (s(bool,X83) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X82))) | s(bool,'T_0') = s(bool,X83) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X82)))))) )),
  inference(superposition,[],[f942,f345])).
fof(f13546,plain,(
  ( ! [X78,X76,X77] : (s(bool,X78) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X77),s('type.num.num',X76))) | s(bool,'T_0') = s(bool,X78) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X76))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X77)))))) )),
  inference(superposition,[],[f942,f412])).
fof(f13543,plain,(
  ( ! [X70,X72,X71] : (s(bool,X72) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X71),s('type.num.num',X70))) | s(bool,'T_0') = s(bool,X72) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X70))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X71)))))) )),
  inference(superposition,[],[f942,f411])).
fof(f13542,plain,(
  ( ! [X68,X69] : (s(bool,X69) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X68))) | s(bool,'T_0') = s(bool,X69) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X68))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f942,f356])).
fof(f13539,plain,(
  ( ! [X64,X62,X63] : (s(bool,X64) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X62),s('type.num.num',X63))) | s(bool,'T_0') = s(bool,X64) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X62))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X63)))))) )),
  inference(superposition,[],[f942,f415])).
fof(f13538,plain,(
  ( ! [X61,X59,X60] : (s(bool,X61) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X59),s('type.num.num',X60))) | s(bool,'T_0') = s(bool,X61) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X59))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X60)))))) )),
  inference(superposition,[],[f942,f413])).
fof(f13537,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X56),s('type.num.num',X57))) | s(bool,'T_0') = s(bool,X58) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X57)))))) )),
  inference(superposition,[],[f942,f414])).
fof(f13536,plain,(
  ( ! [X54,X55,X53] : (s(bool,X55) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X53),s('type.num.num',X54))) | s(bool,'T_0') = s(bool,X55) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X53))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X54)))))) )),
  inference(superposition,[],[f942,f408])).
fof(f13532,plain,(
  ( ! [X47,X46] : (s(bool,X47) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X46),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X47) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f942,f359])).
fof(f13530,plain,(
  ( ! [X39,X41,X42,X40] : (s(bool,X42) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num',X41))) | s(bool,'T_0') = s(bool,X42) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num',X40))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X39),s('type.num.num',X41)))))) )),
  inference(superposition,[],[f942,f470])).
fof(f13529,plain,(
  ( ! [X37,X38,X36] : (s(bool,X38) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X37))) | s(bool,'T_0') = s(bool,X38) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)))) )),
  inference(superposition,[],[f942,f407])).
fof(f13525,plain,(
  ( ! [X30,X28,X29] : (s(bool,X30) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X28),s('type.num.num',X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))))) )),
  inference(superposition,[],[f942,f409])).
fof(f13523,plain,(
  ( ! [X26,X25] : (s(bool,X26) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X25))) | s(bool,'T_0') = s(bool,X26) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f942,f357])).
fof(f13420,plain,(
  ( ! [X0] : (sP4('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f8751,f486])).
fof(f13383,plain,(
  ( ! [X6,X4,X5] : (sP1(X6,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))))))) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1389,f422])).
fof(f13410,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46)))))))) | sP1(X47,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))))))))) | p(s(bool,X47))) )),
  inference(forward_demodulation,[],[f13369,f346])).
fof(f13369,plain,(
  ( ! [X47,X45,X46] : (sP1(X47,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))))))))) | p(s(bool,X47)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46))))))))))) )),
  inference(superposition,[],[f1389,f422])).
fof(f13406,plain,(
  ( ! [X43] : (sP1(X43,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X43))) )),
  inference(subsumption_resolution,[],[f13405,f318])).
fof(f13405,plain,(
  ( ! [X43] : (~p(s(bool,'T_0')) | sP1(X43,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X43))) )),
  inference(forward_demodulation,[],[f13404,f2761])).
fof(f13404,plain,(
  ( ! [X43] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) | sP1(X43,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X43))) )),
  inference(forward_demodulation,[],[f13367,f346])).
fof(f13367,plain,(
  ( ! [X43] : (sP1(X43,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X43)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(superposition,[],[f1389,f1935])).
fof(f13403,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41)))))) | sP1(X42,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))))))) | p(s(bool,X42))) )),
  inference(forward_demodulation,[],[f13402,f346])).
fof(f13402,plain,(
  ( ! [X41,X42] : (sP1(X42,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))))))) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(forward_demodulation,[],[f13366,f353])).
fof(f13366,plain,(
  ( ! [X41,X42] : (sP1(X42,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))))))) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(superposition,[],[f1389,f353])).
fof(f13393,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))))) | sP1(X23,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22))))))))) | p(s(bool,X23))) )),
  inference(forward_demodulation,[],[f13392,f551])).
fof(f13392,plain,(
  ( ! [X23,X21,X22] : (sP1(X23,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22))))))))) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(forward_demodulation,[],[f13358,f4750])).
fof(f13358,plain,(
  ( ! [X23,X21,X22] : (sP1(X23,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))))))))) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(superposition,[],[f1389,f551])).
fof(f13386,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) | sP1(X9,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))))))) | p(s(bool,X9))) )),
  inference(forward_demodulation,[],[f13385,f2786])).
fof(f13385,plain,(
  ( ! [X8,X9] : (sP1(X9,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))))))) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f13353,f353])).
fof(f13353,plain,(
  ( ! [X8,X9] : (sP1(X9,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(superposition,[],[f1389,f2786])).
fof(f13350,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))))) | p(s(bool,X2)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(superposition,[],[f1389,f406])).
fof(f13317,plain,(
  ( ! [X0] : (~sP2('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f567,f561,f8750,f489])).
fof(f13244,plain,(
  ( ! [X6,X4,X5] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))))),X6) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1388,f422])).
fof(f13271,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46)))))))) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46)))))))),X47) | p(s(bool,X47))) )),
  inference(forward_demodulation,[],[f13230,f346])).
fof(f13230,plain,(
  ( ! [X47,X45,X46] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46)))))))),X47) | p(s(bool,X47)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46))))))))))) )),
  inference(superposition,[],[f1388,f422])).
fof(f13267,plain,(
  ( ! [X43] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X43) | p(s(bool,X43))) )),
  inference(subsumption_resolution,[],[f13266,f318])).
fof(f13266,plain,(
  ( ! [X43] : (~p(s(bool,'T_0')) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X43) | p(s(bool,X43))) )),
  inference(forward_demodulation,[],[f13265,f2761])).
fof(f13265,plain,(
  ( ! [X43] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X43) | p(s(bool,X43))) )),
  inference(forward_demodulation,[],[f13228,f346])).
fof(f13228,plain,(
  ( ! [X43] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X43) | p(s(bool,X43)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(superposition,[],[f1388,f1935])).
fof(f13264,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41)))))) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))))))),X42) | p(s(bool,X42))) )),
  inference(forward_demodulation,[],[f13263,f346])).
fof(f13263,plain,(
  ( ! [X41,X42] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))))))),X42) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(forward_demodulation,[],[f13227,f353])).
fof(f13227,plain,(
  ( ! [X41,X42] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))))))),X42) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(superposition,[],[f1388,f353])).
fof(f13254,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))))) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22)))))))),X23) | p(s(bool,X23))) )),
  inference(forward_demodulation,[],[f13253,f551])).
fof(f13253,plain,(
  ( ! [X23,X21,X22] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22)))))))),X23) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(forward_demodulation,[],[f13219,f4750])).
fof(f13219,plain,(
  ( ! [X23,X21,X22] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))))),X23) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(superposition,[],[f1388,f551])).
fof(f13247,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))),X9) | p(s(bool,X9))) )),
  inference(forward_demodulation,[],[f13246,f2786])).
fof(f13246,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))),X9) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f13214,f353])).
fof(f13214,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))),X9) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(superposition,[],[f1388,f2786])).
fof(f13211,plain,(
  ( ! [X2,X0,X1] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),X2) | p(s(bool,X2)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(superposition,[],[f1388,f406])).
fof(f13178,plain,(
  ( ! [X0] : (~sP2('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f565,f559,f8747,f490])).
fof(f13173,plain,(
  sP4('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f7488,f486])).
fof(f13132,plain,(
  ( ! [X6,X4,X5] : (sP1(X6,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))))))) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1311,f422])).
fof(f13158,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46)))))))) | sP1(X47,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))))))))) | p(s(bool,X47))) )),
  inference(forward_demodulation,[],[f13118,f345])).
fof(f13118,plain,(
  ( ! [X47,X45,X46] : (sP1(X47,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46))))))))) | p(s(bool,X47)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46))))))))))) )),
  inference(superposition,[],[f1311,f422])).
fof(f13152,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41)))))) | sP1(X42,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))))))) | p(s(bool,X42))) )),
  inference(forward_demodulation,[],[f13151,f345])).
fof(f13151,plain,(
  ( ! [X41,X42] : (sP1(X42,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))))))) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(forward_demodulation,[],[f13115,f353])).
fof(f13115,plain,(
  ( ! [X41,X42] : (sP1(X42,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))))))))) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(superposition,[],[f1311,f353])).
fof(f13142,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))))) | sP1(X23,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22))))))))) | p(s(bool,X23))) )),
  inference(forward_demodulation,[],[f13141,f551])).
fof(f13141,plain,(
  ( ! [X23,X21,X22] : (sP1(X23,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22))))))))) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(forward_demodulation,[],[f13107,f4750])).
fof(f13107,plain,(
  ( ! [X23,X21,X22] : (sP1(X23,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))))))))) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(superposition,[],[f1311,f551])).
fof(f13135,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) | sP1(X9,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))))))) | p(s(bool,X9))) )),
  inference(forward_demodulation,[],[f13134,f2786])).
fof(f13134,plain,(
  ( ! [X8,X9] : (sP1(X9,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))))))) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f13102,f353])).
fof(f13102,plain,(
  ( ! [X8,X9] : (sP1(X9,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(superposition,[],[f1311,f2786])).
fof(f13099,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))))))) | p(s(bool,X2)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(superposition,[],[f1311,f406])).
fof(f13096,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f574,f561,f7441,f490])).
fof(f13051,plain,(
  ( ! [X6,X4,X5] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))))),X6) | p(s(bool,X6)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1310,f422])).
fof(f13077,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46)))))))) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46)))))))),X47) | p(s(bool,X47))) )),
  inference(forward_demodulation,[],[f13037,f345])).
fof(f13037,plain,(
  ( ! [X47,X45,X46] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46)))))))),X47) | p(s(bool,X47)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X46))))))))))) )),
  inference(superposition,[],[f1310,f422])).
fof(f13071,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41)))))) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))))))),X42) | p(s(bool,X42))) )),
  inference(forward_demodulation,[],[f13070,f345])).
fof(f13070,plain,(
  ( ! [X41,X42] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))))))),X42) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(forward_demodulation,[],[f13034,f353])).
fof(f13034,plain,(
  ( ! [X41,X42] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))))))),X42) | p(s(bool,X42)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X41))))))))) )),
  inference(superposition,[],[f1310,f353])).
fof(f13061,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))))) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22)))))))),X23) | p(s(bool,X23))) )),
  inference(forward_demodulation,[],[f13060,f551])).
fof(f13060,plain,(
  ( ! [X23,X21,X22] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num',X22)))))))),X23) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(forward_demodulation,[],[f13026,f4750])).
fof(f13026,plain,(
  ( ! [X23,X21,X22] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X21),s('type.num.num','const.num.SUC_1'(s('type.num.num',X22)))))))),X23) | p(s(bool,X23)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))),s('type.num.num',X22))))))) )),
  inference(superposition,[],[f1310,f551])).
fof(f13054,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))),X9) | p(s(bool,X9))) )),
  inference(forward_demodulation,[],[f13053,f2786])).
fof(f13053,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))),X9) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(forward_demodulation,[],[f13021,f353])).
fof(f13021,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))),X9) | p(s(bool,X9)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))))))) )),
  inference(superposition,[],[f1310,f2786])).
fof(f13018,plain,(
  ( ! [X2,X0,X1] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),X2) | p(s(bool,X2)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) )),
  inference(superposition,[],[f1310,f406])).
fof(f13015,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f572,f561,f7321,f490])).
fof(f12951,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X2,X3) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))),X2,X3)) )),
  inference(superposition,[],[f920,f406])).
fof(f12948,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f582,f920])).
fof(f12950,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f584,f920])).
fof(f12944,plain,(
  ~sP2('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f566,f561,f7259,f490])).
fof(f12940,plain,(
  sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f7187,f486])).
fof(f12880,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')),X2,X3) | ~sP8('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')),X2,X3)) )),
  inference(superposition,[],[f889,f406])).
fof(f12877,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f582,f889])).
fof(f12879,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f584,f889])).
fof(f12873,plain,(
  ~sP2('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f568,f561,f7138,f490])).
fof(f12809,plain,(
  ( ! [X2,X0,X3,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))),X2,X3) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))),X2,X3)) )),
  inference(superposition,[],[f792,f406])).
fof(f12806,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f582,f792])).
fof(f12808,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f584,f792])).
fof(f12764,plain,(
  ( ! [X54,X52,X53,X51] : (~sP8('const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num',X52)))),X53,X54) | sP9('const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num',X52)))),X53,X54)) )),
  inference(subsumption_resolution,[],[f12739,f318])).
fof(f12739,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num',X52)))),X53,X54) | sP9('const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X51),s('type.num.num',X52)))),X53,X54)) )),
  inference(superposition,[],[f541,f6865])).
fof(f12763,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X49)) | sP9('const.arithmetic.<=_2'(s('type.num.num',X47),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X47),s('type.num.num',X48)))),X49,X50)) )),
  inference(subsumption_resolution,[],[f12738,f318])).
fof(f12738,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X49)) | sP9('const.arithmetic.<=_2'(s('type.num.num',X47),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X47),s('type.num.num',X48)))),X49,X50)) )),
  inference(superposition,[],[f537,f6865])).
fof(f12762,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X45)) | sP9('const.arithmetic.<=_2'(s('type.num.num',X43),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))),X46,X45)) )),
  inference(subsumption_resolution,[],[f12737,f318])).
fof(f12737,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X45)) | sP9('const.arithmetic.<=_2'(s('type.num.num',X43),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))),X46,X45)) )),
  inference(superposition,[],[f536,f6865])).
fof(f12761,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X42)) | sP1(X42,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41)))))) )),
  inference(subsumption_resolution,[],[f12736,f318])).
fof(f12736,plain,(
  ( ! [X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | sP1(X42,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41)))))) )),
  inference(superposition,[],[f453,f6865])).
fof(f12760,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X39)) | sP1('const.arithmetic.<=_2'(s('type.num.num',X37),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))),X39)) )),
  inference(subsumption_resolution,[],[f12735,f318])).
fof(f12735,plain,(
  ( ! [X39,X37,X38] : (~p(s(bool,'T_0')) | p(s(bool,X39)) | sP1('const.arithmetic.<=_2'(s('type.num.num',X37),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X37),s('type.num.num',X38)))),X39)) )),
  inference(superposition,[],[f450,f6865])).
fof(f12759,plain,(
  ( ! [X35,X36,X34] : (sP0('const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num',X35)))),X36)) )),
  inference(subsumption_resolution,[],[f12734,f318])).
fof(f12734,plain,(
  ( ! [X35,X36,X34] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.<=_2'(s('type.num.num',X34),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X34),s('type.num.num',X35)))),X36)) )),
  inference(superposition,[],[f435,f6865])).
fof(f12758,plain,(
  ( ! [X33,X31,X32] : (sP0(X33,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32)))))) )),
  inference(subsumption_resolution,[],[f12733,f318])).
fof(f12733,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'T_0')) | sP0(X33,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X31),s('type.num.num',X32)))))) )),
  inference(superposition,[],[f434,f6865])).
fof(f12749,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X4))))) )),
  inference(subsumption_resolution,[],[f12721,f318])).
fof(f12721,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X4))))) )),
  inference(superposition,[],[f471,f6865])).
fof(f12748,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))))) )),
  inference(forward_demodulation,[],[f12719,f551])).
fof(f12719,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X4))),s('type.num.num',X5)))))) )),
  inference(superposition,[],[f6865,f407])).
fof(f12718,plain,(
  ( ! [X15,X16] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X15),s('type.num.num',X16))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X15))),s('type.num.num',X16)))))) )),
  inference(superposition,[],[f6865,f418])).
fof(f12717,plain,(
  ( ! [X14,X13] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X13),s('type.num.num','const.num.SUC_1'(s('type.num.num',X14)))))))) )),
  inference(superposition,[],[f6865,f551])).
fof(f12716,plain,(
  ( ! [X12] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X12)))))) )),
  inference(superposition,[],[f6865,f550])).
fof(f12714,plain,(
  ( ! [X10,X8,X9] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10)))))))) )),
  inference(superposition,[],[f6865,f469])).
fof(f12713,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.num.SUC_1'(s('type.num.num',X7)))))))) )),
  inference(superposition,[],[f6865,f421])).
fof(f12712,plain,(
  ( ! [X5] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5)))))) )),
  inference(superposition,[],[f6865,f2786])).
fof(f12711,plain,(
  ( ! [X4] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X4)))) )),
  inference(superposition,[],[f6865,f338])).
fof(f12709,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(superposition,[],[f6865,f406])).
fof(f12632,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f545,f6865,f506])).
fof(f12635,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f6865,f504])).
fof(f12640,plain,(
  ( ! [X0,X1] : (~sP5('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f6865,f506])).
fof(f12643,plain,(
  ( ! [X0,X1] : (sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f6865,f504])).
fof(f12530,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f3933,f416])).
fof(f12497,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f3933,f486])).
fof(f12394,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f3932,f416])).
fof(f12293,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f3931,f416])).
fof(f12290,plain,(
  ( ! [X0,X1] : (sP4('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f3931,f486])).
fof(f12157,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f3930,f416])).
fof(f12144,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3859,f422])).
fof(f12143,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3859,f322])).
fof(f12142,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3859,f1935])).
fof(f12141,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3859,f353])).
fof(f12108,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f3859,f486])).
fof(f12035,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3858,f422])).
fof(f12034,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3858,f322])).
fof(f12033,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3858,f1935])).
fof(f12032,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3858,f353])).
fof(f11931,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f3857,f422])).
fof(f11930,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3857,f322])).
fof(f11929,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3857,f1935])).
fof(f11928,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f3857,f353])).
fof(f11925,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f3857,f486])).
fof(f11882,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f3856,f422])).
fof(f11881,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3856,f322])).
fof(f11880,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3856,f1935])).
fof(f11879,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f3856,f353])).
fof(f11874,plain,(
  sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f11759,f486])).
fof(f11869,plain,(
  ~sP2('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f545,f567,f561,f11758,f489])).
fof(f11867,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f3794,f422])).
fof(f11866,plain,(
  sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f3794,f322])).
fof(f11865,plain,(
  sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3794,f1935])).
fof(f11864,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f3794,f353])).
fof(f11831,plain,(
  ( ! [X0] : (sP4('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f3794,f486])).
fof(f11819,plain,(
  ~sP2('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f565,f561,f11755,f490])).
fof(f11818,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f3793,f422])).
fof(f11817,plain,(
  ~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f3793,f322])).
fof(f11816,plain,(
  ~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3793,f1935])).
fof(f11815,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f3793,f353])).
fof(f11782,plain,(
  ( ! [X14] : (s(bool,X14) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(subsumption_resolution,[],[f11775,f545])).
fof(f11775,plain,(
  ( ! [X14] : (p(s(bool,'F_0')) | s(bool,X14) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X14)) )),
  inference(superposition,[],[f1116,f2762])).
fof(f11755,plain,(
  ~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f2762,f506])).
fof(f11756,plain,(
  sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f318,f2762,f504])).
fof(f11758,plain,(
  ~sP5('F_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f545,f2762,f506])).
fof(f11759,plain,(
  sP5('T_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f318,f2762,f504])).
fof(f11687,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f3792,f422])).
fof(f11686,plain,(
  sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f3792,f322])).
fof(f11685,plain,(
  sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3792,f1935])).
fof(f11684,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f3792,f353])).
fof(f11681,plain,(
  ( ! [X0] : (sP4('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f3792,f486])).
fof(f11576,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f3791,f422])).
fof(f11575,plain,(
  ~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f3791,f322])).
fof(f11574,plain,(
  ~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3791,f1935])).
fof(f11573,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f3791,f353])).
fof(f11501,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3734,f422])).
fof(f11500,plain,(
  sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3734,f322])).
fof(f11499,plain,(
  sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3734,f1935])).
fof(f11498,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3734,f353])).
fof(f11465,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f3734,f486])).
fof(f11385,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3733,f422])).
fof(f11384,plain,(
  ~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3733,f322])).
fof(f11383,plain,(
  ~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f3733,f1935])).
fof(f11382,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f3733,f353])).
fof(f11251,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f3732,f422])).
fof(f11250,plain,(
  sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f3732,f322])).
fof(f11249,plain,(
  sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3732,f1935])).
fof(f11248,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f3732,f353])).
fof(f11245,plain,(
  ( ! [X0] : (sP4('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f3732,f486])).
fof(f11135,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f3731,f422])).
fof(f11134,plain,(
  ~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f3731,f322])).
fof(f11133,plain,(
  ~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f3731,f1935])).
fof(f11132,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f3731,f353])).
fof(f11053,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X146),s('type.num.num',X148)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X147),s('type.num.num','const.num.SUC_1'(s('type.num.num',X146)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X147),s('type.num.num',X148))))) )),
  inference(forward_demodulation,[],[f11052,f338])).
fof(f11052,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X146),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X148)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X147),s('type.num.num','const.num.SUC_1'(s('type.num.num',X146)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X147),s('type.num.num',X148))))) )),
  inference(forward_demodulation,[],[f10912,f4655])).
fof(f10912,plain,(
  ( ! [X146,X147,X148] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X147),s('type.num.num','const.num.SUC_1'(s('type.num.num',X146)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X146),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X148)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X147),s('type.num.num',X148))))) )),
  inference(superposition,[],[f471,f2786])).
fof(f11050,plain,(
  ( ! [X125,X124] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X124),s('type.num.num',X125)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X124))))))) )),
  inference(forward_demodulation,[],[f11049,f338])).
fof(f11049,plain,(
  ( ! [X125,X124] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X124),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X125)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X124))))))) )),
  inference(forward_demodulation,[],[f10904,f4655])).
fof(f10904,plain,(
  ( ! [X125,X124] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X124)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X124),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X125))))) )),
  inference(superposition,[],[f459,f2786])).
fof(f11044,plain,(
  ( ! [X116,X117] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X116),s('type.num.num',X117)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X117))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X116))))))))) )),
  inference(forward_demodulation,[],[f11043,f338])).
fof(f11043,plain,(
  ( ! [X116,X117] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X116),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X117)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X117))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X116))))))))) )),
  inference(forward_demodulation,[],[f10900,f4655])).
fof(f10900,plain,(
  ( ! [X116,X117] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X117))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X116)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X116),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X117))))) )),
  inference(superposition,[],[f432,f2786])).
fof(f11041,plain,(
  ( ! [X112,X113] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X112),s('type.num.num',X113)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X113))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X112))))))))) )),
  inference(forward_demodulation,[],[f11040,f338])).
fof(f11040,plain,(
  ( ! [X112,X113] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X112),s('type.num.num','const.num.|48|_0'))),s('type.num.num',X113)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X113))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X112))))))))) )),
  inference(forward_demodulation,[],[f10898,f4655])).
fof(f10898,plain,(
  ( ! [X112,X113] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X113))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X112)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X112),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X113))))) )),
  inference(superposition,[],[f431,f2786])).
fof(f11036,plain,(
  ( ! [X105,X104] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X105))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X104))))) | s('type.num.num',X105) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X104),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f10894,f353])).
fof(f10894,plain,(
  ( ! [X105,X104] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X105))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X104))))) | s('type.num.num',X105) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X104),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f429,f2786])).
fof(f11032,plain,(
  ( ! [X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f10858,f353])).
fof(f10858,plain,(
  ( ! [X36] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X36),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f374,f2786])).
fof(f10832,plain,(
  ( ! [X8,X9] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X8)))))) )),
  inference(superposition,[],[f470,f2786])).
fof(f11029,plain,(
  ( ! [X6,X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))) )),
  inference(forward_demodulation,[],[f10831,f407])).
fof(f10831,plain,(
  ( ! [X6,X7] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X7))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7)))))) )),
  inference(superposition,[],[f470,f2786])).
fof(f11027,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(forward_demodulation,[],[f10829,f353])).
fof(f10829,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f422,f2786])).
fof(f10824,plain,(
  ( ! [X5] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f2786,f551])).
fof(f1375,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f346,f322])).
fof(f1297,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f345,f322])).
fof(f10402,plain,(
  ( ! [X60] : (s(bool,X60) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(subsumption_resolution,[],[f10377,f545])).
fof(f10377,plain,(
  ( ! [X60] : (p(s(bool,'F_0')) | s(bool,X60) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X60)) )),
  inference(superposition,[],[f1116,f608])).
fof(f10398,plain,(
  ( ! [X47,X48] : (s(bool,X48) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X47))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f10371,f545])).
fof(f10371,plain,(
  ( ! [X47,X48] : (p(s(bool,'F_0')) | s(bool,X48) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X47))) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f1116,f331])).
fof(f10392,plain,(
  ( ! [X33,X32] : (s(bool,X33) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(subsumption_resolution,[],[f10365,f545])).
fof(f10365,plain,(
  ( ! [X33,X32] : (p(s(bool,'F_0')) | s(bool,X33) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X33)) )),
  inference(superposition,[],[f1116,f348])).
fof(f10391,plain,(
  ( ! [X30,X31] : (s(bool,X31) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(subsumption_resolution,[],[f10364,f545])).
fof(f10364,plain,(
  ( ! [X30,X31] : (p(s(bool,'F_0')) | s(bool,X31) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X31)) )),
  inference(superposition,[],[f1116,f347])).
fof(f10390,plain,(
  ( ! [X29] : (s(bool,X29) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(subsumption_resolution,[],[f10363,f545])).
fof(f10363,plain,(
  ( ! [X29] : (p(s(bool,'F_0')) | s(bool,X29) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X29)) )),
  inference(superposition,[],[f1116,f3900])).
fof(f10383,plain,(
  ( ! [X6,X7] : (s(bool,X7) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f10352,f545])).
fof(f10352,plain,(
  ( ! [X6,X7] : (p(s(bool,'F_0')) | s(bool,X7) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1116,f3903])).
fof(f10382,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f10351,f545])).
fof(f10351,plain,(
  ( ! [X4,X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1116,f330])).
fof(f10275,plain,(
  ( ! [X12,X13] : (~sP8('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X12,X13) | sP9('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X12,X13)) )),
  inference(subsumption_resolution,[],[f10258,f318])).
fof(f10258,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X12,X13) | sP9('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X12,X13)) )),
  inference(superposition,[],[f541,f8671])).
fof(f10274,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X10,X11)) )),
  inference(subsumption_resolution,[],[f10257,f318])).
fof(f10257,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X10,X11)) )),
  inference(superposition,[],[f537,f8671])).
fof(f10273,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X9,X8)) )),
  inference(subsumption_resolution,[],[f10256,f318])).
fof(f10256,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X9,X8)) )),
  inference(superposition,[],[f536,f8671])).
fof(f10272,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f10255,f318])).
fof(f10255,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1(X7,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f453,f8671])).
fof(f10271,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X6)) )),
  inference(subsumption_resolution,[],[f10254,f318])).
fof(f10254,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X6)) )),
  inference(superposition,[],[f450,f8671])).
fof(f10270,plain,(
  ( ! [X5] : (sP0('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X5)) )),
  inference(subsumption_resolution,[],[f10253,f318])).
fof(f10253,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),X5)) )),
  inference(superposition,[],[f435,f8671])).
fof(f10269,plain,(
  ( ! [X4] : (sP0(X4,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f10252,f318])).
fof(f10252,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(X4,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f434,f8671])).
fof(f10237,plain,(
  ~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f545,f8671,f506])).
fof(f10238,plain,(
  sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f318,f8671,f504])).
fof(f10239,plain,(
  ~sP5('F_0','T_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f545,f8671,f506])).
fof(f10240,plain,(
  sP5('T_0','T_0','const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f318,f8671,f504])).
fof(f10200,plain,(
  ( ! [X12,X13] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12,X13) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12,X13)) )),
  inference(subsumption_resolution,[],[f10180,f318])).
fof(f10180,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12,X13) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X12,X13)) )),
  inference(superposition,[],[f541,f8670])).
fof(f10199,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X10,X11)) )),
  inference(subsumption_resolution,[],[f10179,f318])).
fof(f10179,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X10,X11)) )),
  inference(superposition,[],[f537,f8670])).
fof(f10198,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X9,X8)) )),
  inference(subsumption_resolution,[],[f10178,f318])).
fof(f10178,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X9,X8)) )),
  inference(superposition,[],[f536,f8670])).
fof(f10197,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f10177,f318])).
fof(f10177,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1(X7,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f453,f8670])).
fof(f10196,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6)) )),
  inference(subsumption_resolution,[],[f10176,f318])).
fof(f10176,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X6)) )),
  inference(superposition,[],[f450,f8670])).
fof(f10195,plain,(
  ( ! [X5] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X5)) )),
  inference(subsumption_resolution,[],[f10175,f318])).
fof(f10175,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),X5)) )),
  inference(superposition,[],[f435,f8670])).
fof(f10194,plain,(
  ( ! [X4] : (sP0(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f10174,f318])).
fof(f10174,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(X4,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f434,f8670])).
fof(f10158,plain,(
  ~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f545,f8670,f506])).
fof(f10159,plain,(
  sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'T_0')),
  inference(unit_resulting_resolution,[],[f318,f8670,f504])).
fof(f10160,plain,(
  ~sP5('F_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f545,f8670,f506])).
fof(f10161,plain,(
  sP5('T_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(unit_resulting_resolution,[],[f318,f8670,f504])).
fof(f10086,plain,(
  ~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f3900,f506])).
fof(f10087,plain,(
  sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f318,f3900,f504])).
fof(f10089,plain,(
  ~sP5('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f545,f3900,f506])).
fof(f10090,plain,(
  sP5('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f318,f3900,f504])).
fof(f9881,plain,(
  ( ! [X32] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))) | s('type.num.num',X32) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f429,f2328])).
fof(f9751,plain,(
  ( ! [X6,X4,X5] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))))),X6) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1387,f422])).
fof(f9776,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))))))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))),X45)) )),
  inference(forward_demodulation,[],[f9738,f346])).
fof(f9738,plain,(
  ( ! [X45,X43,X44] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))),X45) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))))))))))) )),
  inference(superposition,[],[f1387,f422])).
fof(f9772,plain,(
  ( ! [X41] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X41)) )),
  inference(subsumption_resolution,[],[f9771,f318])).
fof(f9771,plain,(
  ( ! [X41] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X41)) )),
  inference(forward_demodulation,[],[f9770,f2761])).
fof(f9770,plain,(
  ( ! [X41] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X41)) )),
  inference(forward_demodulation,[],[f9736,f346])).
fof(f9736,plain,(
  ( ! [X41] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),X41) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(superposition,[],[f1387,f1935])).
fof(f9769,plain,(
  ( ! [X39,X40] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))),X40)) )),
  inference(forward_demodulation,[],[f9768,f346])).
fof(f9768,plain,(
  ( ! [X39,X40] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))),X40) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(forward_demodulation,[],[f9735,f353])).
fof(f9735,plain,(
  ( ! [X39,X40] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))),X40) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(superposition,[],[f1387,f353])).
fof(f9759,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20)))))))),X21)) )),
  inference(forward_demodulation,[],[f9758,f551])).
fof(f9758,plain,(
  ( ! [X21,X19,X20] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20)))))))),X21) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(forward_demodulation,[],[f9727,f4750])).
fof(f9727,plain,(
  ( ! [X21,X19,X20] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))))),X21) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(superposition,[],[f1387,f551])).
fof(f9757,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17)))))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))),X18)) )),
  inference(forward_demodulation,[],[f9756,f550])).
fof(f9756,plain,(
  ( ! [X17,X18] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))),X18) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(forward_demodulation,[],[f9726,f353])).
fof(f9726,plain,(
  ( ! [X17,X18] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17)))))),X18) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(superposition,[],[f1387,f550])).
fof(f9721,plain,(
  ( ! [X4,X2,X3] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))),X4) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f1387,f406])).
fof(f9651,plain,(
  ( ! [X6,X4,X5] : (sP0(X6,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1386,f422])).
fof(f9676,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))))))) | sP0(X45,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))))) )),
  inference(forward_demodulation,[],[f9638,f346])).
fof(f9638,plain,(
  ( ! [X45,X43,X44] : (sP0(X45,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))))))))))) )),
  inference(superposition,[],[f1386,f422])).
fof(f9672,plain,(
  ( ! [X41] : (sP0(X41,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(subsumption_resolution,[],[f9671,f318])).
fof(f9671,plain,(
  ( ! [X41] : (~p(s(bool,'T_0')) | sP0(X41,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f9670,f2761])).
fof(f9670,plain,(
  ( ! [X41] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) | sP0(X41,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(forward_demodulation,[],[f9636,f346])).
fof(f9636,plain,(
  ( ! [X41] : (sP0(X41,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) )),
  inference(superposition,[],[f1386,f1935])).
fof(f9669,plain,(
  ( ! [X39,X40] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))) | sP0(X40,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))))) )),
  inference(forward_demodulation,[],[f9668,f346])).
fof(f9668,plain,(
  ( ! [X39,X40] : (sP0(X40,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(forward_demodulation,[],[f9635,f353])).
fof(f9635,plain,(
  ( ! [X39,X40] : (sP0(X40,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(superposition,[],[f1386,f353])).
fof(f9659,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))))) | sP0(X21,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20)))))))))) )),
  inference(forward_demodulation,[],[f9658,f551])).
fof(f9658,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(forward_demodulation,[],[f9627,f4750])).
fof(f9627,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(superposition,[],[f1386,f551])).
fof(f9657,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17)))))) | sP0(X18,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))))) )),
  inference(forward_demodulation,[],[f9656,f550])).
fof(f9656,plain,(
  ( ! [X17,X18] : (sP0(X18,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(forward_demodulation,[],[f9626,f353])).
fof(f9626,plain,(
  ( ! [X17,X18] : (sP0(X18,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(superposition,[],[f1386,f550])).
fof(f9621,plain,(
  ( ! [X4,X2,X3] : (sP0(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f1386,f406])).
fof(f9517,plain,(
  ( ! [X6,X4,X5] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))))),X6) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1309,f422])).
fof(f9541,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))))))) | sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))),X45)) )),
  inference(forward_demodulation,[],[f9504,f345])).
fof(f9504,plain,(
  ( ! [X45,X43,X44] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))),X45) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))))))))))) )),
  inference(superposition,[],[f1309,f422])).
fof(f9535,plain,(
  ( ! [X39,X40] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))) | sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))),X40)) )),
  inference(forward_demodulation,[],[f9534,f345])).
fof(f9534,plain,(
  ( ! [X39,X40] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))),X40) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(forward_demodulation,[],[f9501,f353])).
fof(f9501,plain,(
  ( ! [X39,X40] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))),X40) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(superposition,[],[f1309,f353])).
fof(f9525,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))))) | sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20)))))))),X21)) )),
  inference(forward_demodulation,[],[f9524,f551])).
fof(f9524,plain,(
  ( ! [X21,X19,X20] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20)))))))),X21) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(forward_demodulation,[],[f9493,f4750])).
fof(f9493,plain,(
  ( ! [X21,X19,X20] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))))),X21) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(superposition,[],[f1309,f551])).
fof(f9523,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17)))))) | sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))),X18)) )),
  inference(forward_demodulation,[],[f9522,f550])).
fof(f9522,plain,(
  ( ! [X17,X18] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))),X18) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(forward_demodulation,[],[f9492,f353])).
fof(f9492,plain,(
  ( ! [X17,X18] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17)))))),X18) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(superposition,[],[f1309,f550])).
fof(f9487,plain,(
  ( ! [X4,X2,X3] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))),X4) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f1309,f406])).
fof(f9389,plain,(
  ( ! [X6,X4,X5] : (sP0(X6,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) )),
  inference(superposition,[],[f1308,f422])).
fof(f9413,plain,(
  ( ! [X45,X43,X44] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44)))))))) | sP0(X45,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44)))))))))) )),
  inference(forward_demodulation,[],[f9376,f345])).
fof(f9376,plain,(
  ( ! [X45,X43,X44] : (sP0(X45,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X44))))))))))) )),
  inference(superposition,[],[f1308,f422])).
fof(f9407,plain,(
  ( ! [X39,X40] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39)))))) | sP0(X40,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))))))))) )),
  inference(forward_demodulation,[],[f9406,f345])).
fof(f9406,plain,(
  ( ! [X39,X40] : (sP0(X40,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(forward_demodulation,[],[f9373,f353])).
fof(f9373,plain,(
  ( ! [X39,X40] : (sP0(X40,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X39))))))))) )),
  inference(superposition,[],[f1308,f353])).
fof(f9397,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20)))))))) | sP0(X21,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20)))))))))) )),
  inference(forward_demodulation,[],[f9396,f551])).
fof(f9396,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num',X20))))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(forward_demodulation,[],[f9365,f4750])).
fof(f9365,plain,(
  ( ! [X21,X19,X20] : (sP0(X21,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X20))))))) )),
  inference(superposition,[],[f1308,f551])).
fof(f9395,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17)))))) | sP0(X18,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))))))) )),
  inference(forward_demodulation,[],[f9394,f550])).
fof(f9394,plain,(
  ( ! [X17,X18] : (sP0(X18,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(forward_demodulation,[],[f9364,f353])).
fof(f9364,plain,(
  ( ! [X17,X18] : (sP0(X18,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X17))))))) )),
  inference(superposition,[],[f1308,f550])).
fof(f9359,plain,(
  ( ! [X4,X2,X3] : (sP0(X4,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f1308,f406])).
fof(f9265,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f318,f919])).
fof(f9168,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f918])).
fof(f9071,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f318,f888])).
fof(f9034,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f887])).
fof(f9000,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f318,f791])).
fof(f8961,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f790])).
fof(f8955,plain,(
  sP3('const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))),'F_0')),
  inference(unit_resulting_resolution,[],[f1219,f485])).
fof(f8956,plain,(
  sP2('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))),'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f1219,f484])).
fof(f8929,plain,(
  sP3('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))),'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f1217,f485])).
fof(f8930,plain,(
  sP2('const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))),'F_0')),
  inference(unit_resulting_resolution,[],[f1217,f484])).
fof(f8747,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f3903,f506])).
fof(f8748,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f3903,f504])).
fof(f8750,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f3903,f506])).
fof(f8751,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f3903,f504])).
fof(f8671,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(backward_demodulation,[],[f8670,f3835])).
fof(f8670,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(backward_demodulation,[],[f8391,f3814])).
fof(f8391,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3247,f427])).
fof(f8392,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f3247,f427])).
fof(f8393,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f434])).
fof(f8394,plain,(
  ( ! [X0,X1] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f3247,f435])).
fof(f8395,plain,(
  ( ! [X0] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f3247,f449])).
fof(f8397,plain,(
  ( ! [X0] : (~sP1('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3247,f449])).
fof(f8398,plain,(
  ( ! [X0,X1] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f3247,f449])).
fof(f8399,plain,(
  ( ! [X0] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f3247,f450])).
fof(f8400,plain,(
  ( ! [X0] : (sP1('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f3247,f453])).
fof(f8403,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3247,f487])).
fof(f8404,plain,(
  ( ! [X0,X1] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3247,f545,f3247,f487])).
fof(f8406,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f569,f557,f563,f3247,f488])).
fof(f8408,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f569,f558,f563,f3247,f488])).
fof(f8410,plain,(
  ( ! [X0] : (sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f570,f557,f564,f3247,f488])).
fof(f8412,plain,(
  ( ! [X0] : (sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f570,f558,f564,f3247,f488])).
fof(f8413,plain,(
  ( ! [X0] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3247,f491])).
fof(f8420,plain,(
  ( ! [X0] : (~sP4('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3247,f491])).
fof(f8421,plain,(
  ( ! [X0,X1] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3247,f545,f3247,f491])).
fof(f8427,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f492])).
fof(f8428,plain,(
  ( ! [X0] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3247,f495])).
fof(f8435,plain,(
  ( ! [X0] : (~sP3('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3247,f495])).
fof(f8436,plain,(
  ( ! [X0,X1] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3247,f545,f3247,f495])).
fof(f8442,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f496])).
fof(f8443,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f500])).
fof(f8444,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f3247,f501])).
fof(f8445,plain,(
  ( ! [X2,X0,X1] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f3247,f502])).
fof(f8447,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f334,f3247,f504])).
fof(f8448,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f343,f3247,f504])).
fof(f8449,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f340,f3247,f504])).
fof(f8450,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f591,f3247,f504])).
fof(f8451,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2761,f3247,f504])).
fof(f8452,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f330,f3247,f504])).
fof(f8454,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f347,f3247,f504])).
fof(f8455,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f348,f3247,f504])).
fof(f8457,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f331,f3247,f504])).
fof(f8458,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f608,f3247,f504])).
fof(f8459,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f356,f3247,f504])).
fof(f8460,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f357,f3247,f504])).
fof(f8461,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f407,f3247,f504])).
fof(f8462,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f409,f3247,f504])).
fof(f8463,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f411,f3247,f504])).
fof(f8464,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f359,f3247,f504])).
fof(f8465,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X3)))),'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(unit_resulting_resolution,[],[f470,f3247,f504])).
fof(f8466,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f408,f3247,f504])).
fof(f8467,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f413,f3247,f504])).
fof(f8468,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f414,f3247,f504])).
fof(f8469,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f415,f3247,f504])).
fof(f8470,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f412,f3247,f504])).
fof(f8471,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.ODD_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f345,f3247,f504])).
fof(f8472,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.EVEN_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f346,f3247,f504])).
fof(f8474,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f334,f3247,f504])).
fof(f8475,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f343,f3247,f504])).
fof(f8476,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f340,f3247,f504])).
fof(f8477,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f591,f3247,f504])).
fof(f8478,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f2761,f3247,f504])).
fof(f8479,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f330,f3247,f504])).
fof(f8481,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f347,f3247,f504])).
fof(f8482,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f348,f3247,f504])).
fof(f8484,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f331,f3247,f504])).
fof(f8485,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f608,f3247,f504])).
fof(f8486,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f356,f3247,f504])).
fof(f8487,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f357,f3247,f504])).
fof(f8488,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f407,f3247,f504])).
fof(f8489,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f409,f3247,f504])).
fof(f8490,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f411,f3247,f504])).
fof(f8491,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f359,f3247,f504])).
fof(f8492,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f470,f3247,f504])).
fof(f8493,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f408,f3247,f504])).
fof(f8494,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f413,f3247,f504])).
fof(f8495,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f414,f3247,f504])).
fof(f8496,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f415,f3247,f504])).
fof(f8497,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f412,f3247,f504])).
fof(f8498,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X1)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f345,f3247,f504])).
fof(f8499,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X1)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f346,f3247,f504])).
fof(f8500,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f587,f3247,f505])).
fof(f8503,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f588,f3247,f505])).
fof(f8504,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f587,f3247,f505])).
fof(f8507,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f588,f3247,f505])).
fof(f8508,plain,(
  ( ! [X0,X1] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f3247,f508])).
fof(f8509,plain,(
  ( ! [X0,X1] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f545,f3247,f509])).
fof(f8512,plain,(
  ( ! [X0] : (sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f576,f3247,f510])).
fof(f8518,plain,(
  ( ! [X0] : (sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f576,f3247,f510])).
fof(f8519,plain,(
  ( ! [X0,X1] : (sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f576,f3247,f510])).
fof(f8522,plain,(
  ( ! [X0] : (~sP6('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3247,f514])).
fof(f8529,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3247,f514])).
fof(f8530,plain,(
  ( ! [X0,X1] : (~sP6('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3247,f545,f3247,f514])).
fof(f8536,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f515])).
fof(f8537,plain,(
  ( ! [X0] : (sP7('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f318,f3247,f520])).
fof(f8543,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f318,f3247,f520])).
fof(f8545,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3247,f3247,f520])).
fof(f8549,plain,(
  ( ! [X0] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f318,f3247,f520])).
fof(f8550,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f3247,f318,f3247,f520])).
fof(f8551,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3247,f3247,f520])).
fof(f8552,plain,(
  ( ! [X2,X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f3247,f3247,f520])).
fof(f8553,plain,(
  ( ! [X0] : (~sP7('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f545,f318,f3247,f523])).
fof(f8558,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f318,f3247,f523])).
fof(f8559,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f3247,f3247,f523])).
fof(f8563,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f3247,f524])).
fof(f8568,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f545,f3247,f525])).
fof(f8577,plain,(
  ( ! [X0] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f545,f3247,f531])).
fof(f8582,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f532])).
fof(f8583,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f3247,f533])).
fof(f8584,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3247,f536])).
fof(f8589,plain,(
  ( ! [X2,X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f3247,f536])).
fof(f8590,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f318,f3247,f537])).
fof(f8595,plain,(
  ( ! [X2,X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f3247,f3247,f537])).
fof(f8596,plain,(
  ( ! [X0] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f545,f3247,f538])).
fof(f8597,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f581,f3247,f541])).
fof(f8598,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f582,f3247,f541])).
fof(f8599,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f583,f3247,f541])).
fof(f8600,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f584,f3247,f541])).
fof(f8601,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3247,f674])).
fof(f8602,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f3247,f684])).
fof(f7999,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3246,f427])).
fof(f8000,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f3246,f427])).
fof(f8001,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f434])).
fof(f8002,plain,(
  ( ! [X0,X1] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f3246,f435])).
fof(f8003,plain,(
  ( ! [X0] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f3246,f449])).
fof(f8005,plain,(
  ( ! [X0] : (~sP1('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3246,f449])).
fof(f8006,plain,(
  ( ! [X0,X1] : (~sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f3246,f449])).
fof(f8007,plain,(
  ( ! [X0] : (sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f3246,f450])).
fof(f8008,plain,(
  ( ! [X0] : (sP1('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f3246,f453])).
fof(f8011,plain,(
  ( ! [X0] : (~sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3246,f487])).
fof(f8012,plain,(
  ( ! [X0,X1] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3246,f545,f3246,f487])).
fof(f8014,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f569,f557,f563,f3246,f488])).
fof(f8016,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f569,f558,f563,f3246,f488])).
fof(f8018,plain,(
  ( ! [X0] : (sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f570,f557,f564,f3246,f488])).
fof(f8020,plain,(
  ( ! [X0] : (sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f570,f558,f564,f3246,f488])).
fof(f8021,plain,(
  ( ! [X0] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3246,f491])).
fof(f8028,plain,(
  ( ! [X0] : (~sP4('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3246,f491])).
fof(f8029,plain,(
  ( ! [X0,X1] : (~sP4('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3246,f545,f3246,f491])).
fof(f8035,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f492])).
fof(f8036,plain,(
  ( ! [X0] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3246,f495])).
fof(f8043,plain,(
  ( ! [X0] : (~sP3('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3246,f495])).
fof(f8044,plain,(
  ( ! [X0,X1] : (~sP3('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3246,f545,f3246,f495])).
fof(f8050,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f496])).
fof(f8051,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f500])).
fof(f8052,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f3246,f501])).
fof(f8053,plain,(
  ( ! [X2,X0,X1] : (sP2('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f3246,f502])).
fof(f8055,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f334,f3246,f504])).
fof(f8056,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f343,f3246,f504])).
fof(f8057,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f340,f3246,f504])).
fof(f8058,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f591,f3246,f504])).
fof(f8059,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f2761,f3246,f504])).
fof(f8060,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f330,f3246,f504])).
fof(f8062,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f347,f3246,f504])).
fof(f8063,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f348,f3246,f504])).
fof(f8065,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f331,f3246,f504])).
fof(f8066,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f608,f3246,f504])).
fof(f8067,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f356,f3246,f504])).
fof(f8068,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f357,f3246,f504])).
fof(f8069,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f407,f3246,f504])).
fof(f8070,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f409,f3246,f504])).
fof(f8071,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f411,f3246,f504])).
fof(f8072,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f359,f3246,f504])).
fof(f8073,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X3)))),'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(unit_resulting_resolution,[],[f470,f3246,f504])).
fof(f8074,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f408,f3246,f504])).
fof(f8075,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f413,f3246,f504])).
fof(f8076,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f414,f3246,f504])).
fof(f8077,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f415,f3246,f504])).
fof(f8078,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f412,f3246,f504])).
fof(f8079,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.ODD_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f345,f3246,f504])).
fof(f8080,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.EVEN_1'(s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f346,f3246,f504])).
fof(f8082,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f334,f3246,f504])).
fof(f8083,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f343,f3246,f504])).
fof(f8084,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.>=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f340,f3246,f504])).
fof(f8085,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f591,f3246,f504])).
fof(f8086,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f2761,f3246,f504])).
fof(f8087,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f330,f3246,f504])).
fof(f8089,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f347,f3246,f504])).
fof(f8090,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f348,f3246,f504])).
fof(f8092,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f331,f3246,f504])).
fof(f8093,plain,(
  ( ! [X0] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f608,f3246,f504])).
fof(f8094,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f356,f3246,f504])).
fof(f8095,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f357,f3246,f504])).
fof(f8096,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f407,f3246,f504])).
fof(f8097,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f409,f3246,f504])).
fof(f8098,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f411,f3246,f504])).
fof(f8099,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f359,f3246,f504])).
fof(f8100,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f470,f3246,f504])).
fof(f8101,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f408,f3246,f504])).
fof(f8102,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f413,f3246,f504])).
fof(f8103,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f414,f3246,f504])).
fof(f8104,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f415,f3246,f504])).
fof(f8105,plain,(
  ( ! [X2,X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f412,f3246,f504])).
fof(f8106,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X1)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f345,f3246,f504])).
fof(f8107,plain,(
  ( ! [X0,X1] : (sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X1)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f346,f3246,f504])).
fof(f8108,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f587,f3246,f505])).
fof(f8111,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f588,f3246,f505])).
fof(f8112,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f587,f3246,f505])).
fof(f8115,plain,(
  ( ! [X0] : (~sP5('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f588,f3246,f505])).
fof(f8116,plain,(
  ( ! [X0,X1] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f3246,f508])).
fof(f8117,plain,(
  ( ! [X0,X1] : (~sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f545,f3246,f509])).
fof(f8120,plain,(
  ( ! [X0] : (sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f576,f3246,f510])).
fof(f8126,plain,(
  ( ! [X0] : (sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f576,f3246,f510])).
fof(f8127,plain,(
  ( ! [X0,X1] : (sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f576,f3246,f510])).
fof(f8130,plain,(
  ( ! [X0] : (~sP6('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3246,f514])).
fof(f8137,plain,(
  ( ! [X0] : (~sP6('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f3246,f514])).
fof(f8138,plain,(
  ( ! [X0,X1] : (~sP6('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'F_0')) )),
  inference(unit_resulting_resolution,[],[f3246,f545,f3246,f514])).
fof(f8144,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f515])).
fof(f8145,plain,(
  ( ! [X0] : (sP7('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f318,f3246,f520])).
fof(f8151,plain,(
  ( ! [X0] : (sP7('T_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f318,f3246,f520])).
fof(f8153,plain,(
  ( ! [X0,X1] : (sP7('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3246,f3246,f520])).
fof(f8157,plain,(
  ( ! [X0] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f318,f3246,f520])).
fof(f8158,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f3246,f318,f3246,f520])).
fof(f8159,plain,(
  ( ! [X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3246,f3246,f520])).
fof(f8160,plain,(
  ( ! [X2,X0,X1] : (sP7('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f3246,f3246,f520])).
fof(f8161,plain,(
  ( ! [X0] : (~sP7('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0')) )),
  inference(unit_resulting_resolution,[],[f545,f318,f3246,f523])).
fof(f8166,plain,(
  ( ! [X0] : (~sP7('F_0','T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f318,f3246,f523])).
fof(f8167,plain,(
  ( ! [X0,X1] : (~sP7('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f3246,f3246,f523])).
fof(f8171,plain,(
  ( ! [X0,X1] : (~sP9('F_0',X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f3246,f524])).
fof(f8176,plain,(
  ( ! [X0,X1] : (~sP9('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f545,f3246,f525])).
fof(f8185,plain,(
  ( ! [X0] : (~sP8('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f545,f3246,f531])).
fof(f8190,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f532])).
fof(f8191,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f3246,f533])).
fof(f8192,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f3246,f536])).
fof(f8197,plain,(
  ( ! [X2,X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f3246,f536])).
fof(f8198,plain,(
  ( ! [X0,X1] : (sP9('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f318,f3246,f537])).
fof(f8203,plain,(
  ( ! [X2,X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),X2)) )),
  inference(unit_resulting_resolution,[],[f3246,f3246,f537])).
fof(f8204,plain,(
  ( ! [X0] : (~sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f545,f3246,f538])).
fof(f8205,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f581,f3246,f541])).
fof(f8206,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f582,f3246,f541])).
fof(f8207,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f583,f3246,f541])).
fof(f8208,plain,(
  ( ! [X0,X1] : (sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f584,f3246,f541])).
fof(f8209,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f3246,f674])).
fof(f8210,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))),X1)) )),
  inference(unit_resulting_resolution,[],[f3246,f684])).
fof(f7858,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1452,f486])).
fof(f7813,plain,(
  ( ! [X0] : (~sP2('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f567,f561,f1451,f489])).
fof(f7712,plain,(
  ( ! [X0] : (~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f565,f559,f1448,f490])).
fof(f7624,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1410,f486])).
fof(f7587,plain,(
  ( ! [X0] : (~sP2('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f567,f561,f1409,f489])).
fof(f7502,plain,(
  ( ! [X0] : (~sP2('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f565,f559,f1406,f490])).
fof(f7490,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1366,f422])).
fof(f7489,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f1366,f322])).
fof(f7488,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f1366,f1935])).
fof(f7487,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f1366,f353])).
fof(f7456,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1366,f486])).
fof(f7443,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1365,f422])).
fof(f7442,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f1365,f322])).
fof(f7441,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f1365,f1935])).
fof(f7440,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f1365,f353])).
fof(f7409,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f572,f558,f1365,f490])).
fof(f7398,plain,(
  ~sP3('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f571,f561,f3128,f490])).
fof(f7369,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))),'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1364,f422])).
fof(f7368,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1364,f322])).
fof(f7367,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f1364,f1935])).
fof(f7366,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1364,f353])).
fof(f7363,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1364,f486])).
fof(f7358,plain,(
  sP4('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f3127,f486])).
fof(f7353,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f318,f573,f561,f3126,f488])).
fof(f7323,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))),'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1363,f422])).
fof(f7322,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1363,f322])).
fof(f7321,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f1363,f1935])).
fof(f7320,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1363,f353])).
fof(f7317,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f574,f561,f1363,f490])).
fof(f7308,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1288,f422])).
fof(f7307,plain,(
  sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f1288,f322])).
fof(f7306,plain,(
  sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f1288,f1935])).
fof(f7305,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f1288,f353])).
fof(f7274,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1288,f486])).
fof(f7261,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1287,f422])).
fof(f7260,plain,(
  ~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f1287,f322])).
fof(f7259,plain,(
  ~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f1287,f1935])).
fof(f7258,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f1287,f353])).
fof(f7227,plain,(
  ~sP2('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f568,f561,f1287,f490])).
fof(f7189,plain,(
  ( ! [X2,X1] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))),'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1286,f422])).
fof(f7188,plain,(
  sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1286,f322])).
fof(f7187,plain,(
  sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f1286,f1935])).
fof(f7186,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1286,f353])).
fof(f7183,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f1286,f486])).
fof(f7140,plain,(
  ( ! [X2,X1] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))),'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f1285,f422])).
fof(f7139,plain,(
  ~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1285,f322])).
fof(f7138,plain,(
  ~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f1285,f1935])).
fof(f7137,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))),'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f1285,f353])).
fof(f7134,plain,(
  ~sP2('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f566,f560,f1285,f490])).
fof(f7093,plain,(
  ( ! [X0] : (sP1('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f917])).
fof(f7055,plain,(
  ( ! [X0] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f916])).
fof(f7018,plain,(
  ( ! [X0] : (sP1('F_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f886])).
fof(f6932,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK13(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK13(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f468,f360])).
fof(f6931,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,X2,bool))) )),
  inference(superposition,[],[f468,f360])).
fof(f6864,plain,(
  ( ! [X68,X66,X64,X67,X65] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X65),s('type.num.num',X66)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num',X65))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num',X66)))),X67,X68) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num',X65))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X64),s('type.num.num',X66)))),X67,X68)) )),
  inference(superposition,[],[f541,f470])).
fof(f6863,plain,(
  ( ! [X61,X59,X62,X60,X63] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X60),s('type.num.num',X61)))) | ~p(s(bool,X62)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num',X60))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X59),s('type.num.num',X61)))),X62,X63)) )),
  inference(superposition,[],[f537,f470])).
fof(f6862,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X55),s('type.num.num',X56)))) | ~p(s(bool,X57)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X54),s('type.num.num',X55))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X54),s('type.num.num',X56)))),X58,X57)) )),
  inference(superposition,[],[f536,f470])).
fof(f6861,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X51),s('type.num.num',X52)))) | p(s(bool,X53)) | sP1(X53,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X51))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X52)))))) )),
  inference(superposition,[],[f453,f470])).
fof(f6860,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X47),s('type.num.num',X48)))) | p(s(bool,X49)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X48)))),X49)) )),
  inference(superposition,[],[f450,f470])).
fof(f6859,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X43),s('type.num.num',X44)))) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X44)))),X45)) )),
  inference(superposition,[],[f435,f470])).
fof(f6858,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X39),s('type.num.num',X40)))) | sP0(X41,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X39))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X40)))))) )),
  inference(superposition,[],[f434,f470])).
fof(f6848,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X6))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))) )),
  inference(superposition,[],[f458,f470])).
fof(f6847,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2))),s('type.num.num',X3)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X3))))) )),
  inference(superposition,[],[f471,f470])).
fof(f6846,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X5))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X5)))))) )),
  inference(superposition,[],[f470,f360])).
fof(f6845,plain,(
  ( ! [X24,X23,X22] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num',X23))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23))),s('type.num.num',X24))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num',X23)))))) )),
  inference(superposition,[],[f470,f418])).
fof(f6840,plain,(
  ( ! [X10,X8,X9] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X10))),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9)))))))) )),
  inference(superposition,[],[f470,f421])).
fof(f6838,plain,(
  ( ! [X4,X2,X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num',X3))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) )),
  inference(superposition,[],[f470,f406])).
fof(f6837,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(superposition,[],[f470,f338])).
fof(f6780,plain,(
  ( ! [X24,X23,X22] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num',X24))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num',X23))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23))),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f470,f418])).
fof(f6868,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num',X21))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X21)))))))) )),
  inference(forward_demodulation,[],[f6779,f551])).
fof(f6779,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num',X21))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X19),s('type.num.num','const.num.SUC_1'(s('type.num.num',X20))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X19))),s('type.num.num',X21)))))) )),
  inference(superposition,[],[f470,f551])).
fof(f6867,plain,(
  ( ! [X17,X18] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X17),s('type.num.num',X18))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18)))))) )),
  inference(forward_demodulation,[],[f6778,f550])).
fof(f6778,plain,(
  ( ! [X17,X18] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X17),s('type.num.num',X18))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18)))))) )),
  inference(superposition,[],[f470,f550])).
fof(f6866,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X14)))))))) )),
  inference(forward_demodulation,[],[f6776,f469])).
fof(f6776,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X14))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))),s('type.num.num',X14)))))) )),
  inference(superposition,[],[f470,f469])).
fof(f6775,plain,(
  ( ! [X10,X8,X9] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X10))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f470,f421])).
fof(f6773,plain,(
  ( ! [X4,X2,X3] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X4))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f470,f406])).
fof(f6865,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f6772,f334])).
fof(f6772,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f470,f338])).
fof(f6740,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f545,f470,f506])).
fof(f6741,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X2)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) )),
  inference(unit_resulting_resolution,[],[f318,f470,f504])).
fof(f6742,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f470,f506])).
fof(f6743,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X0))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f470,f504])).
fof(f6719,plain,(
  ( ! [X4,X2,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4)))))) )),
  inference(forward_demodulation,[],[f6593,f422])).
fof(f6593,plain,(
  ( ! [X4,X2,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4)))))))))) )),
  inference(superposition,[],[f422,f469])).
fof(f6591,plain,(
  ( ! [X12,X13,X11] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X13))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X11),s('type.num.num',X12))),s('type.num.num',X13)))))))) )),
  inference(superposition,[],[f469,f421])).
fof(f6589,plain,(
  ( ! [X6,X7,X5] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num',X6)))))) )),
  inference(superposition,[],[f469,f406])).
fof(f6559,plain,(
  ( ! [X24,X23,X22] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X22))),s('type.num.num',X23))),s('type.num.num',X24))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X22),s('type.num.num',X23))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num',X24)))))) )),
  inference(superposition,[],[f469,f418])).
fof(f6715,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X18)))))) )),
  inference(forward_demodulation,[],[f6557,f551])).
fof(f6557,plain,(
  ( ! [X17,X18] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X17))),s('type.num.num',X18))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X18)))))) )),
  inference(superposition,[],[f469,f550])).
fof(f6554,plain,(
  ( ! [X10,X8,X9] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))),s('type.num.num',X10))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X8),s('type.num.num',X9))),s('type.num.num',X10)))))) )),
  inference(superposition,[],[f469,f421])).
fof(f6552,plain,(
  ( ! [X4,X2,X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))),s('type.num.num',X4)))) )),
  inference(superposition,[],[f469,f406])).
fof(f6475,plain,(
  ( ! [X112,X113] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X112),s('type.num.num',X113))))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X112))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X113)))))) )),
  inference(forward_demodulation,[],[f6378,f4298])).
fof(f6378,plain,(
  ( ! [X112,X113] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X112))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X113))))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X112),s('type.num.num',X113))))))))))) )),
  inference(superposition,[],[f379,f422])).
fof(f6351,plain,(
  ( ! [X57,X58,X56] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X58))) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X56))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X57))))) | s('type.num.num',X58) = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X56),s('type.num.num',X57)))))) )),
  inference(superposition,[],[f429,f422])).
fof(f6350,plain,(
  ( ! [X54,X55,X53] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X55),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X53),s('type.num.num',X54))))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X55))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X53))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X54)))))))) )),
  inference(superposition,[],[f420,f422])).
fof(f6349,plain,(
  ( ! [X52,X50,X51] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X50),s('type.num.num',X51))))),s('type.num.num',X52))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X50))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X51))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X52)))))) )),
  inference(superposition,[],[f420,f422])).
fof(f6348,plain,(
  ( ! [X47,X48,X49] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X49),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X47),s('type.num.num',X48))))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X47))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48)))))))) )),
  inference(superposition,[],[f419,f422])).
fof(f6347,plain,(
  ( ! [X45,X46,X44] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X45))))),s('type.num.num',X46))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X46)))))) )),
  inference(superposition,[],[f419,f422])).
fof(f6346,plain,(
  ( ! [X43,X41,X42] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X43),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X41),s('type.num.num',X42))))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42)))))))) )),
  inference(superposition,[],[f417,f422])).
fof(f6345,plain,(
  ( ! [X39,X38,X40] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X38),s('type.num.num',X39))))),s('type.num.num',X40))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40)))))) )),
  inference(superposition,[],[f417,f422])).
fof(f6344,plain,(
  ( ! [X37,X35,X36] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X37),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X35),s('type.num.num',X36))))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36)))))))) )),
  inference(superposition,[],[f414,f422])).
fof(f6343,plain,(
  ( ! [X33,X34,X32] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X32),s('type.num.num',X33))))),s('type.num.num',X34))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f414,f422])).
fof(f6342,plain,(
  ( ! [X30,X31,X29] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))))),s('type.num.num',X31))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30)))))))) )),
  inference(superposition,[],[f412,f422])).
fof(f6341,plain,(
  ( ! [X28,X26,X27] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X28),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X26),s('type.num.num',X27))))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28)))))) )),
  inference(superposition,[],[f412,f422])).
fof(f6340,plain,(
  ( ! [X24,X23,X25] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X23),s('type.num.num',X24))))),s('type.num.num',X25))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24)))))))) )),
  inference(superposition,[],[f411,f422])).
fof(f6339,plain,(
  ( ! [X21,X22,X20] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X22),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X20),s('type.num.num',X21))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))))) )),
  inference(superposition,[],[f411,f422])).
fof(f6338,plain,(
  ( ! [X19,X17,X18] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X19),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X17),s('type.num.num',X18))))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18)))))))) )),
  inference(superposition,[],[f409,f422])).
fof(f6337,plain,(
  ( ! [X14,X15,X16] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X14),s('type.num.num',X15))))),s('type.num.num',X16))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16)))))) )),
  inference(superposition,[],[f409,f422])).
fof(f6336,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X12),s('type.num.num',X13)))))) )),
  inference(superposition,[],[f374,f422])).
fof(f6335,plain,(
  ( ! [X10,X11] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X11))))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f359,f422])).
fof(f6333,plain,(
  ( ! [X6,X7] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X6),s('type.num.num',X7))))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f356,f422])).
fof(f6332,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X5))))))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f349,f422])).
fof(f6331,plain,(
  ( ! [X2,X3] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))))) )),
  inference(superposition,[],[f346,f422])).
fof(f6330,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f345,f422])).
fof(f6471,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12)))))))) )),
  inference(forward_demodulation,[],[f6329,f5882])).
fof(f6329,plain,(
  ( ! [X12,X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X11))),s('type.num.num',X12))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X11),s('type.num.num',X12))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X12)))))) )),
  inference(superposition,[],[f422,f418])).
fof(f6467,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))) )),
  inference(forward_demodulation,[],[f6327,f353])).
fof(f6327,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8)))))) )),
  inference(superposition,[],[f422,f550])).
fof(f6326,plain,(
  ( ! [X7] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X7))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7)))))) )),
  inference(superposition,[],[f422,f339])).
fof(f6466,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6)))))))))) )),
  inference(forward_demodulation,[],[f6325,f6084])).
fof(f6325,plain,(
  ( ! [X6,X5] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X6))))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X5),s('type.num.num',X6)))))))) )),
  inference(superposition,[],[f422,f421])).
fof(f6322,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.numeral.iZ_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f422,f338])).
fof(f6168,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X22),s('type.num.num',X23)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X23))),s('type.num.num',X24)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X22))),s('type.num.num',X24))))) )),
  inference(superposition,[],[f471,f415])).
fof(f6167,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X20)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X20))),s('type.num.num',X21)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X19))),s('type.num.num',X21))))) )),
  inference(superposition,[],[f471,f413])).
fof(f6166,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num',X17)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num',X18)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num',X18))))) )),
  inference(superposition,[],[f471,f414])).
fof(f6165,plain,(
  ( ! [X14,X15,X13] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X13),s('type.num.num',X14)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X14))),s('type.num.num',X15)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X13))),s('type.num.num',X15))))) )),
  inference(superposition,[],[f471,f408])).
fof(f6258,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.ZERO_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num',X8))))) )),
  inference(subsumption_resolution,[],[f6257,f318])).
fof(f6257,plain,(
  ( ! [X8,X7] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.ZERO_0')))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num',X8))))) )),
  inference(forward_demodulation,[],[f6162,f334])).
fof(f6162,plain,(
  ( ! [X8,X7] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X7),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X7))),s('type.num.num',X8))))) )),
  inference(superposition,[],[f471,f359])).
fof(f6234,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))),s('type.num.num',X32)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num',X32))))) )),
  inference(forward_demodulation,[],[f6141,f407])).
fof(f6141,plain,(
  ( ! [X30,X31,X29,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num','const.num.SUC_1'(s('type.num.num',X30)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X29),s('type.num.num',X30))))),s('type.num.num',X32)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X31),s('type.num.num',X32))))) )),
  inference(superposition,[],[f471,f416])).
fof(f6134,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num',X6)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6))))) )),
  inference(superposition,[],[f471,f406])).
fof(f6175,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X20)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X20))))) )),
  inference(forward_demodulation,[],[f6174,f407])).
fof(f6174,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num',X20)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X20))))) )),
  inference(forward_demodulation,[],[f6173,f550])).
fof(f6173,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X18),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))),s('type.num.num',X20))))) )),
  inference(forward_demodulation,[],[f6112,f407])).
fof(f6112,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X18))),s('type.num.num',X19)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X19),s('type.num.num',X20)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18))),s('type.num.num',X20))))) )),
  inference(superposition,[],[f471,f550])).
fof(f6108,plain,(
  ( ! [X6,X4,X5,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num',X3))),s('type.num.num',X5)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X5),s('type.num.num',X6)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X4))),s('type.num.num',X6))))) )),
  inference(superposition,[],[f471,f406])).
fof(f6026,plain,(
  ( ! [X146,X147,X148] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X148))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X146),s('type.num.num','const.num.SUC_1'(s('type.num.num',X147))))))) | s('type.num.num',X148) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X146),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X146),s('type.num.num',X147)))))) )),
  inference(superposition,[],[f429,f421])).
fof(f5994,plain,(
  ( ! [X54,X55] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X54),s('type.num.num','const.num.SUC_1'(s('type.num.num',X55))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X54),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X54),s('type.num.num',X55)))))) )),
  inference(superposition,[],[f374,f421])).
fof(f5968,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f416,f421])).
fof(f5967,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X8),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X8))),s('type.num.num',X9)))))))) )),
  inference(superposition,[],[f421,f551])).
fof(f6084,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X5)))))))) )),
  inference(forward_demodulation,[],[f5964,f5768])).
fof(f5964,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num',X5)))))))) )),
  inference(superposition,[],[f421,f420])).
fof(f5961,plain,(
  ( ! [X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))))) )),
  inference(superposition,[],[f421,f548])).
fof(f5882,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(forward_demodulation,[],[f5770,f5740])).
fof(f5770,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f418,f420])).
fof(f5881,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f5769,f333])).
fof(f5769,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f420,f1935])).
fof(f5768,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f420,f353])).
fof(f5880,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(forward_demodulation,[],[f5767,f548])).
fof(f5767,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f420,f322])).
fof(f5879,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f5741,f332])).
fof(f5741,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f420,f1935])).
fof(f5740,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f420,f353])).
fof(f5632,plain,(
  ( ! [X125,X126,X124] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X126))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X124),s('type.num.num',X125))))))) | s('type.num.num',X126) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X124))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X125)))))) )),
  inference(superposition,[],[f429,f419])).
fof(f5690,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f5689,f419])).
fof(f5689,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(forward_demodulation,[],[f5606,f4298])).
fof(f5606,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X48),s('type.num.num',X49)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f378,f419])).
fof(f5604,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X44),s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f374,f419])).
fof(f5687,plain,(
  ( ! [X3] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f5581,f549])).
fof(f5581,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f419,f1935])).
fof(f5580,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f419,f353])).
fof(f5579,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f419,f322])).
fof(f5554,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f419,f1935])).
fof(f5553,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f419,f353])).
fof(f5552,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f419,f322])).
fof(f5428,plain,(
  ( ! [X127,X125,X126] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X127))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X125))),s('type.num.num',X126))))) | s('type.num.num',X127) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X125),s('type.num.num',X126))),s('type.num.num',X126)))) )),
  inference(superposition,[],[f429,f418])).
fof(f5402,plain,(
  ( ! [X52,X51] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X51))),s('type.num.num',X52))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X51),s('type.num.num',X52))),s('type.num.num',X52)))) )),
  inference(superposition,[],[f374,f418])).
fof(f5377,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X1),s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f416,f418])).
fof(f5374,plain,(
  ( ! [X4,X3] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))),s('type.num.num',X4))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X3),s('type.num.num',X4)))))) )),
  inference(superposition,[],[f418,f406])).
fof(f5349,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X3))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f418,f547])).
fof(f5483,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f5322,f406])).
fof(f5322,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f418,f430])).
fof(f5240,plain,(
  ( ! [X114,X112,X113] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X114))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X112),s('type.num.num',X113))))))) | s('type.num.num',X114) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X112))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X113)))))) )),
  inference(superposition,[],[f429,f417])).
fof(f5299,plain,(
  ( ! [X48,X49] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(forward_demodulation,[],[f5298,f417])).
fof(f5298,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49))),s('type.num.num','const.arithmetic.ZERO_0')))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(forward_demodulation,[],[f5218,f4298])).
fof(f5218,plain,(
  ( ! [X48,X49] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X48),s('type.num.num',X49)))))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X48))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X49)))))) )),
  inference(superposition,[],[f378,f417])).
fof(f5216,plain,(
  ( ! [X45,X44] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X44),s('type.num.num',X45))))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))))) )),
  inference(superposition,[],[f374,f417])).
fof(f5296,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f5193,f335])).
fof(f5193,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X3),s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f417,f1935])).
fof(f5192,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f417,f353])).
fof(f5191,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f417,f322])).
fof(f5295,plain,(
  ( ! [X3] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X3)))))) )),
  inference(forward_demodulation,[],[f5168,f336])).
fof(f5168,plain,(
  ( ! [X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X3))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f417,f1935])).
fof(f5167,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f417,f353])).
fof(f5166,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f417,f322])).
fof(f5134,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(subsumption_resolution,[],[f5095,f545])).
fof(f5095,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f432,f360])).
fof(f5119,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))),s('type.num.num',X17)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))))))) )),
  inference(forward_demodulation,[],[f5080,f407])).
fof(f5080,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16)))))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num',X16))))),s('type.num.num',X17))))) )),
  inference(superposition,[],[f432,f416])).
fof(f5075,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f432,f406])).
fof(f5054,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f432,f406])).
fof(f4979,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f431,f406])).
fof(f4943,plain,(
  ( ! [X134,X132,X133] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X134),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X132),s('type.num.num','const.num.SUC_1'(s('type.num.num',X133)))))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X132),s('type.num.num',X133))),s('type.num.num',X134))))) )),
  inference(forward_demodulation,[],[f4942,f551])).
fof(f4942,plain,(
  ( ! [X134,X132,X133] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X132),s('type.num.num',X133))),s('type.num.num',X134)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X134),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X132))),s('type.num.num',X133))))))) )),
  inference(forward_demodulation,[],[f4886,f4655])).
fof(f4886,plain,(
  ( ! [X134,X132,X133] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X132),s('type.num.num','const.num.SUC_1'(s('type.num.num',X133))))),s('type.num.num',X134)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X134),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X132))),s('type.num.num',X133))))))) )),
  inference(superposition,[],[f458,f551])).
fof(f4829,plain,(
  ( ! [X4,X5] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num',X5))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X5),s('type.num.num','const.num.SUC_1'(s('type.num.num',X4)))))) )),
  inference(superposition,[],[f551,f406])).
fof(f4762,plain,(
  ( ! [X130,X131,X129] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num',X130))),s('type.num.num',X131)))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X131),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130))))))))) )),
  inference(forward_demodulation,[],[f4707,f407])).
fof(f4707,plain,(
  ( ! [X130,X131,X129] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X131),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num','const.num.SUC_1'(s('type.num.num',X130)))))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num',X130))))),s('type.num.num',X131))))) )),
  inference(superposition,[],[f459,f416])).
fof(f4655,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))),s('type.num.num',X2)))) )),
  inference(superposition,[],[f407,f416])).
fof(f4651,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X1)))))) )),
  inference(superposition,[],[f416,f406])).
fof(f4750,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f4609,f353])).
fof(f4609,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1)))))))) )),
  inference(unit_resulting_resolution,[],[f416,f430])).
fof(f4592,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X44),s('type.num.num',X45)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X45)))),X46,X47) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X45)))),X46,X47)) )),
  inference(superposition,[],[f541,f415])).
fof(f4591,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num',X41)))) | ~p(s(bool,X42)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X41)))),X42,X43)) )),
  inference(superposition,[],[f537,f415])).
fof(f4590,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num',X37)))) | ~p(s(bool,X38)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X36))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X37)))),X39,X38)) )),
  inference(superposition,[],[f536,f415])).
fof(f4589,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X33),s('type.num.num',X34)))) | p(s(bool,X35)) | sP1(X35,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f453,f415])).
fof(f4588,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31)))) | p(s(bool,X32)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31)))),X32)) )),
  inference(superposition,[],[f450,f415])).
fof(f4587,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',X28)))) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28)))),X29)) )),
  inference(superposition,[],[f435,f415])).
fof(f4586,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num',X25)))) | sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f434,f415])).
fof(f4576,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))))))) )),
  inference(superposition,[],[f458,f415])).
fof(f4575,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f415,f360])).
fof(f4533,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f545,f415,f506])).
fof(f4534,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f318,f415,f504])).
fof(f4535,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f415,f506])).
fof(f4536,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f415,f504])).
fof(f4506,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X44),s('type.num.num',X45)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))),X46,X47) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))),X46,X47)) )),
  inference(superposition,[],[f541,f414])).
fof(f4505,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num',X41)))) | ~p(s(bool,X42)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))),X42,X43)) )),
  inference(superposition,[],[f537,f414])).
fof(f4504,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num',X37)))) | ~p(s(bool,X38)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37)))),X39,X38)) )),
  inference(superposition,[],[f536,f414])).
fof(f4503,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X33),s('type.num.num',X34)))) | p(s(bool,X35)) | sP1(X35,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f453,f414])).
fof(f4502,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31)))) | p(s(bool,X32)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31)))),X32)) )),
  inference(superposition,[],[f450,f414])).
fof(f4501,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',X28)))) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28)))),X29)) )),
  inference(superposition,[],[f435,f414])).
fof(f4500,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num',X25)))) | sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f434,f414])).
fof(f4489,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f414,f360])).
fof(f4487,plain,(
  ( ! [X2,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f414,f353])).
fof(f4486,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f414,f322])).
fof(f4510,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(forward_demodulation,[],[f4465,f407])).
fof(f4465,plain,(
  ( ! [X2,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f414,f353])).
fof(f4509,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(backward_demodulation,[],[f4508,f4318])).
fof(f4441,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f545,f414,f506])).
fof(f4442,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f318,f414,f504])).
fof(f4443,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f414,f506])).
fof(f4444,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f414,f504])).
fof(f4419,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X44),s('type.num.num',X45)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X45)))),X46,X47) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X45)))),X46,X47)) )),
  inference(superposition,[],[f541,f413])).
fof(f4418,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num',X41)))) | ~p(s(bool,X42)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X41)))),X42,X43)) )),
  inference(superposition,[],[f537,f413])).
fof(f4417,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num',X37)))) | ~p(s(bool,X38)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X36))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X37)))),X39,X38)) )),
  inference(superposition,[],[f536,f413])).
fof(f4416,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X33),s('type.num.num',X34)))) | p(s(bool,X35)) | sP1(X35,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f453,f413])).
fof(f4415,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31)))) | p(s(bool,X32)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X31)))),X32)) )),
  inference(superposition,[],[f450,f413])).
fof(f4414,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',X28)))) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X28)))),X29)) )),
  inference(superposition,[],[f435,f413])).
fof(f4413,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num',X25)))) | sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f434,f413])).
fof(f4403,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) )),
  inference(superposition,[],[f458,f413])).
fof(f4402,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f413,f360])).
fof(f4360,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f545,f413,f506])).
fof(f4361,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f318,f413,f504])).
fof(f4362,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f413,f506])).
fof(f4363,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f413,f504])).
fof(f4337,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X43),s('type.num.num',X42)))) | ~sP8('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43)))),X44,X45) | sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43)))),X44,X45)) )),
  inference(superposition,[],[f541,f412])).
fof(f4336,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X39),s('type.num.num',X38)))) | ~p(s(bool,X40)) | sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))),X40,X41)) )),
  inference(superposition,[],[f537,f412])).
fof(f4335,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X35),s('type.num.num',X34)))) | ~p(s(bool,X36)) | sP9('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35)))),X37,X36)) )),
  inference(superposition,[],[f536,f412])).
fof(f4334,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X32),s('type.num.num',X31)))) | p(s(bool,X33)) | sP1(X33,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))) )),
  inference(superposition,[],[f453,f412])).
fof(f4333,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X29),s('type.num.num',X28)))) | p(s(bool,X30)) | sP1('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))),X30)) )),
  inference(superposition,[],[f450,f412])).
fof(f4332,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X26),s('type.num.num',X25)))) | sP0('const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))),X27)) )),
  inference(superposition,[],[f435,f412])).
fof(f4331,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X23),s('type.num.num',X22)))) | sP0(X24,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))) )),
  inference(superposition,[],[f434,f412])).
fof(f4321,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f412,f360])).
fof(f4338,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(forward_demodulation,[],[f4319,f407])).
fof(f4319,plain,(
  ( ! [X2,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f412,f353])).
fof(f4318,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f412,f322])).
fof(f4297,plain,(
  ( ! [X2,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f412,f353])).
fof(f4296,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f412,f322])).
fof(f4273,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f412,f506])).
fof(f4274,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f412,f504])).
fof(f4275,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f412,f506])).
fof(f4276,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f412,f504])).
fof(f4251,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X43),s('type.num.num',X42)))) | ~sP8('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43)))),X44,X45) | sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X42))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X43)))),X44,X45)) )),
  inference(superposition,[],[f541,f411])).
fof(f4250,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X39),s('type.num.num',X38)))) | ~p(s(bool,X40)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X38))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X39)))),X40,X41)) )),
  inference(superposition,[],[f537,f411])).
fof(f4249,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X35),s('type.num.num',X34)))) | ~p(s(bool,X36)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X35)))),X37,X36)) )),
  inference(superposition,[],[f536,f411])).
fof(f4248,plain,(
  ( ! [X33,X31,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X32),s('type.num.num',X31)))) | p(s(bool,X33)) | sP1(X33,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X32)))))) )),
  inference(superposition,[],[f453,f411])).
fof(f4247,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X29),s('type.num.num',X28)))) | p(s(bool,X30)) | sP1('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29)))),X30)) )),
  inference(superposition,[],[f450,f411])).
fof(f4246,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X26),s('type.num.num',X25)))) | sP0('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X26)))),X27)) )),
  inference(superposition,[],[f435,f411])).
fof(f4245,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X23),s('type.num.num',X22)))) | sP0(X24,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23)))))) )),
  inference(superposition,[],[f434,f411])).
fof(f4235,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X3),s('type.num.num',X2))) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f411,f360])).
fof(f4233,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f411,f353])).
fof(f4232,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f411,f322])).
fof(f4211,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f411,f353])).
fof(f4210,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f411,f322])).
fof(f4187,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f411,f506])).
fof(f4188,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f411,f504])).
fof(f4189,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f411,f506])).
fof(f4190,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f411,f504])).
fof(f4166,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X44),s('type.num.num',X45)))) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))),X46,X47) | sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X45)))),X46,X47)) )),
  inference(superposition,[],[f541,f409])).
fof(f4165,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X40),s('type.num.num',X41)))) | ~p(s(bool,X42)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X41)))),X42,X43)) )),
  inference(superposition,[],[f537,f409])).
fof(f4164,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X37)))) | ~p(s(bool,X38)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X36))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X37)))),X39,X38)) )),
  inference(superposition,[],[f536,f409])).
fof(f4163,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X34)))) | p(s(bool,X35)) | sP1(X35,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f453,f409])).
fof(f4162,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X31)))) | p(s(bool,X32)) | sP1('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31)))),X32)) )),
  inference(superposition,[],[f450,f409])).
fof(f4161,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num',X28)))) | sP0('const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28)))),X29)) )),
  inference(superposition,[],[f435,f409])).
fof(f4160,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X25)))) | sP0(X26,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f434,f409])).
fof(f4149,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f409,f360])).
fof(f4147,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f409,f353])).
fof(f4146,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f409,f322])).
fof(f4125,plain,(
  ( ! [X2,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X2))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))) )),
  inference(superposition,[],[f409,f353])).
fof(f4124,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num',X0))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(superposition,[],[f409,f322])).
fof(f4101,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f545,f409,f506])).
fof(f4102,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f318,f409,f504])).
fof(f4103,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f409,f506])).
fof(f4104,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f409,f504])).
fof(f4080,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X44),s('type.num.num',X45)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X45)))),X46,X47) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X44))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X45)))),X46,X47)) )),
  inference(superposition,[],[f541,f408])).
fof(f4079,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X40),s('type.num.num',X41)))) | ~p(s(bool,X42)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X40))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X41)))),X42,X43)) )),
  inference(superposition,[],[f537,f408])).
fof(f4078,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X36),s('type.num.num',X37)))) | ~p(s(bool,X38)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X36))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X37)))),X39,X38)) )),
  inference(superposition,[],[f536,f408])).
fof(f4077,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X33),s('type.num.num',X34)))) | p(s(bool,X35)) | sP1(X35,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X33))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X34)))))) )),
  inference(superposition,[],[f453,f408])).
fof(f4076,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X30),s('type.num.num',X31)))) | p(s(bool,X32)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X30))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X31)))),X32)) )),
  inference(superposition,[],[f450,f408])).
fof(f4075,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X27),s('type.num.num',X28)))) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X27))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X28)))),X29)) )),
  inference(superposition,[],[f435,f408])).
fof(f4074,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X24),s('type.num.num',X25)))) | sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X24))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X25)))))) )),
  inference(superposition,[],[f434,f408])).
fof(f4064,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))))))) )),
  inference(superposition,[],[f458,f408])).
fof(f4063,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X2))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3)))))) )),
  inference(superposition,[],[f408,f360])).
fof(f4021,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f545,f408,f506])).
fof(f4022,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f318,f408,f504])).
fof(f4023,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f545,f408,f506])).
fof(f4024,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(unit_resulting_resolution,[],[f318,f408,f504])).
fof(f3989,plain,(
  ( ! [X47,X45,X46,X44] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X44),s('type.num.num',X45)))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))),s('type.num.num',X45)),X46,X47) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X44))),s('type.num.num',X45)),X46,X47)) )),
  inference(superposition,[],[f541,f407])).
fof(f3988,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X40),s('type.num.num',X41)))) | ~p(s(bool,X42)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X40))),s('type.num.num',X41)),X42,X43)) )),
  inference(superposition,[],[f537,f407])).
fof(f3987,plain,(
  ( ! [X39,X37,X38,X36] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X36),s('type.num.num',X37)))) | ~p(s(bool,X38)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X36))),s('type.num.num',X37)),X39,X38)) )),
  inference(superposition,[],[f536,f407])).
fof(f3986,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X33),s('type.num.num',X34)))) | p(s(bool,X35)) | sP1(X35,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X33))),s('type.num.num',X34)))) )),
  inference(superposition,[],[f453,f407])).
fof(f3985,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X30),s('type.num.num',X31)))) | p(s(bool,X32)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X30))),s('type.num.num',X31)),X32)) )),
  inference(superposition,[],[f450,f407])).
fof(f3984,plain,(
  ( ! [X28,X29,X27] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X27))),s('type.num.num',X28)),X29) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X27),s('type.num.num',X28))))) )),
  inference(superposition,[],[f435,f407])).
fof(f3983,plain,(
  ( ! [X26,X24,X25] : (sP0(X26,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X24))),s('type.num.num',X25))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X24),s('type.num.num',X25))))) )),
  inference(superposition,[],[f434,f407])).
fof(f3972,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X2),s('type.num.num',X3))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))),s('type.num.num',X3)))) )),
  inference(superposition,[],[f407,f360])).
fof(f3930,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f545,f407,f506])).
fof(f3931,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)),'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f318,f407,f504])).
fof(f3932,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f545,f407,f506])).
fof(f3933,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1)),'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1)))) )),
  inference(unit_resulting_resolution,[],[f318,f407,f504])).
fof(f3899,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X28),s('type.num.num','const.arithmetic.ZERO_0')))) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))),s('type.num.num','const.num.|48|_0')),X29,X30) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28))),s('type.num.num','const.num.|48|_0')),X29,X30)) )),
  inference(superposition,[],[f541,f359])).
fof(f3898,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X25),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,X26)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))),s('type.num.num','const.num.|48|_0')),X26,X27)) )),
  inference(superposition,[],[f537,f359])).
fof(f3897,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X22),s('type.num.num','const.arithmetic.ZERO_0')))) | ~p(s(bool,X23)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22))),s('type.num.num','const.num.|48|_0')),X24,X23)) )),
  inference(superposition,[],[f536,f359])).
fof(f3896,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X20),s('type.num.num','const.arithmetic.ZERO_0')))) | p(s(bool,X21)) | sP1(X21,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f453,f359])).
fof(f3895,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X18),s('type.num.num','const.arithmetic.ZERO_0')))) | p(s(bool,X19)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))),s('type.num.num','const.num.|48|_0')),X19)) )),
  inference(superposition,[],[f450,f359])).
fof(f3894,plain,(
  ( ! [X17,X16] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))),s('type.num.num','const.num.|48|_0')),X17) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X16),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(superposition,[],[f435,f359])).
fof(f3893,plain,(
  ( ! [X14,X15] : (sP0(X15,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))),s('type.num.num','const.num.|48|_0'))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X14),s('type.num.num','const.arithmetic.ZERO_0'))))) )),
  inference(superposition,[],[f434,f359])).
fof(f3882,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num','const.arithmetic.ZERO_0'))) | s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f359,f360])).
fof(f3903,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f3902,f330])).
fof(f3902,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3901,f407])).
fof(f3901,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(forward_demodulation,[],[f3880,f407])).
fof(f3880,plain,(
  ( ! [X0] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f359,f353])).
fof(f3900,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f3879,f348])).
fof(f3879,plain,(
  s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f359,f322])).
fof(f3856,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f359,f506])).
fof(f3857,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')),'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f359,f504])).
fof(f3858,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f359,f506])).
fof(f3859,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.ZERO_0')),'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f359,f504])).
fof(f3834,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X28)))) | ~sP8('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28)))),X29,X30) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X28)))),X29,X30)) )),
  inference(superposition,[],[f541,f357])).
fof(f3833,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X25)))) | ~p(s(bool,X26)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25)))),X26,X27)) )),
  inference(superposition,[],[f537,f357])).
fof(f3832,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X22)))) | ~p(s(bool,X23)) | sP9('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X22)))),X24,X23)) )),
  inference(superposition,[],[f536,f357])).
fof(f3831,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X20)))) | p(s(bool,X21)) | sP1(X21,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20)))))) )),
  inference(superposition,[],[f453,f357])).
fof(f3830,plain,(
  ( ! [X19,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X18)))) | p(s(bool,X19)) | sP1('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18)))),X19)) )),
  inference(superposition,[],[f450,f357])).
fof(f3829,plain,(
  ( ! [X17,X16] : (sP0('const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16)))),X17) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X16))))) )),
  inference(superposition,[],[f435,f357])).
fof(f3828,plain,(
  ( ! [X14,X15] : (sP0(X15,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X14))))) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X14))))) )),
  inference(superposition,[],[f434,f357])).
fof(f3817,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(superposition,[],[f357,f360])).
fof(f3815,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f357,f353])).
fof(f3835,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(backward_demodulation,[],[f3814,f3754])).
fof(f3814,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f357,f322])).
fof(f3791,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f357,f506])).
fof(f3792,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f357,f504])).
fof(f3793,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f357,f506])).
fof(f3794,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'const.prim_rec.<_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f357,f504])).
fof(f3773,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X27)))) | ~sP8('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))),s('type.num.num','const.num.|48|_0')),X28,X29) | sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))),s('type.num.num','const.num.|48|_0')),X28,X29)) )),
  inference(superposition,[],[f541,f356])).
fof(f3772,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X24)))) | ~p(s(bool,X25)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))),s('type.num.num','const.num.|48|_0')),X25,X26)) )),
  inference(superposition,[],[f537,f356])).
fof(f3771,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X21)))) | ~p(s(bool,X22)) | sP9('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))),s('type.num.num','const.num.|48|_0')),X23,X22)) )),
  inference(superposition,[],[f536,f356])).
fof(f3770,plain,(
  ( ! [X19,X20] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X19)))) | p(s(bool,X20)) | sP1(X20,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f453,f356])).
fof(f3769,plain,(
  ( ! [X17,X18] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X17)))) | p(s(bool,X18)) | sP1('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))),s('type.num.num','const.num.|48|_0')),X18)) )),
  inference(superposition,[],[f450,f356])).
fof(f3768,plain,(
  ( ! [X15,X16] : (sP0('const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))),s('type.num.num','const.num.|48|_0')),X16) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X15))))) )),
  inference(superposition,[],[f435,f356])).
fof(f3757,plain,(
  ( ! [X1] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X1))) | s(bool,'T_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f356,f360])).
fof(f3774,plain,(
  s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f3756,f331])).
fof(f3756,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.ZERO_0'))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f356,f1935])).
fof(f3755,plain,(
  ( ! [X0] : (s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))))),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f356,f353])).
fof(f3754,plain,(
  s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f356,f322])).
fof(f3731,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f356,f506])).
fof(f3732,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f356,f504])).
fof(f3733,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f356,f506])).
fof(f3734,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f356,f504])).
fof(f3684,plain,(
  ( ! [X0] : (~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f571,f1200,f557,f545,f490])).
fof(f3685,plain,(
  ( ! [X0] : (~sP3('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f571,f1201,f557,f545,f490])).
fof(f3686,plain,(
  ( ! [X0] : (~sP3('T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f571,f1202,f557,f545,f490])).
fof(f3687,plain,(
  ~sP3('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f574,f1199,f557,f545,f490])).
fof(f3700,plain,(
  ( ! [X0] : (~sP2('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1197,f559,f565,f545,f490])).
fof(f3701,plain,(
  ( ! [X0] : (~sP2('F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1198,f559,f565,f545,f490])).
fof(f3704,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f572,f558,f565,f545,f490])).
fof(f3712,plain,(
  sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f572,f558,f566,f545,f490])).
fof(f3723,plain,(
  sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f574,f561,f568,f545,f490])).
fof(f3725,plain,(
  ~sP2('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f1196,f559,f568,f545,f490])).
fof(f3629,plain,(
  sP5('F_0','F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f574,f561,f567,f545,f489])).
fof(f3631,plain,(
  ( ! [X0] : (~sP2('const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1204,f561,f567,f545,f489])).
fof(f3632,plain,(
  ( ! [X0] : (~sP2('const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f1205,f561,f567,f545,f489])).
fof(f3635,plain,(
  sP5('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f574,f562,f567,f545,f489])).
fof(f3644,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f569,f559,f565,f545,f489])).
fof(f3645,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f570,f559,f565,f545,f489])).
fof(f3649,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f570,f560,f566,f545,f489])).
fof(f3651,plain,(
  ~sP2('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1203,f561,f566,f545,f489])).
fof(f3573,plain,(
  ( ! [X0] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f1193,f561,f318,f488])).
fof(f3574,plain,(
  ( ! [X0] : (~sP3('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f1194,f561,f318,f488])).
fof(f3575,plain,(
  ( ! [X0] : (~sP3('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f573,f1195,f561,f318,f488])).
fof(f3576,plain,(
  ~sP3('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f573,f1192,f561,f318,f488])).
fof(f3586,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f573,f557,f563,f318,f488])).
fof(f3588,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f573,f557,f564,f318,f488])).
fof(f3589,plain,(
  sP5('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f573,f558,f564,f318,f488])).
fof(f3592,plain,(
  sP5('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f573,f562,f567,f318,f488])).
fof(f3593,plain,(
  ~sP4('T_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f573,f1178,f568,f318,f488])).
fof(f3595,plain,(
  sP5('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f573,f562,f568,f318,f488])).
fof(f3273,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f459,f406])).
fof(f3254,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))),s('type.num.num',X4)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f459,f406])).
fof(f3247,plain,(
  ( ! [X3] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3))))))) )),
  inference(subsumption_resolution,[],[f3206,f545])).
fof(f3206,plain,(
  ( ! [X3] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X3))))))) )),
  inference(superposition,[],[f458,f348])).
fof(f3246,plain,(
  ( ! [X2] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))))) )),
  inference(subsumption_resolution,[],[f3205,f545])).
fof(f3205,plain,(
  ( ! [X2] : (p(s(bool,'F_0')) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X2))))))) )),
  inference(superposition,[],[f458,f347])).
fof(f3185,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4))))) )),
  inference(superposition,[],[f458,f406])).
fof(f3211,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10)))))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X11))))) )),
  inference(forward_demodulation,[],[f3210,f550])).
fof(f3210,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X10))))))) )),
  inference(forward_demodulation,[],[f3169,f407])).
fof(f3169,plain,(
  ( ! [X10,X11] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))),s('type.num.num',X11)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X11),s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X10))))))) )),
  inference(superposition,[],[f458,f550])).
fof(f3166,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))),s('type.num.num',X4)))) | p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) )),
  inference(superposition,[],[f458,f406])).
fof(f3164,plain,(
  ( ! [X12,X13] : (~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X12,X13) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X12,X13)) )),
  inference(subsumption_resolution,[],[f3147,f318])).
fof(f3147,plain,(
  ( ! [X12,X13] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X12,X13) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X12,X13)) )),
  inference(superposition,[],[f541,f2761])).
fof(f3163,plain,(
  ( ! [X10,X11] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X10,X11) | ~p(s(bool,X10))) )),
  inference(subsumption_resolution,[],[f3146,f318])).
fof(f3146,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | ~p(s(bool,X10)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X10,X11)) )),
  inference(superposition,[],[f537,f2761])).
fof(f3162,plain,(
  ( ! [X8,X9] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X9,X8) | ~p(s(bool,X8))) )),
  inference(subsumption_resolution,[],[f3145,f318])).
fof(f3145,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | ~p(s(bool,X8)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X9,X8)) )),
  inference(superposition,[],[f536,f2761])).
fof(f3161,plain,(
  ( ! [X7] : (sP1(X7,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | p(s(bool,X7))) )),
  inference(subsumption_resolution,[],[f3144,f318])).
fof(f3144,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,X7)) | sP1(X7,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f453,f2761])).
fof(f3160,plain,(
  ( ! [X6] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X6) | p(s(bool,X6))) )),
  inference(subsumption_resolution,[],[f3143,f318])).
fof(f3143,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,X6)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X6)) )),
  inference(superposition,[],[f450,f2761])).
fof(f3159,plain,(
  ( ! [X5] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X5)) )),
  inference(subsumption_resolution,[],[f3142,f318])).
fof(f3142,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),X5)) )),
  inference(superposition,[],[f435,f2761])).
fof(f3158,plain,(
  ( ! [X4] : (sP0(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(subsumption_resolution,[],[f3141,f318])).
fof(f3141,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | sP0(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f434,f2761])).
fof(f3126,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f545,f2761,f506])).
fof(f3127,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f318,f2761,f504])).
fof(f3128,plain,(
  ~sP5('F_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f545,f2761,f506])).
fof(f3129,plain,(
  sP5('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(unit_resulting_resolution,[],[f318,f2761,f504])).
fof(f3032,plain,(
  ( ! [X4,X2,X3] : (s('type.num.num',X4) != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))))) )),
  inference(superposition,[],[f430,f406])).
fof(f3018,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f406,f430])).
fof(f2918,plain,(
  ( ! [X2,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s('type.num.num',X2) = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f429,f353])).
fof(f2915,plain,(
  ( ! [X34] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X34))) | s('type.num.num',X34) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f429,f321])).
fof(f2914,plain,(
  ( ! [X33,X32] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X33))) | s('type.num.num',X33) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X32)))))))) )),
  inference(superposition,[],[f429,f352])).
fof(f2913,plain,(
  ( ! [X30,X31] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X31))) | s('type.num.num',X31) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X30)))))))) )),
  inference(superposition,[],[f429,f351])).
fof(f2946,plain,(
  ( ! [X28,X29] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X29) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f2912,f353])).
fof(f2912,plain,(
  ( ! [X28,X29] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X29))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X29) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X28),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f429,f549])).
fof(f2911,plain,(
  ( ! [X26,X27] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27))) | s('type.num.num',X27) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X26)))) )),
  inference(superposition,[],[f429,f336])).
fof(f2910,plain,(
  ( ! [X24,X25] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X25))) | s('type.num.num',X25) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X24),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f429,f335])).
fof(f2909,plain,(
  ( ! [X23] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X23))) | s('type.num.num',X23) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(superposition,[],[f429,f1935])).
fof(f2944,plain,(
  ( ! [X20] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(forward_demodulation,[],[f2907,f353])).
fof(f2907,plain,(
  ( ! [X20] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X20))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num',X20) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f429,f322])).
fof(f2906,plain,(
  ( ! [X19,X18] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X18))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))) | s('type.num.num',X19) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X18)))) )),
  inference(superposition,[],[f429,f547])).
fof(f2905,plain,(
  ( ! [X17,X16] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X16))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17))) | s('type.num.num',X17) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X16),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f429,f548])).
fof(f2904,plain,(
  ( ! [X14,X15] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15))) | s('type.num.num',X15) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X14),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f429,f333])).
fof(f2903,plain,(
  ( ! [X12,X13] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))) | s('type.num.num',X13) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)))) )),
  inference(superposition,[],[f429,f332])).
fof(f2943,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))) != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X10)))) )),
  inference(forward_demodulation,[],[f2902,f353])).
fof(f2902,plain,(
  ( ! [X10,X11] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X11))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) | s('type.num.num',X11) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X10)))) )),
  inference(superposition,[],[f429,f550])).
fof(f2901,plain,(
  ( ! [X8,X9] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X8))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9))) | s('type.num.num',X9) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)))) )),
  inference(superposition,[],[f429,f339])).
fof(f2899,plain,(
  ( ! [X4,X2,X3] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))) | s('type.num.num',X4) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3)))) )),
  inference(superposition,[],[f429,f406])).
fof(f2898,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num',X1) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f429,f338])).
fof(f2850,plain,(
  ( ! [X21] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X21)))) )),
  inference(forward_demodulation,[],[f2809,f353])).
fof(f2809,plain,(
  ( ! [X21] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X21))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X21)))) )),
  inference(superposition,[],[f374,f550])).
fof(f2786,plain,(
  ( ! [X1] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f550,f406])).
fof(f2762,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2696,f608])).
fof(f2696,plain,(
  s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f345,f1935])).
fof(f2761,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(forward_demodulation,[],[f2695,f591])).
fof(f2695,plain,(
  s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f346,f1935])).
fof(f2760,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f2693,f321])).
fof(f2693,plain,(
  s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f349,f1935])).
fof(f2594,plain,(
  ( ! [X47,X46] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X47),s('type.num.num',X46))) | p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X46),s('type.num.num',X47))))))) )),
  inference(superposition,[],[f379,f406])).
fof(f2593,plain,(
  ( ! [X45,X44] : (~p(s(bool,'const.arithmetic.>=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X45),s('type.num.num',X44)))))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X44),s('type.num.num',X45)))) )),
  inference(superposition,[],[f378,f406])).
fof(f2592,plain,(
  ( ! [X43,X42] : (s('type.num.num','const.arithmetic.ZERO_0') != s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X43),s('type.num.num',X42))) | s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X42),s('type.num.num',X43)))))) )),
  inference(superposition,[],[f375,f406])).
fof(f2591,plain,(
  ( ! [X41,X40] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X41),s('type.num.num',X40))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X40),s('type.num.num',X41)))) )),
  inference(superposition,[],[f374,f406])).
fof(f2438,plain,(
  ( ! [X3] : (s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X3))))) = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X3)))))))) )),
  inference(superposition,[],[f345,f353])).
fof(f2437,plain,(
  ( ! [X2] : (s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X2))))) = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X2)))))))) )),
  inference(superposition,[],[f346,f353])).
fof(f2436,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.num.SUC_1'(s('type.num.num',X1)))) )),
  inference(superposition,[],[f374,f353])).
fof(f2435,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))))) = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f349,f353])).
fof(f2328,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(superposition,[],[f352,f322])).
fof(f2025,plain,(
  ( ! [X0] : (sP4('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1169,f486])).
fof(f2009,plain,(
  ( ! [X0] : (sP4('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1168,f486])).
fof(f1862,plain,(
  ( ! [X6] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X6))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f374,f335])).
fof(f1860,plain,(
  ( ! [X5] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X5))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X5)))) )),
  inference(superposition,[],[f374,f547])).
fof(f1859,plain,(
  ( ! [X4] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X4))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X4),s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f374,f548])).
fof(f1856,plain,(
  ( ! [X1] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X1)))) )),
  inference(superposition,[],[f374,f339])).
fof(f1855,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f374,f338])).
fof(f1823,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f1159,f486])).
fof(f1766,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1158,f486])).
fof(f1751,plain,(
  ( ! [X0] : (sP4('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f1157,f486])).
fof(f1638,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1629,f526])).
fof(f1629,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f318,f545,f545,f531])).
fof(f1618,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f585,f545,f545,f530])).
fof(f1619,plain,(
  sP9('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f586,f545,f545,f530])).
fof(f1569,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1560,f507])).
fof(f1560,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f318,f545,f318,f514])).
fof(f1550,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f575,f318,f318,f510])).
fof(f1551,plain,(
  sP7('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f576,f318,f318,f510])).
fof(f1546,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f1534,f484])).
fof(f1534,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f545,f545,f545,f499])).
fof(f1527,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f1518,f485])).
fof(f1518,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f318,f545,f318,f495])).
fof(f1501,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f318,f545,f318,f491])).
fof(f1448,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f348,f506])).
fof(f1449,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f348,f504])).
fof(f1451,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f348,f506])).
fof(f1452,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f348,f504])).
fof(f1406,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f347,f506])).
fof(f1407,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f347,f504])).
fof(f1409,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f347,f506])).
fof(f1410,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f347,f504])).
fof(f1392,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X27)))) | ~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X28,X29) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X28,X29)) )),
  inference(superposition,[],[f541,f346])).
fof(f1391,plain,(
  ( ! [X26,X24,X25] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24)))),X25,X26) | ~p(s(bool,X25)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X24))))) )),
  inference(superposition,[],[f537,f346])).
fof(f1390,plain,(
  ( ! [X23,X21,X22] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))),X23,X22) | ~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X21))))) )),
  inference(superposition,[],[f536,f346])).
fof(f1389,plain,(
  ( ! [X19,X20] : (sP1(X20,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))) | p(s(bool,X20)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X19))))) )),
  inference(superposition,[],[f453,f346])).
fof(f1388,plain,(
  ( ! [X17,X18] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X17))))) )),
  inference(superposition,[],[f450,f346])).
fof(f1387,plain,(
  ( ! [X15,X16] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15)))),X16) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X15))))) )),
  inference(superposition,[],[f435,f346])).
fof(f1386,plain,(
  ( ! [X14,X13] : (sP0(X14,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))) | ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num',X13))))) )),
  inference(superposition,[],[f434,f346])).
fof(f1363,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f346,f506])).
fof(f1364,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.EVEN_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f346,f504])).
fof(f1365,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f346,f506])).
fof(f1366,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num',X0)),'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f346,f504])).
fof(f1314,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X27)))) | ~sP8('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X28,X29) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X27)))),X28,X29)) )),
  inference(superposition,[],[f541,f345])).
fof(f1313,plain,(
  ( ! [X26,X24,X25] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X24)))),X25,X26) | ~p(s(bool,X25)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X24))))) )),
  inference(superposition,[],[f537,f345])).
fof(f1312,plain,(
  ( ! [X23,X21,X22] : (sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X21)))),X23,X22) | ~p(s(bool,X22)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X21))))) )),
  inference(superposition,[],[f536,f345])).
fof(f1311,plain,(
  ( ! [X19,X20] : (sP1(X20,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X19))))) | p(s(bool,X20)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X19))))) )),
  inference(superposition,[],[f453,f345])).
fof(f1310,plain,(
  ( ! [X17,X18] : (sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X17)))),X18) | p(s(bool,X18)) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X17))))) )),
  inference(superposition,[],[f450,f345])).
fof(f1309,plain,(
  ( ! [X15,X16] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X15)))),X16) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X15))))) )),
  inference(superposition,[],[f435,f345])).
fof(f1308,plain,(
  ( ! [X14,X13] : (sP0(X14,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X13))))) | ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num',X13))))) )),
  inference(superposition,[],[f434,f345])).
fof(f1285,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f345,f506])).
fof(f1286,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))),'const.arithmetic.ODD_1'(s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f345,f504])).
fof(f1287,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f545,f345,f506])).
fof(f1288,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num',X0)),'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0)))))) )),
  inference(unit_resulting_resolution,[],[f318,f345,f504])).
fof(f1216,plain,(
  ~sP5('T_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))),'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)))),
  inference(unit_resulting_resolution,[],[f318,f317,f505])).
fof(f1217,plain,(
  sP5('F_0','const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))),'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)))),
  inference(unit_resulting_resolution,[],[f545,f317,f503])).
fof(f1218,plain,(
  ~sP5('T_0','const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))))),
  inference(unit_resulting_resolution,[],[f318,f317,f505])).
fof(f1219,plain,(
  sP5('F_0','const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12)),'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))))),
  inference(unit_resulting_resolution,[],[f545,f317,f503])).
fof(f1192,plain,(
  ~sP5('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f545,f591,f506])).
fof(f1193,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f545,f334,f506])).
fof(f1194,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f545,f343,f506])).
fof(f1195,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f545,f340,f506])).
fof(f1196,plain,(
  ~sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f608,f506])).
fof(f1197,plain,(
  ( ! [X0] : (~sP5('F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f330,f506])).
fof(f1198,plain,(
  ( ! [X0] : (~sP5('F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f331,f506])).
fof(f1199,plain,(
  ~sP5('F_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f545,f591,f506])).
fof(f1200,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f334,f506])).
fof(f1201,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f343,f506])).
fof(f1202,plain,(
  ( ! [X0] : (~sP5('F_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f340,f506])).
fof(f1203,plain,(
  ~sP5('F_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f545,f608,f506])).
fof(f1204,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f545,f330,f506])).
fof(f1205,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f331,f506])).
fof(f1177,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f318,f587,f505])).
fof(f1178,plain,(
  ~sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f318,f588,f505])).
fof(f1181,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f318,f587,f505])).
fof(f1182,plain,(
  ~sP5('T_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f318,f588,f505])).
fof(f1156,plain,(
  sP5('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f318,f591,f504])).
fof(f1157,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f334,f504])).
fof(f1158,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f343,f504])).
fof(f1159,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f340,f504])).
fof(f1160,plain,(
  sP5('T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f318,f608,f504])).
fof(f1161,plain,(
  ( ! [X0] : (sP5('T_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f330,f504])).
fof(f1162,plain,(
  ( ! [X0] : (sP5('T_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f331,f504])).
fof(f1163,plain,(
  sP5('T_0','T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f318,f591,f504])).
fof(f1164,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f334,f504])).
fof(f1165,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f343,f504])).
fof(f1166,plain,(
  ( ! [X0] : (sP5('T_0','T_0','const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f340,f504])).
fof(f1167,plain,(
  sP5('T_0','F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f318,f608,f504])).
fof(f1168,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f318,f330,f504])).
fof(f1169,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f318,f331,f504])).
fof(f1146,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f545,f587,f503])).
fof(f1147,plain,(
  sP5('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),'T_0')),
  inference(unit_resulting_resolution,[],[f545,f588,f503])).
fof(f1148,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f545,f587,f503])).
fof(f1149,plain,(
  sP5('F_0','T_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f545,f588,f503])).
fof(f1129,plain,(
  ( ! [X0] : (sP1('const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f885])).
fof(f1116,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(subsumption_resolution,[],[f1098,f545])).
fof(f1098,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X1)) | s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f428,f360])).
fof(f1077,plain,(
  ( ! [X0] : (sP1('F_0','const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(unit_resulting_resolution,[],[f545,f789])).
fof(f1035,plain,(
  ( ! [X0] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f788])).
fof(f942,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f360,f360])).
fof(f920,plain,(
  ( ! [X19,X20,X18] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X18)),X19,X20) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X18)),X19,X20)) )),
  inference(subsumption_resolution,[],[f908,f318])).
fof(f908,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X18)),X19,X20) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X18)),X19,X20)) )),
  inference(superposition,[],[f541,f343])).
fof(f919,plain,(
  ( ! [X17,X15,X16] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X15)),X16,X17) | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f907,f318])).
fof(f907,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X15)),X16,X17)) )),
  inference(superposition,[],[f537,f343])).
fof(f918,plain,(
  ( ! [X14,X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X12)),X14,X13) | ~p(s(bool,X13))) )),
  inference(subsumption_resolution,[],[f906,f318])).
fof(f906,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X12)),X14,X13)) )),
  inference(superposition,[],[f536,f343])).
fof(f917,plain,(
  ( ! [X10,X11] : (sP1(X11,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X10))) | p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f905,f318])).
fof(f905,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | sP1(X11,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f453,f343])).
fof(f916,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8)),X9) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f904,f318])).
fof(f904,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X8)),X9)) )),
  inference(superposition,[],[f450,f343])).
fof(f915,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X6)),X7)) )),
  inference(subsumption_resolution,[],[f903,f318])).
fof(f903,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X6)),X7)) )),
  inference(superposition,[],[f435,f343])).
fof(f914,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X4)))) )),
  inference(subsumption_resolution,[],[f902,f318])).
fof(f902,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f434,f343])).
fof(f889,plain,(
  ( ! [X19,X20,X18] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20) | ~sP8('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20)) )),
  inference(subsumption_resolution,[],[f877,f318])).
fof(f877,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20) | sP9('const.arithmetic.>=_2'(s('type.num.num',X18),s('type.num.num','const.num.|48|_0')),X19,X20)) )),
  inference(superposition,[],[f541,f340])).
fof(f888,plain,(
  ( ! [X17,X15,X16] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16,X17) | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f876,f318])).
fof(f876,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | sP9('const.arithmetic.>=_2'(s('type.num.num',X15),s('type.num.num','const.num.|48|_0')),X16,X17)) )),
  inference(superposition,[],[f537,f340])).
fof(f887,plain,(
  ( ! [X14,X12,X13] : (sP9('const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')),X14,X13) | ~p(s(bool,X13))) )),
  inference(subsumption_resolution,[],[f875,f318])).
fof(f875,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | sP9('const.arithmetic.>=_2'(s('type.num.num',X12),s('type.num.num','const.num.|48|_0')),X14,X13)) )),
  inference(superposition,[],[f536,f340])).
fof(f886,plain,(
  ( ! [X10,X11] : (sP1(X11,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0'))) | p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f874,f318])).
fof(f874,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | sP1(X11,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f453,f340])).
fof(f885,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')),X9) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f873,f318])).
fof(f873,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | sP1('const.arithmetic.>=_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')),X9)) )),
  inference(superposition,[],[f450,f340])).
fof(f884,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')),X7)) )),
  inference(subsumption_resolution,[],[f872,f318])).
fof(f872,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.>=_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')),X7)) )),
  inference(superposition,[],[f435,f340])).
fof(f883,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(subsumption_resolution,[],[f871,f318])).
fof(f871,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.>=_2'(s('type.num.num',X4),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f434,f340])).
fof(f792,plain,(
  ( ! [X19,X20,X18] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20)) )),
  inference(subsumption_resolution,[],[f780,f318])).
fof(f780,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X18)),X19,X20)) )),
  inference(superposition,[],[f541,f334])).
fof(f791,plain,(
  ( ! [X17,X15,X16] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X15)),X16,X17) | ~p(s(bool,X16))) )),
  inference(subsumption_resolution,[],[f779,f318])).
fof(f779,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X16)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X15)),X16,X17)) )),
  inference(superposition,[],[f537,f334])).
fof(f790,plain,(
  ( ! [X14,X12,X13] : (sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)),X14,X13) | ~p(s(bool,X13))) )),
  inference(subsumption_resolution,[],[f778,f318])).
fof(f778,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,'T_0')) | ~p(s(bool,X13)) | sP9('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X12)),X14,X13)) )),
  inference(superposition,[],[f536,f334])).
fof(f789,plain,(
  ( ! [X10,X11] : (sP1(X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))) | p(s(bool,X11))) )),
  inference(subsumption_resolution,[],[f777,f318])).
fof(f777,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | p(s(bool,X11)) | sP1(X11,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))) )),
  inference(superposition,[],[f453,f334])).
fof(f788,plain,(
  ( ! [X8,X9] : (sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X9) | p(s(bool,X9))) )),
  inference(subsumption_resolution,[],[f776,f318])).
fof(f776,plain,(
  ( ! [X8,X9] : (~p(s(bool,'T_0')) | p(s(bool,X9)) | sP1('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X8)),X9)) )),
  inference(superposition,[],[f450,f334])).
fof(f787,plain,(
  ( ! [X6,X7] : (sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)),X7)) )),
  inference(subsumption_resolution,[],[f775,f318])).
fof(f775,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | sP0('const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X6)),X7)) )),
  inference(superposition,[],[f435,f334])).
fof(f786,plain,(
  ( ! [X4,X5] : (sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(subsumption_resolution,[],[f774,f318])).
fof(f774,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | sP0(X5,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X4)))) )),
  inference(superposition,[],[f434,f334])).
fof(f736,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f318,f684])).
fof(f710,plain,(
  ( ! [X0] : (sP7('F_0',X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f578,f545,f542])).
fof(f712,plain,(
  ( ! [X0] : (sP7('F_0','const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f580,f545,f542])).
fof(f702,plain,(
  ( ! [X0,X1] : (~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(subsumption_resolution,[],[f700,f318])).
fof(f700,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(superposition,[],[f541,f591])).
fof(f697,plain,(
  ( ! [X0] : (sP9('T_0',X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f582,f318,f541])).
fof(f699,plain,(
  ( ! [X0] : (sP9('T_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f584,f318,f541])).
fof(f688,plain,(
  ( ! [X0] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f674])).
fof(f684,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f682,f318])).
fof(f682,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(superposition,[],[f537,f591])).
fof(f681,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f318,f318,f537])).
fof(f674,plain,(
  ( ! [X0,X1] : (sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1,X0) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f672,f318])).
fof(f672,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1,X0)) )),
  inference(superposition,[],[f536,f591])).
fof(f671,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f318,f536])).
fof(f669,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f318,f545,f525])).
fof(f668,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f524])).
fof(f661,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f545,f545,f519])).
fof(f658,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f545,f545,f518])).
fof(f655,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f318,f545,f509])).
fof(f652,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f318,f545,f508])).
fof(f651,plain,(
  sP1('F_0','const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f545,f637])).
fof(f646,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f545,f545,f552])).
fof(f637,plain,(
  ( ! [X0] : (sP1(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f635,f318])).
fof(f635,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f453,f591])).
fof(f634,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f545,f318,f453])).
fof(f631,plain,(
  sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),'F_0')),
  inference(unit_resulting_resolution,[],[f545,f628])).
fof(f628,plain,(
  ( ! [X0] : (sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f626,f318])).
fof(f626,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(superposition,[],[f450,f591])).
fof(f625,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f545,f318,f450])).
fof(f622,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f318,f318,f449])).
fof(f621,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f545,f545,f448])).
fof(f608,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f328,f371])).
fof(f591,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f320,f367])).
fof(f587,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f545,f366])).
fof(f588,plain,(
  s(bool,'T_0') != s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')))),
  inference(unit_resulting_resolution,[],[f328,f366])).
fof(f585,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f545,f534])).
fof(f586,plain,(
  ( ! [X0,X1] : (sP8('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f328,f534])).
fof(f583,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f318,f533])).
fof(f584,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f320,f533])).
fof(f581,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f532])).
fof(f582,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f320,f532])).
fof(f579,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f545,f517])).
fof(f580,plain,(
  ( ! [X0,X1] : (sP6('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f328,f517])).
fof(f577,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f545,f516])).
fof(f578,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f328,f516])).
fof(f575,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f515])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f320,f515])).
fof(f573,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f318,f502])).
fof(f574,plain,(
  ( ! [X0,X1] : (sP2('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f320,f502])).
fof(f571,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f318,f501])).
fof(f572,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f320,f501])).
fof(f569,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f500])).
fof(f570,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f320,f500])).
fof(f567,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f545,f498])).
fof(f568,plain,(
  ( ! [X0,X1] : (sP3('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f328,f498])).
fof(f565,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f545,f497])).
fof(f566,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f328,f497])).
fof(f563,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f496])).
fof(f564,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f320,f496])).
fof(f561,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f545,f494])).
fof(f562,plain,(
  ( ! [X0,X1] : (sP4('const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f328,f494])).
fof(f559,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f545,f493])).
fof(f560,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0')),X1)) )),
  inference(unit_resulting_resolution,[],[f328,f493])).
fof(f557,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f492])).
fof(f558,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f320,f492])).
fof(f555,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f318,f435])).
fof(f556,plain,(
  ( ! [X0] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')),X0)) )),
  inference(unit_resulting_resolution,[],[f320,f435])).
fof(f553,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f318,f434])).
fof(f554,plain,(
  ( ! [X0] : (sP0(X0,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(unit_resulting_resolution,[],[f320,f434])).
fof(f535,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f316])).
fof(f316,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f315])).
fof(f315,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f206,f237,f236])).
fof(f236,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f237,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f205])).
fof(f205,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f96])).
fof(f96,axiom,(
  ! [X19,X20,X11] : ((p(s(bool,X11)) <=> (p(s(bool,X20)) | p(s(bool,X19)))) <=> ((p(s(bool,X11)) | ~p(s(bool,X20))) & (p(s(bool,X11)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | p(s(bool,X19)) | ~p(s(bool,X11)))))),
  file('thm.integer.LT_LADD','thm.sat.sat_25_prove_ltb')).
fof(f536,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f537,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f316])).
fof(f538,plain,(
  ( ! [X2,X0,X1] : (~sP9(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f316])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f313])).
fof(f313,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f236])).
fof(f532,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f314])).
fof(f533,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f314])).
fof(f534,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f314])).
fof(f524,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f311,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f310])).
fof(f310,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f309])).
fof(f309,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f237])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f311])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f311])).
fof(f541,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f527])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f311])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f311])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f311])).
fof(f518,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f308])).
fof(f308,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f204,f234,f233])).
fof(f233,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f234,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f204,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f203])).
fof(f203,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f95])).
fof(f95,axiom,(
  ! [X19,X20,X11] : ((p(s(bool,X11)) <=> (p(s(bool,X20)) & p(s(bool,X19)))) <=> ((p(s(bool,X11)) | ~p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X11))) & (p(s(bool,X19)) | ~p(s(bool,X11)))))),
  file('thm.integer.LT_LADD','thm.sat.sat_23_prove_ltb')).
fof(f519,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f308])).
fof(f520,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f308])).
fof(f523,plain,(
  ( ! [X2,X0,X1] : (~sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f308])).
fof(f514,plain,(
  ( ! [X2,X0,X1] : (~sP6(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f306,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f305])).
fof(f305,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f233])).
fof(f515,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f306])).
fof(f516,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f306])).
fof(f517,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f306])).
fof(f507,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f304,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f303])).
fof(f303,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f302])).
fof(f302,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f234])).
fof(f508,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f304])).
fof(f509,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f304])).
fof(f510,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f542,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f513])).
fof(f513,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f304])).
fof(f503,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f202,f231,f230,f229,f228])).
fof(f228,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f229,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f230,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f231,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f202,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f94])).
fof(f94,axiom,(
  ! [X19,X20,X11] : ((p(s(bool,X11)) <=> s(bool,X19) = s(bool,X20)) <=> ((p(s(bool,X11)) | p(s(bool,X20)) | p(s(bool,X19))) & (p(s(bool,X11)) | ~p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X20)) | ~p(s(bool,X19)) | ~p(s(bool,X11))) & (p(s(bool,X19)) | ~p(s(bool,X20)) | ~p(s(bool,X11)))))),
  file('thm.integer.LT_LADD','thm.sat.sat_2|49|_prove_ltb')).
fof(f504,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f301])).
fof(f505,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f301])).
fof(f506,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f301])).
fof(f499,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f228])).
fof(f500,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f300])).
fof(f501,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f300])).
fof(f502,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f300])).
fof(f495,plain,(
  ( ! [X2,X0,X1] : (~sP3(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f297])).
fof(f297,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f229])).
fof(f496,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f298])).
fof(f497,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f298])).
fof(f498,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f491,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f295,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f294])).
fof(f294,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f293])).
fof(f293,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f230])).
fof(f492,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f295])).
fof(f493,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f295])).
fof(f494,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f295])).
fof(f484,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f292,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f291])).
fof(f291,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f290])).
fof(f290,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f231])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f292])).
fof(f487,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f292])).
fof(f488,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | ~p(s(bool,X2)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f489,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f490,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f292])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(flattening,[],[f220])).
fof(f220,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(ennf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X2))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X2)))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X9,X10,X11] : ((p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X10)))) & p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X11))))) => p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X9),s('type.num.num',X11)))))),
  file('thm.integer.LT_LADD','thm.arithmetic.LESS_EQ_TRANS')).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))) )),
  inference(cnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X2))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X2)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X9,X10,X11] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num',X11))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.num.num',X11)))),
  file('thm.integer.LT_LADD','thm.arithmetic.ADD_ASSOC')).
fof(f468,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f219])).
fof(f219,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.integer.LT_LADD',aHL_EXT)).
fof(f448,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f277])).
fof(f277,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f449,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f278])).
fof(f450,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f278])).
fof(f453,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f278])).
fof(f552,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f433,f545])).
fof(f433,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f435,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f271])).
fof(f432,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) | p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f269])).
fof(f430,plain,(
  ( ! [X0,X1] : (s('type.num.num',X0) != s('type.num.num',X1) | s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f268])).
fof(f421,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X9,X10] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.arithmetic.MULT_SUC')).
fof(f419,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X10,X9] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X10),s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|12')).
fof(f418,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))),s('type.num.num',X1))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num',X1)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9,X10] : s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num',X10))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num',X10))),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.arithmetic.MULT|split|1')).
fof(f417,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X10,X9] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X10),s('type.num.num',X9))))) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|8')).
fof(f551,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0)))) )),
  inference(backward_demodulation,[],[f416,f410])).
fof(f415,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_lte|split|6')).
fof(f413,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f86])).
fof(f86,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_lte|split|4')).
fof(f411,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X1),s('type.num.num',X0))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f71])).
fof(f71,axiom,(
  ! [X10,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X9),s('type.num.num',X10))) = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|25')).
fof(f410,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X1))),s('type.num.num',X0))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X1),s('type.num.num',X0)))))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X9] : s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X9))),s('type.num.num',X10))) = s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X9),s('type.num.num',X10)))))),
  file('thm.integer.LT_LADD','thm.arithmetic.ADD_CLAUSES|split|2')).
fof(f409,plain,(
  ( ! [X0,X1] : (s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f68])).
fof(f68,axiom,(
  ! [X10,X9] : s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num',X9))) = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|22')).
fof(f408,plain,(
  ( ! [X0,X1] : (s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X0),s('type.num.num',X1))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X10,X9] : s(bool,'const.arithmetic.<=_2'(s('type.num.num',X10),s('type.num.num',X9))) = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X9)))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_lte|split|3')).
fof(f370,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f244])).
fof(f367,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f242])).
fof(f550,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f358,f322])).
fof(f358,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X0))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : s('type.num.num','const.num.SUC_1'(s('type.num.num',X10))) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.arithmetic.SUC_ONE_ADD')).
fof(f549,plain,(
  ( ! [X0] : (s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f355,f322])).
fof(f355,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X10] : s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|11')).
fof(f352,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10)))))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|9')).
fof(f351,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0)))))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10)))))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|10')).
fof(f348,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f83])).
fof(f83,axiom,(
  ! [X10] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_lte|split|1')).
fof(f347,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))) )),
  inference(cnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(rectify,[],[f84])).
fof(f84,axiom,(
  ! [X10] : s(bool,'F_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X10))),s('type.num.num','const.arithmetic.ZERO_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_lte|split|2')).
fof(f343,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X10] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.ZERO_0'),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_lte|split|0')).
fof(f340,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f75])).
fof(f75,axiom,(
  ! [X10] : s(bool,'T_0') = s(bool,'const.arithmetic.>=_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|29')).
fof(f339,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X10] : s('type.num.num',X10) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.arithmetic.ADD|split|0')).
fof(f338,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X10] : s('type.num.num',X10) = s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|1')).
fof(f336,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f52])).
fof(f52,axiom,(
  ! [X10] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|6')).
fof(f335,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f53])).
fof(f53,axiom,(
  ! [X10] : s('type.num.num',X10) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|7')).
fof(f334,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f72])).
fof(f72,axiom,(
  ! [X10] : s(bool,'T_0') = s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|26')).
fof(f333,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X9),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.arithmetic.MULT_|48|')).
fof(f332,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.*_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X9)))),
  file('thm.integer.LT_LADD','thm.arithmetic.MULT_CLAUSES|split|0')).
fof(f331,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f69])).
fof(f69,axiom,(
  ! [X10] : s(bool,'F_0') = s(bool,'const.arithmetic.>_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10)))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|23')).
fof(f330,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f104])).
fof(f104,plain,(
  ! [X0] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f66])).
fof(f66,axiom,(
  ! [X10] : s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X10),s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|20')).
fof(f329,plain,(
  ( ! [X0] : (p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0))))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X10))))),
  file('thm.integer.LT_LADD','thm.arithmetic.ZERO_LESS_EQ')).
fof(f328,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f80])).
fof(f80,axiom,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|34')).
fof(f321,plain,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f61])).
fof(f61,axiom,(
  s('type.num.num','const.num.|48|_0') = s('type.num.num','const.prim_rec.PRE_1'(s('type.num.num','const.num.|48|_0')))),
  file('thm.integer.LT_LADD','thm.prim_rec.PRE|split|0')).
fof(f320,plain,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(cnf_transformation,[],[f81])).
fof(f81,axiom,(
  p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.LT_LADD','thm.numeral.numeral_distrib|split|35')).
fof(f14597,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14398,f318])).
fof(f14398,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X4))))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f4958])).
fof(f14596,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14595,f14538])).
fof(f14595,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14392,f318])).
fof(f14392,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X4))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X3),s('type.num.num',X2))))))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f4979])).
fof(f14570,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14569,f14538])).
fof(f14569,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14321,f318])).
fof(f14321,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X0))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num',X1))))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f431])).
fof(f14560,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14559,f14538])).
fof(f14559,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14278,f318])).
fof(f14278,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X17))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X15),s('type.num.num','const.num.SUC_1'(s('type.num.num',X16))))))))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f5023])).
fof(f14558,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14557,f14538])).
fof(f14557,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14276,f318])).
fof(f14276,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X4),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f3254])).
fof(f14556,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14555,f14538])).
fof(f14555,plain,(
  ( ! [X130,X131,X129] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num',X130))),s('type.num.num',X131))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14267,f318])).
fof(f14267,plain,(
  ( ! [X130,X131,X129] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X129),s('type.num.num',X130))),s('type.num.num',X131))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f4762])).
fof(f14554,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14553,f14538])).
fof(f14553,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14264,f318])).
fof(f14264,plain,(
  ( ! [X4,X2,X3] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',X2),s('type.num.num',X3))),s('type.num.num',X4))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f3273])).
fof(f14552,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14551,f14538])).
fof(f14551,plain,(
  ( ! [X125,X124] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X124))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14258,f318])).
fof(f14258,plain,(
  ( ! [X125,X124] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X125),s('type.num.num','const.num.SUC_1'(s('type.num.num',X124))))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f11050])).
fof(f14550,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14549,f14538])).
fof(f14549,plain,(
  ( ! [X112,X113] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X113))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X112))))))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14256,f318])).
fof(f14256,plain,(
  ( ! [X112,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num','const.arithmetic.BIT2_1'(s('type.num.num',X113))),s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num',X112))))))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f11041])).
fof(f14534,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14533,f318])).
fof(f14533,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(forward_demodulation,[],[f14532,f14153])).
fof(f14532,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl56),
  inference(subsumption_resolution,[],[f14229,f318])).
fof(f14229,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,'const.arithmetic.<=_2'(s('type.num.num',X1),s('type.num.num',X0))))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f459])).
fof(f14531,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14195,f318])).
fof(f14195,plain,(
  ~p(s(bool,'T_0')) | $spl56),
  inference(backward_demodulation,[],[f14153,f545])).
fof(f14530,plain,(
  $false | $spl56),
  inference(subsumption_resolution,[],[f14529,f14153])).
fof(f14529,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl56),
  inference(subsumption_resolution,[],[f14193,f366])).
fof(f14193,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl56),
  inference(backward_demodulation,[],[f14153,f370])).
fof(f14165,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) ) | $spl64),
  inference(cnf_transformation,[],[f14165_D])).
fof(f14165_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f14153,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl56),
  inference(cnf_transformation,[],[f14153_D])).
fof(f14153_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f14188,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f14187])).
fof(f14187,plain,(
  ( ! [X52,X50,X51,X49] : (~sP8('T_0',X50,X51) | sP9('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f13899,f1116])).
fof(f13899,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~sP8('T_0',X50,X51) | sP9('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f541,f942])).
fof(f14162,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl62),
  inference(cnf_transformation,[],[f14162_D])).
fof(f14162_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f14186,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f14185])).
fof(f14185,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | sP9('T_0',X46,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f13898,f1116])).
fof(f13898,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | sP9('T_0',X46,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f537,f942])).
fof(f14159,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl60),
  inference(cnf_transformation,[],[f14159_D])).
fof(f14159_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f14184,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f14183])).
fof(f14183,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | sP9('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f13897,f1116])).
fof(f13897,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,X42)) | sP9('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f536,f942])).
fof(f14156,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl58),
  inference(cnf_transformation,[],[f14156_D])).
fof(f14156_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f14182,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f14181])).
fof(f14181,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X39)) | sP1(X39,'T_0') | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13896,f1116])).
fof(f13896,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | sP1(X39,'T_0') | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f453,f942])).
fof(f14152,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl54),
  inference(cnf_transformation,[],[f14152_D])).
fof(f14152_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f14180,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f14179])).
fof(f14179,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1('T_0',X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f13895,f1116])).
fof(f13895,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP1('T_0',X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f450,f942])).
fof(f14164,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f14163])).
fof(f14163,plain,(
  ( ! [X52,X50,X51,X49] : (~sP8('T_0',X50,X51) | sP9('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f13822,f1116])).
fof(f13822,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,X49)) | ~sP8('T_0',X50,X51) | sP9('T_0',X50,X51) | s(bool,X49) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f541,f942])).
fof(f14161,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f14160])).
fof(f14160,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X46)) | sP9('T_0',X46,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(subsumption_resolution,[],[f13821,f1116])).
fof(f13821,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | sP9('T_0',X46,X47) | s(bool,X45) = s(bool,X48) | s(bool,'T_0') = s(bool,X48)) )),
  inference(superposition,[],[f537,f942])).
fof(f14158,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f14157])).
fof(f14157,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X42)) | sP9('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(subsumption_resolution,[],[f13820,f1116])).
fof(f13820,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,X42)) | sP9('T_0',X43,X42) | s(bool,X41) = s(bool,X44) | s(bool,'T_0') = s(bool,X44)) )),
  inference(superposition,[],[f536,f942])).
fof(f14155,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f14154])).
fof(f14154,plain,(
  ( ! [X39,X38,X40] : (p(s(bool,X39)) | sP1(X39,'T_0') | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f13819,f1116])).
fof(f13819,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X38)) | p(s(bool,X39)) | sP1(X39,'T_0') | s(bool,X38) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f453,f942])).
fof(f14151,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f14150])).
fof(f14150,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1('T_0',X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f13818,f1116])).
fof(f13818,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP1('T_0',X36) | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f450,f942])).
fof(f10632,plain,(
  $false | ($spl34 | $spl41)),
  inference(subsumption_resolution,[],[f10631,f318])).
fof(f10631,plain,(
  ~p(s(bool,'T_0')) | ($spl34 | $spl41)),
  inference(backward_demodulation,[],[f10630,f10622])).
fof(f10630,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl34),
  inference(backward_demodulation,[],[f10609,f1375])).
fof(f10628,plain,(
  ( ! [X12,X13] : (~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13)) ) | $spl52),
  inference(cnf_transformation,[],[f10628_D])).
fof(f10628_D,plain,(
  ( ! [X12,X13] : (~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f10622,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl41),
  inference(cnf_transformation,[],[f10622_D])).
fof(f10622_D,plain,(
  ~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f10606,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP8('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13)) )),
  inference(superposition,[],[f541,f1375])).
fof(f10627,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10,X11)) ) | $spl50),
  inference(cnf_transformation,[],[f10627_D])).
fof(f10627_D,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10,X11)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f10605,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10,X11)) )),
  inference(superposition,[],[f537,f1375])).
fof(f10626,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X9,X8)) ) | $spl48),
  inference(cnf_transformation,[],[f10626_D])).
fof(f10626_D,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X9,X8)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f10604,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X8)) | sP9('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X9,X8)) )),
  inference(superposition,[],[f536,f1375])).
fof(f10625,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl46),
  inference(cnf_transformation,[],[f10625_D])).
fof(f10625_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f10603,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X7)) | sP1(X7,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f453,f1375])).
fof(f10624,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X6)) ) | $spl44),
  inference(cnf_transformation,[],[f10624_D])).
fof(f10624_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X6)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f10602,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X6)) | sP1('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X6)) )),
  inference(superposition,[],[f450,f1375])).
fof(f10623,plain,(
  ( ! [X5] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X5)) ) | $spl42),
  inference(cnf_transformation,[],[f10623_D])).
fof(f10623_D,plain,(
  ( ! [X5] : (sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X5)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f10601,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | sP0('const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X5)) )),
  inference(superposition,[],[f435,f1375])).
fof(f10620,plain,(
  ( ! [X4] : (sP0(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl38),
  inference(cnf_transformation,[],[f10620_D])).
fof(f10620_D,plain,(
  ( ! [X4] : (sP0(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f10600,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | sP0(X4,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f434,f1375])).
fof(f10609,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl34),
  inference(cnf_transformation,[],[f10609_D])).
fof(f10609_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f10611,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl36),
  inference(cnf_transformation,[],[f10611_D])).
fof(f10611_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f10591,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f360,f1375])).
fof(f10590,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.EVEN_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1375,f360])).
fof(f10486,plain,(
  $false | ($spl14 | $spl21)),
  inference(subsumption_resolution,[],[f10485,f318])).
fof(f10485,plain,(
  ~p(s(bool,'T_0')) | ($spl14 | $spl21)),
  inference(backward_demodulation,[],[f10484,f10476])).
fof(f10484,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl14),
  inference(backward_demodulation,[],[f10463,f1297])).
fof(f10482,plain,(
  ( ! [X12,X13] : (~sP8('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13)) ) | $spl32),
  inference(cnf_transformation,[],[f10482_D])).
fof(f10482_D,plain,(
  ( ! [X12,X13] : (~sP8('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f10476,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | $spl21),
  inference(cnf_transformation,[],[f10476_D])).
fof(f10476_D,plain,(
  ~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f10460,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~sP8('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X12,X13)) )),
  inference(superposition,[],[f541,f1297])).
fof(f10481,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10,X11)) ) | $spl30),
  inference(cnf_transformation,[],[f10481_D])).
fof(f10481_D,plain,(
  ( ! [X10,X11] : (~p(s(bool,X10)) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10,X11)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f10459,plain,(
  ( ! [X10,X11] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X10)) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X10,X11)) )),
  inference(superposition,[],[f537,f1297])).
fof(f10480,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X9,X8)) ) | $spl28),
  inference(cnf_transformation,[],[f10480_D])).
fof(f10480_D,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X9,X8)) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f10458,plain,(
  ( ! [X8,X9] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X8)) | sP9('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X9,X8)) )),
  inference(superposition,[],[f536,f1297])).
fof(f10479,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl26),
  inference(cnf_transformation,[],[f10479_D])).
fof(f10479_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | sP1(X7,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10457,plain,(
  ( ! [X7] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X7)) | sP1(X7,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f453,f1297])).
fof(f10478,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X6)) ) | $spl24),
  inference(cnf_transformation,[],[f10478_D])).
fof(f10478_D,plain,(
  ( ! [X6] : (p(s(bool,X6)) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X6)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f10456,plain,(
  ( ! [X6] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X6)) | sP1('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X6)) )),
  inference(superposition,[],[f450,f1297])).
fof(f10477,plain,(
  ( ! [X5] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X5)) ) | $spl22),
  inference(cnf_transformation,[],[f10477_D])).
fof(f10477_D,plain,(
  ( ! [X5] : (sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X5)) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f10455,plain,(
  ( ! [X5] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | sP0('const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),X5)) )),
  inference(superposition,[],[f435,f1297])).
fof(f10474,plain,(
  ( ! [X4] : (sP0(X4,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) | $spl18),
  inference(cnf_transformation,[],[f10474_D])).
fof(f10474_D,plain,(
  ( ! [X4] : (sP0(X4,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f10454,plain,(
  ( ! [X4] : (~p(s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0')))))) | sP0(X4,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))) )),
  inference(superposition,[],[f434,f1297])).
fof(f10463,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl14),
  inference(cnf_transformation,[],[f10463_D])).
fof(f10463_D,plain,(
  s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f10465,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | $spl16),
  inference(cnf_transformation,[],[f10465_D])).
fof(f10465_D,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f10445,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f360,f1297])).
fof(f10444,plain,(
  s(bool,'F_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'T_0') = s(bool,'const.arithmetic.ODD_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f1297,f360])).
fof(f1879,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) | $spl10),
  inference(cnf_transformation,[],[f1879_D])).
fof(f1879_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1881,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | $spl13),
  inference(cnf_transformation,[],[f1881_D])).
fof(f1881_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f1866,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))),
  inference(superposition,[],[f374,f322])).
fof(f1876,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) | $spl8),
  inference(cnf_transformation,[],[f1876_D])).
fof(f1876_D,plain,(
  ( ! [X8] : (s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f1873,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | $spl7),
  inference(cnf_transformation,[],[f1873_D])).
fof(f1873_D,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f1875,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1864,f353])).
fof(f1864,plain,(
  ( ! [X8] : (s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.EXP_2'(s('type.num.num',X8),s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f374,f549])).
fof(f1871,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) | $spl4),
  inference(cnf_transformation,[],[f1871_D])).
fof(f1871_D,plain,(
  s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1869,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.num.SUC_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(forward_demodulation,[],[f1861,f353])).
fof(f1861,plain,(
  s('type.num.num','const.num.|48|_0') != s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.num.SUC_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.num.num','const.arithmetic.ZERO_0') = s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))),
  inference(superposition,[],[f374,f322])).
fof(f1222,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))))) | $spl0),
  inference(cnf_transformation,[],[f1222_D])).
fof(f1222_D,plain,(
  s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1224,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) | $spl3),
  inference(cnf_transformation,[],[f1224_D])).
fof(f1224_D,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f1220,plain,(
  s(bool,'F_0') != s(bool,'const.prim_rec.<_2'(s('type.num.num',sK11),s('type.num.num',sK12))) | s(bool,'T_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK11))),s('type.num.num','const.arithmetic.+_2'(s('type.num.num',sK10),s('type.num.num',sK12)))))),
  inference(superposition,[],[f317,f360])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1019
% Passive clauses: 3320
% Generated clauses: 38063
% Final active clauses: 652
% Final passive clauses: 1336
% Input formulas: 96
% Initial clauses: 224
% 
% Duplicate literals: 128
% Trivial inequalities: 64
% Fw subsumption resolutions: 500
% Fw demodulations: 5491
% Bw demodulations: 869
% Condensations: 130
% Global subsumptions: 1
% 
% Simple tautologies: 86
% Equational tautologies: 62
% Forward subsumptions: 28020
% Fw demodulations to eq. taut.: 194
% Bw demodulations to eq. taut.: 21
% 
% Unit resulting resolution: 4833
% Factoring: 1
% Forward superposition: 16798
% Backward superposition: 8300
% Self superposition: 77
% Equality factoring: 28
% Equality resolution: 9
% 
% Split clauses: 39
% Split components: 78
% SAT solver clauses: 3429
% SAT solver unit clauses: 1581
% SAT solver binary clauses: 1286
% 
% Sat splits: 39
% Sat splitting refutations: 28
% 
% Memory used [KB]: 19189
% Time elapsed: 0.846 s
% ------------------------------
% ------------------------------
% Success in time 2.166 s
