Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (14531)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 977
% Passive clauses: 9441
% Generated clauses: 23756
% Final active clauses: 960
% Final passive clauses: 8121
% Input formulas: 133
% Initial clauses: 390
% 
% Duplicate literals: 201
% Trivial inequalities: 175
% Fw subsumption resolutions: 1528
% Fw demodulations: 3209
% Bw demodulations: 10
% 
% Simple tautologies: 342
% Equational tautologies: 18
% Forward subsumptions: 8678
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 8489
% Factoring: 611
% Forward superposition: 4840
% Backward superposition: 4098
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 26
% 
% Split clauses: 344
% Split components: 690
% SAT solver clauses: 493
% SAT solver unit clauses: 7
% SAT solver binary clauses: 93
% 
% Sat splits: 344
% 
% TWLsolver clauses: 493
% TWLsolver calls for satisfiability: 330
% 
% Memory used [KB]: 15223
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (14548)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 503
% Passive clauses: 3207
% Generated clauses: 14757
% Final active clauses: 441
% Final passive clauses: 2638
% Input formulas: 133
% Initial clauses: 390
% 
% Duplicate literals: 89
% Trivial inequalities: 14
% Fw subsumption resolutions: 589
% Fw demodulations: 2343
% Bw demodulations: 108
% Condensations: 55
% Global subsumptions: 23
% 
% Simple tautologies: 278
% Equational tautologies: 9
% Forward subsumptions: 4598
% Fw demodulations to eq. taut.: 2
% 
% Binary resolution: 2771
% Factoring: 71
% Forward superposition: 3117
% Backward superposition: 1903
% Self superposition: 32
% Equality factoring: 2
% Equality resolution: 5
% 
% Split clauses: 50
% Split components: 100
% SAT solver clauses: 3903
% SAT solver unit clauses: 376
% SAT solver binary clauses: 135
% SAT solver learnt clauses: 46
% SAT solver learnt literals: 12
% 
% Sat splits: 50
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 3811
% TWLsolver calls for satisfiability: 2689
% 
% Memory used [KB]: 20084
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 137 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (14565)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2360
% Passive clauses: 20459
% Generated clauses: 112539
% Final active clauses: 1848
% Final passive clauses: 5013
% Input formulas: 133
% Initial clauses: 390
% Discarded non-redundant clauses: 12905
% 
% Duplicate literals: 393
% Trivial inequalities: 77
% Fw subsumption resolutions: 2768
% Fw demodulations: 27219
% Bw demodulations: 4597
% Condensations: 921
% Global subsumptions: 956
% 
% Simple tautologies: 185
% Equational tautologies: 140
% Forward subsumptions: 51524
% Fw demodulations to eq. taut.: 3443
% Bw demodulations to eq. taut.: 54
% 
% Unit resulting resolution: 39503
% Factoring: 1
% Forward superposition: 20429
% Backward superposition: 11220
% Self superposition: 105
% Equality factoring: 26
% Equality resolution: 18
% 
% Split clauses: 187
% Split components: 380
% SAT solver clauses: 23617
% SAT solver unit clauses: 4233
% SAT solver binary clauses: 11008
% 
% Sat splits: 187
% Sat splitting refutations: 1825
% 
% Memory used [KB]: 135861
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f51485,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f559,f913,f558,f915,f911,f557,f916,f910,f925,f926,f927,f928,f929,f931,f932,f1030,f1044,f1093,f1049,f1095,f1101,f1113,f1266,f1094,f1118,f1268,f1120,f1270,f1269,f1122,f1124,f1125,f1223,f1411,f1410,f1405,f1528,f1529,f1530,f1531,f1532,f1534,f1535,f1545,f1546,f1547,f1548,f1549,f1551,f1552,f1650,f1693,f1765,f1764,f1905,f1904,f1903,f1898,f2022,f2023,f2024,f2025,f2026,f2028,f2029,f2076,f2077,f2078,f2090,f2105,f2121,f2191,f2190,f2194,f2206,f2220,f2235,f2251,f2267,f2299,f2300,f2301,f2302,f2303,f2305,f2306,f2404,f2443,f2452,f2453,f2454,f2455,f2456,f2458,f2459,f2552,f2582,f2594,f2595,f2596,f2597,f1679,f2598,f2600,f2601,f2699,f1678,f2740,f2754,f2768,f2802,f2803,f2804,f2805,f2806,f2808,f2935,f2809,f2902,f2934,f2944,f2958,f2971,f2985,f3000,f3026,f3027,f3028,f3029,f3030,f3032,f3165,f3033,f3126,f3164,f3175,f3228,f3229,f3230,f3231,f3232,f3234,f1078,f3235,f3328,f3376,f3388,f3398,f3403,f3413,f3420,f3430,f3437,f3447,f3843,f3845,f3873,f3874,f3890,f3891,f1768,f4158,f4173,f4197,f4211,f4315,f4316,f4317,f4318,f4453,f4319,f4321,f4322,f4420,f4452,f4426,f4460,f4471,f4487,f4497,f4516,f4526,f4541,f4552,f4563,f4572,f4587,f4596,f4611,f4637,f4638,f4639,f4640,f4641,f4643,f2435,f4644,f4737,f4769,f4780,f4792,f4801,f4813,f4823,f4845,f4856,f4866,f4893,f4894,f4895,f4896,f4897,f4899,f5026,f4900,f4993,f5038,f5060,f5071,f5081,f5090,f5101,f5127,f5128,f5129,f5130,f5131,f5133,f1079,f5260,f5134,f5227,f5270,f5281,f5289,f5305,f5314,f5330,f5442,f5444,f5446,f5448,f5449,f5547,f5662,f5671,f5682,f5691,f5702,f5712,f5723,f5732,f5743,f5752,f5762,f5772,f5782,f5793,f5803,f5814,f5824,f5835,f6705,f6719,f6796,f6810,f7036,f7038,f7040,f7042,f7043,f7141,f7199,f7201,f7203,f7205,f7206,f7304,f7992,f8004,f8012,f8029,f8073,f8075,f8077,f8079,f8080,f8178,f8208,f8415,f8423,f8429,f8437,f8621,f8622,f8911,f8912,f8913,f8914,f8929,f8915,f8931,f8951,f8952,f8953,f8954,f8969,f8955,f8971,f8993,f8995,f8997,f8999,f9000,f9098,f7171,f9160,f9162,f9163,f9164,f9167,f9260,f9292,f9342,f9344,f9346,f9348,f9349,f9447,f7334,f9484,f9681,f9680,f9679,f9636,f9622,f9621,f9610,f9609,f9608,f9762,f9763,f9764,f9765,f9766,f9767,f9772,f9773,f9774,f9775,f9984,f9986,f9988,f9990,f9991,f10089,f5577,f10158,f8965,f10159,f8925,f10441,f8967,f10442,f8927,f10625,f8963,f10626,f8923,f10682,f10695,f10721,f10734,f10745,f10758,f10771,f10772,f10773,f10784,f10796,f10807,f10819,f10831,f10841,f10854,f10864,f10877,f10888,f10902,f10912,f10926,f10932,f10954,f10938,f10955,f10967,f10979,f11012,f10997,f10998,f10999,f11000,f11024,f11002,f11013,f11074,f11059,f11060,f11061,f11062,f11078,f11064,f11075,f11111,f11122,f11134,f11146,f11180,f11166,f11167,f11168,f11169,f11184,f11171,f11181,f11253,f11254,f11289,f11275,f11276,f11277,f11278,f11293,f11280,f11290,f11369,f11404,f11390,f11391,f11392,f11393,f11408,f11395,f11405,f11478,f11514,f11499,f11500,f11501,f11502,f11518,f11504,f11515,f11566,f11579,f11593,f11627,f11628,f11664,f11649,f11650,f11651,f11652,f11668,f11654,f11665,f11716,f11717,f11761,f11804,f11839,f11824,f11825,f11826,f11827,f11843,f11829,f11840,f11888,f11902,f11915,f11928,f11942,f11978,f11962,f11963,f11964,f11965,f11982,f11967,f11979,f12029,f12063,f12096,f12135,f12119,f12120,f12121,f12122,f12139,f12124,f12136,f12191,f12216,f12230,f12243,f12267,f12275,f12308,f12309,f12291,f12292,f12293,f12294,f12313,f12296,f12310,f12348,f12361,f12372,f12384,f12418,f12405,f12406,f12407,f11021,f12408,f12422,f11023,f12410,f12419,f12449,f12462,f12474,f12486,f12861,f12863,f12865,f12867,f12868,f12966,f13024,f13026,f13027,f13028,f13031,f13124,f13223,f13224,f13225,f13226,f13228,f13382,f13229,f13230,f13233,f13234,f13235,f13236,f13246,f13248,f13249,f13250,f13253,f13346,f13389,f13423,f13424,f13425,f13426,f13429,f13383,f13430,f13433,f13434,f13435,f13436,f13446,f13448,f13449,f13450,f13453,f13546,f13578,f13699,f13700,f13729,f13730,f13731,f13732,f13736,f13384,f13739,f13740,f13741,f13742,f13752,f13754,f13755,f13756,f13759,f13852,f13884,f14033,f14034,f14035,f14036,f14043,f13385,f14044,f14045,f14046,f14056,f14058,f14059,f14060,f14063,f14156,f14188,f14380,f14443,f14455,f14459,f14460,f14461,f14462,f14470,f13386,f14471,f14472,f14482,f14484,f14485,f14486,f14489,f14582,f14614,f14632,f14644,f14658,f14670,f14673,f13378,f14674,f14675,f14676,f14685,f14686,f14696,f14698,f14699,f14700,f14703,f14796,f14828,f14839,f15222,f15456,f15484,f15554,f15590,f15656,f15672,f15832,f15848,f15863,f15879,f15895,f15911,f16098,f16112,f16128,f16143,f16157,f16173,f16191,f16220,f17950,f17952,f17953,f17954,f17957,f18050,f18089,f18741,f18997,f19022,f19041,f19068,f19087,f19109,f19128,f19144,f19198,f19223,f19242,f19267,f19318,f19338,f19356,f19376,f19392,f19419,f19438,f19453,f19506,f19533,f19552,f19567,f19595,f19659,f19675,f19694,f19710,f19729,f23785,f23811,f23985,f24016,f24191,f24218,f24238,f24265,f24287,f24313,f24336,f24362,f24387,f24411,f24431,f24458,f24584,f24608,f24659,f24686,f25630,f25663,f26405,f26438,f26528,f26551,f26574,f26597,f1703,f26604,f26798,f26607,f26800,f26802,f26930,f26932,f26934,f26936,f26937,f27035,f27063,f27080,f27082,f27083,f27084,f27087,f27180,f27210,f27423,f27438,f29014,f29016,f29017,f29018,f29021,f29114,f29153,f30643,f30645,f30646,f30647,f30650,f30743,f30806,f32399,f32398,f32397,f32383,f32372,f32371,f32370,f32538,f32539,f32540,f32541,f32542,f32543,f32548,f32549,f32550,f32551,f32552,f32553,f36134,f36136,f36137,f36138,f36141,f36234,f36274,f36407,f36409,f36411,f36413,f36414,f36512,f36551,f37416,f37439,f39473,f39475,f39477,f39479,f39480,f39578,f39644,f39693,f39695,f39696,f39697,f39700,f39793,f39912,f39914,f39915,f39916,f39919,f40012,f40206,f40018,f13379,f40019,f40020,f40029,f40030,f40210,f40040,f40042,f40043,f40044,f40047,f40140,f40213,f40329,f13380,f40330,f40339,f40340,f40645,f40350,f40352,f40353,f40354,f40357,f40450,f40648,f40479,f40481,f40482,f40483,f40486,f40579,f40652,f40760,f40779,f40807,f40809,f40810,f40811,f40814,f40907,f41102,f40915,f13381,f40924,f40925,f41106,f40935,f40937,f40938,f40939,f40942,f41035,f41109,f41173,f41175,f41176,f41177,f41180,f41273,f41307,f41362,f13387,f41363,f41669,f41373,f41375,f41376,f41377,f41380,f41473,f41672,f41502,f41504,f41505,f41506,f41509,f41602,f41676,f41708,f41710,f41711,f41712,f41715,f41808,f41909,f41911,f41912,f41913,f41916,f42009,f42205,f42204,f42211,f42209,f42028,f42212,f13388,f42038,f42040,f42041,f42042,f42045,f42138,f42215,f42247,f42249,f42250,f42251,f42254,f42347,f42381,f42448,f42450,f42451,f42452,f42455,f42548,f42586,f42835,f39598,f42843,f42842,f42964,f42966,f42968,f42970,f42971,f43069,f13007,f43125,f43127,f43129,f43131,f43132,f43230,f1076,f43306,f43308,f43309,f43310,f43313,f43406,f13168,f43527,f43529,f43531,f43533,f43534,f43632,f1267,f43700,f43721,f43723,f43725,f43727,f43728,f43826,f43858,f44105,f44126,f44106,f44128,f44207,f44209,f44211,f44213,f44214,f44312,f2434,f44501,f13726,f44502,f13727,f44711,f44713,f44714,f44715,f44718,f44811,f44843,f45025,f14406,f45342,f10783,f45343,f45363,f45500,f45502,f45503,f45504,f45507,f45600,f5025,f45725,f45727,f45728,f45729,f45732,f45825,f5259,f46755,f10780,f46966,f10782,f47288,f47303,f47543,f47558,f48128,f48129,f48133,f1704,f48146,f48165,f48194,f48195,f48199,f2099,f48232,f48233,f48237,f2115,f48272,f48273,f48277,f2130,f48311,f48312,f48316,f2215,f48333,f12307,f48334,f48372,f48356,f48357,f48361,f2230,f48401,f48402,f48406,f2246,f48617,f48618,f48622,f2262,f49281,f11017,f49282,f11019,f49286,f2278,f11026,f50276,f50279,f50278,f50288,f50290,f50291,f50292,f50295,f50388,f39827,f50738,f50740,f50741,f1077,f50742,f1082,f50745,f1080,f50838,f1083,f41842,f51171,f51114,f51113,f51108,f51243,f51244,f51245,f51246,f51247,f51249,f51250,f51474])).
fof(f51474,plain,(
  $false | ($spl1 | $spl6)),
  inference(subsumption_resolution,[],[f51320,f1098])).
fof(f1098,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) | $spl1),
  inference(subsumption_resolution,[],[f1047,f886])).
fof(f886,plain,(
  ( ! [X0] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.NIL_0'))))) )),
  inference(cnf_transformation,[],[f279])).
fof(f279,plain,(
  ! [X0] : p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.NIL_0'))))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X6] : p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.NIL_0'))))),
  file('thm.list.NULL_FILTER','thm.list.NULL|split|0')).
fof(f1047,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) | $spl1),
  inference(superposition,[],[f911,f594])).
fof(f594,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X2,X1,X0))))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f407,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK35(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK35(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK35])],[f406])).
fof(f406,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f405])).
fof(f405,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0] : ! [X1,X2] : (~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X11] : (~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X11))) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X11)))))) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))))),
  file('thm.list.NULL_FILTER','thm.list.FILTER_NEQ_NIL')).
fof(f51320,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) | ($spl1 | $spl6)),
  inference(resolution,[],[f43289,f916])).
fof(f43289,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl1),
  inference(subsumption_resolution,[],[f43246,f886])).
fof(f43246,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl1),
  inference(superposition,[],[f911,f593])).
fof(f593,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK35(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f407])).
fof(f51250,plain,(
  $false | ($spl3 | $spl6 | $spl68)),
  inference(subsumption_resolution,[],[f51129,f50984])).
fof(f50984,plain,(
  ( ! [X35] : (p(s(bool,X35))) ) | ($spl3 | $spl68)),
  inference(subsumption_resolution,[],[f50983,f855])).
fof(f855,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f544])).
fof(f544,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f247])).
fof(f247,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f246])).
fof(f246,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.EQ_CLAUSES|split|3')).
fof(f50983,plain,(
  ( ! [X35] : (s(bool,'F_0') != s(bool,X35) | p(s(bool,X35))) ) | ($spl3 | $spl68)),
  inference(forward_demodulation,[],[f2934,f43820])).
fof(f43820,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | $spl3),
  inference(resolution,[],[f912,f853])).
fof(f853,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f543])).
fof(f543,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f245])).
fof(f245,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f244])).
fof(f244,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.EQ_CLAUSES|split|2')).
fof(f912,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | $spl3),
  inference(cnf_transformation,[],[f912_D])).
fof(f912_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f51129,plain,(
  ( ! [X373] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X373))))) ) | ($spl3 | $spl6 | $spl68)),
  inference(resolution,[],[f50984,f916])).
fof(f51249,plain,(
  $false | ($spl3 | $spl68)),
  inference(subsumption_resolution,[],[f51123,f50984])).
fof(f51123,plain,(
  ( ! [X352,X351] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X351),X352))))) ) | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f628])).
fof(f628,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) )),
  inference(cnf_transformation,[],[f442])).
fof(f442,plain,(
  ! [X0,X1] : ((p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK52(X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) & (! [X2] : ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK52])],[f441])).
fof(f441,plain,(
  ! [X0,X1] : ((p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) | ? [X3] : p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) & (! [X2] : ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f440])).
fof(f440,plain,(
  ! [X0,X1] : ((p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) | ? [X2] : p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1))))))) & (! [X2] : ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1))))))),
  inference(nnf_transformation,[],[f184])).
fof(f184,plain,(
  ! [X0,X1] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) <=> ? [X2] : p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X0] : ! [X1] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) <=> ? [X2] : p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X11] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),X11)))) <=> ? [X12] : p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X11)))))))),
  file('thm.list.NULL_FILTER','thm.list.NOT_NULL_MEM')).
fof(f51247,plain,(
  $false | ($spl3 | $spl68)),
  inference(subsumption_resolution,[],[f51112,f50984])).
fof(f51112,plain,(
  ( ! [X313,X312,X314] : (~p(s(bool,'const.bool.IN_2'(s(X312,X313),s(fun(X312,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X312),X314))))))) ) | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f628])).
fof(f51246,plain,(
  $false | ($spl3 | $spl68)),
  inference(subsumption_resolution,[],[f51110,f50984])).
fof(f51110,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f875])).
fof(f875,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f548])).
fof(f548,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f265])).
fof(f265,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f13])).
fof(f13,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.list.NULL_FILTER','thm.bool.NOT_CLAUSES|split|2')).
fof(f51245,plain,(
  $false | ($spl3 | $spl68)),
  inference(subsumption_resolution,[],[f51109,f50984])).
fof(f51109,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f873])).
fof(f873,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f547])).
fof(f547,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f264])).
fof(f264,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f12])).
fof(f12,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.NULL_FILTER','thm.bool.NOT_CLAUSES|split|1')).
fof(f51244,plain,(
  $false | ($spl3 | $spl68)),
  inference(subsumption_resolution,[],[f51106,f50984])).
fof(f51106,plain,(
  ( ! [X308] : (~p(s(bool,X308))) ) | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f867])).
fof(f867,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f345])).
fof(f345,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f257])).
fof(f257,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f256])).
fof(f256,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f88])).
fof(f88,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) => ~p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_F')).
fof(f51243,plain,(
  $false | ($spl3 | $spl68)),
  inference(subsumption_resolution,[],[f51100,f50984])).
fof(f51100,plain,(
  ( ! [X302] : (~p(s(bool,X302))) ) | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f808])).
fof(f808,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f514])).
fof(f514,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f513])).
fof(f513,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f335])).
fof(f335,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f229])).
fof(f229,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f228])).
fof(f228,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_CLAUSES|split|4')).
fof(f51108,plain,(
  $false | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f871])).
fof(f871,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.NULL_FILTER',aHL_FALSITY)).
fof(f51113,plain,(
  $false | ($spl3 | $spl68)),
  inference(resolution,[],[f50984,f880])).
fof(f880,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f271])).
fof(f271,plain,(
  ! [X0,X1,X2] : ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X0] : ! [X1,X2] : ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X6] : ! [X13,X10] : ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),X10))))))),
  file('thm.list.NULL_FILTER','thm.list.NULL|split|1')).
fof(f51114,plain,(
  $false | ($spl1 | $spl3 | $spl68)),
  inference(resolution,[],[f50984,f911])).
fof(f51171,plain,(
  $false | ($spl3 | $spl29 | $spl68)),
  inference(resolution,[],[f50984,f1266])).
fof(f41842,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | $spl502),
  inference(cnf_transformation,[],[f41842_D])).
fof(f41842_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f1083,plain,(
  ( ! [X32] : (p(s(bool,X32))) ) | $spl20),
  inference(cnf_transformation,[],[f1083_D])).
fof(f1083_D,plain,(
  ( ! [X32] : (p(s(bool,X32))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f50838,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | $spl487),
  inference(resolution,[],[f50281,f901])).
fof(f901,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f650])).
fof(f650,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f444])).
fof(f444,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f443])).
fof(f443,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f313])).
fof(f313,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f312])).
fof(f312,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f122])).
fof(f122,axiom,(
  ! [X9,X51,X20,X52] : ((s(bool,X9) = s(bool,X51) & (p(s(bool,X51)) => s(bool,X20) = s(bool,X52))) => ((p(s(bool,X9)) => p(s(bool,X20))) <=> (p(s(bool,X51)) => p(s(bool,X52)))))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_CONG')).
fof(f50281,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) | $spl487),
  inference(resolution,[],[f50270,f622])).
fof(f622,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f436])).
fof(f436,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK50(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK50(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK50])],[f435])).
fof(f435,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f434])).
fof(f434,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f179])).
fof(f179,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X6] : ! [X7,X11] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X11)))) <=> ? [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X11)))))) & p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12))))))),
  file('thm.list.NULL_FILTER','thm.list.EXISTS_MEM')).
fof(f50270,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl487),
  inference(trivial_inequality_removal,[],[f50263])).
fof(f50263,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl487),
  inference(superposition,[],[f40779,f855])).
fof(f1080,plain,(
  ( ! [X22,X20] : (s(bool,X20) != s(bool,X22) | ~p(s(bool,X20))) ) | $spl16),
  inference(cnf_transformation,[],[f1080_D])).
fof(f1080_D,plain,(
  ( ! [X22,X20] : (s(bool,X20) != s(bool,X22) | ~p(s(bool,X20))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f50745,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl487),
  inference(resolution,[],[f50281,f653])).
fof(f653,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f444])).
fof(f1082,plain,(
  ( ! [X34] : (~p(s(bool,X34))) ) | $spl18),
  inference(cnf_transformation,[],[f1082_D])).
fof(f1082_D,plain,(
  ( ! [X34] : (~p(s(bool,X34))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f50742,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | $spl487),
  inference(resolution,[],[f50281,f652])).
fof(f652,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f444])).
fof(f1077,plain,(
  ( ! [X14,X16] : (s(bool,X14) != s(bool,X16) | p(s(bool,X14))) ) | $spl10),
  inference(cnf_transformation,[],[f1077_D])).
fof(f1077_D,plain,(
  ( ! [X14,X16] : (s(bool,X14) != s(bool,X16) | p(s(bool,X14))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f50741,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | $spl487),
  inference(resolution,[],[f50281,f651])).
fof(f651,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f444])).
fof(f50740,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl487),
  inference(resolution,[],[f50281,f649])).
fof(f649,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f444])).
fof(f50738,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | $spl487),
  inference(resolution,[],[f50281,f647])).
fof(f647,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f444])).
fof(f39827,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl472),
  inference(cnf_transformation,[],[f39827_D])).
fof(f39827_D,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f50388,plain,(
  ( ! [X189,X188] : (p(s(bool,X188)) | p(s(bool,X189)) | s(bool,X189) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl487),
  inference(resolution,[],[f50270,f901])).
fof(f50295,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X40) != s(bool,X41)) ) | $spl487),
  inference(resolution,[],[f50270,f653])).
fof(f50292,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl487),
  inference(resolution,[],[f50270,f652])).
fof(f50291,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl487),
  inference(resolution,[],[f50270,f651])).
fof(f50290,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X25) != s(bool,X26)) ) | $spl487),
  inference(resolution,[],[f50270,f649])).
fof(f50288,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl487),
  inference(resolution,[],[f50270,f647])).
fof(f50278,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl538),
  inference(cnf_transformation,[],[f50278_D])).
fof(f50278_D,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) <=> ~$spl538),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl538])])).
fof(f50279,plain,(
  ( ! [X9] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X9))) ) | $spl487),
  inference(subsumption_resolution,[],[f50259,f853])).
fof(f50259,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | p(s(bool,X9))) ) | $spl487),
  inference(superposition,[],[f40779,f795])).
fof(f795,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f334])).
fof(f334,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f333])).
fof(f333,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f221])).
fof(f221,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X17,X18] : ((p(s(bool,X17)) => p(s(bool,X18))) => ((p(s(bool,X18)) => p(s(bool,X17))) => s(bool,X17) = s(bool,X18)))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_ANTISYM_AX')).
fof(f50276,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl487),
  inference(subsumption_resolution,[],[f50258,f853])).
fof(f50258,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) ) | $spl487),
  inference(superposition,[],[f40779,f795])).
fof(f11026,plain,(
  s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) | $spl244),
  inference(cnf_transformation,[],[f11026_D])).
fof(f11026_D,plain,(
  s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)))) <=> ~$spl244),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl244])])).
fof(f2278,plain,(
  ( ! [X0,X1] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl60),
  inference(cnf_transformation,[],[f2278_D])).
fof(f2278_D,plain,(
  ( ! [X0,X1] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f49286,plain,(
  ( ! [X12,X13] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X12,X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f986,f900])).
fof(f900,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK54(X1,X0)),s('type.list.list'(X0),sK55(X1,X0)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) )),
  inference(cnf_transformation,[],[f556])).
fof(f556,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK54(X1,X0)),s('type.list.list'(X0),sK55(X1,X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK54,sK55])],[f297])).
fof(f297,plain,(
  ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(flattening,[],[f296])).
fof(f296,plain,(
  ! [X0] : ! [X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ? [X2,X3] : s('type.list.list'(X0),X1) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))),
  inference(rectify,[],[f112])).
fof(f112,axiom,(
  ! [X6] : ! [X11] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X11) | ? [X13,X10] : s('type.list.list'(X6),X11) = s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),X10))))),
  file('thm.list.NULL_FILTER','thm.list.list_CASES')).
fof(f986,plain,(
  ( ! [X144,X145] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X144,X145)) ) | $spl1),
  inference(resolution,[],[f911,f719])).
fof(f719,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f476])).
fof(f476,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f475])).
fof(f475,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(flattening,[],[f474])).
fof(f474,plain,(
  ! [X1,X0,X2] : ((~sP7(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP7(X1,X0,X2)))),
  inference(nnf_transformation,[],[f361])).
fof(f361,plain,(
  ! [X1,X0,X2] : (sP7(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f11019,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl236),
  inference(cnf_transformation,[],[f11019_D])).
fof(f11019_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f49282,plain,(
  ( ! [X4,X5] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X4,X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f986,f593])).
fof(f11017,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) | $spl235),
  inference(cnf_transformation,[],[f11017_D])).
fof(f11017_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15)))) <=> ~$spl235),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl235])])).
fof(f49281,plain,(
  ( ! [X2,X3] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2,X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f986,f889])).
fof(f889,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f550])).
fof(f550,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2))))) & (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f282])).
fof(f282,plain,(
  ! [X0] : ! [X1,X2] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X0,bool),X1))),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X6] : ! [X7,X11] : (s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X11))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(X6,bool),X7))),s('type.list.list'(X6),X11)))))),
  file('thm.list.NULL_FILTER','thm.list.FILTER_EQ_NIL')).
fof(f2262,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl58),
  inference(cnf_transformation,[],[f2262_D])).
fof(f2262_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f48622,plain,(
  ( ! [X12,X13] : (sP7(X12,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f985,f900])).
fof(f985,plain,(
  ( ! [X142,X143] : (sP7(X142,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X143)) ) | $spl1),
  inference(resolution,[],[f911,f718])).
fof(f718,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f476])).
fof(f48618,plain,(
  ( ! [X4,X5] : (sP7(X4,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f985,f593])).
fof(f48617,plain,(
  ( ! [X2,X3] : (sP7(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f985,f889])).
fof(f2246,plain,(
  ( ! [X0,X1] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl56),
  inference(cnf_transformation,[],[f2246_D])).
fof(f2246_D,plain,(
  ( ! [X0,X1] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f48406,plain,(
  ( ! [X12,X13] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X12,X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f983,f900])).
fof(f983,plain,(
  ( ! [X138,X139] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X138,X139)) ) | $spl1),
  inference(resolution,[],[f911,f715])).
fof(f715,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f473])).
fof(f473,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f472])).
fof(f472,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f471])).
fof(f471,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f362])).
fof(f362,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f48402,plain,(
  ( ! [X4,X5] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X4,X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f983,f593])).
fof(f48401,plain,(
  ( ! [X2,X3] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2,X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f983,f889])).
fof(f2230,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl54),
  inference(cnf_transformation,[],[f2230_D])).
fof(f2230_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f48361,plain,(
  ( ! [X12,X13] : (sP8(X12,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f982,f900])).
fof(f982,plain,(
  ( ! [X136,X137] : (sP8(X136,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X137)) ) | $spl1),
  inference(resolution,[],[f911,f714])).
fof(f714,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f473])).
fof(f48357,plain,(
  ( ! [X4,X5] : (sP8(X4,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f982,f593])).
fof(f48356,plain,(
  ( ! [X2,X3] : (sP8(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f982,f889])).
fof(f48372,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | $spl537),
  inference(cnf_transformation,[],[f48372_D])).
fof(f48372_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) <=> ~$spl537),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl537])])).
fof(f48334,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X9))) ) | ($spl1 | $spl25)),
  inference(resolution,[],[f982,f8903])).
fof(f8903,plain,(
  ( ! [X14,X15] : (~sP8(X14,X15,'F_0') | s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X14))) ) | $spl25),
  inference(subsumption_resolution,[],[f8902,f721])).
fof(f721,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f478])).
fof(f478,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f477])).
fof(f477,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f360])).
fof(f360,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f8902,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl25),
  inference(subsumption_resolution,[],[f8901,f808])).
fof(f8901,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP6(X15,'F_0',X14)) ) | $spl25),
  inference(subsumption_resolution,[],[f8893,f719])).
fof(f8893,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | s(bool,'T_0') != s(bool,X15) | p(s(bool,'F_0')) | ~sP8(X14,X15,'F_0') | ~sP7('F_0',X15,X14) | ~sP6(X15,'F_0',X14)) ) | $spl25),
  inference(resolution,[],[f4136,f710])).
fof(f710,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | p(s(bool,X1)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f470])).
fof(f470,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | (sP6(X2,X1,X0) & sP7(X1,X2,X0) & sP8(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X2,X1,X0) | ~sP7(X1,X2,X0) | ~sP8(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f469])).
fof(f469,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP9(X2,X1,X0)))),
  inference(flattening,[],[f468])).
fof(f468,plain,(
  ! [X2,X1,X0] : ((~sP9(X2,X1,X0) | (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | ~sP7(X1,X0,X2) | ~sP8(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP9(X2,X1,X0)))),
  inference(nnf_transformation,[],[f363])).
fof(f363,plain,(
  ! [X2,X1,X0] : (sP9(X2,X1,X0) <=> (sP6(X0,X1,X2) & sP7(X1,X0,X2) & sP8(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f4136,plain,(
  ( ! [X8,X9] : (~sP9(X9,'F_0',X8) | ~p(s(bool,X9)) | s(bool,'T_0') != s(bool,X8)) ) | $spl25),
  inference(superposition,[],[f4096,f726])).
fof(f726,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f479])).
fof(f479,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP9(X2,X1,X0)) & (~sP9(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f364])).
fof(f364,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP9(X2,X1,X0))),
  inference(definition_folding,[],[f197,f363,f362,f361,f360])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f90])).
fof(f90,axiom,(
  ! [X34,X35,X36] : ((p(s(bool,X36)) <=> s(bool,X34) = s(bool,X35)) <=> ((p(s(bool,X36)) | p(s(bool,X35)) | p(s(bool,X34))) & (p(s(bool,X36)) | ~p(s(bool,X34)) | ~p(s(bool,X35))) & (p(s(bool,X35)) | ~p(s(bool,X34)) | ~p(s(bool,X36))) & (p(s(bool,X34)) | ~p(s(bool,X35)) | ~p(s(bool,X36)))))),
  file('thm.list.NULL_FILTER','thm.sat.sat_2|49|_prove_ltb')).
fof(f4096,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl25),
  inference(superposition,[],[f2695,f2693])).
fof(f2693,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | $spl25),
  inference(resolution,[],[f1095,f853])).
fof(f2695,plain,(
  s(bool,'T_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | $spl25),
  inference(resolution,[],[f1095,f856])).
fof(f856,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f545])).
fof(f545,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f248])).
fof(f248,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.EQ_CLAUSES|split|0')).
fof(f12307,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | $spl302),
  inference(cnf_transformation,[],[f12307_D])).
fof(f12307_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f48333,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X8))) ) | ($spl1 | $spl25)),
  inference(resolution,[],[f982,f8824])).
fof(f8824,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl25),
  inference(subsumption_resolution,[],[f8817,f721])).
fof(f8817,plain,(
  ( ! [X2,X3] : (~sP8(X2,X3,'F_0') | ~sP6(X3,'F_0',X2) | s(bool,'F_0') = s(bool,X3) | ~p(s(bool,X2))) ) | $spl25),
  inference(resolution,[],[f8753,f726])).
fof(f8753,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl25),
  inference(subsumption_resolution,[],[f8751,f871])).
fof(f8751,plain,(
  ( ! [X4,X5] : (sP9(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4)) ) | $spl25),
  inference(resolution,[],[f8748,f710])).
fof(f8748,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5)) ) | $spl25),
  inference(subsumption_resolution,[],[f8734,f717])).
fof(f717,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f476])).
fof(f8734,plain,(
  ( ! [X4,X5] : (sP7('F_0',X4,X5) | p(s(bool,X5))) ) | $spl25),
  inference(resolution,[],[f8731,f715])).
fof(f8731,plain,(
  ( ! [X6,X7] : (~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl25),
  inference(subsumption_resolution,[],[f8726,f718])).
fof(f8726,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP8(X7,X6,'F_0') | sP7('F_0',X6,X7)) ) | $spl25),
  inference(resolution,[],[f8720,f706])).
fof(f706,plain,(
  ( ! [X2,X0,X1] : (sP7(X1,X2,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f470])).
fof(f8720,plain,(
  ( ! [X2,X3] : (sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl25),
  inference(subsumption_resolution,[],[f8719,f857])).
fof(f857,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f545])).
fof(f8719,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0')) ) | $spl25),
  inference(subsumption_resolution,[],[f8716,f8706])).
fof(f8706,plain,(
  ( ! [X4,X5] : (sP6(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl25),
  inference(subsumption_resolution,[],[f8700,f721])).
fof(f8700,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP6(X4,'F_0',X5)) ) | $spl25),
  inference(resolution,[],[f4134,f705])).
fof(f705,plain,(
  ( ! [X2,X0,X1] : (sP6(X2,X1,X0) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f470])).
fof(f4134,plain,(
  ( ! [X4,X5] : (sP9(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4) | p(s(bool,X5))) ) | $spl25),
  inference(superposition,[],[f4096,f724])).
fof(f724,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f479])).
fof(f8716,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP9(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP8(X3,X2,'F_0') | ~sP6(X2,'F_0',X3)) ) | $spl25),
  inference(resolution,[],[f8707,f709])).
fof(f709,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~p(s(bool,X2)) | ~sP8(X0,X2,X1) | ~sP7(X1,X2,X0) | ~sP6(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f470])).
fof(f8707,plain,(
  ( ! [X6,X7] : (sP7('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl25),
  inference(subsumption_resolution,[],[f8701,f717])).
fof(f8701,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP7('F_0',X6,X7)) ) | $spl25),
  inference(resolution,[],[f4134,f706])).
fof(f2215,plain,(
  ( ! [X0,X1] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl52),
  inference(cnf_transformation,[],[f2215_D])).
fof(f2215_D,plain,(
  ( ! [X0,X1] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f48316,plain,(
  ( ! [X12,X13] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X12,X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f970,f900])).
fof(f970,plain,(
  ( ! [X112,X113] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X112,X113)) ) | $spl1),
  inference(resolution,[],[f911,f698])).
fof(f698,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f465])).
fof(f465,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(flattening,[],[f464])).
fof(f464,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(nnf_transformation,[],[f357])).
fof(f357,plain,(
  ! [X0,X1,X2] : (sP4(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f48312,plain,(
  ( ! [X4,X5] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X4,X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f970,f593])).
fof(f48311,plain,(
  ( ! [X2,X3] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2,X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f970,f889])).
fof(f2130,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl48),
  inference(cnf_transformation,[],[f2130_D])).
fof(f2130_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f48277,plain,(
  ( ! [X12,X13] : (sP4(X12,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f969,f900])).
fof(f969,plain,(
  ( ! [X111,X110] : (sP4(X110,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X111)) ) | $spl1),
  inference(resolution,[],[f911,f697])).
fof(f697,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f465])).
fof(f48273,plain,(
  ( ! [X4,X5] : (sP4(X4,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f969,f593])).
fof(f48272,plain,(
  ( ! [X2,X3] : (sP4(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f969,f889])).
fof(f2115,plain,(
  ( ! [X0,X1] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl46),
  inference(cnf_transformation,[],[f2115_D])).
fof(f2115_D,plain,(
  ( ! [X0,X1] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f48237,plain,(
  ( ! [X12,X13] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X12,X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f956,f900])).
fof(f956,plain,(
  ( ! [X85,X84] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X84,X85)) ) | $spl1),
  inference(resolution,[],[f911,f681])).
fof(f681,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f458])).
fof(f458,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f457])).
fof(f457,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f456])).
fof(f456,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f354])).
fof(f354,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f48233,plain,(
  ( ! [X4,X5] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X4,X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f956,f593])).
fof(f48232,plain,(
  ( ! [X2,X3] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2,X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f956,f889])).
fof(f2099,plain,(
  ( ! [X0,X1] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl44),
  inference(cnf_transformation,[],[f2099_D])).
fof(f2099_D,plain,(
  ( ! [X0,X1] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f48199,plain,(
  ( ! [X12,X13] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X12,X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f940,f900])).
fof(f940,plain,(
  ( ! [X52,X53] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X52,X53)) ) | $spl1),
  inference(resolution,[],[f911,f664])).
fof(f664,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f450])).
fof(f450,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(rectify,[],[f449])).
fof(f449,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(flattening,[],[f448])).
fof(f448,plain,(
  ! [X2,X0,X1] : ((~sP0(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP0(X2,X0,X1)))),
  inference(nnf_transformation,[],[f351])).
fof(f351,plain,(
  ! [X2,X0,X1] : (sP0(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f48195,plain,(
  ( ! [X4,X5] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X4,X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f940,f593])).
fof(f48194,plain,(
  ( ! [X2,X3] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2,X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f940,f889])).
fof(f48165,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl40),
  inference(subsumption_resolution,[],[f48147,f886])).
fof(f48147,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,X3)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl40),
  inference(resolution,[],[f1704,f661])).
fof(f661,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP0(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f450])).
fof(f48146,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl40 | $spl452)),
  inference(resolution,[],[f1704,f35627])).
fof(f35627,plain,(
  ( ! [X56,X55] : (~sP0(X55,X56,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X55)) | p(s(bool,X56))) ) | $spl452),
  inference(forward_demodulation,[],[f27094,f27072])).
fof(f27072,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),sK55(sK15,sK13)) | $spl452),
  inference(resolution,[],[f26801,f882])).
fof(f882,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f549])).
fof(f549,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1)) & (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X1) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1))))))),
  inference(nnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1))),
  inference(flattening,[],[f274])).
fof(f274,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1)))) <=> s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X1))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X6] : ! [X11] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),X11)))) <=> s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X11))),
  file('thm.list.NULL_FILTER','thm.list.NULL_EQ')).
fof(f26801,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) | $spl452),
  inference(cnf_transformation,[],[f26801_D])).
fof(f26801_D,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f27094,plain,(
  ( ! [X56,X55] : (~p(s(bool,X55)) | p(s(bool,X56)) | ~sP0(X55,X56,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl452),
  inference(resolution,[],[f26801,f661])).
fof(f1704,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl40),
  inference(cnf_transformation,[],[f1704_D])).
fof(f1704_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f48133,plain,(
  ( ! [X12,X13] : (sP0(X12,X13,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl1),
  inference(superposition,[],[f939,f900])).
fof(f939,plain,(
  ( ! [X50,X51] : (sP0(X50,X51,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f662])).
fof(f662,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f450])).
fof(f48129,plain,(
  ( ! [X4,X5] : (sP0(X4,X5,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl1),
  inference(superposition,[],[f939,f593])).
fof(f48128,plain,(
  ( ! [X2,X3] : (sP0(X2,X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl1),
  inference(superposition,[],[f939,f889])).
fof(f47558,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13)))) | $spl535),
  inference(cnf_transformation,[],[f47558_D])).
fof(f47558_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13)))) <=> ~$spl535),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl535])])).
fof(f47543,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) ) | ($spl25 | $spl508)),
  inference(resolution,[],[f42311,f8825])).
fof(f8825,plain,(
  ( ! [X4,X5] : (~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl25),
  inference(subsumption_resolution,[],[f8818,f715])).
fof(f8818,plain,(
  ( ! [X4,X5] : (~sP8(X4,X5,'F_0') | ~sP6(X5,'F_0',X4) | p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | $spl25),
  inference(resolution,[],[f8753,f727])).
fof(f727,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP9(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f479])).
fof(f42311,plain,(
  ( ! [X171,X170] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))),X170,X171)) ) | $spl508),
  inference(resolution,[],[f42211,f723])).
fof(f723,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f478])).
fof(f47303,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))) | $spl533),
  inference(cnf_transformation,[],[f47303_D])).
fof(f47303_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))) <=> ~$spl533),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl533])])).
fof(f47288,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) ) | ($spl1 | $spl25)),
  inference(resolution,[],[f44775,f8825])).
fof(f44775,plain,(
  ( ! [X171,X170] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))),X170,X171)) ) | $spl1),
  inference(resolution,[],[f1098,f723])).
fof(f10782,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13)))) | $spl211),
  inference(cnf_transformation,[],[f10782_D])).
fof(f10782_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13)))) <=> ~$spl211),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl211])])).
fof(f46966,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | ($spl1 | $spl25)),
  inference(resolution,[],[f45789,f8825])).
fof(f45789,plain,(
  ( ! [X171,X170] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))),X170,X171)) ) | $spl1),
  inference(resolution,[],[f44703,f723])).
fof(f44703,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) | $spl1),
  inference(resolution,[],[f1098,f607])).
fof(f607,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X1,X0)))))) )),
  inference(cnf_transformation,[],[f420])).
fof(f420,plain,(
  ! [X0,X1] : ((p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK42(X1,X0))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK43(X1,X0))))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK42,sK43])],[f419])).
fof(f419,plain,(
  ! [X0,X1] : ((? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f418])).
fof(f418,plain,(
  ! [X0,X1] : ((? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f302])).
fof(f302,plain,(
  ! [X0,X1] : (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f167])).
fof(f167,plain,(
  ! [X0] : ! [X1] : (~? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6] : ! [X7] : (~? [X9] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) <=> ! [X9] : ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))))),
  file('thm.list.NULL_FILTER','thm.bool.NOT_EXISTS_THM')).
fof(f10780,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13)))) | $spl209),
  inference(cnf_transformation,[],[f10780_D])).
fof(f10780_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13)))) <=> ~$spl209),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl209])])).
fof(f46755,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | ($spl1 | $spl25)),
  inference(resolution,[],[f45564,f8825])).
fof(f45564,plain,(
  ( ! [X171,X170] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))),X170,X171)) ) | $spl1),
  inference(resolution,[],[f44702,f723])).
fof(f44702,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) | $spl1),
  inference(resolution,[],[f1098,f606])).
fof(f606,plain,(
  ( ! [X4,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK42(X1,X0)))))) )),
  inference(cnf_transformation,[],[f420])).
fof(f5259,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl130),
  inference(cnf_transformation,[],[f5259_D])).
fof(f5259_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f45825,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44703,f901])).
fof(f45732,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl1),
  inference(resolution,[],[f44703,f653])).
fof(f45729,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44703,f652])).
fof(f45728,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44703,f651])).
fof(f45727,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl1),
  inference(resolution,[],[f44703,f649])).
fof(f45725,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44703,f647])).
fof(f5025,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl120),
  inference(cnf_transformation,[],[f5025_D])).
fof(f5025_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f45600,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44702,f901])).
fof(f45507,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl1),
  inference(resolution,[],[f44702,f653])).
fof(f45504,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44702,f652])).
fof(f45503,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44702,f651])).
fof(f45502,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl1),
  inference(resolution,[],[f44702,f649])).
fof(f45500,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f44702,f647])).
fof(f45363,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | $spl531),
  inference(cnf_transformation,[],[f45363_D])).
fof(f45363_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) <=> ~$spl531),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl531])])).
fof(f45343,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | ~p(s(bool,X9))) ) | ($spl25 | $spl29)),
  inference(resolution,[],[f2356,f8903])).
fof(f2356,plain,(
  ( ! [X146,X147] : (sP8(X146,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))),X147)) ) | $spl29),
  inference(resolution,[],[f1266,f714])).
fof(f10783,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | $spl212),
  inference(cnf_transformation,[],[f10783_D])).
fof(f10783_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f45342,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | ~p(s(bool,X8))) ) | ($spl25 | $spl29)),
  inference(resolution,[],[f2356,f8824])).
fof(f14406,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X15),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl358),
  inference(cnf_transformation,[],[f14406_D])).
fof(f14406_D,plain,(
  ( ! [X15] : (s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X15),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f45025,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X15),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | ($spl25 | $spl452)),
  inference(resolution,[],[f43370,f8825])).
fof(f43370,plain,(
  ( ! [X229,X227,X228] : (sP6('const.list.EVERY_2'(s(fun(sK13,bool),X227),s('type.list.list'(sK13),'const.list.NIL_0')),X228,X229)) ) | $spl452),
  inference(resolution,[],[f27760,f723])).
fof(f27760,plain,(
  ( ! [X81] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X81),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl452),
  inference(forward_demodulation,[],[f27759,f27072])).
fof(f27759,plain,(
  ( ! [X81] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X81),s('type.list.list'(sK13),sK55(sK15,sK13)))))) ) | $spl452),
  inference(subsumption_resolution,[],[f27758,f869])).
fof(f869,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f348])).
fof(f348,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X10] : (p(s(bool,'F_0')) => p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.FALSITY')).
fof(f27758,plain,(
  ( ! [X81] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X81),s('type.list.list'(sK13),sK55(sK15,sK13)))))) ) | $spl452),
  inference(forward_demodulation,[],[f27555,f629])).
fof(f629,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))) )),
  inference(cnf_transformation,[],[f186])).
fof(f186,plain,(
  ! [X0,X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(flattening,[],[f185])).
fof(f185,plain,(
  ! [X0] : ! [X1] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.NIL_0')))))),
  inference(rectify,[],[f109])).
fof(f109,axiom,(
  ! [X6] : ! [X9] : s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.NIL_0')))))),
  file('thm.list.NULL_FILTER','thm.list.MEM|split|0')).
fof(f27555,plain,(
  ( ! [X81] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK51(sK55(sK15,sK13),X81,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X81),s('type.list.list'(sK13),sK55(sK15,sK13)))))) ) | $spl452),
  inference(superposition,[],[f625,f27072])).
fof(f625,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f439])).
fof(f439,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & ((p(s(bool,'const.bool.IN_2'(s(X0,sK51(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK51])],[f438])).
fof(f438,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(rectify,[],[f437])).
fof(f437,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f307])).
fof(f307,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f181])).
fof(f181,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X7,X11] : (p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X11)))) <=> ! [X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X11)))))) => p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X12))))))),
  file('thm.list.NULL_FILTER','thm.list.EVERY_MEM')).
fof(f44843,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) ) | $spl528),
  inference(cnf_transformation,[],[f44843_D])).
fof(f44843_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) ) <=> ~$spl528),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl528])])).
fof(f44811,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f1098,f901])).
fof(f44718,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl1),
  inference(resolution,[],[f1098,f653])).
fof(f44715,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f1098,f652])).
fof(f44714,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f1098,f651])).
fof(f44713,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl1),
  inference(resolution,[],[f1098,f649])).
fof(f44711,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) ) | $spl1),
  inference(resolution,[],[f1098,f647])).
fof(f13727,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl352),
  inference(cnf_transformation,[],[f13727_D])).
fof(f13727_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl352),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl352])])).
fof(f44502,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X15))) ) | ($spl25 | $spl452)),
  inference(resolution,[],[f43021,f8903])).
fof(f43021,plain,(
  ( ! [X198,X196,X197] : (sP8(X196,'const.list.EXISTS_2'(s(fun(sK13,bool),X197),s('type.list.list'(sK13),'const.list.NIL_0')),X198)) ) | $spl452),
  inference(resolution,[],[f27741,f714])).
fof(f27741,plain,(
  ( ! [X70] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X70),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl452),
  inference(forward_demodulation,[],[f27740,f27072])).
fof(f27740,plain,(
  ( ! [X70] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X70),s('type.list.list'(sK13),sK55(sK15,sK13)))))) ) | $spl452),
  inference(subsumption_resolution,[],[f27739,f867])).
fof(f27739,plain,(
  ( ! [X70] : (p(s(bool,'F_0')) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X70),s('type.list.list'(sK13),sK55(sK15,sK13)))))) ) | $spl452),
  inference(forward_demodulation,[],[f27548,f629])).
fof(f27548,plain,(
  ( ! [X70] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK50(sK55(sK15,sK13),X70,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X70),s('type.list.list'(sK13),sK55(sK15,sK13)))))) ) | $spl452),
  inference(superposition,[],[f621,f27072])).
fof(f621,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK50(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f436])).
fof(f13726,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X12),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl350),
  inference(cnf_transformation,[],[f13726_D])).
fof(f13726_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X12),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f44501,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X12),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X13))) ) | ($spl25 | $spl452)),
  inference(resolution,[],[f43021,f8824])).
fof(f2434,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl62),
  inference(cnf_transformation,[],[f2434_D])).
fof(f2434_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f44312,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f901])).
fof(f44214,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl29),
  inference(resolution,[],[f1266,f653])).
fof(f44213,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f652])).
fof(f44211,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f651])).
fof(f44209,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl29),
  inference(resolution,[],[f1266,f649])).
fof(f44207,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f647])).
fof(f44128,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | $spl527),
  inference(cnf_transformation,[],[f44128_D])).
fof(f44128_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) <=> ~$spl527),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl527])])).
fof(f44106,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | ~p(s(bool,X9))) ) | ($spl3 | $spl25)),
  inference(resolution,[],[f43778,f8903])).
fof(f43778,plain,(
  ( ! [X146,X147] : (sP8(X146,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)),X147)) ) | $spl3),
  inference(resolution,[],[f912,f714])).
fof(f44126,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | $spl524),
  inference(cnf_transformation,[],[f44126_D])).
fof(f44126_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) <=> ~$spl524),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl524])])).
fof(f44105,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | ~p(s(bool,X8))) ) | ($spl3 | $spl25)),
  inference(resolution,[],[f43778,f8824])).
fof(f43858,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl522),
  inference(cnf_transformation,[],[f43858_D])).
fof(f43858_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) <=> ~$spl522),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl522])])).
fof(f43826,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl3),
  inference(resolution,[],[f912,f901])).
fof(f43728,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | s(bool,X45) != s(bool,X47)) ) | $spl3),
  inference(resolution,[],[f912,f653])).
fof(f43727,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl3),
  inference(resolution,[],[f912,f652])).
fof(f43725,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl3),
  inference(resolution,[],[f912,f651])).
fof(f43723,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | s(bool,X30) != s(bool,X32)) ) | $spl3),
  inference(resolution,[],[f912,f649])).
fof(f43721,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl3),
  inference(resolution,[],[f912,f647])).
fof(f43700,plain,(
  $false | ($spl2 | $spl4 | $spl6)),
  inference(subsumption_resolution,[],[f43699,f877])).
fof(f877,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.NULL_FILTER','thm.bool.TRUTH')).
fof(f43699,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl4 | $spl6)),
  inference(forward_demodulation,[],[f43698,f2899])).
fof(f2899,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | $spl2),
  inference(resolution,[],[f913,f857])).
fof(f43698,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | ($spl4 | $spl6)),
  inference(subsumption_resolution,[],[f43660,f906])).
fof(f906,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f803])).
fof(f803,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f511])).
fof(f511,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f225])).
fof(f225,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f224])).
fof(f224,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6] : ! [X9] : (s(X6,X9) = s(X6,X9) <=> p(s(bool,'T_0')))),
  file('thm.list.NULL_FILTER','thm.bool.REFL_CLAUSE')).
fof(f43660,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | ($spl4 | $spl6)),
  inference(superposition,[],[f916,f3123])).
fof(f3123,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl4),
  inference(resolution,[],[f915,f857])).
fof(f1267,plain,(
  ( ! [X30,X32] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,X30),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl30),
  inference(cnf_transformation,[],[f1267_D])).
fof(f1267_D,plain,(
  ( ! [X30,X32] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | ~p(s(bool,X32)) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,X30),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f43632,plain,(
  ( ! [X313,X315,X314] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X313)))) | p(s(bool,X314)) | ~p(s(bool,X315)) | s(bool,X315) != s(bool,'const.bool.IN_2'(s(sK13,X313),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f901])).
fof(f43534,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X58)))) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.bool.IN_2'(s(sK13,X58),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X59) != s(bool,X61)) ) | $spl6),
  inference(resolution,[],[f916,f653])).
fof(f43533,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X54)))) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,'const.bool.IN_2'(s(sK13,X54),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f652])).
fof(f43531,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X46)))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,X46),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f651])).
fof(f43529,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X38)))) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,X38),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X39) != s(bool,X41)) ) | $spl6),
  inference(resolution,[],[f916,f649])).
fof(f43527,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,X30),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f647])).
fof(f13168,plain,(
  ( ! [X31,X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X31),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl322),
  inference(cnf_transformation,[],[f13168_D])).
fof(f13168_D,plain,(
  ( ! [X31,X29] : (p(s(bool,X29)) | s(bool,X29) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X31),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f43406,plain,(
  ( ! [X290,X289,X291] : (p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X290) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X291),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27760,f901])).
fof(f43313,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X57),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X56) != s(bool,X58)) ) | $spl452),
  inference(resolution,[],[f27760,f653])).
fof(f43310,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X47),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27760,f652])).
fof(f43309,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X43),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27760,f651])).
fof(f43308,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X38),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X37) != s(bool,X39)) ) | $spl452),
  inference(resolution,[],[f27760,f649])).
fof(f43306,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X31),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27760,f647])).
fof(f1076,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl8),
  inference(cnf_transformation,[],[f1076_D])).
fof(f1076_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f43230,plain,(
  ( ! [X192,X193] : (p(s(bool,X192)) | ~p(s(bool,X193)) | s(bool,X193) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f901])).
fof(f43132,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X35) != s(bool,X37)) ) | $spl1),
  inference(resolution,[],[f911,f653])).
fof(f43131,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f652])).
fof(f43129,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f651])).
fof(f43127,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X20) != s(bool,X22)) ) | $spl1),
  inference(resolution,[],[f911,f649])).
fof(f43125,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f647])).
fof(f13007,plain,(
  ( ! [X21,X19] : (~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X21),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl320),
  inference(cnf_transformation,[],[f13007_D])).
fof(f13007_D,plain,(
  ( ! [X21,X19] : (~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X21),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f43069,plain,(
  ( ! [X292,X290,X291] : (p(s(bool,X290)) | ~p(s(bool,X291)) | s(bool,X291) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X292),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27741,f901])).
fof(f42971,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X47),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X45) != s(bool,X48)) ) | $spl452),
  inference(resolution,[],[f27741,f653])).
fof(f42970,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X44),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27741,f652])).
fof(f42968,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X37),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27741,f651])).
fof(f42966,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X28),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X26) != s(bool,X29)) ) | $spl452),
  inference(resolution,[],[f27741,f649])).
fof(f42964,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X21),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl452),
  inference(resolution,[],[f27741,f647])).
fof(f42842,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | $spl518),
  inference(cnf_transformation,[],[f42842_D])).
fof(f42842_D,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) <=> ~$spl518),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl518])])).
fof(f42843,plain,(
  ( ! [X23] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X23)))))) | s('type.list.list'(sK13),X23) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(X23,sK13)),s('type.list.list'(sK13),sK55(X23,sK13))))) ) | $spl520),
  inference(cnf_transformation,[],[f42843_D])).
fof(f42843_D,plain,(
  ( ! [X23] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X23)))))) | s('type.list.list'(sK13),X23) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(X23,sK13)),s('type.list.list'(sK13),sK55(X23,sK13))))) ) <=> ~$spl520),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl520])])).
fof(f39598,plain,(
  ( ! [X23] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X23)))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | s('type.list.list'(sK13),X23) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(X23,sK13)),s('type.list.list'(sK13),sK55(X23,sK13))))) ) | $spl2),
  inference(superposition,[],[f2795,f900])).
fof(f2795,plain,(
  ( ! [X15] : (~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X15)))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X15))))) ) | $spl2),
  inference(resolution,[],[f913,f623])).
fof(f623,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.list.EXISTS_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f436])).
fof(f42835,plain,(
  $false | ($spl0 | $spl2 | $spl4)),
  inference(subsumption_resolution,[],[f42834,f877])).
fof(f42834,plain,(
  ~p(s(bool,'T_0')) | ($spl0 | $spl2 | $spl4)),
  inference(forward_demodulation,[],[f42827,f2899])).
fof(f42827,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | ($spl0 | $spl4)),
  inference(trivial_inequality_removal,[],[f42807])).
fof(f42807,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),'const.list.NIL_0') | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | ($spl0 | $spl4)),
  inference(superposition,[],[f3013,f3220])).
fof(f3220,plain,(
  s('type.list.list'(sK13),'const.list.NIL_0') = s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl0),
  inference(resolution,[],[f910,f882])).
fof(f3013,plain,(
  ( ! [X2] : (s('type.list.list'(sK13),'const.list.NIL_0') != s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),X2),s('type.list.list'(sK13),sK15))) | ~p(s(bool,happ(s(fun(sK13,bool),X2),s(sK13,sK16))))) ) | $spl4),
  inference(resolution,[],[f915,f595])).
fof(f595,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2))))))) )),
  inference(cnf_transformation,[],[f407])).
fof(f42586,plain,(
  ( ! [X61,X63] : (~p(s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK40(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK41(sK14,X61,sK13)))) | p(s(bool,X63))) ) | $spl516),
  inference(cnf_transformation,[],[f42586_D])).
fof(f42586_D,plain,(
  ( ! [X61,X63] : (~p(s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK40(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK41(sK14,X61,sK13)))) | p(s(bool,X63))) ) <=> ~$spl516),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl516])])).
fof(f42548,plain,(
  ( ! [X333,X331,X332] : (~p(s(bool,happ(s(fun(sK13,bool),X331),s(sK13,sK40(sK14,X331,sK13))))) | p(s(bool,X332)) | p(s(bool,X333)) | s(bool,X333) != s(bool,happ(s(fun(sK13,bool),X331),s(sK13,sK41(sK14,X331,sK13))))) ) | $spl2),
  inference(resolution,[],[f2792,f901])).
fof(f2792,plain,(
  ( ! [X14] : (p(s(bool,happ(s(fun(sK13,bool),X14),s(sK13,sK41(sK14,X14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X14),s(sK13,sK40(sK14,X14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f913,f605])).
fof(f605,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK41(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK40(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f417])).
fof(f417,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK40(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK40(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK41(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK40,sK41])],[f416])).
fof(f416,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f301])).
fof(f301,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f300])).
fof(f300,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6] : ! [X25,X7] : (! [X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) => p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))) => (? [X9] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) => ? [X9] : p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))))),
  file('thm.list.NULL_FILTER','thm.bool.MONO_EXISTS')).
fof(f42455,plain,(
  ( ! [X92,X90,X91,X89] : (~p(s(bool,happ(s(fun(sK13,bool),X89),s(sK13,sK40(sK14,X89,sK13))))) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,happ(s(fun(sK13,bool),X89),s(sK13,sK41(sK14,X89,sK13)))) | s(bool,X91) != s(bool,X92)) ) | $spl2),
  inference(resolution,[],[f2792,f653])).
fof(f42452,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,happ(s(fun(sK13,bool),X77),s(sK13,sK40(sK14,X77,sK13))))) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,happ(s(fun(sK13,bool),X77),s(sK13,sK41(sK14,X77,sK13))))) ) | $spl2),
  inference(resolution,[],[f2792,f652])).
fof(f42451,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,happ(s(fun(sK13,bool),X73),s(sK13,sK40(sK14,X73,sK13))))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X75) != s(bool,happ(s(fun(sK13,bool),X73),s(sK13,sK41(sK14,X73,sK13))))) ) | $spl2),
  inference(resolution,[],[f2792,f651])).
fof(f42450,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,happ(s(fun(sK13,bool),X69),s(sK13,sK40(sK14,X69,sK13))))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,happ(s(fun(sK13,bool),X69),s(sK13,sK41(sK14,X69,sK13)))) | s(bool,X71) != s(bool,X72)) ) | $spl2),
  inference(resolution,[],[f2792,f649])).
fof(f42448,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK40(sK14,X61,sK13))))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK41(sK14,X61,sK13))))) ) | $spl2),
  inference(resolution,[],[f2792,f647])).
fof(f42381,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) ) | $spl514),
  inference(cnf_transformation,[],[f42381_D])).
fof(f42381_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) ) <=> ~$spl514),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl514])])).
fof(f42347,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) ) | $spl508),
  inference(resolution,[],[f42211,f901])).
fof(f42254,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl508),
  inference(resolution,[],[f42211,f653])).
fof(f42251,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) ) | $spl508),
  inference(resolution,[],[f42211,f652])).
fof(f42250,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) ) | $spl508),
  inference(resolution,[],[f42211,f651])).
fof(f42249,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl508),
  inference(resolution,[],[f42211,f649])).
fof(f42247,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) ) | $spl508),
  inference(resolution,[],[f42211,f647])).
fof(f42215,plain,(
  ( ! [X61,X59] : (p(s(bool,happ(s(fun(sK13,bool),X59),s(sK13,sK41(sK14,X59,sK13))))) | p(s(bool,X61)) | s(bool,X61) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X59,sK13))))) ) | $spl512),
  inference(cnf_transformation,[],[f42215_D])).
fof(f42215_D,plain,(
  ( ! [X61,X59] : (p(s(bool,happ(s(fun(sK13,bool),X59),s(sK13,sK41(sK14,X59,sK13))))) | p(s(bool,X61)) | s(bool,X61) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X59,sK13))))) ) <=> ~$spl512),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl512])])).
fof(f42138,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,happ(s(fun(sK13,bool),X329),s(sK13,sK41(sK14,X329,sK13))))) | p(s(bool,X330)) | p(s(bool,X331)) | s(bool,X331) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X329,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f901])).
fof(f2791,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK13,bool),X13),s(sK13,sK41(sK14,X13,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X13,sK13)))))) ) | $spl2),
  inference(resolution,[],[f913,f604])).
fof(f604,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK41(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK40(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f417])).
fof(f42045,plain,(
  ( ! [X90,X88,X87,X89] : (p(s(bool,happ(s(fun(sK13,bool),X87),s(sK13,sK41(sK14,X87,sK13))))) | p(s(bool,X88)) | ~p(s(bool,X89)) | s(bool,X88) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X87,sK13)))) | s(bool,X89) != s(bool,X90)) ) | $spl2),
  inference(resolution,[],[f2791,f653])).
fof(f42042,plain,(
  ( ! [X78,X76,X77,X75] : (p(s(bool,happ(s(fun(sK13,bool),X75),s(sK13,sK41(sK14,X75,sK13))))) | p(s(bool,X76)) | ~p(s(bool,X77)) | p(s(bool,X78)) | s(bool,X78) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X75,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f652])).
fof(f42041,plain,(
  ( ! [X74,X72,X71,X73] : (p(s(bool,happ(s(fun(sK13,bool),X71),s(sK13,sK41(sK14,X71,sK13))))) | p(s(bool,X72)) | p(s(bool,X73)) | s(bool,X72) != s(bool,X74) | s(bool,X73) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X71,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f651])).
fof(f42040,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,happ(s(fun(sK13,bool),X67),s(sK13,sK41(sK14,X67,sK13))))) | p(s(bool,X68)) | ~p(s(bool,X69)) | s(bool,X68) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X67,sK13)))) | s(bool,X69) != s(bool,X70)) ) | $spl2),
  inference(resolution,[],[f2791,f649])).
fof(f42038,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,happ(s(fun(sK13,bool),X59),s(sK13,sK41(sK14,X59,sK13))))) | p(s(bool,X60)) | p(s(bool,X61)) | s(bool,X60) != s(bool,X62) | s(bool,X61) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X59,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f647])).
fof(f13388,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl344),
  inference(cnf_transformation,[],[f13388_D])).
fof(f13388_D,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) <=> ~$spl344),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl344])])).
fof(f42212,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13)))))) ) | $spl510),
  inference(cnf_transformation,[],[f42212_D])).
fof(f42212_D,plain,(
  ( ! [X29] : (p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13)))))) ) <=> ~$spl510),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl510])])).
fof(f42028,plain,(
  ( ! [X30,X29] : (p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X30),s(sK13,sK41(sK14,X30,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X30),s(sK13,sK40(sK14,X30,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2791,f605])).
fof(f42209,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X0))))) ) | $spl506),
  inference(cnf_transformation,[],[f42209_D])).
fof(f42209_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X0))))) ) <=> ~$spl506),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl506])])).
fof(f42211,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) | $spl508),
  inference(cnf_transformation,[],[f42211_D])).
fof(f42211_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) <=> ~$spl508),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl508])])).
fof(f42204,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X0))))) ) | $spl2),
  inference(duplicate_literal_removal,[],[f42013])).
fof(f42013,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK41(sK14,sK14,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X0))))) ) | $spl2),
  inference(resolution,[],[f2791,f605])).
fof(f42205,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK41(sK14,X61,sK13)))) | p(s(bool,X63))) ) | $spl504),
  inference(cnf_transformation,[],[f42205_D])).
fof(f42205_D,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK41(sK14,X61,sK13)))) | p(s(bool,X63))) ) <=> ~$spl504),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl504])])).
fof(f42009,plain,(
  ( ! [X333,X331,X332] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X331,sK13))))) | p(s(bool,X332)) | p(s(bool,X333)) | s(bool,X333) != s(bool,happ(s(fun(sK13,bool),X331),s(sK13,sK41(sK14,X331,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f901])).
fof(f41916,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X89,sK13))))) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,happ(s(fun(sK13,bool),X89),s(sK13,sK41(sK14,X89,sK13)))) | s(bool,X91) != s(bool,X92)) ) | $spl2),
  inference(resolution,[],[f2791,f653])).
fof(f41913,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X77,sK13))))) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,happ(s(fun(sK13,bool),X77),s(sK13,sK41(sK14,X77,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f652])).
fof(f41912,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X73,sK13))))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X75) != s(bool,happ(s(fun(sK13,bool),X73),s(sK13,sK41(sK14,X73,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f651])).
fof(f41911,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X69,sK13))))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,happ(s(fun(sK13,bool),X69),s(sK13,sK41(sK14,X69,sK13)))) | s(bool,X71) != s(bool,X72)) ) | $spl2),
  inference(resolution,[],[f2791,f649])).
fof(f41909,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X61,sK13))))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK41(sK14,X61,sK13))))) ) | $spl2),
  inference(resolution,[],[f2791,f647])).
fof(f41808,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39686,f901])).
fof(f39686,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f622])).
fof(f39660,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | ($spl2 | $spl4 | $spl452)),
  inference(forward_demodulation,[],[f39659,f27495])).
fof(f27495,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | ($spl4 | $spl452)),
  inference(backward_demodulation,[],[f27072,f3168])).
fof(f3168,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl4),
  inference(subsumption_resolution,[],[f3167,f871])).
fof(f3167,plain,(
  p(s(bool,'F_0')) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl4),
  inference(forward_demodulation,[],[f3134,f629])).
fof(f3134,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl4),
  inference(superposition,[],[f915,f900])).
fof(f39659,plain,(
  p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0')))))) | ($spl2 | $spl4 | $spl452)),
  inference(subsumption_resolution,[],[f39658,f906])).
fof(f39658,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0')))))) | ($spl2 | $spl4 | $spl452)),
  inference(forward_demodulation,[],[f39586,f3123])).
fof(f39586,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0')))))) | ($spl2 | $spl4 | $spl452)),
  inference(superposition,[],[f2795,f27495])).
fof(f41715,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39686,f653])).
fof(f41712,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39686,f652])).
fof(f41711,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39686,f651])).
fof(f41710,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39686,f649])).
fof(f41708,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK50(sK15,sK14,sK13))))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39686,f647])).
fof(f41676,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK23(sK14,X61,sK13)))) | p(s(bool,X63))) ) | $spl500),
  inference(cnf_transformation,[],[f41676_D])).
fof(f41676_D,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK23(sK14,X61,sK13)))) | p(s(bool,X63))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f41602,plain,(
  ( ! [X333,X331,X332] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X331,sK13))))) | p(s(bool,X332)) | p(s(bool,X333)) | s(bool,X333) != s(bool,happ(s(fun(sK13,bool),X331),s(sK13,sK23(sK14,X331,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f901])).
fof(f2782,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X4,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X4),s(sK13,sK23(sK14,X4,sK13)))))) ) | $spl2),
  inference(resolution,[],[f913,f570])).
fof(f570,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f382,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK21,sK22,sK23])],[f381])).
fof(f381,plain,(
  ! [X0,X1,X2] : ((! [X8] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) | ? [X7] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X7)))) | ? [X6] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f380])).
fof(f380,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f379])).
fof(f379,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((! [X5] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f67])).
fof(f67,axiom,(
  ! [X6] : ! [X7,X25] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) | p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))))),
  file('thm.list.NULL_FILTER','thm.bool.EXISTS_OR_THM')).
fof(f41509,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X89,sK13))))) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,happ(s(fun(sK13,bool),X89),s(sK13,sK23(sK14,X89,sK13)))) | s(bool,X91) != s(bool,X92)) ) | $spl2),
  inference(resolution,[],[f2782,f653])).
fof(f41506,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X77,sK13))))) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,happ(s(fun(sK13,bool),X77),s(sK13,sK23(sK14,X77,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f652])).
fof(f41505,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X73,sK13))))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X75) != s(bool,happ(s(fun(sK13,bool),X73),s(sK13,sK23(sK14,X73,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f651])).
fof(f41504,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X69,sK13))))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,happ(s(fun(sK13,bool),X69),s(sK13,sK23(sK14,X69,sK13)))) | s(bool,X71) != s(bool,X72)) ) | $spl2),
  inference(resolution,[],[f2782,f649])).
fof(f41502,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X61,sK13))))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK23(sK14,X61,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f647])).
fof(f41672,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK23(sK14,X58,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X58,sK13))))) ) | $spl498),
  inference(cnf_transformation,[],[f41672_D])).
fof(f41672_D,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK23(sK14,X58,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X58,sK13))))) ) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f41473,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,happ(s(fun(sK13,bool),X328),s(sK13,sK23(sK14,X328,sK13))))) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X328,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f901])).
fof(f41380,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK23(sK14,X86,sK13))))) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X86,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2782,f653])).
fof(f41377,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK23(sK14,X74,sK13))))) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X74,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f652])).
fof(f41376,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK23(sK14,X70,sK13))))) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X70,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f651])).
fof(f41375,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,happ(s(fun(sK13,bool),X66),s(sK13,sK23(sK14,X66,sK13))))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X66,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2782,f649])).
fof(f41373,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK23(sK14,X58,sK13))))) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X58,sK13))))) ) | $spl2),
  inference(resolution,[],[f2782,f647])).
fof(f41669,plain,(
  ( ! [X26] : (p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK23(sK14,X26,sK13)))))) ) | $spl496),
  inference(cnf_transformation,[],[f41669_D])).
fof(f41669_D,plain,(
  ( ! [X26] : (p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK23(sK14,X26,sK13)))))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f41363,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(sK13,bool),X28),s(sK13,sK23(sK14,X28,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2782,f605])).
fof(f13387,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl342),
  inference(cnf_transformation,[],[f13387_D])).
fof(f13387_D,plain,(
  ( ! [X27] : (p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f41362,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK23(sK14,X26,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2782,f604])).
fof(f41307,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13))))) ) | $spl494),
  inference(cnf_transformation,[],[f41307_D])).
fof(f41307_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13))))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f41273,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f41039,f901])).
fof(f41039,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13))))) | $spl2),
  inference(factoring,[],[f2781])).
fof(f2781,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK23(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X3,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f913,f569])).
fof(f569,plain,(
  ( ! [X2,X0,X5,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK23(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK23(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f41180,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl2),
  inference(resolution,[],[f41039,f653])).
fof(f41177,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f41039,f652])).
fof(f41176,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f41039,f651])).
fof(f41175,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl2),
  inference(resolution,[],[f41039,f649])).
fof(f41173,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f41039,f647])).
fof(f41109,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK23(X58,sK14,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X58,sK14,sK13))))) ) | $spl492),
  inference(cnf_transformation,[],[f41109_D])).
fof(f41109_D,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK23(X58,sK14,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X58,sK14,sK13))))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f41035,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,happ(s(fun(sK13,bool),X328),s(sK13,sK23(X328,sK14,sK13))))) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X328,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f901])).
fof(f40942,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK23(X86,sK14,sK13))))) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X86,sK14,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2781,f653])).
fof(f40939,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK23(X74,sK14,sK13))))) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X74,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f652])).
fof(f40938,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK23(X70,sK14,sK13))))) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X70,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f651])).
fof(f40937,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,happ(s(fun(sK13,bool),X66),s(sK13,sK23(X66,sK14,sK13))))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X66,sK14,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2781,f649])).
fof(f40935,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK23(X58,sK14,sK13))))) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X58,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f647])).
fof(f41106,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK23(X8,sK14,sK13)))))) ) | $spl490),
  inference(cnf_transformation,[],[f41106_D])).
fof(f41106_D,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK23(X8,sK14,sK13)))))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f40925,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(sK13,bool),X28),s(sK13,sK23(X28,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2781,f605])).
fof(f40924,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK23(X26,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2781,f604])).
fof(f13381,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl330),
  inference(cnf_transformation,[],[f13381_D])).
fof(f13381_D,plain,(
  ( ! [X9] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f40915,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK23(X8,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2781,f570])).
fof(f41102,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X61,sK14,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK23(X61,sK14,sK13)))) | p(s(bool,X63))) ) | $spl488),
  inference(cnf_transformation,[],[f41102_D])).
fof(f41102_D,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X61,sK14,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK23(X61,sK14,sK13)))) | p(s(bool,X63))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f40907,plain,(
  ( ! [X333,X331,X332] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X331,sK14,sK13))))) | p(s(bool,X332)) | p(s(bool,X333)) | s(bool,X333) != s(bool,happ(s(fun(sK13,bool),X331),s(sK13,sK23(X331,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f901])).
fof(f40814,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X89,sK14,sK13))))) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,happ(s(fun(sK13,bool),X89),s(sK13,sK23(X89,sK14,sK13)))) | s(bool,X91) != s(bool,X92)) ) | $spl2),
  inference(resolution,[],[f2781,f653])).
fof(f40811,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X77,sK14,sK13))))) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,happ(s(fun(sK13,bool),X77),s(sK13,sK23(X77,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f652])).
fof(f40810,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X73,sK14,sK13))))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X75) != s(bool,happ(s(fun(sK13,bool),X73),s(sK13,sK23(X73,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f651])).
fof(f40809,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X69,sK14,sK13))))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,happ(s(fun(sK13,bool),X69),s(sK13,sK23(X69,sK14,sK13)))) | s(bool,X71) != s(bool,X72)) ) | $spl2),
  inference(resolution,[],[f2781,f649])).
fof(f40807,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X61,sK14,sK13))))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK23(X61,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2781,f647])).
fof(f40779,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl487),
  inference(cnf_transformation,[],[f40779_D])).
fof(f40779_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl487),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl487])])).
fof(f40760,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | ($spl2 | $spl4 | $spl25 | $spl452)),
  inference(resolution,[],[f39757,f8825])).
fof(f39757,plain,(
  ( ! [X154,X155] : (sP6('const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X154,X155)) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f723])).
fof(f40652,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK21(sK14,X61,sK13)))) | p(s(bool,X63))) ) | $spl484),
  inference(cnf_transformation,[],[f40652_D])).
fof(f40652_D,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X61,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK21(sK14,X61,sK13)))) | p(s(bool,X63))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f40579,plain,(
  ( ! [X333,X331,X332] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X331,sK13))))) | p(s(bool,X332)) | p(s(bool,X333)) | s(bool,X333) != s(bool,happ(s(fun(sK13,bool),X331),s(sK13,sK21(sK14,X331,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f901])).
fof(f2780,plain,(
  ( ! [X2] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X2,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X2),s(sK13,sK21(sK14,X2,sK13)))))) ) | $spl2),
  inference(resolution,[],[f913,f568])).
fof(f568,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f40486,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X89,sK13))))) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,happ(s(fun(sK13,bool),X89),s(sK13,sK21(sK14,X89,sK13)))) | s(bool,X91) != s(bool,X92)) ) | $spl2),
  inference(resolution,[],[f2780,f653])).
fof(f40483,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X77,sK13))))) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,happ(s(fun(sK13,bool),X77),s(sK13,sK21(sK14,X77,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f652])).
fof(f40482,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X73,sK13))))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X75) != s(bool,happ(s(fun(sK13,bool),X73),s(sK13,sK21(sK14,X73,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f651])).
fof(f40481,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X69,sK13))))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,happ(s(fun(sK13,bool),X69),s(sK13,sK21(sK14,X69,sK13)))) | s(bool,X71) != s(bool,X72)) ) | $spl2),
  inference(resolution,[],[f2780,f649])).
fof(f40479,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X61,sK13))))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK21(sK14,X61,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f647])).
fof(f40648,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK21(sK14,X58,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X58,sK13))))) ) | $spl482),
  inference(cnf_transformation,[],[f40648_D])).
fof(f40648_D,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK21(sK14,X58,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X58,sK13))))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f40450,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,happ(s(fun(sK13,bool),X328),s(sK13,sK21(sK14,X328,sK13))))) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X328,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f901])).
fof(f40357,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK21(sK14,X86,sK13))))) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X86,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2780,f653])).
fof(f40354,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK21(sK14,X74,sK13))))) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X74,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f652])).
fof(f40353,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK21(sK14,X70,sK13))))) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X70,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f651])).
fof(f40352,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,happ(s(fun(sK13,bool),X66),s(sK13,sK21(sK14,X66,sK13))))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X66,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2780,f649])).
fof(f40350,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK21(sK14,X58,sK13))))) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X58,sK13))))) ) | $spl2),
  inference(resolution,[],[f2780,f647])).
fof(f40645,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(sK13,bool),X6),s(sK13,sK21(sK14,X6,sK13)))))) ) | $spl480),
  inference(cnf_transformation,[],[f40645_D])).
fof(f40645_D,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(sK13,bool),X6),s(sK13,sK21(sK14,X6,sK13)))))) ) <=> ~$spl480),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl480])])).
fof(f40340,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(sK13,bool),X28),s(sK13,sK21(sK14,X28,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2780,f605])).
fof(f40339,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK21(sK14,X26,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2780,f604])).
fof(f40330,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK21(sK14,X8,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2780,f570])).
fof(f13380,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl328),
  inference(cnf_transformation,[],[f13380_D])).
fof(f13380_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f40329,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(sK13,bool),X6),s(sK13,sK21(sK14,X6,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2780,f569])).
fof(f40213,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK22(X58,sK14,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X58,sK14,sK13))))) ) | $spl478),
  inference(cnf_transformation,[],[f40213_D])).
fof(f40213_D,plain,(
  ( ! [X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK22(X58,sK14,sK13))))) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X58,sK14,sK13))))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f40140,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,happ(s(fun(sK13,bool),X328),s(sK13,sK22(X328,sK14,sK13))))) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X328,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f901])).
fof(f2779,plain,(
  ( ! [X1] : (p(s(bool,happ(s(fun(sK13,bool),X1),s(sK13,sK22(X1,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X1,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f913,f567])).
fof(f567,plain,(
  ( ! [X2,X0,X8,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK22(X2,X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK21(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X8))))) )),
  inference(cnf_transformation,[],[f382])).
fof(f40047,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,happ(s(fun(sK13,bool),X86),s(sK13,sK22(X86,sK14,sK13))))) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X86,sK14,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2779,f653])).
fof(f40044,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,happ(s(fun(sK13,bool),X74),s(sK13,sK22(X74,sK14,sK13))))) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X74,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f652])).
fof(f40043,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,happ(s(fun(sK13,bool),X70),s(sK13,sK22(X70,sK14,sK13))))) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X70,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f651])).
fof(f40042,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,happ(s(fun(sK13,bool),X66),s(sK13,sK22(X66,sK14,sK13))))) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X66,sK14,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2779,f649])).
fof(f40040,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,happ(s(fun(sK13,bool),X58),s(sK13,sK22(X58,sK14,sK13))))) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X58,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f647])).
fof(f40210,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),X4),s(sK13,sK22(X4,sK14,sK13)))))) ) | $spl476),
  inference(cnf_transformation,[],[f40210_D])).
fof(f40210_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK13,bool),X4),s(sK13,sK22(X4,sK14,sK13)))))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f40030,plain,(
  ( ! [X28,X29] : (p(s(bool,happ(s(fun(sK13,bool),X28),s(sK13,sK22(X28,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2779,f605])).
fof(f40029,plain,(
  ( ! [X26,X27] : (p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK22(X26,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2779,f604])).
fof(f40020,plain,(
  ( ! [X8,X9] : (p(s(bool,happ(s(fun(sK13,bool),X8),s(sK13,sK22(X8,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2779,f570])).
fof(f40019,plain,(
  ( ! [X6,X7] : (p(s(bool,happ(s(fun(sK13,bool),X6),s(sK13,sK22(X6,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2779,f569])).
fof(f13379,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl326),
  inference(cnf_transformation,[],[f13379_D])).
fof(f13379_D,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f40018,plain,(
  ( ! [X4,X5] : (p(s(bool,happ(s(fun(sK13,bool),X4),s(sK13,sK22(X4,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2779,f568])).
fof(f40206,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X61,sK14,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK22(X61,sK14,sK13)))) | p(s(bool,X63))) ) | $spl474),
  inference(cnf_transformation,[],[f40206_D])).
fof(f40206_D,plain,(
  ( ! [X61,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X61,sK14,sK13))))) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK22(X61,sK14,sK13)))) | p(s(bool,X63))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f40012,plain,(
  ( ! [X333,X331,X332] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X331,sK14,sK13))))) | p(s(bool,X332)) | p(s(bool,X333)) | s(bool,X333) != s(bool,happ(s(fun(sK13,bool),X331),s(sK13,sK22(X331,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f901])).
fof(f39919,plain,(
  ( ! [X92,X90,X91,X89] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X89,sK14,sK13))))) | p(s(bool,X90)) | ~p(s(bool,X91)) | s(bool,X90) != s(bool,happ(s(fun(sK13,bool),X89),s(sK13,sK22(X89,sK14,sK13)))) | s(bool,X91) != s(bool,X92)) ) | $spl2),
  inference(resolution,[],[f2779,f653])).
fof(f39916,plain,(
  ( ! [X80,X78,X79,X77] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X77,sK14,sK13))))) | p(s(bool,X78)) | ~p(s(bool,X79)) | p(s(bool,X80)) | s(bool,X80) != s(bool,happ(s(fun(sK13,bool),X77),s(sK13,sK22(X77,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f652])).
fof(f39915,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X73,sK14,sK13))))) | p(s(bool,X74)) | p(s(bool,X75)) | s(bool,X74) != s(bool,X76) | s(bool,X75) != s(bool,happ(s(fun(sK13,bool),X73),s(sK13,sK22(X73,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f651])).
fof(f39914,plain,(
  ( ! [X70,X72,X71,X69] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X69,sK14,sK13))))) | p(s(bool,X70)) | ~p(s(bool,X71)) | s(bool,X70) != s(bool,happ(s(fun(sK13,bool),X69),s(sK13,sK22(X69,sK14,sK13)))) | s(bool,X71) != s(bool,X72)) ) | $spl2),
  inference(resolution,[],[f2779,f649])).
fof(f39912,plain,(
  ( ! [X61,X64,X62,X63] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X61,sK14,sK13))))) | p(s(bool,X62)) | p(s(bool,X63)) | s(bool,X62) != s(bool,X64) | s(bool,X63) != s(bool,happ(s(fun(sK13,bool),X61),s(sK13,sK22(X61,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2779,f647])).
fof(f39793,plain,(
  ( ! [X189,X188] : (p(s(bool,X188)) | p(s(bool,X189)) | s(bool,X189) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f901])).
fof(f39700,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X40) != s(bool,X41)) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f653])).
fof(f39697,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | p(s(bool,X32)) | s(bool,X32) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f652])).
fof(f39696,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f651])).
fof(f39695,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X25) != s(bool,X26)) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f649])).
fof(f39693,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | ($spl2 | $spl4 | $spl452)),
  inference(resolution,[],[f39660,f647])).
fof(f39644,plain,(
  ( ! [X33,X31] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X31)))) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X31)))))) ) | $spl470),
  inference(cnf_transformation,[],[f39644_D])).
fof(f39644_D,plain,(
  ( ! [X33,X31] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X31)))) | ~p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X31)))))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f39578,plain,(
  ( ! [X315,X316,X314] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X314)))) | p(s(bool,X315)) | ~p(s(bool,X316)) | s(bool,X316) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X314)))))) ) | $spl2),
  inference(resolution,[],[f2795,f901])).
fof(f39480,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X59)))) | ~p(s(bool,X60)) | ~p(s(bool,X61)) | s(bool,X61) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X59))))) | s(bool,X60) != s(bool,X62)) ) | $spl2),
  inference(resolution,[],[f2795,f653])).
fof(f39479,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X55)))) | p(s(bool,X56)) | ~p(s(bool,X57)) | ~p(s(bool,X58)) | s(bool,X57) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X55)))))) ) | $spl2),
  inference(resolution,[],[f2795,f652])).
fof(f39477,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X47)))) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X50) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X47)))))) ) | $spl2),
  inference(resolution,[],[f2795,f651])).
fof(f39475,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X39)))) | ~p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X39))))) | s(bool,X40) != s(bool,X42)) ) | $spl2),
  inference(resolution,[],[f2795,f649])).
fof(f39473,plain,(
  ( ! [X33,X31,X34,X32] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),X31)))) | p(s(bool,X32)) | ~p(s(bool,X33)) | s(bool,X32) != s(bool,X34) | s(bool,X33) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),X31)))))) ) | $spl2),
  inference(resolution,[],[f2795,f647])).
fof(f37439,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl469),
  inference(cnf_transformation,[],[f37439_D])).
fof(f37439_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl469),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl469])])).
fof(f37416,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl25),
  inference(resolution,[],[f29078,f8825])).
fof(f29078,plain,(
  ( ! [X167,X166] : (sP6('const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X166,X167)) ) | $spl25),
  inference(resolution,[],[f2589,f723])).
fof(f2589,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl25),
  inference(resolution,[],[f1095,f627])).
fof(f627,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK52(X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X1)))))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),X1))))) )),
  inference(cnf_transformation,[],[f442])).
fof(f36551,plain,(
  ( ! [X22,X20] : (p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK16)))) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X20),s('type.list.list'(sK13),sK15)))) ) | $spl466),
  inference(cnf_transformation,[],[f36551_D])).
fof(f36551_D,plain,(
  ( ! [X22,X20] : (p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK16)))) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X20),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f36512,plain,(
  ( ! [X304,X303,X305] : (p(s(bool,happ(s(fun(sK13,bool),X303),s(sK13,sK16)))) | p(s(bool,X304)) | ~p(s(bool,X305)) | s(bool,X305) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X303),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3018,f901])).
fof(f3018,plain,(
  ( ! [X11] : (~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X11),s('type.list.list'(sK13),sK15)))) | p(s(bool,happ(s(fun(sK13,bool),X11),s(sK13,sK16))))) ) | $spl4),
  inference(resolution,[],[f915,f624])).
fof(f624,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f439])).
fof(f36414,plain,(
  ( ! [X50,X48,X51,X49] : (p(s(bool,happ(s(fun(sK13,bool),X48),s(sK13,sK16)))) | ~p(s(bool,X49)) | ~p(s(bool,X50)) | s(bool,X50) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X48),s('type.list.list'(sK13),sK15))) | s(bool,X49) != s(bool,X51)) ) | $spl4),
  inference(resolution,[],[f3018,f653])).
fof(f36413,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,happ(s(fun(sK13,bool),X44),s(sK13,sK16)))) | p(s(bool,X45)) | ~p(s(bool,X46)) | ~p(s(bool,X47)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X44),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3018,f652])).
fof(f36411,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,happ(s(fun(sK13,bool),X36),s(sK13,sK16)))) | p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X39) | s(bool,X38) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X36),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3018,f651])).
fof(f36409,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,happ(s(fun(sK13,bool),X28),s(sK13,sK16)))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X28),s('type.list.list'(sK13),sK15))) | s(bool,X29) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f3018,f649])).
fof(f36407,plain,(
  ( ! [X23,X21,X22,X20] : (p(s(bool,happ(s(fun(sK13,bool),X20),s(sK13,sK16)))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,X23) | s(bool,X22) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X20),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3018,f647])).
fof(f36274,plain,(
  ( ! [X28,X26] : (~p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK16)))) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X26),s('type.list.list'(sK13),sK15)))) ) | $spl464),
  inference(cnf_transformation,[],[f36274_D])).
fof(f36274_D,plain,(
  ( ! [X28,X26] : (~p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK16)))) | p(s(bool,X28)) | s(bool,X28) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X26),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f36234,plain,(
  ( ! [X296,X298,X297] : (~p(s(bool,happ(s(fun(sK13,bool),X296),s(sK13,sK16)))) | p(s(bool,X297)) | p(s(bool,X298)) | s(bool,X298) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X296),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3017,f901])).
fof(f3017,plain,(
  ( ! [X10] : (p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X10),s('type.list.list'(sK13),sK15)))) | ~p(s(bool,happ(s(fun(sK13,bool),X10),s(sK13,sK16))))) ) | $spl4),
  inference(resolution,[],[f915,f623])).
fof(f36141,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,happ(s(fun(sK13,bool),X54),s(sK13,sK16)))) | p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X54),s('type.list.list'(sK13),sK15))) | s(bool,X56) != s(bool,X57)) ) | $spl4),
  inference(resolution,[],[f3017,f653])).
fof(f36138,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,happ(s(fun(sK13,bool),X42),s(sK13,sK16)))) | p(s(bool,X43)) | ~p(s(bool,X44)) | p(s(bool,X45)) | s(bool,X45) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X42),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3017,f652])).
fof(f36137,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun(sK13,bool),X38),s(sK13,sK16)))) | p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X38),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3017,f651])).
fof(f36136,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,happ(s(fun(sK13,bool),X34),s(sK13,sK16)))) | p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X34),s('type.list.list'(sK13),sK15))) | s(bool,X36) != s(bool,X37)) ) | $spl4),
  inference(resolution,[],[f3017,f649])).
fof(f36134,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,happ(s(fun(sK13,bool),X26),s(sK13,sK16)))) | p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X27) != s(bool,X29) | s(bool,X28) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X26),s('type.list.list'(sK13),sK15)))) ) | $spl4),
  inference(resolution,[],[f3017,f647])).
fof(f32553,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32450,f32246])).
fof(f32246,plain,(
  ( ! [X15] : (~p(s(bool,X15))) ) | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32245,f859])).
fof(f859,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f546])).
fof(f546,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.EQ_CLAUSES|split|1')).
fof(f32245,plain,(
  ( ! [X15] : (s(bool,'T_0') != s(bool,X15) | ~p(s(bool,X15))) ) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f32244,f885])).
fof(f885,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f278])).
fof(f278,plain,(
  ! [X0] : s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X6] : s(bool,'T_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.NULL_FILTER','thm.list.NULL_DEF|split|0')).
fof(f32244,plain,(
  ( ! [X15] : (s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X15))) ) | ($spl0 | $spl8)),
  inference(forward_demodulation,[],[f1076,f3220])).
fof(f32450,plain,(
  ( ! [X526,X525,X527] : (p(s(bool,'const.bool.IN_2'(s(X525,sK51(X526,X527,X525)),s(fun(X525,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X525),X526))))))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f625])).
fof(f32552,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f32447,f32246])).
fof(f32447,plain,(
  ( ! [X517] : (p(s(bool,X517))) ) | ($spl0 | $spl2 | $spl8)),
  inference(resolution,[],[f32246,f2790])).
fof(f2790,plain,(
  ( ! [X12] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X12,sK13))))) | p(s(bool,X12))) ) | $spl2),
  inference(resolution,[],[f913,f599])).
fof(f599,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f411,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK37(X2,X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK36,sK37])],[f410])).
fof(f410,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X6] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | ? [X5] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f409])).
fof(f409,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f408])).
fof(f408,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X1)) & ! [X3] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X4] : (~p(s(bool,X1)) & ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | (p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X0] : ! [X1,X2] : ((p(s(bool,X1)) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> ? [X4] : (p(s(bool,X1)) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X6] : ! [X7,X25] : ((p(s(bool,X7)) | ? [X9] : p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))) <=> ? [X9] : (p(s(bool,X7)) | p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))))),
  file('thm.list.NULL_FILTER','thm.bool.RIGHT_OR_EXISTS_THM')).
fof(f32551,plain,(
  $false | ($spl0 | $spl2 | $spl8)),
  inference(subsumption_resolution,[],[f32446,f32246])).
fof(f32446,plain,(
  ( ! [X516] : (p(s(bool,X516))) ) | ($spl0 | $spl2 | $spl8)),
  inference(resolution,[],[f32246,f2789])).
fof(f2789,plain,(
  ( ! [X11] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X11,sK13))))) | p(s(bool,X11))) ) | $spl2),
  inference(resolution,[],[f913,f597])).
fof(f597,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK36(X2,X1,X0))))) | p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f411])).
fof(f32550,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32396,f32246])).
fof(f32396,plain,(
  ( ! [X366,X365] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X365),X366))))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f627])).
fof(f32549,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32395,f32246])).
fof(f32395,plain,(
  ( ! [X364,X362,X363] : (p(s(bool,'const.list.EVERY_2'(s(fun(X362,bool),X363),s('type.list.list'(X362),X364))))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f625])).
fof(f32548,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32382,f32246])).
fof(f32382,plain,(
  ( ! [X318,X317] : (p(s(bool,'const.bool.IN_2'(s(X317,sK52(X318,X317)),s(fun(X317,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X317),X318))))))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f627])).
fof(f32543,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32369,f32246])).
fof(f32369,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f874])).
fof(f874,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f548])).
fof(f32542,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32368,f32246])).
fof(f32368,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f872])).
fof(f872,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f547])).
fof(f32541,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32362,f32246])).
fof(f32362,plain,(
  ( ! [X297] : (p(s(bool,X297))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f819])).
fof(f819,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f522])).
fof(f522,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f521])).
fof(f521,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f339])).
fof(f339,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f233])).
fof(f233,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_CLAUSES|split|3')).
fof(f32540,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32361,f32246])).
fof(f32361,plain,(
  ( ! [X296] : (p(s(bool,X296))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f816])).
fof(f816,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f520])).
fof(f520,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f519])).
fof(f519,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f338])).
fof(f338,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f232])).
fof(f232,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X10] : ((p(s(bool,'T_0')) => p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_CLAUSES|split|0')).
fof(f32539,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32359,f32246])).
fof(f32359,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f813])).
fof(f813,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f518])).
fof(f518,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f517])).
fof(f517,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f337])).
fof(f337,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f231])).
fof(f231,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_CLAUSES|split|2')).
fof(f32538,plain,(
  $false | ($spl0 | $spl8)),
  inference(subsumption_resolution,[],[f32358,f32246])).
fof(f32358,plain,(
  ( ! [X294] : (p(s(bool,X294))) ) | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f810])).
fof(f810,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f516])).
fof(f516,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f515])).
fof(f515,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f336])).
fof(f336,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f230])).
fof(f230,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.NULL_FILTER','thm.bool.IMP_CLAUSES|split|1')).
fof(f32370,plain,(
  $false | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f876])).
fof(f876,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.NULL_FILTER',aHL_TRUTH)).
fof(f32371,plain,(
  $false | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f877])).
fof(f32372,plain,(
  $false | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f906])).
fof(f32383,plain,(
  $false | ($spl0 | $spl8)),
  inference(resolution,[],[f32246,f886])).
fof(f32397,plain,(
  $false | ($spl0 | $spl8 | $spl25)),
  inference(resolution,[],[f32246,f3380])).
fof(f3380,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | ($spl0 | $spl25)),
  inference(subsumption_resolution,[],[f3348,f1095])).
fof(f3348,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl0),
  inference(superposition,[],[f910,f590])).
fof(f590,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,sK34(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f404])).
fof(f404,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK34(X2,X1,X0)),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK34])],[f403])).
fof(f403,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f402])).
fof(f402,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f157])).
fof(f157,plain,(
  ! [X0] : ! [X1,X2] : (~s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7,X11] : (~s('type.list.list'(X6),X11) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X11))) <=> ? [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X11)))))) & ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))))),
  file('thm.list.NULL_FILTER','thm.list.FILTER_NEQ_ID')).
fof(f32398,plain,(
  $false | ($spl0 | $spl8 | $spl25)),
  inference(resolution,[],[f32246,f2589])).
fof(f32399,plain,(
  $false | ($spl0 | $spl4 | $spl8)),
  inference(resolution,[],[f32246,f3016])).
fof(f3016,plain,(
  ( ! [X9] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X9),s('type.list.list'(sK13),sK15))))))))) ) | $spl4),
  inference(resolution,[],[f915,f620])).
fof(f620,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) )),
  inference(cnf_transformation,[],[f433])).
fof(f433,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))))),
  inference(flattening,[],[f432])).
fof(f432,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3)))))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3))))))))))),
  inference(nnf_transformation,[],[f178])).
fof(f178,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X3)))))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X3))))))))),
  inference(rectify,[],[f110])).
fof(f110,axiom,(
  ! [X6] : ! [X9,X13,X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),X10)))))))) <=> (s(X6,X9) = s(X6,X13) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X10))))))))),
  file('thm.list.NULL_FILTER','thm.list.MEM|split|1')).
fof(f30806,plain,(
  ( ! [X47,X49] : (p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X49),s('type.list.list'(sK13),sK15)))))))) ) | $spl462),
  inference(cnf_transformation,[],[f30806_D])).
fof(f30806_D,plain,(
  ( ! [X47,X49] : (p(s(bool,X47)) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X49),s('type.list.list'(sK13),sK15)))))))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f30743,plain,(
  ( ! [X308,X309,X307] : (p(s(bool,X307)) | p(s(bool,X308)) | s(bool,X308) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X309),s('type.list.list'(sK13),sK15)))))))) ) | $spl4),
  inference(resolution,[],[f3016,f901])).
fof(f30650,plain,(
  ( ! [X76,X74,X75,X73] : (p(s(bool,X73)) | ~p(s(bool,X74)) | s(bool,X73) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X75),s('type.list.list'(sK13),sK15))))))) | s(bool,X74) != s(bool,X76)) ) | $spl4),
  inference(resolution,[],[f3016,f653])).
fof(f30647,plain,(
  ( ! [X64,X62,X65,X63] : (p(s(bool,X62)) | ~p(s(bool,X63)) | p(s(bool,X64)) | s(bool,X64) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X65),s('type.list.list'(sK13),sK15)))))))) ) | $spl4),
  inference(resolution,[],[f3016,f652])).
fof(f30646,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X58) != s(bool,X60) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X61),s('type.list.list'(sK13),sK15)))))))) ) | $spl4),
  inference(resolution,[],[f3016,f651])).
fof(f30645,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X54)) | ~p(s(bool,X55)) | s(bool,X54) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X56),s('type.list.list'(sK13),sK15))))))) | s(bool,X55) != s(bool,X57)) ) | $spl4),
  inference(resolution,[],[f3016,f649])).
fof(f30643,plain,(
  ( ! [X47,X48,X46,X49] : (p(s(bool,X46)) | p(s(bool,X47)) | s(bool,X46) != s(bool,X48) | s(bool,X47) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,X49),s('type.list.list'(sK13),sK15)))))))) ) | $spl4),
  inference(resolution,[],[f3016,f647])).
fof(f29153,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl460),
  inference(cnf_transformation,[],[f29153_D])).
fof(f29153_D,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f29114,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl25),
  inference(resolution,[],[f2589,f901])).
fof(f29021,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X52) != s(bool,X53)) ) | $spl25),
  inference(resolution,[],[f2589,f653])).
fof(f29018,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl25),
  inference(resolution,[],[f2589,f652])).
fof(f29017,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl25),
  inference(resolution,[],[f2589,f651])).
fof(f29016,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X37) != s(bool,X38)) ) | $spl25),
  inference(resolution,[],[f2589,f649])).
fof(f29014,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK52(sK15,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl25),
  inference(resolution,[],[f2589,f647])).
fof(f27438,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl459),
  inference(cnf_transformation,[],[f27438_D])).
fof(f27438_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13)))) <=> ~$spl459),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl459])])).
fof(f27423,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl452)),
  inference(resolution,[],[f27144,f8825])).
fof(f27144,plain,(
  ( ! [X156,X155] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))),X155,X156)) ) | $spl452),
  inference(resolution,[],[f26801,f723])).
fof(f27210,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl456),
  inference(cnf_transformation,[],[f27210_D])).
fof(f27210_D,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f27180,plain,(
  ( ! [X189,X190] : (p(s(bool,X189)) | p(s(bool,X190)) | s(bool,X190) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl452),
  inference(resolution,[],[f26801,f901])).
fof(f27087,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl452),
  inference(resolution,[],[f26801,f653])).
fof(f27084,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl452),
  inference(resolution,[],[f26801,f652])).
fof(f27083,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl452),
  inference(resolution,[],[f26801,f651])).
fof(f27082,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13)))) | s(bool,X26) != s(bool,X27)) ) | $spl452),
  inference(resolution,[],[f26801,f649])).
fof(f27080,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl452),
  inference(resolution,[],[f26801,f647])).
fof(f27063,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl454),
  inference(cnf_transformation,[],[f27063_D])).
fof(f27063_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f27035,plain,(
  ( ! [X192,X193] : (p(s(bool,X192)) | ~p(s(bool,X193)) | s(bool,X193) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl453),
  inference(resolution,[],[f26802,f901])).
fof(f26937,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13)))) | s(bool,X35) != s(bool,X37)) ) | $spl453),
  inference(resolution,[],[f26802,f653])).
fof(f26936,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl453),
  inference(resolution,[],[f26802,f652])).
fof(f26934,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl453),
  inference(resolution,[],[f26802,f651])).
fof(f26932,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13)))) | s(bool,X20) != s(bool,X22)) ) | $spl453),
  inference(resolution,[],[f26802,f649])).
fof(f26930,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl453),
  inference(resolution,[],[f26802,f647])).
fof(f26802,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) | $spl453),
  inference(cnf_transformation,[],[f26802_D])).
fof(f26802_D,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) <=> ~$spl453),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl453])])).
fof(f26800,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl450),
  inference(cnf_transformation,[],[f26800_D])).
fof(f26800_D,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f26607,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK55(sK15,sK13))))) | $spl4),
  inference(superposition,[],[f3168,f882])).
fof(f26798,plain,(
  s('type.list.list'(sK13),sK55(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK55(sK15,sK13),sK13)),s('type.list.list'(sK13),sK55(sK55(sK15,sK13),sK13)))) | $spl448),
  inference(cnf_transformation,[],[f26798_D])).
fof(f26798_D,plain,(
  s('type.list.list'(sK13),sK55(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK55(sK15,sK13),sK13)),s('type.list.list'(sK13),sK55(sK55(sK15,sK13),sK13)))) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f26604,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK55(sK15,sK13)) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK55(sK15,sK13),sK13)),s('type.list.list'(sK13),sK55(sK55(sK15,sK13),sK13)))) | $spl4),
  inference(superposition,[],[f3168,f900])).
fof(f1703,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl38),
  inference(cnf_transformation,[],[f1703_D])).
fof(f1703_D,plain,(
  s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f26597,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0))) ) | $spl446),
  inference(cnf_transformation,[],[f26597_D])).
fof(f26597_D,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f26574,plain,(
  ( ! [X0,X1] : (~sP10(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3094,f900])).
fof(f3094,plain,(
  ( ! [X175,X174] : (~sP10(X174,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X175) | p(s(bool,X174))) ) | $spl4),
  inference(resolution,[],[f915,f729])).
fof(f729,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f482])).
fof(f482,plain,(
  ! [X0,X1,X2] : ((~sP10(X0,X1,X2) | ((~p(s(bool,X2)) | p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X2)) & ~p(s(bool,X0))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | sP10(X0,X1,X2)))),
  inference(rectify,[],[f481])).
fof(f481,plain,(
  ! [X2,X1,X0] : ((~sP10(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & ((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2))) | sP10(X2,X1,X0)))),
  inference(flattening,[],[f480])).
fof(f480,plain,(
  ! [X2,X1,X0] : ((~sP10(X2,X1,X0) | ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2))))) & (((p(s(bool,X0)) & ~p(s(bool,X2))) | (p(s(bool,X1)) & ~p(s(bool,X2)))) | sP10(X2,X1,X0)))),
  inference(nnf_transformation,[],[f365])).
fof(f365,plain,(
  ! [X2,X1,X0] : (sP10(X2,X1,X0) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP10])])).
fof(f26551,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0))) ) | $spl444),
  inference(cnf_transformation,[],[f26551_D])).
fof(f26551_D,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f26528,plain,(
  ( ! [X0,X1] : (~sP10(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3093,f900])).
fof(f3093,plain,(
  ( ! [X173,X172] : (~sP10(X172,X173,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X172))) ) | $spl4),
  inference(resolution,[],[f915,f728])).
fof(f728,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X2)) | ~sP10(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f482])).
fof(f26438,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X0))) ) | $spl442),
  inference(cnf_transformation,[],[f26438_D])).
fof(f26438_D,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X0))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f26405,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3064,f900])).
fof(f3064,plain,(
  ( ! [X114,X115] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X114,X115) | p(s(bool,X114))) ) | $spl4),
  inference(resolution,[],[f915,f690])).
fof(f690,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f463])).
fof(f463,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP4(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP4(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f462])).
fof(f462,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP5(X2,X0,X1)))),
  inference(flattening,[],[f461])).
fof(f461,plain,(
  ! [X2,X0,X1] : ((~sP5(X2,X0,X1) | (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP4(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP5(X2,X0,X1)))),
  inference(nnf_transformation,[],[f358])).
fof(f358,plain,(
  ! [X2,X0,X1] : (sP5(X2,X0,X1) <=> (sP4(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f25663,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1))) ) | $spl440),
  inference(cnf_transformation,[],[f25663_D])).
fof(f25663_D,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f25630,plain,(
  ( ! [X0,X1] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3063,f900])).
fof(f3063,plain,(
  ( ! [X112,X113] : (~sP5('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X113,X112) | p(s(bool,X112))) ) | $spl4),
  inference(resolution,[],[f915,f689])).
fof(f689,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f463])).
fof(f24686,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0))) ) | $spl438),
  inference(cnf_transformation,[],[f24686_D])).
fof(f24686_D,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f24659,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3059,f900])).
fof(f3059,plain,(
  ( ! [X105,X104] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X104,X105) | ~p(s(bool,X104))) ) | $spl4),
  inference(resolution,[],[f915,f684])).
fof(f684,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f460])).
fof(f460,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f459])).
fof(f459,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP3(X2,X0,X1)) & (~sP3(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f356])).
fof(f356,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP3(X2,X0,X1))),
  inference(definition_folding,[],[f193,f355,f354])).
fof(f355,plain,(
  ! [X2,X0,X1] : (sP3(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f193,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f82])).
fof(f82,axiom,(
  ! [X34,X35,X36] : ((p(s(bool,X36)) <=> (p(s(bool,X35)) | p(s(bool,X34)))) <=> ((p(s(bool,X36)) | ~p(s(bool,X35))) & (p(s(bool,X36)) | ~p(s(bool,X34))) & (p(s(bool,X35)) | p(s(bool,X34)) | ~p(s(bool,X36)))))),
  file('thm.list.NULL_FILTER','thm.sat.sat_25_prove_ltb')).
fof(f24608,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0))) ) | $spl436),
  inference(cnf_transformation,[],[f24608_D])).
fof(f24608_D,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f24584,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3058,f900])).
fof(f3058,plain,(
  ( ! [X103,X102] : (sP3(X102,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X103) | ~p(s(bool,X102))) ) | $spl4),
  inference(resolution,[],[f915,f684])).
fof(f24458,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X1))) ) | $spl434),
  inference(cnf_transformation,[],[f24458_D])).
fof(f24458_D,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f24431,plain,(
  ( ! [X0,X1] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3057,f900])).
fof(f3057,plain,(
  ( ! [X101,X100] : (sP3('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X100,X101) | ~p(s(bool,X101))) ) | $spl4),
  inference(resolution,[],[f915,f683])).
fof(f683,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f460])).
fof(f24411,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0))) ) | $spl432),
  inference(cnf_transformation,[],[f24411_D])).
fof(f24411_D,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f24387,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3056,f900])).
fof(f3056,plain,(
  ( ! [X99,X98] : (sP3(X98,X99,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X98))) ) | $spl4),
  inference(resolution,[],[f915,f683])).
fof(f24362,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0))) ) | $spl430),
  inference(cnf_transformation,[],[f24362_D])).
fof(f24362_D,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f24336,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3050,f900])).
fof(f3050,plain,(
  ( ! [X87,X86] : (~sP3(X86,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X87) | p(s(bool,X86))) ) | $spl4),
  inference(resolution,[],[f915,f672])).
fof(f672,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f455])).
fof(f455,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP2(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP2(X0,X1,X2) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f454])).
fof(f454,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1) | sP3(X2,X0,X1)))),
  inference(flattening,[],[f453])).
fof(f453,plain,(
  ! [X2,X0,X1] : ((~sP3(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP2(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP2(X2,X0,X1)) | sP3(X2,X0,X1)))),
  inference(nnf_transformation,[],[f355])).
fof(f24313,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0))) ) | $spl428),
  inference(cnf_transformation,[],[f24313_D])).
fof(f24313_D,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f24287,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3049,f900])).
fof(f3049,plain,(
  ( ! [X85,X84] : (~sP3(X84,X85,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X84))) ) | $spl4),
  inference(resolution,[],[f915,f671])).
fof(f671,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X0)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f455])).
fof(f24265,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0))) ) | $spl426),
  inference(cnf_transformation,[],[f24265_D])).
fof(f24265_D,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f24238,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3045,f900])).
fof(f3045,plain,(
  ( ! [X76,X77] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X76,X77) | ~p(s(bool,X76))) ) | $spl4),
  inference(resolution,[],[f915,f667])).
fof(f667,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f452])).
fof(f452,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f451])).
fof(f451,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP1(X2,X0,X1)) & (~sP1(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f353])).
fof(f353,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP1(X2,X0,X1))),
  inference(definition_folding,[],[f314,f352,f351])).
fof(f352,plain,(
  ! [X2,X0,X1] : (sP1(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f314,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f85])).
fof(f85,axiom,(
  ! [X34,X35,X36] : ((p(s(bool,X36)) <=> (p(s(bool,X35)) => p(s(bool,X34)))) <=> ((p(s(bool,X36)) | p(s(bool,X35))) & (p(s(bool,X36)) | ~p(s(bool,X34))) & (~p(s(bool,X35)) | p(s(bool,X34)) | ~p(s(bool,X36)))))),
  file('thm.list.NULL_FILTER','thm.sat.sat_27_prove_ltb')).
fof(f24218,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0))) ) | $spl424),
  inference(cnf_transformation,[],[f24218_D])).
fof(f24218_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0))) ) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f24191,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3044,f900])).
fof(f3044,plain,(
  ( ! [X74,X75] : (sP1(X74,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X75) | ~p(s(bool,X74))) ) | $spl4),
  inference(resolution,[],[f915,f667])).
fof(f24016,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1))) ) | $spl422),
  inference(cnf_transformation,[],[f24016_D])).
fof(f24016_D,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1))) ) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f23985,plain,(
  ( ! [X0,X1] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3043,f900])).
fof(f3043,plain,(
  ( ! [X72,X73] : (sP1('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X72,X73) | p(s(bool,X73))) ) | $spl4),
  inference(resolution,[],[f915,f666])).
fof(f666,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f452])).
fof(f23811,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0))) ) | $spl420),
  inference(cnf_transformation,[],[f23811_D])).
fof(f23811_D,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0))) ) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f23785,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3034,f900])).
fof(f3034,plain,(
  ( ! [X54,X55] : (~sP1(X54,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X55) | p(s(bool,X54))) ) | $spl4),
  inference(resolution,[],[f915,f655])).
fof(f655,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f447])).
fof(f447,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP0(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP0(X0,X1,X2) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f446])).
fof(f446,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1) | sP1(X2,X0,X1)))),
  inference(flattening,[],[f445])).
fof(f445,plain,(
  ! [X2,X0,X1] : ((~sP1(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP0(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP0(X2,X0,X1)) | sP1(X2,X0,X1)))),
  inference(nnf_transformation,[],[f352])).
fof(f19729,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1))) ) | $spl418),
  inference(cnf_transformation,[],[f19729_D])).
fof(f19729_D,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f19710,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4385,f900])).
fof(f4385,plain,(
  ( ! [X161,X160] : (sP10(X160,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X161) | p(s(bool,X161))) ) | $spl23),
  inference(resolution,[],[f1093,f730])).
fof(f730,plain,(
  ( ! [X2,X0,X1] : (sP10(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f482])).
fof(f19694,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1))) ) | $spl416),
  inference(cnf_transformation,[],[f19694_D])).
fof(f19694_D,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1))) ) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f19675,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4384,f900])).
fof(f4384,plain,(
  ( ! [X158,X159] : (sP10(X158,X159,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | p(s(bool,X159))) ) | $spl23),
  inference(resolution,[],[f1093,f730])).
fof(f19659,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl0 | $spl23)),
  inference(subsumption_resolution,[],[f19646,f13095])).
fof(f13095,plain,(
  ( ! [X249,X250,X248] : (sP10('const.list.EVERY_2'(s(fun(sK13,bool),X248),s('type.list.list'(sK13),'const.list.NIL_0')),X249,X250)) ) | $spl0),
  inference(resolution,[],[f12778,f735])).
fof(f735,plain,(
  ( ! [X2,X0,X1] : (sP10(X2,X1,X0) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f484])).
fof(f484,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP10(X2,X1,X0)) & (~sP10(X2,X1,X0) | (~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))))),
  inference(flattening,[],[f483])).
fof(f483,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X0)) | p(s(bool,X1))) & ~p(s(bool,X2))) | sP10(X2,X1,X0)) & (~sP10(X2,X1,X0) | ((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f366])).
fof(f366,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> sP10(X2,X1,X0))),
  inference(definition_folding,[],[f315,f365])).
fof(f315,plain,(
  ! [X0,X1,X2] : (((~p(s(bool,X0)) & ~p(s(bool,X1))) | p(s(bool,X2))) <=> ((~p(s(bool,X0)) | p(s(bool,X2))) & (~p(s(bool,X1)) | p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f198])).
fof(f198,plain,(
  ! [X0,X1,X2] : (((p(s(bool,X0)) | p(s(bool,X1))) => p(s(bool,X2))) <=> ((p(s(bool,X0)) => p(s(bool,X2))) & (p(s(bool,X1)) => p(s(bool,X2)))))),
  inference(rectify,[],[f123])).
fof(f123,axiom,(
  ! [X7,X25,X53] : (((p(s(bool,X7)) | p(s(bool,X25))) => p(s(bool,X53))) <=> ((p(s(bool,X7)) => p(s(bool,X53))) & (p(s(bool,X25)) => p(s(bool,X53)))))),
  file('thm.list.NULL_FILTER','thm.bool.DISJ_IMP_THM')).
fof(f12778,plain,(
  ( ! [X87] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X87),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f12777,f3220])).
fof(f12777,plain,(
  ( ! [X87] : (p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X87),s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12776,f869])).
fof(f12776,plain,(
  ( ! [X87] : (p(s(bool,'F_0')) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X87),s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(forward_demodulation,[],[f12573,f629])).
fof(f12573,plain,(
  ( ! [X87] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK51('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X87,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) | p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X87),s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f625,f3220])).
fof(f19646,plain,(
  ( ! [X0,X1] : (~sP10('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4383,f900])).
fof(f4383,plain,(
  ( ! [X156,X157] : (~sP10('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X156,X157) | ~p(s(bool,X156))) ) | $spl23),
  inference(resolution,[],[f1093,f729])).
fof(f19595,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl0 | $spl23)),
  inference(subsumption_resolution,[],[f19582,f13095])).
fof(f19582,plain,(
  ( ! [X0,X1] : (~sP10('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4382,f900])).
fof(f4382,plain,(
  ( ! [X154,X155] : (~sP10('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X155,X154) | ~p(s(bool,X154))) ) | $spl23),
  inference(resolution,[],[f1093,f728])).
fof(f19567,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0))) ) | $spl414),
  inference(cnf_transformation,[],[f19567_D])).
fof(f19567_D,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0))) ) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f19552,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4364,f900])).
fof(f4364,plain,(
  ( ! [X118,X119] : (sP5(X118,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X119) | p(s(bool,X118))) ) | $spl23),
  inference(resolution,[],[f1093,f700])).
fof(f700,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f467])).
fof(f467,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f466])).
fof(f466,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP5(X2,X0,X1)) & (~sP5(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f359])).
fof(f359,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP5(X2,X0,X1))),
  inference(definition_folding,[],[f195,f358,f357])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f93])).
fof(f93,axiom,(
  ! [X34,X35,X36] : ((p(s(bool,X36)) <=> (p(s(bool,X35)) & p(s(bool,X34)))) <=> ((p(s(bool,X36)) | ~p(s(bool,X35)) | ~p(s(bool,X34))) & (p(s(bool,X35)) | ~p(s(bool,X36))) & (p(s(bool,X34)) | ~p(s(bool,X36)))))),
  file('thm.list.NULL_FILTER','thm.sat.sat_23_prove_ltb')).
fof(f19533,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0))) ) | $spl412),
  inference(cnf_transformation,[],[f19533_D])).
fof(f19533_D,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0))) ) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f19506,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4363,f900])).
fof(f4363,plain,(
  ( ! [X116,X117] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X116,X117) | p(s(bool,X116))) ) | $spl23),
  inference(resolution,[],[f1093,f700])).
fof(f19453,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl410),
  inference(cnf_transformation,[],[f19453_D])).
fof(f19453_D,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f19438,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4362,f900])).
fof(f4362,plain,(
  ( ! [X114,X115] : (sP5(X114,X115,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | p(s(bool,X114))) ) | $spl23),
  inference(resolution,[],[f1093,f699])).
fof(f699,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f467])).
fof(f19419,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) | $spl408),
  inference(cnf_transformation,[],[f19419_D])).
fof(f19419_D,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f19392,plain,(
  ( ! [X0,X1] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4361,f900])).
fof(f4361,plain,(
  ( ! [X112,X113] : (sP5('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X112,X113) | p(s(bool,X113))) ) | $spl23),
  inference(resolution,[],[f1093,f699])).
fof(f19376,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0))) ) | $spl406),
  inference(cnf_transformation,[],[f19376_D])).
fof(f19376_D,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0))) ) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f19356,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4355,f900])).
fof(f4355,plain,(
  ( ! [X101,X100] : (~sP5(X100,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X101) | ~p(s(bool,X100))) ) | $spl23),
  inference(resolution,[],[f1093,f690])).
fof(f19338,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) | $spl404),
  inference(cnf_transformation,[],[f19338_D])).
fof(f19338_D,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f19318,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4354,f900])).
fof(f4354,plain,(
  ( ! [X99,X98] : (~sP5(X98,X99,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X98))) ) | $spl23),
  inference(resolution,[],[f1093,f689])).
fof(f19267,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl402),
  inference(cnf_transformation,[],[f19267_D])).
fof(f19267_D,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f19242,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4339,f900])).
fof(f4339,plain,(
  ( ! [X68,X69] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X68,X69) | ~p(s(bool,X68))) ) | $spl23),
  inference(resolution,[],[f1093,f672])).
fof(f19223,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1))) ) | $spl400),
  inference(cnf_transformation,[],[f19223_D])).
fof(f19223_D,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f19198,plain,(
  ( ! [X0,X1] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4338,f900])).
fof(f4338,plain,(
  ( ! [X66,X67] : (~sP3('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X67,X66) | ~p(s(bool,X66))) ) | $spl23),
  inference(resolution,[],[f1093,f671])).
fof(f19144,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) | $spl398),
  inference(cnf_transformation,[],[f19144_D])).
fof(f19144_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f19128,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4333,f900])).
fof(f4333,plain,(
  ( ! [X57,X56] : (sP1(X56,X57,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X56))) ) | $spl23),
  inference(resolution,[],[f1093,f666])).
fof(f19109,plain,(
  ( ! [X0,X1] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl396),
  inference(cnf_transformation,[],[f19109_D])).
fof(f19109_D,plain,(
  ( ! [X0,X1] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl396),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl396])])).
fof(f19087,plain,(
  ( ! [X0,X1] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4325,f900])).
fof(f4325,plain,(
  ( ! [X41,X40] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X40,X41) | ~p(s(bool,X40))) ) | $spl23),
  inference(resolution,[],[f1093,f655])).
fof(f19068,plain,(
  ( ! [X0,X1] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) | $spl394),
  inference(cnf_transformation,[],[f19068_D])).
fof(f19068_D,plain,(
  ( ! [X0,X1] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl394),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl394])])).
fof(f19041,plain,(
  ( ! [X0,X1] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4324,f900])).
fof(f4324,plain,(
  ( ! [X39,X38] : (~sP1('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X39,X38) | p(s(bool,X38))) ) | $spl23),
  inference(resolution,[],[f1093,f654])).
fof(f654,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | p(s(bool,X0)) | ~sP1(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f447])).
fof(f19022,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl392),
  inference(cnf_transformation,[],[f19022_D])).
fof(f19022_D,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl392),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl392])])).
fof(f18997,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4323,f900])).
fof(f4323,plain,(
  ( ! [X37,X36] : (~sP1(X36,X37,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | p(s(bool,X36))) ) | $spl23),
  inference(resolution,[],[f1093,f654])).
fof(f18741,plain,(
  $false | ($spl0 | $spl12)),
  inference(equality_resolution,[],[f18682])).
fof(f18682,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | ($spl0 | $spl12)),
  inference(forward_demodulation,[],[f1078,f3220])).
fof(f18089,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl390),
  inference(cnf_transformation,[],[f18089_D])).
fof(f18089_D,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f18050,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f3380,f901])).
fof(f17957,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X52) != s(bool,X53)) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f3380,f653])).
fof(f17954,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f3380,f652])).
fof(f17953,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f3380,f651])).
fof(f17952,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X37) != s(bool,X38)) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f3380,f649])).
fof(f17950,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f3380,f647])).
fof(f16220,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl388),
  inference(cnf_transformation,[],[f16220_D])).
fof(f16220_D,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | p(s(bool,X1))) ) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f16191,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f12199,f900])).
fof(f12199,plain,(
  ( ! [X6,X7] : (~sP9(X6,X7,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X7)) | p(s(bool,X6))) ) | $spl25),
  inference(resolution,[],[f2658,f705])).
fof(f2658,plain,(
  ( ! [X151,X150] : (~sP6('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X150,X151) | p(s(bool,X151)) | p(s(bool,X150))) ) | $spl25),
  inference(resolution,[],[f1095,f720])).
fof(f720,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f478])).
fof(f16173,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl386),
  inference(cnf_transformation,[],[f16173_D])).
fof(f16173_D,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0)) | p(s(bool,X1))) ) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f16157,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f12162,f900])).
fof(f12162,plain,(
  ( ! [X6,X7] : (~sP9(X6,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X7) | p(s(bool,X7)) | p(s(bool,X6))) ) | $spl25),
  inference(resolution,[],[f2657,f705])).
fof(f2657,plain,(
  ( ! [X149,X148] : (~sP6(X148,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X149) | p(s(bool,X149)) | p(s(bool,X148))) ) | $spl25),
  inference(resolution,[],[f1095,f720])).
fof(f16143,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl384),
  inference(cnf_transformation,[],[f16143_D])).
fof(f16143_D,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) ) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f16128,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f12070,f900])).
fof(f12070,plain,(
  ( ! [X8,X9] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X8,X9) | p(s(bool,X9)) | p(s(bool,X8))) ) | $spl25),
  inference(resolution,[],[f2656,f705])).
fof(f2656,plain,(
  ( ! [X146,X147] : (~sP6(X146,X147,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X147)) | p(s(bool,X146))) ) | $spl25),
  inference(resolution,[],[f1095,f720])).
fof(f16112,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl382),
  inference(cnf_transformation,[],[f16112_D])).
fof(f16112_D,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f16098,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f12036,f900])).
fof(f12036,plain,(
  ( ! [X8,X9] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9,X8) | ~p(s(bool,X9)) | ~p(s(bool,X8))) ) | $spl25),
  inference(resolution,[],[f2653,f706])).
fof(f2653,plain,(
  ( ! [X140,X141] : (~sP7(X140,X141,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X141)) | ~p(s(bool,X140))) ) | $spl25),
  inference(resolution,[],[f1095,f716])).
fof(f716,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f476])).
fof(f15911,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0))) ) | $spl380),
  inference(cnf_transformation,[],[f15911_D])).
fof(f15911_D,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f15895,plain,(
  ( ! [X0] : (sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3103,f900])).
fof(f3103,plain,(
  ( ! [X189] : (sP11(X189,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | p(s(bool,X189))) ) | $spl4),
  inference(resolution,[],[f915,f763])).
fof(f763,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f495])).
fof(f495,plain,(
  ! [X0,X1] : ((~sP11(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP11(X0,X1)))),
  inference(rectify,[],[f494])).
fof(f494,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP11(X1,X0)))),
  inference(flattening,[],[f493])).
fof(f493,plain,(
  ! [X1,X0] : ((~sP11(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP11(X1,X0)))),
  inference(nnf_transformation,[],[f367])).
fof(f367,plain,(
  ! [X1,X0] : (sP11(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP11])])).
fof(f15879,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | p(s(bool,X0))) ) | $spl378),
  inference(cnf_transformation,[],[f15879_D])).
fof(f15879_D,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | p(s(bool,X0))) ) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f15863,plain,(
  ( ! [X0] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3102,f900])).
fof(f3102,plain,(
  ( ! [X188] : (sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X188) | p(s(bool,X188))) ) | $spl4),
  inference(resolution,[],[f915,f760])).
fof(f760,plain,(
  ( ! [X0,X1] : (sP11(X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f495])).
fof(f15848,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0))) ) | $spl376),
  inference(cnf_transformation,[],[f15848_D])).
fof(f15848_D,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0))) ) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f15832,plain,(
  ( ! [X0] : (~sP11(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3101,f900])).
fof(f3101,plain,(
  ( ! [X187] : (~sP11(X187,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | ~p(s(bool,X187))) ) | $spl4),
  inference(resolution,[],[f915,f759])).
fof(f759,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f495])).
fof(f15672,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | ~p(s(bool,X0))) ) | $spl374),
  inference(cnf_transformation,[],[f15672_D])).
fof(f15672_D,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | ~p(s(bool,X0))) ) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f15656,plain,(
  ( ! [X0] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3100,f900])).
fof(f3100,plain,(
  ( ! [X186] : (~sP11('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X186) | ~p(s(bool,X186))) ) | $spl4),
  inference(resolution,[],[f915,f759])).
fof(f15590,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl0 | $spl23)),
  inference(subsumption_resolution,[],[f15575,f13107])).
fof(f13107,plain,(
  ( ! [X275,X276] : (sP12('const.list.EVERY_2'(s(fun(sK13,bool),X275),s('type.list.list'(sK13),'const.list.NIL_0')),X276)) ) | $spl0),
  inference(resolution,[],[f12778,f776])).
fof(f776,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f500])).
fof(f500,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(flattening,[],[f499])).
fof(f499,plain,(
  ! [X0,X1] : ((~sP12(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP12(X0,X1)))),
  inference(nnf_transformation,[],[f369])).
fof(f369,plain,(
  ! [X0,X1] : (sP12(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP12])])).
fof(f15575,plain,(
  ( ! [X0] : (~sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4456,f900])).
fof(f4456,plain,(
  ( ! [X183] : (~sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X183) | p(s(bool,X183))) ) | $spl23),
  inference(subsumption_resolution,[],[f4401,f826])).
fof(f826,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f526])).
fof(f526,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f525])).
fof(f525,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f63])).
fof(f63,axiom,(
  ! [X10] : ((p(s(bool,X10)) | p(s(bool,'F_0'))) <=> p(s(bool,X10)))),
  file('thm.list.NULL_FILTER','thm.bool.OR_CLAUSES|split|3')).
fof(f4401,plain,(
  ( ! [X183] : (p(s(bool,'F_0')) | p(s(bool,X183)) | ~sP12('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X183)) ) | $spl23),
  inference(resolution,[],[f1093,f774])).
fof(f774,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP12(X0,X1)) )),
  inference(cnf_transformation,[],[f500])).
fof(f15554,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl0 | $spl23)),
  inference(subsumption_resolution,[],[f15538,f13106])).
fof(f13106,plain,(
  ( ! [X273,X274] : (sP12(X273,'const.list.EVERY_2'(s(fun(sK13,bool),X274),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12778,f775])).
fof(f775,plain,(
  ( ! [X0,X1] : (sP12(X0,X1) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f500])).
fof(f15538,plain,(
  ( ! [X0] : (~sP12(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4455,f900])).
fof(f4455,plain,(
  ( ! [X182] : (~sP12(X182,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | p(s(bool,X182))) ) | $spl23),
  inference(subsumption_resolution,[],[f4400,f826])).
fof(f4400,plain,(
  ( ! [X182] : (p(s(bool,'F_0')) | p(s(bool,X182)) | ~sP12(X182,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl23),
  inference(resolution,[],[f1093,f774])).
fof(f15484,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl373),
  inference(cnf_transformation,[],[f15484_D])).
fof(f15484_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) <=> ~$spl373),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl373])])).
fof(f15456,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl23),
  inference(superposition,[],[f4416,f900])).
fof(f4416,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl23),
  inference(resolution,[],[f1093,f856])).
fof(f15222,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | $spl370),
  inference(cnf_transformation,[],[f15222_D])).
fof(f15222_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f14839,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl23),
  inference(superposition,[],[f4414,f900])).
fof(f4414,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl23),
  inference(resolution,[],[f1093,f853])).
fof(f14828,plain,(
  ( ! [X60,X58] : (p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X58,sK13))))) ) | $spl368),
  inference(cnf_transformation,[],[f14828_D])).
fof(f14828_D,plain,(
  ( ! [X60,X58] : (p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X58,sK13))))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f14796,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X328,sK13))))) ) | $spl2),
  inference(resolution,[],[f2790,f901])).
fof(f14703,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X86,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2790,f653])).
fof(f14700,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X74,sK13))))) ) | $spl2),
  inference(resolution,[],[f2790,f652])).
fof(f14699,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X70,sK13))))) ) | $spl2),
  inference(resolution,[],[f2790,f651])).
fof(f14698,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X66,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2790,f649])).
fof(f14696,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X58,sK13))))) ) | $spl2),
  inference(resolution,[],[f2790,f647])).
fof(f14686,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2790,f605])).
fof(f14685,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2790,f604])).
fof(f14676,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2790,f570])).
fof(f14675,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2790,f569])).
fof(f14674,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2790,f568])).
fof(f13378,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) | $spl324),
  inference(cnf_transformation,[],[f13378_D])).
fof(f13378_D,plain,(
  ( ! [X3] : (p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f14673,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2790,f567])).
fof(f14670,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0))) ) | $spl366),
  inference(cnf_transformation,[],[f14670_D])).
fof(f14670_D,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f14658,plain,(
  ( ! [X0] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4395,f900])).
fof(f4395,plain,(
  ( ! [X177] : (sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X177) | ~p(s(bool,X177))) ) | $spl23),
  inference(resolution,[],[f1093,f763])).
fof(f14644,plain,(
  ( ! [X0] : (sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) | $spl364),
  inference(cnf_transformation,[],[f14644_D])).
fof(f14644_D,plain,(
  ( ! [X0] : (sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f14632,plain,(
  ( ! [X0] : (sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4394,f900])).
fof(f4394,plain,(
  ( ! [X176] : (sP11(X176,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X176))) ) | $spl23),
  inference(resolution,[],[f1093,f760])).
fof(f14614,plain,(
  ( ! [X60,X58] : (p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X58,sK13))))) ) | $spl362),
  inference(cnf_transformation,[],[f14614_D])).
fof(f14614_D,plain,(
  ( ! [X60,X58] : (p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X58,sK13))))) ) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f14582,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,X328)) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X328,sK13))))) ) | $spl2),
  inference(resolution,[],[f2789,f901])).
fof(f14489,plain,(
  ( ! [X88,X87,X89,X86] : (p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X86,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2789,f653])).
fof(f14486,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X74,sK13))))) ) | $spl2),
  inference(resolution,[],[f2789,f652])).
fof(f14485,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X70,sK13))))) ) | $spl2),
  inference(resolution,[],[f2789,f651])).
fof(f14484,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X66,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2789,f649])).
fof(f14482,plain,(
  ( ! [X61,X59,X60,X58] : (p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X58,sK13))))) ) | $spl2),
  inference(resolution,[],[f2789,f647])).
fof(f14472,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2789,f605])).
fof(f14471,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2789,f604])).
fof(f13386,plain,(
  ( ! [X25] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X25,sK13))))) | p(s(bool,X25))) ) | $spl340),
  inference(cnf_transformation,[],[f13386_D])).
fof(f13386_D,plain,(
  ( ! [X25] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X25,sK13))))) | p(s(bool,X25))) ) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f14470,plain,(
  ( ! [X24,X25] : (p(s(bool,X24)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X25,sK13))))) | p(s(bool,X25))) ) | $spl2),
  inference(resolution,[],[f2789,f599])).
fof(f14462,plain,(
  ( ! [X8,X9] : (p(s(bool,X8)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2789,f570])).
fof(f14461,plain,(
  ( ! [X6,X7] : (p(s(bool,X6)) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2789,f569])).
fof(f14460,plain,(
  ( ! [X4,X5] : (p(s(bool,X4)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2789,f568])).
fof(f14459,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2789,f567])).
fof(f14455,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) | $spl360),
  inference(cnf_transformation,[],[f14455_D])).
fof(f14455_D,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f14443,plain,(
  ( ! [X0] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4393,f900])).
fof(f4393,plain,(
  ( ! [X175] : (~sP11('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X175) | p(s(bool,X175))) ) | $spl23),
  inference(resolution,[],[f1093,f758])).
fof(f758,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | ~sP11(X0,X1)) )),
  inference(cnf_transformation,[],[f495])).
fof(f14380,plain,(
  ( ! [X14,X15] : (p(s(bool,X14)) | s(bool,'F_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X15),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f13088,f8825])).
fof(f13088,plain,(
  ( ! [X229,X227,X228] : (sP6('const.list.EVERY_2'(s(fun(sK13,bool),X227),s('type.list.list'(sK13),'const.list.NIL_0')),X228,X229)) ) | $spl0),
  inference(resolution,[],[f12778,f723])).
fof(f14188,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X58,sK13))))) ) | $spl356),
  inference(cnf_transformation,[],[f14188_D])).
fof(f14188_D,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X58,sK13))))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f14156,plain,(
  ( ! [X329,X328,X330] : (~p(s(bool,X328)) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X328,sK13))))) ) | $spl2),
  inference(resolution,[],[f2786,f901])).
fof(f2786,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X8,sK13))))) | ~p(s(bool,X8))) ) | $spl2),
  inference(resolution,[],[f913,f578])).
fof(f578,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f390])).
fof(f390,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0))))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK27(X2,X1,X0))))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK26,sK27])],[f389])).
fof(f389,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6))))) | (p(s(bool,X1)) & ? [X5] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(rectify,[],[f388])).
fof(f388,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(flattening,[],[f387])).
fof(f387,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,X1)) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & ((~p(s(bool,X1)) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))) | ? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(flattening,[],[f149])).
fof(f149,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,X1)) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) <=> (p(s(bool,X1)) & ? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4))))))),
  inference(rectify,[],[f114])).
fof(f114,axiom,(
  ! [X6] : ! [X7,X25] : (? [X9] : (p(s(bool,X7)) & p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))) <=> (p(s(bool,X7)) & ? [X9] : p(s(bool,happ(s(fun(X6,bool),X25),s(X6,X9))))))),
  file('thm.list.NULL_FILTER','thm.bool.RIGHT_EXISTS_AND_THM')).
fof(f14063,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X86,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2786,f653])).
fof(f14060,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X74,sK13))))) ) | $spl2),
  inference(resolution,[],[f2786,f652])).
fof(f14059,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X70,sK13))))) ) | $spl2),
  inference(resolution,[],[f2786,f651])).
fof(f14058,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X66,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2786,f649])).
fof(f14056,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X58,sK13))))) ) | $spl2),
  inference(resolution,[],[f2786,f647])).
fof(f14046,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2786,f605])).
fof(f14045,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2786,f604])).
fof(f14044,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X25,sK13))))) | p(s(bool,X25))) ) | $spl2),
  inference(resolution,[],[f2786,f599])).
fof(f13385,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X23,sK13))))) | p(s(bool,X23))) ) | $spl338),
  inference(cnf_transformation,[],[f13385_D])).
fof(f13385_D,plain,(
  ( ! [X23] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X23,sK13))))) | p(s(bool,X23))) ) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f14043,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X23,sK13))))) | p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f2786,f597])).
fof(f14036,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2786,f570])).
fof(f14035,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2786,f569])).
fof(f14034,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2786,f568])).
fof(f14033,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2786,f567])).
fof(f13884,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X58,sK13))))) ) | $spl354),
  inference(cnf_transformation,[],[f13884_D])).
fof(f13884_D,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X58,sK13))))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f13852,plain,(
  ( ! [X329,X328,X330] : (~p(s(bool,X328)) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X328,sK13))))) ) | $spl2),
  inference(resolution,[],[f2785,f901])).
fof(f2785,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X7,sK13))))) | ~p(s(bool,X7))) ) | $spl2),
  inference(resolution,[],[f913,f576])).
fof(f576,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK26(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X6)))) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f390])).
fof(f13759,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X86,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2785,f653])).
fof(f13756,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X74,sK13))))) ) | $spl2),
  inference(resolution,[],[f2785,f652])).
fof(f13755,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X70,sK13))))) ) | $spl2),
  inference(resolution,[],[f2785,f651])).
fof(f13754,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X66,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2785,f649])).
fof(f13752,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X58,sK13))))) ) | $spl2),
  inference(resolution,[],[f2785,f647])).
fof(f13742,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2785,f605])).
fof(f13741,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2785,f604])).
fof(f13740,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X25,sK13))))) | p(s(bool,X25))) ) | $spl2),
  inference(resolution,[],[f2785,f599])).
fof(f13739,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X23,sK13))))) | p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f2785,f597])).
fof(f13384,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X17,sK13))))) | ~p(s(bool,X17))) ) | $spl336),
  inference(cnf_transformation,[],[f13384_D])).
fof(f13384_D,plain,(
  ( ! [X17] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X17,sK13))))) | ~p(s(bool,X17))) ) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f13736,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X17,sK13))))) | ~p(s(bool,X17))) ) | $spl2),
  inference(resolution,[],[f2785,f578])).
fof(f13732,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2785,f570])).
fof(f13731,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2785,f569])).
fof(f13730,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2785,f568])).
fof(f13729,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2785,f567])).
fof(f13700,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X14),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X15))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f12918,f8903])).
fof(f12918,plain,(
  ( ! [X198,X196,X197] : (sP8(X196,'const.list.EXISTS_2'(s(fun(sK13,bool),X197),s('type.list.list'(sK13),'const.list.NIL_0')),X198)) ) | $spl0),
  inference(resolution,[],[f12768,f714])).
fof(f12768,plain,(
  ( ! [X76] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X76),s('type.list.list'(sK13),'const.list.NIL_0'))))) ) | $spl0),
  inference(forward_demodulation,[],[f12767,f3220])).
fof(f12767,plain,(
  ( ! [X76] : (~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X76),s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(subsumption_resolution,[],[f12766,f867])).
fof(f12766,plain,(
  ( ! [X76] : (p(s(bool,'F_0')) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X76),s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(forward_demodulation,[],[f12566,f629])).
fof(f12566,plain,(
  ( ! [X76] : (p(s(bool,'const.bool.IN_2'(s(sK13,sK50('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X76,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) | ~p(s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X76),s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f621,f3220])).
fof(f13699,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X12),s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X13))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f12918,f8824])).
fof(f13578,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X58,sK14,sK13))))) ) | $spl348),
  inference(cnf_transformation,[],[f13578_D])).
fof(f13578_D,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X58,sK14,sK13))))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f13546,plain,(
  ( ! [X329,X328,X330] : (~p(s(bool,X328)) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X328,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2784,f901])).
fof(f2784,plain,(
  ( ! [X6] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X6,sK14,sK13))))) | ~p(s(bool,X6))) ) | $spl2),
  inference(resolution,[],[f913,f573])).
fof(f573,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) )),
  inference(cnf_transformation,[],[f386])).
fof(f386,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK24(X2,X1,X0))))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK25(X2,X1,X0))))) & p(s(bool,X2)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK24,sK25])],[f385])).
fof(f385,plain,(
  ! [X0,X1,X2] : ((! [X6] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6)))) | ~p(s(bool,X2))) | (? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(rectify,[],[f384])).
fof(f384,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2)) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(flattening,[],[f383])).
fof(f383,plain,(
  ! [X0,X1,X2] : ((! [X3] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,X2))) | (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2)))) & ((! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | ~p(s(bool,X2))) | ? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2)))))),
  inference(nnf_transformation,[],[f148])).
fof(f148,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(flattening,[],[f147])).
fof(f147,plain,(
  ! [X0] : ! [X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) & p(s(bool,X2))) <=> (? [X4] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) & p(s(bool,X2))))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X6] : ! [X7,X25] : (? [X9] : (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) & p(s(bool,X25))) <=> (? [X9] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) & p(s(bool,X25))))),
  file('thm.list.NULL_FILTER','thm.bool.LEFT_EXISTS_AND_THM')).
fof(f13453,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X86,sK14,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2784,f653])).
fof(f13450,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X74,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2784,f652])).
fof(f13449,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X70,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2784,f651])).
fof(f13448,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X66,sK14,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2784,f649])).
fof(f13446,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X58,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2784,f647])).
fof(f13436,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2784,f605])).
fof(f13435,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2784,f604])).
fof(f13434,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X25,sK13))))) | p(s(bool,X25))) ) | $spl2),
  inference(resolution,[],[f2784,f599])).
fof(f13433,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X23,sK13))))) | p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f2784,f597])).
fof(f13430,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X17,sK13))))) | ~p(s(bool,X17))) ) | $spl2),
  inference(resolution,[],[f2784,f578])).
fof(f13383,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X15,sK13))))) | ~p(s(bool,X15))) ) | $spl334),
  inference(cnf_transformation,[],[f13383_D])).
fof(f13383_D,plain,(
  ( ! [X15] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X15,sK13))))) | ~p(s(bool,X15))) ) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f13429,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X15,sK13))))) | ~p(s(bool,X15))) ) | $spl2),
  inference(resolution,[],[f2784,f576])).
fof(f13426,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2784,f570])).
fof(f13425,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2784,f569])).
fof(f13424,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2784,f568])).
fof(f13423,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2784,f567])).
fof(f13389,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X58,sK14,sK13))))) ) | $spl346),
  inference(cnf_transformation,[],[f13389_D])).
fof(f13389_D,plain,(
  ( ! [X60,X58] : (~p(s(bool,X58)) | p(s(bool,X60)) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X58,sK14,sK13))))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f13346,plain,(
  ( ! [X329,X328,X330] : (~p(s(bool,X328)) | p(s(bool,X329)) | p(s(bool,X330)) | s(bool,X330) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X328,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2783,f901])).
fof(f2783,plain,(
  ( ! [X5] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X5,sK14,sK13))))) | ~p(s(bool,X5))) ) | $spl2),
  inference(resolution,[],[f913,f571])).
fof(f571,plain,(
  ( ! [X6,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK24(X2,X1,X0))))) | ~p(s(bool,X2)) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X6))))) )),
  inference(cnf_transformation,[],[f386])).
fof(f13253,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,X88)) | s(bool,X87) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X86,sK14,sK13)))) | s(bool,X88) != s(bool,X89)) ) | $spl2),
  inference(resolution,[],[f2783,f653])).
fof(f13250,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | ~p(s(bool,X76)) | p(s(bool,X77)) | s(bool,X77) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X74,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2783,f652])).
fof(f13249,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | p(s(bool,X72)) | s(bool,X71) != s(bool,X73) | s(bool,X72) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X70,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2783,f651])).
fof(f13248,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~p(s(bool,X68)) | s(bool,X67) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X66,sK14,sK13)))) | s(bool,X68) != s(bool,X69)) ) | $spl2),
  inference(resolution,[],[f2783,f649])).
fof(f13246,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | p(s(bool,X59)) | p(s(bool,X60)) | s(bool,X59) != s(bool,X61) | s(bool,X60) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK24(X58,sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2783,f647])).
fof(f13236,plain,(
  ( ! [X28,X29] : (~p(s(bool,X28)) | p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK41(sK14,X29,sK13))))) | ~p(s(bool,happ(s(fun(sK13,bool),X29),s(sK13,sK40(sK14,X29,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2783,f605])).
fof(f13235,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,happ(s(fun(sK13,bool),X27),s(sK13,sK41(sK14,X27,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK40(sK14,X27,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2783,f604])).
fof(f13234,plain,(
  ( ! [X24,X25] : (~p(s(bool,X24)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK37(sK14,X25,sK13))))) | p(s(bool,X25))) ) | $spl2),
  inference(resolution,[],[f2783,f599])).
fof(f13233,plain,(
  ( ! [X23,X22] : (~p(s(bool,X22)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK36(sK14,X23,sK13))))) | p(s(bool,X23))) ) | $spl2),
  inference(resolution,[],[f2783,f597])).
fof(f13230,plain,(
  ( ! [X17,X16] : (~p(s(bool,X16)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK27(sK14,X17,sK13))))) | ~p(s(bool,X17))) ) | $spl2),
  inference(resolution,[],[f2783,f578])).
fof(f13229,plain,(
  ( ! [X14,X15] : (~p(s(bool,X14)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK26(sK14,X15,sK13))))) | ~p(s(bool,X15))) ) | $spl2),
  inference(resolution,[],[f2783,f576])).
fof(f13382,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X13,sK14,sK13))))) | ~p(s(bool,X13))) ) | $spl332),
  inference(cnf_transformation,[],[f13382_D])).
fof(f13382_D,plain,(
  ( ! [X13] : (p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X13,sK14,sK13))))) | ~p(s(bool,X13))) ) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f13228,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK25(X13,sK14,sK13))))) | ~p(s(bool,X13))) ) | $spl2),
  inference(resolution,[],[f2783,f573])).
fof(f13226,plain,(
  ( ! [X8,X9] : (~p(s(bool,X8)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(sK14,X9,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X9),s(sK13,sK23(sK14,X9,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2783,f570])).
fof(f13225,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | p(s(bool,happ(s(fun(sK13,bool),X7),s(sK13,sK23(X7,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK23(X7,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2783,f569])).
fof(f13224,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK22(sK14,X5,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),X5),s(sK13,sK21(sK14,X5,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2783,f568])).
fof(f13223,plain,(
  ( ! [X2,X3] : (~p(s(bool,X2)) | p(s(bool,happ(s(fun(sK13,bool),X3),s(sK13,sK22(X3,sK14,sK13))))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK21(X3,sK14,sK13)))))) ) | $spl2),
  inference(resolution,[],[f2783,f567])).
fof(f13124,plain,(
  ( ! [X290,X289,X291] : (p(s(bool,X289)) | p(s(bool,X290)) | s(bool,X290) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X291),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12778,f901])).
fof(f13031,plain,(
  ( ! [X57,X58,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X57),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X56) != s(bool,X58)) ) | $spl0),
  inference(resolution,[],[f12778,f653])).
fof(f13028,plain,(
  ( ! [X47,X45,X46,X44] : (p(s(bool,X44)) | ~p(s(bool,X45)) | p(s(bool,X46)) | s(bool,X46) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X47),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12778,f652])).
fof(f13027,plain,(
  ( ! [X43,X41,X42,X40] : (p(s(bool,X40)) | p(s(bool,X41)) | s(bool,X40) != s(bool,X42) | s(bool,X41) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X43),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12778,f651])).
fof(f13026,plain,(
  ( ! [X39,X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X38),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X37) != s(bool,X39)) ) | $spl0),
  inference(resolution,[],[f12778,f649])).
fof(f13024,plain,(
  ( ! [X30,X28,X31,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),X31),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12778,f647])).
fof(f12966,plain,(
  ( ! [X292,X290,X291] : (p(s(bool,X290)) | ~p(s(bool,X291)) | s(bool,X291) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X292),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12768,f901])).
fof(f12868,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X47),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X45) != s(bool,X48)) ) | $spl0),
  inference(resolution,[],[f12768,f653])).
fof(f12867,plain,(
  ( ! [X43,X41,X44,X42] : (p(s(bool,X41)) | ~p(s(bool,X42)) | ~p(s(bool,X43)) | s(bool,X42) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X44),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12768,f652])).
fof(f12865,plain,(
  ( ! [X37,X35,X36,X34] : (p(s(bool,X34)) | ~p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X37),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12768,f651])).
fof(f12863,plain,(
  ( ! [X28,X26,X29,X27] : (~p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X28),s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X26) != s(bool,X29)) ) | $spl0),
  inference(resolution,[],[f12768,f649])).
fof(f12861,plain,(
  ( ! [X21,X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,X20) | s(bool,X19) != s(bool,'const.list.EXISTS_2'(s(fun(sK13,bool),X21),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl0),
  inference(resolution,[],[f12768,f647])).
fof(f12486,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl318),
  inference(cnf_transformation,[],[f12486_D])).
fof(f12486_D,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f12474,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4392,f900])).
fof(f4392,plain,(
  ( ! [X174] : (~sP11(X174,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | p(s(bool,X174))) ) | $spl23),
  inference(resolution,[],[f1093,f758])).
fof(f12462,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl316),
  inference(cnf_transformation,[],[f12462_D])).
fof(f12462_D,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) ) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f12449,plain,(
  ( ! [X0,X1] : (~sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2712,f900])).
fof(f2712,plain,(
  ( ! [X6,X7] : (~sP9(X7,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X6) | ~p(s(bool,X7)) | ~p(s(bool,X6))) ) | $spl25),
  inference(superposition,[],[f1095,f726])).
fof(f12419,plain,(
  ( ! [X0,X1] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl312),
  inference(cnf_transformation,[],[f12419_D])).
fof(f12419_D,plain,(
  ( ! [X0,X1] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl312),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl312])])).
fof(f12410,plain,(
  ( ! [X12,X13] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X12,X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3299,f900])).
fof(f3299,plain,(
  ( ! [X169,X170] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X169,X170)) ) | $spl0),
  inference(resolution,[],[f910,f735])).
fof(f11023,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl240),
  inference(cnf_transformation,[],[f11023_D])).
fof(f11023_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f12422,plain,(
  ( ! [X8,X9] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X8,X9)) ) | $spl314),
  inference(cnf_transformation,[],[f12422_D])).
fof(f12422_D,plain,(
  ( ! [X8,X9] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X8,X9)) ) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f12408,plain,(
  ( ! [X8,X9] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X8,X9) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3299,f590])).
fof(f11021,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) | $spl238),
  inference(cnf_transformation,[],[f11021_D])).
fof(f11021_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13))))) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f12407,plain,(
  ( ! [X6,X7] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X6,X7) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3299,f594])).
fof(f12406,plain,(
  ( ! [X4,X5] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X4,X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3299,f593])).
fof(f12405,plain,(
  ( ! [X2,X3] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2,X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3299,f889])).
fof(f12418,plain,(
  ( ! [X0,X1] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f12396,f892])).
fof(f892,plain,(
  ( ! [X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f287])).
fof(f287,plain,(
  ! [X0,X1] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(flattening,[],[f286])).
fof(f286,plain,(
  ! [X0] : ! [X1] : s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),'const.list.NIL_0')))),
  inference(rectify,[],[f127])).
fof(f127,axiom,(
  ! [X6] : ! [X7] : s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),'const.list.NIL_0')))),
  file('thm.list.NULL_FILTER','thm.list.FILTER|split|0')).
fof(f12396,plain,(
  ( ! [X0,X1] : (sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3299,f900])).
fof(f12384,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl310),
  inference(cnf_transformation,[],[f12384_D])).
fof(f12384_D,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f12372,plain,(
  ( ! [X0,X1] : (sP9(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2710,f900])).
fof(f2710,plain,(
  ( ! [X2,X3] : (sP9(X3,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X2) | ~p(s(bool,X2)) | p(s(bool,X3))) ) | $spl25),
  inference(superposition,[],[f1095,f724])).
fof(f12361,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl308),
  inference(cnf_transformation,[],[f12361_D])).
fof(f12361_D,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f12348,plain,(
  ( ! [X0,X1] : (sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2709,f900])).
fof(f2709,plain,(
  ( ! [X0,X1] : (sP9(X1,X0,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X0)) | p(s(bool,X1))) ) | $spl25),
  inference(superposition,[],[f1095,f724])).
fof(f12310,plain,(
  ( ! [X0,X1] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl304),
  inference(cnf_transformation,[],[f12310_D])).
fof(f12310_D,plain,(
  ( ! [X0,X1] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f12296,plain,(
  ( ! [X12,X13] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X12,X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3292,f900])).
fof(f3292,plain,(
  ( ! [X156,X155] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X155,X156)) ) | $spl0),
  inference(resolution,[],[f910,f723])).
fof(f12313,plain,(
  ( ! [X8,X9] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X8,X9)) ) | $spl306),
  inference(cnf_transformation,[],[f12313_D])).
fof(f12313_D,plain,(
  ( ! [X8,X9] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X8,X9)) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f12294,plain,(
  ( ! [X8,X9] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X8,X9) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3292,f590])).
fof(f12293,plain,(
  ( ! [X6,X7] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X6,X7) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3292,f594])).
fof(f12292,plain,(
  ( ! [X4,X5] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X4,X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3292,f593])).
fof(f12291,plain,(
  ( ! [X2,X3] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2,X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3292,f889])).
fof(f12309,plain,(
  ( ! [X0,X1] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f12282,f892])).
fof(f12282,plain,(
  ( ! [X0,X1] : (sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3292,f900])).
fof(f12308,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | $spl303),
  inference(cnf_transformation,[],[f12308_D])).
fof(f12308_D,plain,(
  s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) <=> ~$spl303),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl303])])).
fof(f12275,plain,(
  ( ! [X9] : (p(s(bool,X9)) | s(bool,'F_0') != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | ($spl0 | $spl25)),
  inference(resolution,[],[f3292,f8825])).
fof(f12267,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl300),
  inference(cnf_transformation,[],[f12267_D])).
fof(f12267_D,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f12243,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2700,f900])).
fof(f2700,plain,(
  ( ! [X194,X195] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X194,X195) | ~sP4(X194,X195,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl25),
  inference(resolution,[],[f1095,f904])).
fof(f904,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f694])).
fof(f694,plain,(
  ( ! [X2,X0,X1] : (sP5(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP4(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f463])).
fof(f12230,plain,(
  ( ! [X0,X1] : (~sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) ) | $spl298),
  inference(cnf_transformation,[],[f12230_D])).
fof(f12230_D,plain,(
  ( ! [X0,X1] : (~sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f12216,plain,(
  ( ! [X0,X1] : (~sP6('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2658,f900])).
fof(f12191,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl294)),
  inference(subsumption_resolution,[],[f12177,f12136])).
fof(f12177,plain,(
  ( ! [X0,X1] : (~sP6(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2657,f900])).
fof(f12136,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl294),
  inference(cnf_transformation,[],[f12136_D])).
fof(f12136_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f12124,plain,(
  ( ! [X12,X13] : (sP6(X12,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3291,f900])).
fof(f3291,plain,(
  ( ! [X154,X153] : (sP6(X153,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X154)) ) | $spl0),
  inference(resolution,[],[f910,f722])).
fof(f722,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f478])).
fof(f12139,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9)) ) | $spl296),
  inference(cnf_transformation,[],[f12139_D])).
fof(f12139_D,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9)) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f12122,plain,(
  ( ! [X8,X9] : (sP6(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3291,f590])).
fof(f12121,plain,(
  ( ! [X6,X7] : (sP6(X6,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X7) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3291,f594])).
fof(f12120,plain,(
  ( ! [X4,X5] : (sP6(X4,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3291,f593])).
fof(f12119,plain,(
  ( ! [X2,X3] : (sP6(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3291,f889])).
fof(f12135,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f12110,f892])).
fof(f12110,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3291,f900])).
fof(f12096,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl290)),
  inference(subsumption_resolution,[],[f12083,f11979])).
fof(f12083,plain,(
  ( ! [X0,X1] : (~sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2656,f900])).
fof(f12063,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl282)),
  inference(subsumption_resolution,[],[f12052,f11840])).
fof(f12052,plain,(
  ( ! [X0,X1] : (~sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2653,f900])).
fof(f12029,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl276)),
  inference(subsumption_resolution,[],[f12018,f11665])).
fof(f12018,plain,(
  ( ! [X0,X1] : (~sP8(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2650,f900])).
fof(f2650,plain,(
  ( ! [X134,X135] : (~sP8(X134,X135,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X135)) | ~p(s(bool,X134))) ) | $spl25),
  inference(resolution,[],[f1095,f712])).
fof(f712,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f473])).
fof(f11979,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl290),
  inference(cnf_transformation,[],[f11979_D])).
fof(f11979_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f11967,plain,(
  ( ! [X12,X13] : (sP6(X12,X13,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3290,f900])).
fof(f3290,plain,(
  ( ! [X152,X151] : (sP6(X151,X152,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f721])).
fof(f11982,plain,(
  ( ! [X8,X9] : (sP6(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl292),
  inference(cnf_transformation,[],[f11982_D])).
fof(f11982_D,plain,(
  ( ! [X8,X9] : (sP6(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f11965,plain,(
  ( ! [X8,X9] : (sP6(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3290,f590])).
fof(f11964,plain,(
  ( ! [X6,X7] : (sP6(X6,X7,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3290,f594])).
fof(f11963,plain,(
  ( ! [X4,X5] : (sP6(X4,X5,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3290,f593])).
fof(f11962,plain,(
  ( ! [X2,X3] : (sP6(X2,X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3290,f889])).
fof(f11978,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11953,f892])).
fof(f11953,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3290,f900])).
fof(f11942,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl288),
  inference(cnf_transformation,[],[f11942_D])).
fof(f11942_D,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f11928,plain,(
  ( ! [X0,X1] : (~sP9(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2647,f900])).
fof(f2647,plain,(
  ( ! [X128,X129] : (~sP9(X128,X129,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X129)) | ~p(s(bool,X128))) ) | $spl25),
  inference(resolution,[],[f1095,f708])).
fof(f708,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f470])).
fof(f11915,plain,(
  ( ! [X0,X1] : (~sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl286),
  inference(cnf_transformation,[],[f11915_D])).
fof(f11915_D,plain,(
  ( ! [X0,X1] : (~sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f11902,plain,(
  ( ! [X0,X1] : (~sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2646,f900])).
fof(f2646,plain,(
  ( ! [X127,X126] : (~sP5('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X126,X127) | ~p(s(bool,X127)) | ~p(s(bool,X126))) ) | $spl25),
  inference(resolution,[],[f1095,f704])).
fof(f704,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP5(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f467])).
fof(f11888,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl268)),
  inference(subsumption_resolution,[],[f11877,f11515])).
fof(f11877,plain,(
  ( ! [X0,X1] : (~sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2637,f900])).
fof(f2637,plain,(
  ( ! [X109,X108] : (~sP4(X108,X109,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X109)) | ~p(s(bool,X108))) ) | $spl25),
  inference(resolution,[],[f1095,f695])).
fof(f695,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f465])).
fof(f11840,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl282),
  inference(cnf_transformation,[],[f11840_D])).
fof(f11840_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f11829,plain,(
  ( ! [X12,X13] : (sP7(X12,X13,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3289,f900])).
fof(f3289,plain,(
  ( ! [X149,X150] : (sP7(X149,X150,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f717])).
fof(f11843,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl284),
  inference(cnf_transformation,[],[f11843_D])).
fof(f11843_D,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f11827,plain,(
  ( ! [X8,X9] : (sP7(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3289,f590])).
fof(f11826,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3289,f594])).
fof(f11825,plain,(
  ( ! [X4,X5] : (sP7(X4,X5,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3289,f593])).
fof(f11824,plain,(
  ( ! [X2,X3] : (sP7(X2,X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3289,f889])).
fof(f11839,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11815,f892])).
fof(f11815,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3289,f900])).
fof(f11804,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl264)),
  inference(subsumption_resolution,[],[f11778,f11411])).
fof(f11411,plain,(
  ( ! [X2,X3] : (sP3(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,X2))) ) | $spl264),
  inference(resolution,[],[f11405,f903])).
fof(f903,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f674])).
fof(f674,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~sP2(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f455])).
fof(f11778,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0)) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2630,f900])).
fof(f2630,plain,(
  ( ! [X94,X95] : (~sP3(X95,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X94) | ~p(s(bool,X95)) | p(s(bool,X94))) ) | $spl25),
  inference(resolution,[],[f1095,f685])).
fof(f685,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP3(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f460])).
fof(f11761,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl260)),
  inference(subsumption_resolution,[],[f11735,f11296])).
fof(f11296,plain,(
  ( ! [X2,X3] : (sP3(X2,X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X2))) ) | $spl260),
  inference(resolution,[],[f11290,f903])).
fof(f11735,plain,(
  ( ! [X0,X1] : (~sP3(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2629,f900])).
fof(f2629,plain,(
  ( ! [X92,X93] : (~sP3(X93,X92,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X93)) | p(s(bool,X92))) ) | $spl25),
  inference(resolution,[],[f1095,f685])).
fof(f11717,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1))) ) | $spl280),
  inference(cnf_transformation,[],[f11717_D])).
fof(f11717_D,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1))) ) <=> ~$spl280),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl280])])).
fof(f11716,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl264)),
  inference(subsumption_resolution,[],[f11695,f11411])).
fof(f11695,plain,(
  ( ! [X0,X1] : (sP3(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2628,f900])).
fof(f2628,plain,(
  ( ! [X90,X91] : (sP3(X90,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X91) | p(s(bool,X91)) | p(s(bool,X90))) ) | $spl25),
  inference(resolution,[],[f1095,f682])).
fof(f682,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f460])).
fof(f11665,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl276),
  inference(cnf_transformation,[],[f11665_D])).
fof(f11665_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl276),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl276])])).
fof(f11654,plain,(
  ( ! [X12,X13] : (sP8(X12,X13,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3286,f900])).
fof(f3286,plain,(
  ( ! [X144,X143] : (sP8(X143,X144,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f713])).
fof(f713,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f473])).
fof(f11668,plain,(
  ( ! [X8,X9] : (sP8(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl278),
  inference(cnf_transformation,[],[f11668_D])).
fof(f11668_D,plain,(
  ( ! [X8,X9] : (sP8(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl278),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl278])])).
fof(f11652,plain,(
  ( ! [X8,X9] : (sP8(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3286,f590])).
fof(f11651,plain,(
  ( ! [X6,X7] : (sP8(X6,X7,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3286,f594])).
fof(f11650,plain,(
  ( ! [X4,X5] : (sP8(X4,X5,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3286,f593])).
fof(f11649,plain,(
  ( ! [X2,X3] : (sP8(X2,X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3286,f889])).
fof(f11664,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11640,f892])).
fof(f11640,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3286,f900])).
fof(f11628,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1))) ) | $spl274),
  inference(cnf_transformation,[],[f11628_D])).
fof(f11628_D,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f11627,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl260)),
  inference(subsumption_resolution,[],[f11606,f11296])).
fof(f11606,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2627,f900])).
fof(f2627,plain,(
  ( ! [X88,X89] : (sP3(X88,X89,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X89)) | p(s(bool,X88))) ) | $spl25),
  inference(resolution,[],[f1095,f682])).
fof(f11593,plain,(
  ( ! [X0,X1] : (sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl272),
  inference(cnf_transformation,[],[f11593_D])).
fof(f11593_D,plain,(
  ( ! [X0,X1] : (sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f11579,plain,(
  ( ! [X0,X1] : (sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2626,f900])).
fof(f2626,plain,(
  ( ! [X87,X86] : (sP3('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X86,X87) | p(s(bool,X86)) | p(s(bool,X87))) ) | $spl25),
  inference(resolution,[],[f1095,f682])).
fof(f11566,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl264)),
  inference(subsumption_resolution,[],[f11554,f11405])).
fof(f11554,plain,(
  ( ! [X0,X1] : (~sP2(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2624,f900])).
fof(f2624,plain,(
  ( ! [X83,X82] : (~sP2(X82,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X83) | p(s(bool,X83)) | ~p(s(bool,X82))) ) | $spl25),
  inference(resolution,[],[f1095,f678])).
fof(f678,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f458])).
fof(f11515,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl268),
  inference(cnf_transformation,[],[f11515_D])).
fof(f11515_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f11504,plain,(
  ( ! [X12,X13] : (sP4(X12,X13,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3273,f900])).
fof(f3273,plain,(
  ( ! [X118,X117] : (sP4(X117,X118,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f696])).
fof(f696,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f465])).
fof(f11518,plain,(
  ( ! [X8,X9] : (sP4(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl270),
  inference(cnf_transformation,[],[f11518_D])).
fof(f11518_D,plain,(
  ( ! [X8,X9] : (sP4(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f11502,plain,(
  ( ! [X8,X9] : (sP4(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3273,f590])).
fof(f11501,plain,(
  ( ! [X6,X7] : (sP4(X6,X7,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3273,f594])).
fof(f11500,plain,(
  ( ! [X4,X5] : (sP4(X4,X5,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3273,f593])).
fof(f11499,plain,(
  ( ! [X2,X3] : (sP4(X2,X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3273,f889])).
fof(f11514,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11490,f892])).
fof(f11490,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3273,f900])).
fof(f11478,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl260)),
  inference(subsumption_resolution,[],[f11466,f11290])).
fof(f11466,plain,(
  ( ! [X0,X1] : (~sP2(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2623,f900])).
fof(f2623,plain,(
  ( ! [X80,X81] : (~sP2(X80,X81,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X81)) | ~p(s(bool,X80))) ) | $spl25),
  inference(resolution,[],[f1095,f678])).
fof(f11405,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl264),
  inference(cnf_transformation,[],[f11405_D])).
fof(f11405_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f11395,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3257,f900])).
fof(f3257,plain,(
  ( ! [X85,X86] : (sP2(X85,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X86)) ) | $spl0),
  inference(resolution,[],[f910,f680])).
fof(f680,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f458])).
fof(f11408,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9)) ) | $spl266),
  inference(cnf_transformation,[],[f11408_D])).
fof(f11408_D,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9)) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f11393,plain,(
  ( ! [X8,X9] : (sP2(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3257,f590])).
fof(f11392,plain,(
  ( ! [X6,X7] : (sP2(X6,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X7) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3257,f594])).
fof(f11391,plain,(
  ( ! [X4,X5] : (sP2(X4,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3257,f593])).
fof(f11390,plain,(
  ( ! [X2,X3] : (sP2(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3257,f889])).
fof(f11404,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11381,f892])).
fof(f11381,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3257,f900])).
fof(f11369,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl254)),
  inference(subsumption_resolution,[],[f11346,f11187])).
fof(f11187,plain,(
  ( ! [X2,X3] : (sP1(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,X2))) ) | $spl254),
  inference(resolution,[],[f11181,f902])).
fof(f902,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(duplicate_literal_removal,[],[f657])).
fof(f657,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~sP0(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f447])).
fof(f11346,plain,(
  ( ! [X0,X1] : (~sP1(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2613,f900])).
fof(f2613,plain,(
  ( ! [X61,X60] : (~sP1(X61,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X60) | ~p(s(bool,X61)) | ~p(s(bool,X60))) ) | $spl25),
  inference(resolution,[],[f1095,f668])).
fof(f668,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP1(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f452])).
fof(f11290,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl260),
  inference(cnf_transformation,[],[f11290_D])).
fof(f11290_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f11280,plain,(
  ( ! [X12,X13] : (sP2(X12,X13,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3256,f900])).
fof(f3256,plain,(
  ( ! [X83,X84] : (sP2(X83,X84,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f679])).
fof(f679,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f458])).
fof(f11293,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl262),
  inference(cnf_transformation,[],[f11293_D])).
fof(f11293_D,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f11278,plain,(
  ( ! [X8,X9] : (sP2(X8,X9,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3256,f590])).
fof(f11277,plain,(
  ( ! [X6,X7] : (sP2(X6,X7,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3256,f594])).
fof(f11276,plain,(
  ( ! [X4,X5] : (sP2(X4,X5,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3256,f593])).
fof(f11275,plain,(
  ( ! [X2,X3] : (sP2(X2,X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3256,f889])).
fof(f11289,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11266,f892])).
fof(f11266,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3256,f900])).
fof(f11254,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1))) ) | $spl258),
  inference(cnf_transformation,[],[f11254_D])).
fof(f11254_D,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1))) ) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f11253,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | ($spl25 | $spl254)),
  inference(subsumption_resolution,[],[f11235,f11187])).
fof(f11235,plain,(
  ( ! [X0,X1] : (sP1(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2611,f900])).
fof(f2611,plain,(
  ( ! [X57,X56] : (sP1(X56,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X57) | ~p(s(bool,X57)) | p(s(bool,X56))) ) | $spl25),
  inference(resolution,[],[f1095,f665])).
fof(f665,plain,(
  ( ! [X2,X0,X1] : (sP1(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f452])).
fof(f11181,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl254),
  inference(cnf_transformation,[],[f11181_D])).
fof(f11181_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl254),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl254])])).
fof(f11171,plain,(
  ( ! [X12,X13] : (sP0(X12,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X13) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3243,f900])).
fof(f3243,plain,(
  ( ! [X57,X58] : (sP0(X57,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X58)) ) | $spl0),
  inference(resolution,[],[f910,f663])).
fof(f663,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f450])).
fof(f11184,plain,(
  ( ! [X8,X9] : (sP0(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9)) ) | $spl256),
  inference(cnf_transformation,[],[f11184_D])).
fof(f11184_D,plain,(
  ( ! [X8,X9] : (sP0(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9)) ) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f11169,plain,(
  ( ! [X8,X9] : (sP0(X8,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X9) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3243,f590])).
fof(f11168,plain,(
  ( ! [X6,X7] : (sP0(X6,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X7) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3243,f594])).
fof(f11167,plain,(
  ( ! [X4,X5] : (sP0(X4,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X5) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3243,f593])).
fof(f11166,plain,(
  ( ! [X2,X3] : (sP0(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3243,f889])).
fof(f11180,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11157,f892])).
fof(f11157,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3243,f900])).
fof(f11146,plain,(
  ( ! [X0,X1] : (sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) | $spl252),
  inference(cnf_transformation,[],[f11146_D])).
fof(f11146_D,plain,(
  ( ! [X0,X1] : (sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0))) ) <=> ~$spl252),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl252])])).
fof(f11134,plain,(
  ( ! [X0,X1] : (sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2610,f900])).
fof(f2610,plain,(
  ( ! [X54,X55] : (sP1('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X54,X55) | p(s(bool,X54)) | ~p(s(bool,X55))) ) | $spl25),
  inference(resolution,[],[f1095,f665])).
fof(f11122,plain,(
  ( ! [X0,X1] : (~sP0(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl250),
  inference(cnf_transformation,[],[f11122_D])).
fof(f11122_D,plain,(
  ( ! [X0,X1] : (~sP0(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) ) <=> ~$spl250),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl250])])).
fof(f11111,plain,(
  ( ! [X0,X1] : (~sP0(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X1)) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2607,f900])).
fof(f2607,plain,(
  ( ! [X48,X49] : (~sP0(X48,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X49) | ~p(s(bool,X49)) | ~p(s(bool,X48))) ) | $spl25),
  inference(resolution,[],[f1095,f661])).
fof(f11075,plain,(
  ( ! [X0] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0)) ) | $spl246),
  inference(cnf_transformation,[],[f11075_D])).
fof(f11075_D,plain,(
  ( ! [X0] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0)) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f11064,plain,(
  ( ! [X6] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X6) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3311,f900])).
fof(f3311,plain,(
  ( ! [X184] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))),X184)) ) | $spl0),
  inference(resolution,[],[f910,f776])).
fof(f11078,plain,(
  ( ! [X4] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X4)) ) | $spl248),
  inference(cnf_transformation,[],[f11078_D])).
fof(f11078_D,plain,(
  ( ! [X4] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X4)) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f11062,plain,(
  ( ! [X4] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X4) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3311,f590])).
fof(f11061,plain,(
  ( ! [X3] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X3) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3311,f594])).
fof(f11060,plain,(
  ( ! [X2] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X2) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3311,f593])).
fof(f11059,plain,(
  ( ! [X1] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3311,f889])).
fof(f11074,plain,(
  ( ! [X0] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f11050,f892])).
fof(f11050,plain,(
  ( ! [X0] : (sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3311,f900])).
fof(f11013,plain,(
  ( ! [X0] : (sP12(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl232),
  inference(cnf_transformation,[],[f11013_D])).
fof(f11013_D,plain,(
  ( ! [X0] : (sP12(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f11002,plain,(
  ( ! [X6] : (sP12(X6,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13)),s('type.list.list'(sK13),sK55('const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),sK13))))) ) | $spl0),
  inference(superposition,[],[f3310,f900])).
fof(f3310,plain,(
  ( ! [X183] : (sP12(X183,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f775])).
fof(f11024,plain,(
  ( ! [X4] : (sP12(X4,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl242),
  inference(cnf_transformation,[],[f11024_D])).
fof(f11024_D,plain,(
  ( ! [X4] : (sP12(X4,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl242),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl242])])).
fof(f11000,plain,(
  ( ! [X4] : (sP12(X4,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK34(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3310,f590])).
fof(f10999,plain,(
  ( ! [X3] : (sP12(X3,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK35(sK15,sK14,sK13)))))) ) | $spl0),
  inference(superposition,[],[f3310,f594])).
fof(f10998,plain,(
  ( ! [X2] : (sP12(X2,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,'const.bool.IN_2'(s(sK13,sK35(sK15,sK14,sK13)),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))))) ) | $spl0),
  inference(superposition,[],[f3310,f593])).
fof(f10997,plain,(
  ( ! [X1] : (sP12(X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),'_dst_x0x1_2'(s(fun(bool,bool),'~_0'),s(fun(sK13,bool),sK14))),s('type.list.list'(sK13),sK15))))) ) | $spl0),
  inference(superposition,[],[f3310,f889])).
fof(f11012,plain,(
  ( ! [X0] : (sP12(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(forward_demodulation,[],[f10988,f892])).
fof(f10988,plain,(
  ( ! [X0] : (sP12(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl0),
  inference(superposition,[],[f3310,f900])).
fof(f10979,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1)) ) | $spl230),
  inference(cnf_transformation,[],[f10979_D])).
fof(f10979_D,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1)) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f10967,plain,(
  ( ! [X0,X1] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3097,f900])).
fof(f3097,plain,(
  ( ! [X180,X181] : (sP10('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X180,X181)) ) | $spl4),
  inference(resolution,[],[f915,f735])).
fof(f10955,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1)) ) | $spl228),
  inference(cnf_transformation,[],[f10955_D])).
fof(f10955_D,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1)) ) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f10938,plain,(
  ( ! [X0,X1] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3090,f900])).
fof(f3090,plain,(
  ( ! [X167,X166] : (sP6('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X166,X167)) ) | $spl4),
  inference(resolution,[],[f915,f723])).
fof(f10954,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | $spl227),
  inference(cnf_transformation,[],[f10954_D])).
fof(f10954_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) <=> ~$spl227),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl227])])).
fof(f10932,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | ($spl4 | $spl25)),
  inference(resolution,[],[f3090,f8825])).
fof(f10926,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1)) ) | $spl224),
  inference(cnf_transformation,[],[f10926_D])).
fof(f10926_D,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1)) ) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f10912,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3089,f900])).
fof(f3089,plain,(
  ( ! [X165,X164] : (sP6(X164,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X165)) ) | $spl4),
  inference(resolution,[],[f915,f722])).
fof(f10902,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) | $spl222),
  inference(cnf_transformation,[],[f10902_D])).
fof(f10902_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f10888,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3088,f900])).
fof(f3088,plain,(
  ( ! [X163,X162] : (sP6(X162,X163,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f721])).
fof(f10877,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) | $spl220),
  inference(cnf_transformation,[],[f10877_D])).
fof(f10877_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f10864,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3087,f900])).
fof(f3087,plain,(
  ( ! [X161,X160] : (sP7(X160,X161,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f717])).
fof(f10854,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) | $spl218),
  inference(cnf_transformation,[],[f10854_D])).
fof(f10854_D,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f10841,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3084,f900])).
fof(f3084,plain,(
  ( ! [X154,X155] : (sP8(X154,X155,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f713])).
fof(f10831,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) | $spl216),
  inference(cnf_transformation,[],[f10831_D])).
fof(f10831_D,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f10819,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3071,f900])).
fof(f3071,plain,(
  ( ! [X128,X129] : (sP4(X128,X129,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f696])).
fof(f10807,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X0)) ) | $spl214),
  inference(cnf_transformation,[],[f10807_D])).
fof(f10807_D,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X0)) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f10796,plain,(
  ( ! [X0] : (~sP6(X0,'F_0','const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,'F_0') != s(bool,X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f8827,f900])).
fof(f8827,plain,(
  ( ! [X14] : (~sP6(X14,'F_0','const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,'F_0') != s(bool,X14)) ) | $spl25),
  inference(subsumption_resolution,[],[f8823,f2652])).
fof(f2652,plain,(
  ( ! [X138,X139] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X138,X139)) ) | $spl25),
  inference(resolution,[],[f1095,f715])).
fof(f8823,plain,(
  ( ! [X14] : (~sP8('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X14,'F_0') | ~sP6(X14,'F_0','const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,'F_0') != s(bool,X14)) ) | $spl25),
  inference(resolution,[],[f8753,f2660])).
fof(f2660,plain,(
  ( ! [X154,X155] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X155,X154) | s(bool,X154) != s(bool,X155)) ) | $spl25),
  inference(resolution,[],[f1095,f727])).
fof(f10784,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | $spl213),
  inference(cnf_transformation,[],[f10784_D])).
fof(f10784_D,plain,(
  s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f10773,plain,(
  ( ! [X18] : (p(s(bool,X18)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | ($spl25 | $spl28)),
  inference(resolution,[],[f8825,f4701])).
fof(f4701,plain,(
  ( ! [X171,X170] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))),X170,X171)) ) | $spl28),
  inference(resolution,[],[f1265,f723])).
fof(f1265,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) | $spl28),
  inference(cnf_transformation,[],[f1265_D])).
fof(f1265_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f10772,plain,(
  ( ! [X17] : (p(s(bool,X17)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | ($spl2 | $spl25)),
  inference(resolution,[],[f8825,f5191])).
fof(f5191,plain,(
  ( ! [X171,X170] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))),X170,X171)) ) | $spl2),
  inference(resolution,[],[f2794,f723])).
fof(f2794,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) | $spl2),
  inference(resolution,[],[f913,f607])).
fof(f10771,plain,(
  ( ! [X16] : (p(s(bool,X16)) | s(bool,'F_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | ($spl2 | $spl25)),
  inference(resolution,[],[f8825,f4957])).
fof(f4957,plain,(
  ( ! [X171,X170] : (sP6(happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))),X170,X171)) ) | $spl2),
  inference(resolution,[],[f2793,f723])).
fof(f2793,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) | $spl2),
  inference(resolution,[],[f913,f606])).
fof(f10758,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1)) ) | $spl206),
  inference(cnf_transformation,[],[f10758_D])).
fof(f10758_D,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1)) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f10745,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3055,f900])).
fof(f3055,plain,(
  ( ! [X97,X96] : (sP2(X96,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X97)) ) | $spl4),
  inference(resolution,[],[f915,f680])).
fof(f10734,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) | $spl204),
  inference(cnf_transformation,[],[f10734_D])).
fof(f10734_D,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f10721,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3054,f900])).
fof(f3054,plain,(
  ( ! [X94,X95] : (sP2(X94,X95,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f679])).
fof(f10695,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1)) ) | $spl202),
  inference(cnf_transformation,[],[f10695_D])).
fof(f10695_D,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1)) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f10682,plain,(
  ( ! [X0,X1] : (sP0(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3041,f900])).
fof(f3041,plain,(
  ( ! [X68,X69] : (sP0(X68,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X69)) ) | $spl4),
  inference(resolution,[],[f915,f663])).
fof(f8923,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | $spl181),
  inference(cnf_transformation,[],[f8923_D])).
fof(f8923_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) <=> ~$spl181),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl181])])).
fof(f10626,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | ~p(s(bool,X9))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f5499,f8903])).
fof(f5499,plain,(
  ( ! [X146,X147] : (sP8(X146,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))),X147)) ) | $spl27),
  inference(resolution,[],[f1101,f714])).
fof(f8963,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | $spl190),
  inference(cnf_transformation,[],[f8963_D])).
fof(f8963_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f10625,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | ~p(s(bool,X8))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f5499,f8824])).
fof(f8927,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | $spl185),
  inference(cnf_transformation,[],[f8927_D])).
fof(f8927_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) <=> ~$spl185),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl185])])).
fof(f10442,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | ~p(s(bool,X9))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f7256,f8903])).
fof(f7256,plain,(
  ( ! [X146,X147] : (sP8(X146,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))),X147)) ) | $spl27),
  inference(resolution,[],[f5436,f714])).
fof(f5436,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) | $spl27),
  inference(resolution,[],[f1101,f609])).
fof(f609,plain,(
  ( ! [X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK45(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f423])).
fof(f423,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X1,X0)))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK45(X1,X0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK44,sK45])],[f422])).
fof(f422,plain,(
  ! [X0,X1] : ((! [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))) | ? [X4] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(rectify,[],[f421])).
fof(f421,plain,(
  ! [X0,X1] : ((! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))))),
  inference(nnf_transformation,[],[f303])).
fof(f303,plain,(
  ! [X0,X1] : (? [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(ennf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0] : ! [X1] : (~! [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) <=> ? [X3] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6] : ! [X7] : (~! [X9] : p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))) <=> ? [X9] : ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9)))))),
  file('thm.list.NULL_FILTER','thm.bool.NOT_FORALL_THM')).
fof(f8967,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | $spl194),
  inference(cnf_transformation,[],[f8967_D])).
fof(f8967_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f10441,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | ~p(s(bool,X8))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f7256,f8824])).
fof(f8925,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | $spl183),
  inference(cnf_transformation,[],[f8925_D])).
fof(f8925_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) <=> ~$spl183),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl183])])).
fof(f10159,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | ~p(s(bool,X9))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f7093,f8903])).
fof(f7093,plain,(
  ( ! [X146,X147] : (sP8(X146,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))),X147)) ) | $spl27),
  inference(resolution,[],[f5435,f714])).
fof(f5435,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) | $spl27),
  inference(resolution,[],[f1101,f608])).
fof(f608,plain,(
  ( ! [X0,X5,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK44(X1,X0))))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f423])).
fof(f8965,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | $spl192),
  inference(cnf_transformation,[],[f8965_D])).
fof(f8965_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f10158,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | ~p(s(bool,X8))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f7093,f8824])).
fof(f5577,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl140),
  inference(cnf_transformation,[],[f5577_D])).
fof(f5577_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f10089,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f901])).
fof(f9991,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl27),
  inference(resolution,[],[f1101,f653])).
fof(f9990,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f652])).
fof(f9988,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f651])).
fof(f9986,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl27),
  inference(resolution,[],[f1101,f649])).
fof(f9984,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f647])).
fof(f9775,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9684,f1082])).
fof(f9684,plain,(
  ( ! [X520,X521,X519] : (p(s(bool,'const.bool.IN_2'(s(X519,sK51(X520,X521,X519)),s(fun(X519,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X519),X520))))))) ) | $spl18),
  inference(resolution,[],[f1082,f625])).
fof(f9774,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9635,f1082])).
fof(f9635,plain,(
  ( ! [X366,X365] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X365),X366))))) ) | $spl18),
  inference(resolution,[],[f1082,f627])).
fof(f9773,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9634,f1082])).
fof(f9634,plain,(
  ( ! [X364,X362,X363] : (p(s(bool,'const.list.EVERY_2'(s(fun(X362,bool),X363),s('type.list.list'(X362),X364))))) ) | $spl18),
  inference(resolution,[],[f1082,f625])).
fof(f9772,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9620,f1082])).
fof(f9620,plain,(
  ( ! [X318,X317] : (p(s(bool,'const.bool.IN_2'(s(X317,sK52(X318,X317)),s(fun(X317,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X317),X318))))))) ) | $spl18),
  inference(resolution,[],[f1082,f627])).
fof(f9767,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9607,f1082])).
fof(f9607,plain,(
  p(s(bool,'F_0')) | $spl18),
  inference(resolution,[],[f1082,f874])).
fof(f9766,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9606,f1082])).
fof(f9606,plain,(
  p(s(bool,'F_0')) | $spl18),
  inference(resolution,[],[f1082,f872])).
fof(f9765,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9600,f1082])).
fof(f9600,plain,(
  ( ! [X297] : (p(s(bool,X297))) ) | $spl18),
  inference(resolution,[],[f1082,f819])).
fof(f9764,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9599,f1082])).
fof(f9599,plain,(
  ( ! [X296] : (p(s(bool,X296))) ) | $spl18),
  inference(resolution,[],[f1082,f816])).
fof(f9763,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9597,f1082])).
fof(f9597,plain,(
  p(s(bool,'F_0')) | $spl18),
  inference(resolution,[],[f1082,f813])).
fof(f9762,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f9596,f1082])).
fof(f9596,plain,(
  ( ! [X294] : (p(s(bool,X294))) ) | $spl18),
  inference(resolution,[],[f1082,f810])).
fof(f9608,plain,(
  $false | $spl18),
  inference(resolution,[],[f1082,f876])).
fof(f9609,plain,(
  $false | $spl18),
  inference(resolution,[],[f1082,f877])).
fof(f9610,plain,(
  $false | $spl18),
  inference(resolution,[],[f1082,f906])).
fof(f9621,plain,(
  $false | $spl18),
  inference(resolution,[],[f1082,f886])).
fof(f9622,plain,(
  $false | ($spl0 | $spl18)),
  inference(resolution,[],[f1082,f910])).
fof(f9636,plain,(
  $false | ($spl4 | $spl18)),
  inference(resolution,[],[f1082,f915])).
fof(f9679,plain,(
  $false | ($spl2 | $spl18)),
  inference(resolution,[],[f1082,f2793])).
fof(f9680,plain,(
  $false | ($spl2 | $spl18)),
  inference(resolution,[],[f1082,f2794])).
fof(f9681,plain,(
  $false | ($spl18 | $spl28)),
  inference(resolution,[],[f1082,f1265])).
fof(f9484,plain,(
  $false | ($spl0 | $spl25 | $spl26)),
  inference(subsumption_resolution,[],[f3381,f1100])).
fof(f1100,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) | $spl26),
  inference(cnf_transformation,[],[f1100_D])).
fof(f1100_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f3381,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) | ($spl0 | $spl25)),
  inference(subsumption_resolution,[],[f3349,f1095])).
fof(f3349,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) | $spl0),
  inference(superposition,[],[f910,f591])).
fof(f591,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK34(X2,X1,X0))))) | s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) )),
  inference(cnf_transformation,[],[f404])).
fof(f7334,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl166),
  inference(cnf_transformation,[],[f7334_D])).
fof(f7334_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f9447,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8067,f901])).
fof(f8067,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) | $spl23),
  inference(resolution,[],[f4310,f609])).
fof(f4310,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))))) | $spl23),
  inference(resolution,[],[f1093,f626])).
fof(f626,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK51(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f439])).
fof(f9349,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl23),
  inference(resolution,[],[f8067,f653])).
fof(f9348,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8067,f652])).
fof(f9346,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8067,f651])).
fof(f9344,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl23),
  inference(resolution,[],[f8067,f649])).
fof(f9342,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8067,f647])).
fof(f9292,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl200),
  inference(cnf_transformation,[],[f9292_D])).
fof(f9292_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f9260,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl26),
  inference(resolution,[],[f1100,f901])).
fof(f9167,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl26),
  inference(resolution,[],[f1100,f653])).
fof(f9164,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl26),
  inference(resolution,[],[f1100,f652])).
fof(f9163,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl26),
  inference(resolution,[],[f1100,f651])).
fof(f9162,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl26),
  inference(resolution,[],[f1100,f649])).
fof(f9160,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl26),
  inference(resolution,[],[f1100,f647])).
fof(f7171,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl164),
  inference(cnf_transformation,[],[f7171_D])).
fof(f7171_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f9098,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8066,f901])).
fof(f8066,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) | $spl23),
  inference(resolution,[],[f4310,f608])).
fof(f9000,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl23),
  inference(resolution,[],[f8066,f653])).
fof(f8999,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8066,f652])).
fof(f8997,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8066,f651])).
fof(f8995,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl23),
  inference(resolution,[],[f8066,f649])).
fof(f8993,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f8066,f647])).
fof(f8971,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl198),
  inference(cnf_transformation,[],[f8971_D])).
fof(f8971_D,plain,(
  s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f8955,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X13))) ) | ($spl23 | $spl25)),
  inference(resolution,[],[f8824,f4372])).
fof(f4372,plain,(
  ( ! [X134,X135] : (sP8(X134,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X135)) ) | $spl23),
  inference(resolution,[],[f1093,f714])).
fof(f8969,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) | $spl196),
  inference(cnf_transformation,[],[f8969_D])).
fof(f8969_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f8954,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) | ~p(s(bool,X12))) ) | ($spl23 | $spl25)),
  inference(resolution,[],[f8824,f8130])).
fof(f8130,plain,(
  ( ! [X146,X147] : (sP8(X146,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))),X147)) ) | $spl23),
  inference(resolution,[],[f4310,f714])).
fof(f8953,plain,(
  ( ! [X11] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | ~p(s(bool,X11))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f8824,f7256])).
fof(f8952,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | ~p(s(bool,X10))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f8824,f7093])).
fof(f8951,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | ~p(s(bool,X9))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f8824,f5499])).
fof(f8931,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | $spl189),
  inference(cnf_transformation,[],[f8931_D])).
fof(f8931_D,plain,(
  s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f8915,plain,(
  ( ! [X13] : (s(bool,'T_0') != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | ~p(s(bool,X13))) ) | ($spl23 | $spl25)),
  inference(resolution,[],[f8903,f4372])).
fof(f8929,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) | $spl187),
  inference(cnf_transformation,[],[f8929_D])).
fof(f8929_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) <=> ~$spl187),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl187])])).
fof(f8914,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) | ~p(s(bool,X12))) ) | ($spl23 | $spl25)),
  inference(resolution,[],[f8903,f8130])).
fof(f8913,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | ~p(s(bool,X11))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f8903,f7256])).
fof(f8912,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | ~p(s(bool,X10))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f8903,f7093])).
fof(f8911,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | ~p(s(bool,X9))) ) | ($spl25 | $spl27)),
  inference(resolution,[],[f8903,f5499])).
fof(f8622,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) | $spl178),
  inference(cnf_transformation,[],[f8622_D])).
fof(f8622_D,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f8621,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl25),
  inference(duplicate_literal_removal,[],[f8615])).
fof(f8615,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl25),
  inference(resolution,[],[f4133,f727])).
fof(f4133,plain,(
  ( ! [X2,X3] : (sP9(X3,X2,'F_0') | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl25),
  inference(superposition,[],[f4096,f724])).
fof(f8437,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl176),
  inference(cnf_transformation,[],[f8437_D])).
fof(f8437_D,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f8429,plain,(
  ( ! [X0,X1] : (sP7(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X0) = s(bool,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f7989,f900])).
fof(f7989,plain,(
  ( ! [X6,X7] : (sP7(X6,X7,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X6) = s(bool,X7)) ) | $spl25),
  inference(resolution,[],[f2659,f706])).
fof(f2659,plain,(
  ( ! [X152,X153] : (sP9('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X152,X153) | s(bool,X152) = s(bool,X153)) ) | $spl25),
  inference(resolution,[],[f1095,f724])).
fof(f8423,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X0) = s(bool,X1)) ) | $spl174),
  inference(cnf_transformation,[],[f8423_D])).
fof(f8423_D,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f8415,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s(bool,X0) = s(bool,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f7988,f900])).
fof(f7988,plain,(
  ( ! [X4,X5] : (sP6(X5,X4,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X4) = s(bool,X5)) ) | $spl25),
  inference(resolution,[],[f2659,f705])).
fof(f8208,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))))) ) | $spl172),
  inference(cnf_transformation,[],[f8208_D])).
fof(f8208_D,plain,(
  ( ! [X25] : (~p(s(bool,X25)) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f8178,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f4310,f901])).
fof(f8080,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl23),
  inference(resolution,[],[f4310,f653])).
fof(f8079,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f4310,f652])).
fof(f8077,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f4310,f651])).
fof(f8075,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl23),
  inference(resolution,[],[f4310,f649])).
fof(f8073,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK51(sK15,sK14,sK13))))) ) | $spl23),
  inference(resolution,[],[f4310,f647])).
fof(f8029,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s(bool,X0) != s(bool,X1)) ) | $spl170),
  inference(cnf_transformation,[],[f8029_D])).
fof(f8029_D,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s(bool,X0) != s(bool,X1)) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f8012,plain,(
  ( ! [X0,X1] : (~sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s(bool,X0) != s(bool,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2660,f900])).
fof(f8004,plain,(
  ( ! [X0,X1] : (sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s(bool,X0) = s(bool,X1)) ) | $spl168),
  inference(cnf_transformation,[],[f8004_D])).
fof(f8004_D,plain,(
  ( ! [X0,X1] : (sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s(bool,X0) = s(bool,X1)) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f7992,plain,(
  ( ! [X0,X1] : (sP9('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s(bool,X0) = s(bool,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2659,f900])).
fof(f7304,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5436,f901])).
fof(f7206,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl27),
  inference(resolution,[],[f5436,f653])).
fof(f7205,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5436,f652])).
fof(f7203,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5436,f651])).
fof(f7201,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl27),
  inference(resolution,[],[f5436,f649])).
fof(f7199,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK45(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5436,f647])).
fof(f7141,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5435,f901])).
fof(f7043,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl27),
  inference(resolution,[],[f5435,f653])).
fof(f7042,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5435,f652])).
fof(f7040,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5435,f651])).
fof(f7038,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl27),
  inference(resolution,[],[f5435,f649])).
fof(f7036,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK44(sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f5435,f647])).
fof(f6810,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0)) ) | $spl162),
  inference(cnf_transformation,[],[f6810_D])).
fof(f6810_D,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0)) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f6796,plain,(
  ( ! [X0] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))),X0) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3109,f900])).
fof(f3109,plain,(
  ( ! [X195] : (sP12('const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))),X195)) ) | $spl4),
  inference(resolution,[],[f915,f776])).
fof(f6719,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) | $spl160),
  inference(cnf_transformation,[],[f6719_D])).
fof(f6719_D,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0')))))) ) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f6705,plain,(
  ( ! [X0] : (sP12(X0,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),'const.list.NIL_0'))))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl4),
  inference(superposition,[],[f3108,f900])).
fof(f3108,plain,(
  ( ! [X194] : (sP12(X194,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f775])).
fof(f5835,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl158),
  inference(cnf_transformation,[],[f5835_D])).
fof(f5835_D,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f5824,plain,(
  ( ! [X0,X1] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4376,f900])).
fof(f4376,plain,(
  ( ! [X142,X143] : (sP7('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X142,X143)) ) | $spl23),
  inference(resolution,[],[f1093,f719])).
fof(f5814,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl156),
  inference(cnf_transformation,[],[f5814_D])).
fof(f5814_D,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f5803,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4375,f900])).
fof(f4375,plain,(
  ( ! [X140,X141] : (sP7(X140,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X141)) ) | $spl23),
  inference(resolution,[],[f1093,f718])).
fof(f5793,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl154),
  inference(cnf_transformation,[],[f5793_D])).
fof(f5793_D,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f5782,plain,(
  ( ! [X0,X1] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4373,f900])).
fof(f4373,plain,(
  ( ! [X136,X137] : (sP8('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X136,X137)) ) | $spl23),
  inference(resolution,[],[f1093,f715])).
fof(f5772,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl152),
  inference(cnf_transformation,[],[f5772_D])).
fof(f5772_D,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f5762,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4372,f900])).
fof(f5752,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl150),
  inference(cnf_transformation,[],[f5752_D])).
fof(f5752_D,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f5743,plain,(
  ( ! [X0,X1] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4360,f900])).
fof(f4360,plain,(
  ( ! [X111,X110] : (sP4('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X110,X111)) ) | $spl23),
  inference(resolution,[],[f1093,f698])).
fof(f5732,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) | $spl148),
  inference(cnf_transformation,[],[f5732_D])).
fof(f5732_D,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1)) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f5723,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4359,f900])).
fof(f4359,plain,(
  ( ! [X109,X108] : (sP4(X108,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X109)) ) | $spl23),
  inference(resolution,[],[f1093,f697])).
fof(f5712,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl146),
  inference(cnf_transformation,[],[f5712_D])).
fof(f5712_D,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f5702,plain,(
  ( ! [X0,X1] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4346,f900])).
fof(f4346,plain,(
  ( ! [X83,X82] : (sP2('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X82,X83)) ) | $spl23),
  inference(resolution,[],[f1093,f681])).
fof(f5691,plain,(
  ( ! [X0,X1] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) | $spl144),
  inference(cnf_transformation,[],[f5691_D])).
fof(f5691_D,plain,(
  ( ! [X0,X1] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1)) ) <=> ~$spl144),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl144])])).
fof(f5682,plain,(
  ( ! [X0,X1] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4330,f900])).
fof(f4330,plain,(
  ( ! [X50,X51] : (sP0('const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)),X50,X51)) ) | $spl23),
  inference(resolution,[],[f1093,f664])).
fof(f5671,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) | $spl142),
  inference(cnf_transformation,[],[f5671_D])).
fof(f5671_D,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) ) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f5662,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl23),
  inference(superposition,[],[f4329,f900])).
fof(f4329,plain,(
  ( ! [X48,X49] : (sP0(X48,X49,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl23),
  inference(resolution,[],[f1093,f662])).
fof(f5547,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f901])).
fof(f5449,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl27),
  inference(resolution,[],[f1101,f653])).
fof(f5448,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f652])).
fof(f5446,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f651])).
fof(f5444,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl27),
  inference(resolution,[],[f1101,f649])).
fof(f5442,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) ) | $spl27),
  inference(resolution,[],[f1101,f647])).
fof(f5330,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1))) ) | $spl138),
  inference(cnf_transformation,[],[f5330_D])).
fof(f5330_D,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1))) ) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f5314,plain,(
  ( ! [X0,X1] : (sP10(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2664,f900])).
fof(f2664,plain,(
  ( ! [X163,X162] : (sP10(X162,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X163) | p(s(bool,X163))) ) | $spl25),
  inference(resolution,[],[f1095,f730])).
fof(f5305,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1))) ) | $spl136),
  inference(cnf_transformation,[],[f5305_D])).
fof(f5305_D,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1))) ) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f5289,plain,(
  ( ! [X0,X1] : (sP10(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2663,f900])).
fof(f2663,plain,(
  ( ! [X161,X160] : (sP10(X160,X161,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X161))) ) | $spl25),
  inference(resolution,[],[f1095,f730])).
fof(f5281,plain,(
  ( ! [X0,X1] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl134),
  inference(cnf_transformation,[],[f5281_D])).
fof(f5281_D,plain,(
  ( ! [X0,X1] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f5270,plain,(
  ( ! [X0,X1] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2662,f900])).
fof(f2662,plain,(
  ( ! [X158,X159] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X158,X159) | ~p(s(bool,X158))) ) | $spl25),
  inference(resolution,[],[f1095,f729])).
fof(f5227,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2794,f901])).
fof(f5134,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl2),
  inference(resolution,[],[f2794,f653])).
fof(f5260,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl132),
  inference(cnf_transformation,[],[f5260_D])).
fof(f5260_D,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f1079,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,X18) | ~p(s(bool,X17)) | p(s(bool,X18))) ) | $spl14),
  inference(cnf_transformation,[],[f1079_D])).
fof(f1079_D,plain,(
  ( ! [X17,X18] : (s(bool,X17) != s(bool,X18) | ~p(s(bool,X17)) | p(s(bool,X18))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f5133,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2794,f653])).
fof(f5131,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2794,f652])).
fof(f5130,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2794,f651])).
fof(f5129,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl2),
  inference(resolution,[],[f2794,f649])).
fof(f5128,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2794,f649])).
fof(f5127,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK43(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2794,f647])).
fof(f5101,plain,(
  ( ! [X0,X1] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1))) ) | $spl128),
  inference(cnf_transformation,[],[f5101_D])).
fof(f5101_D,plain,(
  ( ! [X0,X1] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f5090,plain,(
  ( ! [X0,X1] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2661,f900])).
fof(f2661,plain,(
  ( ! [X156,X157] : (~sP10('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X157,X156) | ~p(s(bool,X156))) ) | $spl25),
  inference(resolution,[],[f1095,f728])).
fof(f5081,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0))) ) | $spl126),
  inference(cnf_transformation,[],[f5081_D])).
fof(f5081_D,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f5071,plain,(
  ( ! [X0,X1] : (sP5(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2643,f900])).
fof(f2643,plain,(
  ( ! [X121,X120] : (sP5(X120,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X121) | p(s(bool,X120))) ) | $spl25),
  inference(resolution,[],[f1095,f700])).
fof(f5060,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0))) ) | $spl124),
  inference(cnf_transformation,[],[f5060_D])).
fof(f5060_D,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f5038,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2642,f900])).
fof(f2642,plain,(
  ( ! [X118,X119] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X118,X119) | p(s(bool,X118))) ) | $spl25),
  inference(resolution,[],[f1095,f700])).
fof(f4993,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2793,f901])).
fof(f4900,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl2),
  inference(resolution,[],[f2793,f653])).
fof(f5026,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl122),
  inference(cnf_transformation,[],[f5026_D])).
fof(f5026_D,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f4899,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2793,f653])).
fof(f4897,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2793,f652])).
fof(f4896,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2793,f651])).
fof(f4895,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl2),
  inference(resolution,[],[f2793,f649])).
fof(f4894,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2793,f649])).
fof(f4893,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK42(sK14,sK13))))) ) | $spl2),
  inference(resolution,[],[f2793,f647])).
fof(f4866,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl118),
  inference(cnf_transformation,[],[f4866_D])).
fof(f4866_D,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f4856,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2641,f900])).
fof(f2641,plain,(
  ( ! [X116,X117] : (sP5(X116,X117,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X116))) ) | $spl25),
  inference(resolution,[],[f1095,f699])).
fof(f4845,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) | $spl116),
  inference(cnf_transformation,[],[f4845_D])).
fof(f4845_D,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f4823,plain,(
  ( ! [X0,X1] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2640,f900])).
fof(f2640,plain,(
  ( ! [X114,X115] : (sP5('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X114,X115) | p(s(bool,X115))) ) | $spl25),
  inference(resolution,[],[f1095,f699])).
fof(f4813,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0))) ) | $spl114),
  inference(cnf_transformation,[],[f4813_D])).
fof(f4813_D,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f4801,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2634,f900])).
fof(f2634,plain,(
  ( ! [X103,X102] : (~sP5(X102,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X103) | ~p(s(bool,X102))) ) | $spl25),
  inference(resolution,[],[f1095,f690])).
fof(f4792,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) | $spl112),
  inference(cnf_transformation,[],[f4792_D])).
fof(f4792_D,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f4780,plain,(
  ( ! [X0,X1] : (~sP5(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2633,f900])).
fof(f2633,plain,(
  ( ! [X101,X100] : (~sP5(X100,X101,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X100))) ) | $spl25),
  inference(resolution,[],[f1095,f689])).
fof(f4769,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl110),
  inference(cnf_transformation,[],[f4769_D])).
fof(f4769_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f4737,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl28),
  inference(resolution,[],[f1265,f901])).
fof(f4644,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X56) != s(bool,X57)) ) | $spl28),
  inference(resolution,[],[f1265,f653])).
fof(f2435,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl64),
  inference(cnf_transformation,[],[f2435_D])).
fof(f2435_D,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4643,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl28),
  inference(resolution,[],[f1265,f653])).
fof(f4641,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl28),
  inference(resolution,[],[f1265,f652])).
fof(f4640,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl28),
  inference(resolution,[],[f1265,f651])).
fof(f4639,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X41) != s(bool,X42)) ) | $spl28),
  inference(resolution,[],[f1265,f649])).
fof(f4638,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl28),
  inference(resolution,[],[f1265,f649])).
fof(f4637,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl28),
  inference(resolution,[],[f1265,f647])).
fof(f4611,plain,(
  ( ! [X0,X1] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl108),
  inference(cnf_transformation,[],[f4611_D])).
fof(f4611_D,plain,(
  ( ! [X0,X1] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f4596,plain,(
  ( ! [X0,X1] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2618,f900])).
fof(f2618,plain,(
  ( ! [X70,X71] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X70,X71) | ~p(s(bool,X70))) ) | $spl25),
  inference(resolution,[],[f1095,f672])).
fof(f4587,plain,(
  ( ! [X0,X1] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1))) ) | $spl106),
  inference(cnf_transformation,[],[f4587_D])).
fof(f4587_D,plain,(
  ( ! [X0,X1] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f4572,plain,(
  ( ! [X0,X1] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2617,f900])).
fof(f2617,plain,(
  ( ! [X68,X69] : (~sP3('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X69,X68) | ~p(s(bool,X68))) ) | $spl25),
  inference(resolution,[],[f1095,f671])).
fof(f4563,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) | $spl104),
  inference(cnf_transformation,[],[f4563_D])).
fof(f4563_D,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f4552,plain,(
  ( ! [X0,X1] : (sP1(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2612,f900])).
fof(f2612,plain,(
  ( ! [X59,X58] : (sP1(X58,X59,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X58))) ) | $spl25),
  inference(resolution,[],[f1095,f666])).
fof(f4541,plain,(
  ( ! [X0,X1] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) | $spl102),
  inference(cnf_transformation,[],[f4541_D])).
fof(f4541_D,plain,(
  ( ! [X0,X1] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f4526,plain,(
  ( ! [X0,X1] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2604,f900])).
fof(f2604,plain,(
  ( ! [X43,X42] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X42,X43) | ~p(s(bool,X42))) ) | $spl25),
  inference(resolution,[],[f1095,f655])).
fof(f4516,plain,(
  ( ! [X0,X1] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) | $spl100),
  inference(cnf_transformation,[],[f4516_D])).
fof(f4516_D,plain,(
  ( ! [X0,X1] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4497,plain,(
  ( ! [X0,X1] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | p(s(bool,X1)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2603,f900])).
fof(f2603,plain,(
  ( ! [X41,X40] : (~sP1('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X41,X40) | p(s(bool,X40))) ) | $spl25),
  inference(resolution,[],[f1095,f654])).
fof(f4487,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl98),
  inference(cnf_transformation,[],[f4487_D])).
fof(f4487_D,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f4471,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2602,f900])).
fof(f2602,plain,(
  ( ! [X39,X38] : (~sP1(X38,X39,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X38))) ) | $spl25),
  inference(resolution,[],[f1095,f654])).
fof(f4460,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | $spl97),
  inference(cnf_transformation,[],[f4460_D])).
fof(f4460_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f4426,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl23),
  inference(superposition,[],[f1093,f900])).
fof(f4452,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl92),
  inference(cnf_transformation,[],[f4452_D])).
fof(f4452_D,plain,(
  ( ! [X13] : (~p(s(bool,X13)) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4420,plain,(
  ( ! [X191,X190] : (p(s(bool,X190)) | ~p(s(bool,X191)) | s(bool,X191) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl23),
  inference(resolution,[],[f1093,f901])).
fof(f4322,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X34) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X33) != s(bool,X35)) ) | $spl23),
  inference(resolution,[],[f1093,f653])).
fof(f4321,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | ~p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl23),
  inference(resolution,[],[f1093,f652])).
fof(f4319,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl23),
  inference(resolution,[],[f1093,f651])).
fof(f4453,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl94),
  inference(cnf_transformation,[],[f4453_D])).
fof(f4453_D,plain,(
  ( ! [X17] : (s(bool,X17) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4318,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X23) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X21) != s(bool,X22)) ) | $spl23),
  inference(resolution,[],[f1093,f651])).
fof(f4317,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X19) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X18) != s(bool,X20)) ) | $spl23),
  inference(resolution,[],[f1093,f649])).
fof(f4316,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X17) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))) | s(bool,X15) != s(bool,X16)) ) | $spl23),
  inference(resolution,[],[f1093,f647])).
fof(f4315,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,X14) | s(bool,X13) != s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) ) | $spl23),
  inference(resolution,[],[f1093,f647])).
fof(f4211,plain,(
  ( ! [X0] : (~sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) | $spl90),
  inference(cnf_transformation,[],[f4211_D])).
fof(f4211_D,plain,(
  ( ! [X0] : (~sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f4197,plain,(
  ( ! [X0] : (~sP12('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2729,f900])).
fof(f2729,plain,(
  ( ! [X185] : (~sP12('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X185) | p(s(bool,X185))) ) | $spl25),
  inference(subsumption_resolution,[],[f2680,f826])).
fof(f2680,plain,(
  ( ! [X185] : (p(s(bool,'F_0')) | p(s(bool,X185)) | ~sP12('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X185)) ) | $spl25),
  inference(resolution,[],[f1095,f774])).
fof(f4173,plain,(
  ( ! [X0] : (~sP12(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl88),
  inference(cnf_transformation,[],[f4173_D])).
fof(f4173_D,plain,(
  ( ! [X0] : (~sP12(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f4158,plain,(
  ( ! [X0] : (~sP12(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2728,f900])).
fof(f2728,plain,(
  ( ! [X184] : (~sP12(X184,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X184))) ) | $spl25),
  inference(subsumption_resolution,[],[f2679,f826])).
fof(f2679,plain,(
  ( ! [X184] : (p(s(bool,'F_0')) | p(s(bool,X184)) | ~sP12(X184,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl25),
  inference(resolution,[],[f1095,f774])).
fof(f1768,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | $spl42),
  inference(cnf_transformation,[],[f1768_D])).
fof(f1768_D,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f3891,plain,(
  ( ! [X121,X122,X120] : (s(bool,'F_0') != s(bool,X120) | p(s(bool,X121)) | ~p(s(bool,X122)) | s(bool,X121) != s(bool,X122)) ) | $spl25),
  inference(subsumption_resolution,[],[f3530,f651])).
fof(f3530,plain,(
  ( ! [X121,X122,X120] : (s(bool,'F_0') != s(bool,X120) | p(s(bool,X121)) | ~p(s(bool,X120)) | ~p(s(bool,X122)) | s(bool,X121) != s(bool,X122)) ) | $spl25),
  inference(superposition,[],[f653,f2693])).
fof(f3890,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,X119)) | s(bool,X118) != s(bool,X119)) ) | $spl25),
  inference(subsumption_resolution,[],[f3529,f901])).
fof(f3529,plain,(
  ( ! [X118,X119,X117] : (s(bool,'F_0') != s(bool,X117) | p(s(bool,X118)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | ~p(s(bool,X119)) | s(bool,X118) != s(bool,X119)) ) | $spl25),
  inference(superposition,[],[f653,f2693])).
fof(f3874,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') != s(bool,X60) | p(s(bool,X61)) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,X62)) ) | $spl25),
  inference(subsumption_resolution,[],[f3510,f901])).
fof(f3510,plain,(
  ( ! [X61,X62,X60] : (s(bool,'F_0') != s(bool,X60) | p(s(bool,X61)) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | ~p(s(bool,X62)) | s(bool,X61) != s(bool,X62)) ) | $spl25),
  inference(superposition,[],[f649,f2693])).
fof(f3873,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,X59)) ) | $spl25),
  inference(subsumption_resolution,[],[f3509,f651])).
fof(f3509,plain,(
  ( ! [X59,X57,X58] : (s(bool,'F_0') != s(bool,X57) | p(s(bool,X58)) | ~p(s(bool,X57)) | ~p(s(bool,X59)) | s(bool,X58) != s(bool,X59)) ) | $spl25),
  inference(superposition,[],[f649,f2693])).
fof(f3845,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl86),
  inference(cnf_transformation,[],[f3845_D])).
fof(f3845_D,plain,(
  s(bool,'T_0') = s(bool,'F_0') <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f3843,plain,(
  s(bool,'T_0') = s(bool,'F_0') | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl25),
  inference(forward_demodulation,[],[f3463,f885])).
fof(f3463,plain,(
  s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13)))) | $spl25),
  inference(superposition,[],[f2693,f900])).
fof(f3447,plain,(
  ( ! [X0] : (sP11('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0))) ) | $spl84),
  inference(cnf_transformation,[],[f3447_D])).
fof(f3447_D,plain,(
  ( ! [X0] : (sP11('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f3437,plain,(
  ( ! [X0] : (sP11('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2674,f900])).
fof(f2674,plain,(
  ( ! [X179] : (sP11('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X179) | ~p(s(bool,X179))) ) | $spl25),
  inference(resolution,[],[f1095,f763])).
fof(f3430,plain,(
  ( ! [X0] : (sP11(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) | $spl82),
  inference(cnf_transformation,[],[f3430_D])).
fof(f3430_D,plain,(
  ( ! [X0] : (sP11(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f3420,plain,(
  ( ! [X0] : (sP11(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | ~p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2673,f900])).
fof(f2673,plain,(
  ( ! [X178] : (sP11(X178,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | ~p(s(bool,X178))) ) | $spl25),
  inference(resolution,[],[f1095,f760])).
fof(f3413,plain,(
  ( ! [X0] : (~sP11('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) | $spl80),
  inference(cnf_transformation,[],[f3413_D])).
fof(f3413_D,plain,(
  ( ! [X0] : (~sP11('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f3403,plain,(
  ( ! [X0] : (~sP11('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2672,f900])).
fof(f2672,plain,(
  ( ! [X177] : (~sP11('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X177) | p(s(bool,X177))) ) | $spl25),
  inference(resolution,[],[f1095,f758])).
fof(f3398,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) | $spl78),
  inference(cnf_transformation,[],[f3398_D])).
fof(f3398_D,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f3388,plain,(
  ( ! [X0] : (~sP11(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | p(s(bool,X0)) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2671,f900])).
fof(f2671,plain,(
  ( ! [X176] : (~sP11(X176,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | p(s(bool,X176))) ) | $spl25),
  inference(resolution,[],[f1095,f758])).
fof(f3376,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl76),
  inference(cnf_transformation,[],[f3376_D])).
fof(f3376_D,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f3328,plain,(
  ( ! [X189,X190] : (p(s(bool,X189)) | p(s(bool,X190)) | s(bool,X190) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f901])).
fof(f3235,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X41) != s(bool,X42)) ) | $spl0),
  inference(resolution,[],[f910,f653])).
fof(f1078,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl12),
  inference(cnf_transformation,[],[f1078_D])).
fof(f1078_D,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f3234,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f653])).
fof(f3232,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f652])).
fof(f3231,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f651])).
fof(f3230,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X26) != s(bool,X27)) ) | $spl0),
  inference(resolution,[],[f910,f649])).
fof(f3229,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f649])).
fof(f3228,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl0),
  inference(resolution,[],[f910,f647])).
fof(f3175,plain,(
  ( ! [X0,X1] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2655,f900])).
fof(f2655,plain,(
  ( ! [X144,X145] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X144,X145)) ) | $spl25),
  inference(resolution,[],[f1095,f719])).
fof(f3164,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl72),
  inference(cnf_transformation,[],[f3164_D])).
fof(f3164_D,plain,(
  ( ! [X31] : (p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f3126,plain,(
  ( ! [X200,X201] : (p(s(bool,X200)) | p(s(bool,X201)) | s(bool,X201) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f901])).
fof(f3033,plain,(
  ( ! [X52,X53,X51] : (p(s(bool,X51)) | ~p(s(bool,X52)) | s(bool,X51) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X52) != s(bool,X53)) ) | $spl4),
  inference(resolution,[],[f915,f653])).
fof(f3165,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl74),
  inference(cnf_transformation,[],[f3165_D])).
fof(f3165_D,plain,(
  ( ! [X35] : (s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f3032,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X48) != s(bool,X49) | s(bool,X50) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f653])).
fof(f3030,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X44) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f652])).
fof(f3029,plain,(
  ( ! [X39,X41,X40] : (p(s(bool,X39)) | p(s(bool,X40)) | s(bool,X39) != s(bool,X41) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f651])).
fof(f3028,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X37) != s(bool,X38)) ) | $spl4),
  inference(resolution,[],[f915,f649])).
fof(f3027,plain,(
  ( ! [X35,X33,X34] : (p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,X34) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f649])).
fof(f3026,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X30) != s(bool,X32) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl4),
  inference(resolution,[],[f915,f647])).
fof(f3000,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2654,f900])).
fof(f2654,plain,(
  ( ! [X142,X143] : (sP7(X142,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X143)) ) | $spl25),
  inference(resolution,[],[f1095,f718])).
fof(f2985,plain,(
  ( ! [X0,X1] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2652,f900])).
fof(f2971,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2651,f900])).
fof(f2651,plain,(
  ( ! [X136,X137] : (sP8(X136,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X137)) ) | $spl25),
  inference(resolution,[],[f1095,f714])).
fof(f2958,plain,(
  ( ! [X0,X1] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2639,f900])).
fof(f2639,plain,(
  ( ! [X112,X113] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X112,X113)) ) | $spl25),
  inference(resolution,[],[f1095,f698])).
fof(f2944,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2638,f900])).
fof(f2638,plain,(
  ( ! [X111,X110] : (sP4(X110,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X111)) ) | $spl25),
  inference(resolution,[],[f1095,f697])).
fof(f2934,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl68),
  inference(cnf_transformation,[],[f2934_D])).
fof(f2934_D,plain,(
  ( ! [X35] : (p(s(bool,X35)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2902,plain,(
  ( ! [X204,X205] : (p(s(bool,X204)) | p(s(bool,X205)) | s(bool,X205) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl2),
  inference(resolution,[],[f913,f901])).
fof(f2809,plain,(
  ( ! [X57,X56,X55] : (p(s(bool,X55)) | ~p(s(bool,X56)) | s(bool,X55) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | s(bool,X56) != s(bool,X57)) ) | $spl2),
  inference(resolution,[],[f913,f653])).
fof(f2935,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl70),
  inference(cnf_transformation,[],[f2935_D])).
fof(f2935_D,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2808,plain,(
  ( ! [X54,X52,X53] : (p(s(bool,X52)) | ~p(s(bool,X53)) | s(bool,X52) != s(bool,X53) | s(bool,X54) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl2),
  inference(resolution,[],[f913,f653])).
fof(f2806,plain,(
  ( ! [X47,X48,X46] : (p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | s(bool,X48) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl2),
  inference(resolution,[],[f913,f652])).
fof(f2805,plain,(
  ( ! [X45,X43,X44] : (p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X43) != s(bool,X45) | s(bool,X44) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl2),
  inference(resolution,[],[f913,f651])).
fof(f2804,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))) | s(bool,X41) != s(bool,X42)) ) | $spl2),
  inference(resolution,[],[f913,f649])).
fof(f2803,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl2),
  inference(resolution,[],[f913,f649])).
fof(f2802,plain,(
  ( ! [X35,X36,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X34) != s(bool,X36) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) ) | $spl2),
  inference(resolution,[],[f913,f647])).
fof(f2768,plain,(
  ( ! [X0,X1] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2625,f900])).
fof(f2625,plain,(
  ( ! [X85,X84] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X84,X85)) ) | $spl25),
  inference(resolution,[],[f1095,f681])).
fof(f2754,plain,(
  ( ! [X0,X1] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2609,f900])).
fof(f2609,plain,(
  ( ! [X52,X53] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X52,X53)) ) | $spl25),
  inference(resolution,[],[f1095,f664])).
fof(f2740,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl25),
  inference(superposition,[],[f2608,f900])).
fof(f2608,plain,(
  ( ! [X50,X51] : (sP0(X50,X51,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl25),
  inference(resolution,[],[f1095,f662])).
fof(f1678,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl34),
  inference(cnf_transformation,[],[f1678_D])).
fof(f1678_D,plain,(
  ( ! [X15] : (~p(s(bool,X15)) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2699,plain,(
  ( ! [X192,X193] : (p(s(bool,X192)) | ~p(s(bool,X193)) | s(bool,X193) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl25),
  inference(resolution,[],[f1095,f901])).
fof(f2601,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X35) != s(bool,X37)) ) | $spl25),
  inference(resolution,[],[f1095,f653])).
fof(f2600,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl25),
  inference(resolution,[],[f1095,f652])).
fof(f2598,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl25),
  inference(resolution,[],[f1095,f651])).
fof(f1679,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl36),
  inference(cnf_transformation,[],[f1679_D])).
fof(f1679_D,plain,(
  ( ! [X19] : (s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2597,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X25) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl25),
  inference(resolution,[],[f1095,f651])).
fof(f2596,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X20) != s(bool,X22)) ) | $spl25),
  inference(resolution,[],[f1095,f649])).
fof(f2595,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X17) != s(bool,X18)) ) | $spl25),
  inference(resolution,[],[f1095,f647])).
fof(f2594,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl25),
  inference(resolution,[],[f1095,f647])).
fof(f2582,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl66),
  inference(cnf_transformation,[],[f2582_D])).
fof(f2582_D,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2552,plain,(
  ( ! [X189,X190] : (p(s(bool,X189)) | p(s(bool,X190)) | s(bool,X190) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl24),
  inference(resolution,[],[f1094,f901])).
fof(f2459,plain,(
  ( ! [X41,X42,X40] : (p(s(bool,X40)) | ~p(s(bool,X41)) | s(bool,X40) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X41) != s(bool,X42)) ) | $spl24),
  inference(resolution,[],[f1094,f653])).
fof(f2458,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | ~p(s(bool,X38)) | s(bool,X37) != s(bool,X38) | s(bool,X39) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl24),
  inference(resolution,[],[f1094,f653])).
fof(f2456,plain,(
  ( ! [X33,X31,X32] : (p(s(bool,X31)) | ~p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl24),
  inference(resolution,[],[f1094,f652])).
fof(f2455,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X28) != s(bool,X30) | s(bool,X29) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl24),
  inference(resolution,[],[f1094,f651])).
fof(f2454,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X26) != s(bool,X27)) ) | $spl24),
  inference(resolution,[],[f1094,f649])).
fof(f2453,plain,(
  ( ! [X24,X23,X22] : (p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,X23) | s(bool,X24) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl24),
  inference(resolution,[],[f1094,f649])).
fof(f2452,plain,(
  ( ! [X21,X19,X20] : (p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl24),
  inference(resolution,[],[f1094,f647])).
fof(f2443,plain,(
  $false | ($spl1 | $spl24)),
  inference(subsumption_resolution,[],[f1094,f1091])).
fof(f1091,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | $spl1),
  inference(subsumption_resolution,[],[f1090,f886])).
fof(f1090,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | $spl1),
  inference(forward_demodulation,[],[f1039,f892])).
fof(f1039,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),'const.list.NIL_0')))))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | $spl1),
  inference(superposition,[],[f911,f882])).
fof(f2404,plain,(
  ( ! [X202,X203] : (p(s(bool,X202)) | ~p(s(bool,X203)) | s(bool,X203) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f901])).
fof(f2306,plain,(
  ( ! [X47,X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46)) | s(bool,X46) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X45) != s(bool,X47)) ) | $spl29),
  inference(resolution,[],[f1266,f653])).
fof(f2305,plain,(
  ( ! [X43,X44,X42] : (p(s(bool,X42)) | ~p(s(bool,X43)) | ~p(s(bool,X44)) | s(bool,X43) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f652])).
fof(f2303,plain,(
  ( ! [X37,X38,X36] : (p(s(bool,X36)) | ~p(s(bool,X37)) | s(bool,X36) != s(bool,X38) | s(bool,X37) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f651])).
fof(f2302,plain,(
  ( ! [X35,X33,X34] : (~p(s(bool,X33)) | p(s(bool,X34)) | s(bool,X35) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X33) != s(bool,X34)) ) | $spl29),
  inference(resolution,[],[f1266,f651])).
fof(f2301,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,X30)) | ~p(s(bool,X31)) | s(bool,X31) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X30) != s(bool,X32)) ) | $spl29),
  inference(resolution,[],[f1266,f649])).
fof(f2300,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,X27)) | p(s(bool,X28)) | s(bool,X29) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13)))) | s(bool,X27) != s(bool,X28)) ) | $spl29),
  inference(resolution,[],[f1266,f647])).
fof(f2299,plain,(
  ( ! [X26,X24,X25] : (p(s(bool,X24)) | ~p(s(bool,X25)) | s(bool,X24) != s(bool,X26) | s(bool,X25) != s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) ) | $spl29),
  inference(resolution,[],[f1266,f647])).
fof(f2267,plain,(
  ( ! [X0,X1] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1606,f900])).
fof(f1606,plain,(
  ( ! [X144,X145] : (sP7('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X144,X145)) ) | $spl1),
  inference(resolution,[],[f1091,f719])).
fof(f2251,plain,(
  ( ! [X0,X1] : (sP7(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1605,f900])).
fof(f1605,plain,(
  ( ! [X142,X143] : (sP7(X142,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X143)) ) | $spl1),
  inference(resolution,[],[f1091,f718])).
fof(f2235,plain,(
  ( ! [X0,X1] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1603,f900])).
fof(f1603,plain,(
  ( ! [X138,X139] : (sP8('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X138,X139)) ) | $spl1),
  inference(resolution,[],[f1091,f715])).
fof(f2220,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1602,f900])).
fof(f1602,plain,(
  ( ! [X136,X137] : (sP8(X136,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X137)) ) | $spl1),
  inference(resolution,[],[f1091,f714])).
fof(f2206,plain,(
  ( ! [X0,X1] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1590,f900])).
fof(f1590,plain,(
  ( ! [X112,X113] : (sP4('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X112,X113)) ) | $spl1),
  inference(resolution,[],[f1091,f698])).
fof(f2194,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33)) ) | $spl50),
  inference(cnf_transformation,[],[f2194_D])).
fof(f2194_D,plain,(
  ( ! [X33] : (s(bool,'T_0') != s(bool,X33)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2190,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2142])).
fof(f2142,plain,(
  ( ! [X35,X34] : (s(bool,'T_0') != s(bool,X35) | ~p(s(bool,X34)) | ~p(s(bool,X34))) ) | $spl16),
  inference(superposition,[],[f1080,f859])).
fof(f2191,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32))) ) | $spl16),
  inference(duplicate_literal_removal,[],[f2141])).
fof(f2141,plain,(
  ( ! [X33,X32] : (s(bool,'T_0') != s(bool,X33) | ~p(s(bool,X32)) | ~p(s(bool,X32))) ) | $spl16),
  inference(superposition,[],[f1080,f857])).
fof(f2121,plain,(
  ( ! [X0,X1] : (sP4(X0,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1589,f900])).
fof(f1589,plain,(
  ( ! [X111,X110] : (sP4(X110,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X111)) ) | $spl1),
  inference(resolution,[],[f1091,f697])).
fof(f2105,plain,(
  ( ! [X0,X1] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1576,f900])).
fof(f1576,plain,(
  ( ! [X85,X84] : (sP2('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X84,X85)) ) | $spl1),
  inference(resolution,[],[f1091,f681])).
fof(f2090,plain,(
  ( ! [X0,X1] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0')),X0,X1) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1560,f900])).
fof(f1560,plain,(
  ( ! [X52,X53] : (sP0('const.list.NULL_1'(s('type.list.list'(sK13),sK15)),X52,X53)) ) | $spl1),
  inference(resolution,[],[f1091,f664])).
fof(f2078,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2063])).
fof(f2063,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(superposition,[],[f1768,f881])).
fof(f881,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))) )),
  inference(cnf_transformation,[],[f273])).
fof(f273,plain,(
  ! [X0,X1,X2] : s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f272])).
fof(f272,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X6] : ! [X13,X10] : s(bool,'F_0') = s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X13),s('type.list.list'(X6),X10)))))),
  file('thm.list.NULL_FILTER','thm.list.NULL_DEF|split|1')).
fof(f2077,plain,(
  $false | $spl42),
  inference(trivial_inequality_removal,[],[f2065])).
fof(f2065,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl42),
  inference(superposition,[],[f1768,f629])).
fof(f2076,plain,(
  $false | $spl42),
  inference(equality_resolution,[],[f1768])).
fof(f2029,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f1920,f1763])).
fof(f1763,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl10),
  inference(equality_resolution,[],[f1077])).
fof(f1920,plain,(
  ( ! [X373] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X373))))) ) | ($spl6 | $spl10)),
  inference(resolution,[],[f1763,f916])).
fof(f2028,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1914,f1763])).
fof(f1914,plain,(
  ( ! [X352,X351] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X351),X352))))) ) | $spl10),
  inference(resolution,[],[f1763,f628])).
fof(f2026,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1902,f1763])).
fof(f1902,plain,(
  ( ! [X313,X312,X314] : (~p(s(bool,'const.bool.IN_2'(s(X312,X313),s(fun(X312,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X312),X314))))))) ) | $spl10),
  inference(resolution,[],[f1763,f628])).
fof(f2025,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1900,f1763])).
fof(f1900,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(resolution,[],[f1763,f875])).
fof(f2024,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1899,f1763])).
fof(f1899,plain,(
  ~p(s(bool,'T_0')) | $spl10),
  inference(resolution,[],[f1763,f873])).
fof(f2023,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1896,f1763])).
fof(f1896,plain,(
  ( ! [X308] : (~p(s(bool,X308))) ) | $spl10),
  inference(resolution,[],[f1763,f867])).
fof(f2022,plain,(
  $false | $spl10),
  inference(subsumption_resolution,[],[f1890,f1763])).
fof(f1890,plain,(
  ( ! [X302] : (~p(s(bool,X302))) ) | $spl10),
  inference(resolution,[],[f1763,f808])).
fof(f1898,plain,(
  $false | $spl10),
  inference(resolution,[],[f1763,f871])).
fof(f1903,plain,(
  $false | $spl10),
  inference(resolution,[],[f1763,f880])).
fof(f1904,plain,(
  $false | ($spl1 | $spl10)),
  inference(resolution,[],[f1763,f911])).
fof(f1905,plain,(
  $false | ($spl1 | $spl10)),
  inference(resolution,[],[f1763,f1091])).
fof(f1764,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f1714])).
fof(f1714,plain,(
  ( ! [X30,X31] : (s(bool,'F_0') != s(bool,X31) | p(s(bool,X30)) | p(s(bool,X30))) ) | $spl10),
  inference(superposition,[],[f1077,f855])).
fof(f1765,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28))) ) | $spl10),
  inference(duplicate_literal_removal,[],[f1713])).
fof(f1713,plain,(
  ( ! [X28,X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X28)) | p(s(bool,X28))) ) | $spl10),
  inference(superposition,[],[f1077,f853])).
fof(f1693,plain,(
  ( ! [X0,X1] : (sP0(X0,X1,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.NIL_0'))) | s('type.list.list'(sK13),sK15) = s('type.list.list'(sK13),'const.list.CONS_2'(s(sK13,sK54(sK15,sK13)),s('type.list.list'(sK13),sK55(sK15,sK13))))) ) | $spl1),
  inference(superposition,[],[f1559,f900])).
fof(f1559,plain,(
  ( ! [X50,X51] : (sP0(X50,X51,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1091,f662])).
fof(f1650,plain,(
  ( ! [X192,X193] : (p(s(bool,X192)) | ~p(s(bool,X193)) | s(bool,X193) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1091,f901])).
fof(f1552,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X35) != s(bool,X37)) ) | $spl1),
  inference(resolution,[],[f1091,f653])).
fof(f1551,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1091,f652])).
fof(f1549,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1091,f651])).
fof(f1548,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X25) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X23) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f1091,f651])).
fof(f1547,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X20) != s(bool,X22)) ) | $spl1),
  inference(resolution,[],[f1091,f649])).
fof(f1546,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15))) | s(bool,X17) != s(bool,X18)) ) | $spl1),
  inference(resolution,[],[f1091,f647])).
fof(f1545,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) ) | $spl1),
  inference(resolution,[],[f1091,f647])).
fof(f1535,plain,(
  $false | ($spl6 | $spl20)),
  inference(subsumption_resolution,[],[f1426,f1083])).
fof(f1426,plain,(
  ( ! [X373] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X373))))) ) | ($spl6 | $spl20)),
  inference(resolution,[],[f1083,f916])).
fof(f1534,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1420,f1083])).
fof(f1420,plain,(
  ( ! [X352,X351] : (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X351),X352))))) ) | $spl20),
  inference(resolution,[],[f1083,f628])).
fof(f1532,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1409,f1083])).
fof(f1409,plain,(
  ( ! [X313,X312,X314] : (~p(s(bool,'const.bool.IN_2'(s(X312,X313),s(fun(X312,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X312),X314))))))) ) | $spl20),
  inference(resolution,[],[f1083,f628])).
fof(f1531,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1407,f1083])).
fof(f1407,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(resolution,[],[f1083,f875])).
fof(f1530,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1406,f1083])).
fof(f1406,plain,(
  ~p(s(bool,'T_0')) | $spl20),
  inference(resolution,[],[f1083,f873])).
fof(f1529,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1403,f1083])).
fof(f1403,plain,(
  ( ! [X308] : (~p(s(bool,X308))) ) | $spl20),
  inference(resolution,[],[f1083,f867])).
fof(f1528,plain,(
  $false | $spl20),
  inference(subsumption_resolution,[],[f1397,f1083])).
fof(f1397,plain,(
  ( ! [X302] : (~p(s(bool,X302))) ) | $spl20),
  inference(resolution,[],[f1083,f808])).
fof(f1405,plain,(
  $false | $spl20),
  inference(resolution,[],[f1083,f871])).
fof(f1410,plain,(
  $false | $spl20),
  inference(resolution,[],[f1083,f880])).
fof(f1411,plain,(
  $false | ($spl1 | $spl20)),
  inference(resolution,[],[f1083,f911])).
fof(f1223,plain,(
  ( ! [X313,X315,X314] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X313)))) | p(s(bool,X314)) | ~p(s(bool,X315)) | s(bool,X315) != s(bool,'const.bool.IN_2'(s(sK13,X313),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f901])).
fof(f1125,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X58)))) | ~p(s(bool,X59)) | ~p(s(bool,X60)) | s(bool,X60) != s(bool,'const.bool.IN_2'(s(sK13,X58),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X59) != s(bool,X61)) ) | $spl6),
  inference(resolution,[],[f916,f653])).
fof(f1124,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X54)))) | p(s(bool,X55)) | ~p(s(bool,X56)) | ~p(s(bool,X57)) | s(bool,X56) != s(bool,'const.bool.IN_2'(s(sK13,X54),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f652])).
fof(f1122,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X46)))) | p(s(bool,X47)) | ~p(s(bool,X48)) | s(bool,X47) != s(bool,X49) | s(bool,X48) != s(bool,'const.bool.IN_2'(s(sK13,X46),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f651])).
fof(f1269,plain,(
  ( ! [X37,X34] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,X34),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl32),
  inference(cnf_transformation,[],[f1269_D])).
fof(f1269_D,plain,(
  ( ! [X37,X34] : (s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,X34),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1270,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X45) != s(bool,'const.bool.IN_2'(s(sK13,X42),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X43) != s(bool,X44)) ) | $spl6),
  inference(subsumption_resolution,[],[f1121,f901])).
fof(f1121,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X42)))) | ~p(s(bool,X43)) | p(s(bool,X44)) | s(bool,X45) != s(bool,'const.bool.IN_2'(s(sK13,X42),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X43) != s(bool,X44)) ) | $spl6),
  inference(resolution,[],[f916,f651])).
fof(f1120,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X38)))) | ~p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X40) != s(bool,'const.bool.IN_2'(s(sK13,X38),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X39) != s(bool,X41)) ) | $spl6),
  inference(resolution,[],[f916,f649])).
fof(f1268,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,X34),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X35) != s(bool,X36)) ) | $spl6),
  inference(subsumption_resolution,[],[f1119,f901])).
fof(f1119,plain,(
  ( ! [X37,X35,X36,X34] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X34)))) | ~p(s(bool,X35)) | p(s(bool,X36)) | s(bool,X37) != s(bool,'const.bool.IN_2'(s(sK13,X34),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15))))) | s(bool,X35) != s(bool,X36)) ) | $spl6),
  inference(resolution,[],[f916,f647])).
fof(f1118,plain,(
  ( ! [X30,X33,X31,X32] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X30)))) | p(s(bool,X31)) | ~p(s(bool,X32)) | s(bool,X31) != s(bool,X33) | s(bool,X32) != s(bool,'const.bool.IN_2'(s(sK13,X30),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) ) | $spl6),
  inference(resolution,[],[f916,f647])).
fof(f1094,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | $spl24),
  inference(cnf_transformation,[],[f1094_D])).
fof(f1094_D,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1266,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) | $spl29),
  inference(cnf_transformation,[],[f1266_D])).
fof(f1266_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f1113,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK52(sK15,sK13))))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | $spl6),
  inference(resolution,[],[f916,f627])).
fof(f1101,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) | $spl27),
  inference(cnf_transformation,[],[f1101_D])).
fof(f1101_D,plain,(
  ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1095,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | $spl25),
  inference(cnf_transformation,[],[f1095_D])).
fof(f1095_D,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f1049,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK34(sK15,sK14,sK13))))) | $spl1),
  inference(superposition,[],[f911,f591])).
fof(f1093,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl23),
  inference(cnf_transformation,[],[f1093_D])).
fof(f1093_D,plain,(
  ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f1044,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),sK15)))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))) | $spl1),
  inference(superposition,[],[f911,f891])).
fof(f891,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | ~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f551])).
fof(f551,plain,(
  ! [X0,X1,X2] : ((s('type.list.list'(X0),X2) != s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) | p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))))) & (~p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(nnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0] : ! [X1,X2] : (s('type.list.list'(X0),X2) = s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))),
  inference(rectify,[],[f74])).
fof(f74,axiom,(
  ! [X6] : ! [X7,X11] : (s('type.list.list'(X6),X11) = s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X11))) <=> p(s(bool,'const.list.EVERY_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X11)))))),
  file('thm.list.NULL_FILTER','thm.list.FILTER_EQ_ID')).
fof(f1030,plain,(
  ( ! [X192,X193] : (p(s(bool,X192)) | ~p(s(bool,X193)) | s(bool,X193) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f901])).
fof(f932,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | ~p(s(bool,X36)) | s(bool,X36) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X35) != s(bool,X37)) ) | $spl1),
  inference(resolution,[],[f911,f653])).
fof(f931,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X32)) | ~p(s(bool,X33)) | ~p(s(bool,X34)) | s(bool,X33) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f652])).
fof(f929,plain,(
  ( ! [X28,X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X26) != s(bool,X28) | s(bool,X27) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f651])).
fof(f928,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | p(s(bool,X24)) | s(bool,X25) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X23) != s(bool,X24)) ) | $spl1),
  inference(resolution,[],[f911,f651])).
fof(f927,plain,(
  ( ! [X21,X22,X20] : (~p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X21) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X20) != s(bool,X22)) ) | $spl1),
  inference(resolution,[],[f911,f649])).
fof(f926,plain,(
  ( ! [X19,X17,X18] : (~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X19) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))) | s(bool,X17) != s(bool,X18)) ) | $spl1),
  inference(resolution,[],[f911,f647])).
fof(f925,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,X16) | s(bool,X15) != s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) ) | $spl1),
  inference(resolution,[],[f911,f647])).
fof(f910,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) | $spl0),
  inference(cnf_transformation,[],[f910_D])).
fof(f910_D,plain,(
  p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f916,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X4),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4))))) ) | $spl6),
  inference(cnf_transformation,[],[f916_D])).
fof(f916_D,plain,(
  ( ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X4),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f557,plain,(
  ( ! [X4] : (~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4)))) | ~p(s(bool,'const.bool.IN_2'(s(sK13,X4),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))) )),
  inference(cnf_transformation,[],[f373])).
fof(f373,plain,(
  (p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(sK13,X4),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,X4)))))) & (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) | (p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) & p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13,sK14,sK15,sK16])],[f372])).
fof(f372,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))) | ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))))) & (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(rectify,[],[f371])).
fof(f371,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))) | ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))) & (~p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) & p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))))),
  inference(nnf_transformation,[],[f298])).
fof(f298,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))) <~> ! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(ennf_transformation,[],[f136])).
fof(f136,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X0),'const.list.FILTER_2'(s(fun(X0,bool),X1),s('type.list.list'(X0),X2)))))) <=> ! [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X0),X2)))))) => ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) => ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.list.NULL_1'(s('type.list.list'(X6),'const.list.FILTER_2'(s(fun(X6,bool),X7),s('type.list.list'(X6),X8)))))) <=> ! [X9] : (p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(X6),X8)))))) => ~p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X9))))))),
  file('thm.list.NULL_FILTER','thm.list.NULL_FILTER')).
fof(f911,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) | $spl1),
  inference(cnf_transformation,[],[f911_D])).
fof(f911_D,plain,(
  ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f915,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | $spl4),
  inference(cnf_transformation,[],[f915_D])).
fof(f915_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f558,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK13,sK16),s(fun(sK13,bool),'const.list.LIST_TO_SET_1'(s('type.list.list'(sK13),sK15)))))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))),
  inference(cnf_transformation,[],[f373])).
fof(f913,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | $spl2),
  inference(cnf_transformation,[],[f913_D])).
fof(f913_D,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f559,plain,(
  p(s(bool,happ(s(fun(sK13,bool),sK14),s(sK13,sK16)))) | ~p(s(bool,'const.list.NULL_1'(s('type.list.list'(sK13),'const.list.FILTER_2'(s(fun(sK13,bool),sK14),s('type.list.list'(sK13),sK15))))))),
  inference(cnf_transformation,[],[f373])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1148
% Passive clauses: 17611
% Generated clauses: 50654
% Final active clauses: 755
% Final passive clauses: 8067
% Input formulas: 133
% Initial clauses: 344
% 
% Selected by SInE selection: 100
% SInE iterations: 1
% 
% Duplicate literals: 1153
% Trivial inequalities: 113
% Fw subsumption resolutions: 5496
% Fw demodulations: 6091
% Bw demodulations: 50
% 
% Simple tautologies: 2285
% Equational tautologies: 3
% Forward subsumptions: 17775
% 
% Binary resolution: 17738
% Factoring: 2
% Forward superposition: 12342
% Backward superposition: 7306
% Equality resolution: 7
% 
% Split clauses: 769
% Split components: 1611
% SAT solver clauses: 855
% SAT solver unit clauses: 33
% SAT solver binary clauses: 38
% 
% Sat splits: 769
% Sat splitting refutations: 79
% 
% Memory used [KB]: 17782
% Time elapsed: 1.091 s
% ------------------------------
% ------------------------------
% Success in time 12.81 s
