Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (51495)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 151
% Passive clauses: 3216
% Generated clauses: 22966
% Final active clauses: 127
% Final passive clauses: 2763
% Input formulas: 29
% Initial clauses: 63
% 
% Duplicate literals: 857
% Trivial inequalities: 212
% Fw subsumption resolutions: 3385
% Fw demodulations: 3645
% Bw demodulations: 22
% 
% Simple tautologies: 388
% Equational tautologies: 45
% Forward subsumptions: 10123
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 954
% Factoring: 4
% Forward superposition: 7248
% Backward superposition: 6151
% Self superposition: 168
% Equality factoring: 100
% Equality resolution: 28
% 
% Split clauses: 554
% Split components: 1130
% SAT solver clauses: 1351
% SAT solver unit clauses: 21
% SAT solver binary clauses: 241
% 
% Sat splits: 554
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1346
% TWLsolver calls for satisfiability: 200
% 
% Memory used [KB]: 11129
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (51506)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 108
% Passive clauses: 1134
% Generated clauses: 12381
% Final active clauses: 105
% Final passive clauses: 1011
% Input formulas: 29
% Initial clauses: 63
% 
% Duplicate literals: 616
% Trivial inequalities: 76
% Fw subsumption resolutions: 1135
% Fw demodulations: 1157
% Bw demodulations: 13
% Condensations: 880
% Global subsumptions: 5
% 
% Simple tautologies: 173
% Equational tautologies: 37
% Forward subsumptions: 4921
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 800
% Factoring: 2
% Forward superposition: 3703
% Backward superposition: 2602
% Self superposition: 103
% Equality factoring: 72
% Equality resolution: 11
% 
% Split clauses: 119
% Split components: 247
% SAT solver clauses: 1779
% SAT solver unit clauses: 54
% SAT solver binary clauses: 35
% SAT solver learnt clauses: 10
% 
% Sat splits: 119
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1759
% TWLsolver calls for satisfiability: 270
% 
% Memory used [KB]: 10490
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (51522)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 685
% Passive clauses: 8554
% Generated clauses: 151955
% Final active clauses: 106
% Final passive clauses: 1007
% Input formulas: 29
% Initial clauses: 63
% 
% Duplicate literals: 7470
% Trivial inequalities: 984
% Fw subsumption resolutions: 13647
% Fw demodulations: 20453
% Bw demodulations: 5275
% Condensations: 17435
% Global subsumptions: 493
% 
% Simple tautologies: 1261
% Equational tautologies: 1943
% Forward subsumptions: 77154
% Fw demodulations to eq. taut.: 62
% Bw demodulations to eq. taut.: 453
% 
% Unit resulting resolution: 9044
% Factoring: 17
% Forward superposition: 38203
% Backward superposition: 30156
% Self superposition: 996
% Equality factoring: 592
% Equality resolution: 123
% 
% Split clauses: 373
% Split components: 770
% SAT solver clauses: 10744
% SAT solver unit clauses: 138
% SAT solver binary clauses: 752
% 
% Sat splits: 373
% Sat splitting refutations: 949
% 
% Memory used [KB]: 54114
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f2530,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f215,f216,f227,f217,f218,f228,f219,f220,f221,f222,f225,f223,f232,f234,f224,f230,f236,f262,f263,f462,f464,f226,f494,f495,f651,f652,f663,f653,f654,f664,f655,f656,f657,f658,f661,f659,f668,f670,f660,f666,f672,f698,f699,f834,f835,f463,f662,f906,f907,f908,f2526])).
fof(f2526,plain,(
  $false | $spl0),
  inference(subsumption_resolution,[],[f2525,f157])).
fof(f157,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.COMPL_CLAUSES|split|0',aHL_TRUTH)).
fof(f2525,plain,(
  ~p(s(bool,'T_0')) | $spl0),
  inference(forward_demodulation,[],[f2495,f2337])).
fof(f2337,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),sK1))) | $spl0),
  inference(resolution,[],[f2290,f115])).
fof(f115,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f78])).
fof(f78,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.bool.EQ_CLAUSES|split|1')).
fof(f2290,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),sK1)))) | $spl0),
  inference(resolution,[],[f226,f120])).
fof(f120,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))) | p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(flattening,[],[f44])).
fof(f44,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X7,X10,X11] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X10)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))) & p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.pred_set.IN_INTER')).
fof(f2495,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),sK1)))) | $spl0),
  inference(resolution,[],[f2289,f122])).
fof(f122,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))))) & (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2))))))))),
  inference(nnf_transformation,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(flattening,[],[f46])).
fof(f46,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X2)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6] : ! [X11,X7] : (p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7)))))) <=> ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7)))))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.pred_set.IN_COMPL')).
fof(f2289,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1)))))) | $spl0),
  inference(resolution,[],[f226,f119])).
fof(f119,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))) )),
  inference(cnf_transformation,[],[f81])).
fof(f908,plain,(
  $false | ($spl6 | $spl18 | $spl22)),
  inference(trivial_inequality_removal,[],[f903])).
fof(f903,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl6 | $spl18 | $spl22)),
  inference(superposition,[],[f874,f229])).
fof(f229,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl6),
  inference(cnf_transformation,[],[f229_D])).
fof(f229_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f874,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | ($spl18 | $spl22)),
  inference(subsumption_resolution,[],[f873,f112])).
fof(f112,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f40])).
fof(f40,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f873,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2))) ) | ($spl18 | $spl22)),
  inference(forward_demodulation,[],[f663,f665])).
fof(f665,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl22),
  inference(cnf_transformation,[],[f665_D])).
fof(f665_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f907,plain,(
  $false | ($spl18 | $spl22)),
  inference(trivial_inequality_removal,[],[f905])).
fof(f905,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl18 | $spl22)),
  inference(superposition,[],[f874,f665])).
fof(f906,plain,(
  $false | ($spl18 | $spl22)),
  inference(equality_resolution,[],[f874])).
fof(f662,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl16),
  inference(cnf_transformation,[],[f662_D])).
fof(f662_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f463,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | $spl14),
  inference(cnf_transformation,[],[f463_D])).
fof(f463_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f835,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl22),
  inference(subsumption_resolution,[],[f822,f110])).
fof(f110,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f38])).
fof(f38,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.bool.EQ_CLAUSES|split|2')).
fof(f822,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) ) | $spl22),
  inference(superposition,[],[f700,f102])).
fof(f102,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f33])).
fof(f33,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X8,X9] : ((p(s(bool,X8)) => p(s(bool,X9))) => ((p(s(bool,X9)) => p(s(bool,X8))) => s(bool,X8) = s(bool,X9)))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.bool.IMP_ANTISYM_AX')).
fof(f700,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl22),
  inference(backward_demodulation,[],[f665,f379])).
fof(f379,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f165])).
fof(f165,plain,(
  ( ! [X1] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X1) | s(bool,'const.bool.IN_2'(s(sK0,sK2(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),X1))) != s(bool,'const.bool.IN_2'(s(sK0,sK2(X1,'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f101,f107])).
fof(f107,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK2(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK2])],[f73])).
fof(f73,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f35])).
fof(f35,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f34])).
fof(f34,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6] : ! [X7,X10] : (s(fun(X6,bool),X7) = s(fun(X6,bool),X10) <=> ! [X11] : s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X7))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),X10))))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.pred_set.EXTENSION')).
fof(f101,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),
  inference(cnf_transformation,[],[f71])).
fof(f71,plain,(
  s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f65])).
fof(f65,plain,(
  ? [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(ennf_transformation,[],[f32])).
fof(f32,plain,(
  ~! [X0,X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f31])).
fof(f31,plain,(
  ~! [X0] : ! [X1] : s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),'const.pred_set.INTER_2'(s(fun(X0,bool),'const.pred_set.COMPL_1'(s(fun(X0,bool),X1))),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),'const.pred_set.INTER_2'(s(fun(X6,bool),'const.pred_set.COMPL_1'(s(fun(X6,bool),X7))),s(fun(X6,bool),X7)))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.pred_set.COMPL_CLAUSES|split|0')).
fof(f834,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl22),
  inference(subsumption_resolution,[],[f821,f110])).
fof(f821,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl22),
  inference(superposition,[],[f700,f102])).
fof(f699,plain,(
  $false | $spl23),
  inference(subsumption_resolution,[],[f692,f158])).
fof(f158,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X6] : ! [X11] : ~p(s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.pred_set.NOT_IN_EMPTY')).
fof(f692,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(trivial_inequality_removal,[],[f686])).
fof(f686,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(superposition,[],[f666,f113])).
fof(f113,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f698,plain,(
  $false | $spl23),
  inference(subsumption_resolution,[],[f693,f158])).
fof(f693,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(trivial_inequality_removal,[],[f685])).
fof(f685,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl23),
  inference(superposition,[],[f666,f111])).
fof(f111,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f76])).
fof(f672,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl28),
  inference(cnf_transformation,[],[f672_D])).
fof(f672_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f666,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl23),
  inference(cnf_transformation,[],[f666_D])).
fof(f666_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f660,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f379,f118])).
fof(f118,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.COMPL_CLAUSES|split|0',aHL_BOOL_CASES)).
fof(f670,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl26),
  inference(cnf_transformation,[],[f670_D])).
fof(f670_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f668,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl25),
  inference(cnf_transformation,[],[f668_D])).
fof(f668_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f659,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f379,f118])).
fof(f661,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl17),
  inference(cnf_transformation,[],[f661_D])).
fof(f661_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f658,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f379,f117])).
fof(f117,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.pred_set.COMPL_CLAUSES|split|0','thm.bool.EQ_CLAUSES|split|0')).
fof(f657,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f379,f115])).
fof(f656,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f379,f113])).
fof(f655,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f379,f111])).
fof(f664,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) | $spl20),
  inference(cnf_transformation,[],[f664_D])).
fof(f664_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f654,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f379,f105])).
fof(f105,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f67])).
fof(f653,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f379,f105])).
fof(f663,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f663_D])).
fof(f663_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f652,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f379,f102])).
fof(f651,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.EMPTY_0','const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f379,f102])).
fof(f495,plain,(
  $false | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f493])).
fof(f493,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6)),
  inference(superposition,[],[f440,f229])).
fof(f440,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2)) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f439,f112])).
fof(f439,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f227,f229])).
fof(f494,plain,(
  $false | ($spl2 | $spl6)),
  inference(equality_resolution,[],[f440])).
fof(f226,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl0),
  inference(cnf_transformation,[],[f226_D])).
fof(f226_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f464,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f450,f110])).
fof(f450,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f265,f102])).
fof(f265,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl6),
  inference(backward_demodulation,[],[f229,f206])).
fof(f206,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(equality_resolution,[],[f164])).
fof(f164,plain,(
  ( ! [X0] : (s(fun(sK0,bool),'const.pred_set.EMPTY_0') != s(fun(sK0,bool),X0) | s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),X0,sK0)),s(fun(sK0,bool),X0))) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),X0,sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) )),
  inference(superposition,[],[f101,f107])).
fof(f462,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl6),
  inference(subsumption_resolution,[],[f449,f110])).
fof(f449,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) ) | $spl6),
  inference(superposition,[],[f265,f102])).
fof(f263,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f256,f158])).
fof(f256,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(trivial_inequality_removal,[],[f250])).
fof(f250,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(superposition,[],[f230,f113])).
fof(f262,plain,(
  $false | $spl7),
  inference(subsumption_resolution,[],[f257,f158])).
fof(f257,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(trivial_inequality_removal,[],[f249])).
fof(f249,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0')))) | $spl7),
  inference(superposition,[],[f230,f111])).
fof(f236,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl12),
  inference(cnf_transformation,[],[f236_D])).
fof(f236_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f230,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl7),
  inference(cnf_transformation,[],[f230_D])).
fof(f230_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f224,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f206,f118])).
fof(f234,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) | $spl10),
  inference(cnf_transformation,[],[f234_D])).
fof(f234_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f232,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | $spl9),
  inference(cnf_transformation,[],[f232_D])).
fof(f232_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f223,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))),
  inference(superposition,[],[f206,f118])).
fof(f225,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | $spl1),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f222,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f206,f117])).
fof(f221,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f206,f115])).
fof(f220,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f206,f113])).
fof(f219,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))),
  inference(superposition,[],[f206,f111])).
fof(f228,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) | $spl4),
  inference(cnf_transformation,[],[f228_D])).
fof(f228_D,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X4))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f218,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,X5)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f206,f105])).
fof(f217,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | ~p(s(bool,X4))) )),
  inference(superposition,[],[f206,f105])).
fof(f227,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) | $spl2),
  inference(cnf_transformation,[],[f227_D])).
fof(f227_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f216,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)))))) | p(s(bool,X3))) )),
  inference(superposition,[],[f206,f102])).
fof(f215,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.EMPTY_0'))) | p(s(bool,X2)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK2('const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1)),'const.pred_set.EMPTY_0',sK0)),s(fun(sK0,bool),'const.pred_set.INTER_2'(s(fun(sK0,bool),'const.pred_set.COMPL_1'(s(fun(sK0,bool),sK1))),s(fun(sK0,bool),sK1))))))) )),
  inference(superposition,[],[f206,f102])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 99
% Passive clauses: 732
% Generated clauses: 2359
% Final active clauses: 81
% Final passive clauses: 573
% Input formulas: 29
% Initial clauses: 58
% 
% Selected by SInE selection: 26
% SInE iterations: 1
% 
% Duplicate literals: 23
% Trivial inequalities: 84
% Fw subsumption resolutions: 212
% Fw demodulations: 309
% Bw demodulations: 14
% 
% Simple tautologies: 84
% Forward subsumptions: 936
% 
% Binary resolution: 138
% Forward superposition: 987
% Backward superposition: 565
% Equality resolution: 11
% 
% Split clauses: 24
% Split components: 48
% SAT solver clauses: 35
% SAT solver unit clauses: 5
% SAT solver binary clauses: 24
% 
% Sat splits: 24
% Sat splitting refutations: 10
% 
% Memory used [KB]: 2046
% Time elapsed: 0.060 s
% ------------------------------
% ------------------------------
% Success in time 11.831 s
