Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (53132)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 202
% Passive clauses: 2272
% Generated clauses: 25950
% Final active clauses: 173
% Final passive clauses: 1795
% Input formulas: 32
% Initial clauses: 73
% 
% Duplicate literals: 1204
% Trivial inequalities: 649
% Fw subsumption resolutions: 4076
% Fw demodulations: 4392
% Bw demodulations: 25
% 
% Simple tautologies: 586
% Equational tautologies: 31
% Forward subsumptions: 11494
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 201
% Forward superposition: 8656
% Backward superposition: 8874
% Self superposition: 217
% Equality factoring: 65
% Equality resolution: 68
% 
% Split clauses: 924
% Split components: 1865
% SAT solver clauses: 1526
% SAT solver unit clauses: 20
% SAT solver binary clauses: 112
% 
% Sat splits: 924
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 1518
% TWLsolver calls for satisfiability: 201
% 
% Memory used [KB]: 14072
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (53171)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 114
% Passive clauses: 268
% Generated clauses: 14725
% Final active clauses: 105
% Final passive clauses: 143
% Input formulas: 32
% Initial clauses: 73
% 
% Duplicate literals: 626
% Trivial inequalities: 172
% Fw subsumption resolutions: 1790
% Fw demodulations: 2188
% Bw demodulations: 16
% Condensations: 1644
% Global subsumptions: 8
% 
% Simple tautologies: 255
% Equational tautologies: 27
% Forward subsumptions: 6712
% 
% Binary resolution: 29
% Forward superposition: 4228
% Backward superposition: 3412
% Self superposition: 158
% Equality factoring: 51
% Equality resolution: 58
% 
% Split clauses: 40
% Split components: 80
% SAT solver clauses: 596
% SAT solver unit clauses: 84
% SAT solver binary clauses: 32
% SAT solver learnt clauses: 16
% 
% Sat splits: 40
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 564
% TWLsolver calls for satisfiability: 299
% 
% Memory used [KB]: 9210
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f8760,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f8731])).
fof(f8731,plain,(
  s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))) != s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1)))))),
  inference(superposition,[],[f120,f7416])).
fof(f7416,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f7415,f1202])).
fof(f1202,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),X0))) = s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),X0)))) )),
  inference(forward_demodulation,[],[f1188,f125])).
fof(f125,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),X0))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_SET_0'),s(fun('type.num.num',bool),X0)))) )),
  inference(cnf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0] : s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),X0))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_SET_0'),s(fun('type.num.num',bool),X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X21] : s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),X21))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_SET_0'),s(fun('type.num.num',bool),X21)))),
  file('thm.pred_set.PROD_SET_THM|split|1','ARITY_0_1_const.pred_set.PROD_SET')).
fof(f1188,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),X0))) = s('type.num.num',happ(s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_SET_0'),s(fun('type.num.num',bool),X0)))) )),
  inference(superposition,[],[f170,f123])).
fof(f123,plain,(
  s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_SET_0') = s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_IMAGE_1'(s(fun('type.num.num','type.num.num'),'const.combin.I_0')))),
  inference(cnf_transformation,[],[f30])).
fof(f30,axiom,(
  s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_SET_0') = s(fun(fun('type.num.num',bool),'type.num.num'),'const.pred_set.PROD_IMAGE_1'(s(fun('type.num.num','type.num.num'),'const.combin.I_0')))),
  file('thm.pred_set.PROD_SET_THM|split|1','thm.pred_set.PROD_SET_DEF')).
fof(f170,plain,(
  ( ! [X2,X0,X1] : (s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))) = s('type.num.num',happ(s(fun(fun(X0,bool),'type.num.num'),'const.pred_set.PROD_IMAGE_1'(s(fun(X0,'type.num.num'),X1))),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))) = s('type.num.num',happ(s(fun(fun(X0,bool),'type.num.num'),'const.pred_set.PROD_IMAGE_1'(s(fun(X0,'type.num.num'),X1))),s(fun(X0,bool),X2)))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),X2))) = s('type.num.num',happ(s(fun(fun(X0,bool),'type.num.num'),'const.pred_set.PROD_IMAGE_1'(s(fun(X0,'type.num.num'),X1))),s(fun(X0,bool),X2)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X12] : ! [X21,X22] : s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X12,'type.num.num'),X21),s(fun(X12,bool),X22))) = s('type.num.num',happ(s(fun(fun(X12,bool),'type.num.num'),'const.pred_set.PROD_IMAGE_1'(s(fun(X12,'type.num.num'),X21))),s(fun(X12,bool),X22)))),
  file('thm.pred_set.PROD_SET_THM|split|1','ARITY_1_2_const.pred_set.PROD_IMAGE')).
fof(f7415,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X0)))))))) )),
  inference(forward_demodulation,[],[f7358,f1202])).
fof(f7358,plain,(
  ( ! [X0] : (s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),'const.combin.I_0'),s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X0)))))))) )),
  inference(superposition,[],[f1261,f210])).
fof(f210,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(forward_demodulation,[],[f166,f165])).
fof(f165,plain,(
  ( ! [X0,X1] : (s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))) )),
  inference(cnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : ! [X1] : s(X0,X1) = s(X0,'const.combin.I_1'(s(X0,X1)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X12] : ! [X6] : s(X12,X6) = s(X12,'const.combin.I_1'(s(X12,X6)))),
  file('thm.pred_set.PROD_SET_THM|split|1','thm.combin.I_THM')).
fof(f166,plain,(
  ( ! [X0,X1] : (s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0,X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(flattening,[],[f56])).
fof(f56,plain,(
  ! [X0] : ! [X1] : s(X0,'const.combin.I_1'(s(X0,X1))) = s(X0,happ(s(fun(X0,X0),'const.combin.I_0'),s(X0,X1)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X12] : ! [X21] : s(X12,'const.combin.I_1'(s(X12,X21))) = s(X12,happ(s(fun(X12,X12),'const.combin.I_0'),s(X12,X21)))),
  file('thm.pred_set.PROD_SET_THM|split|1','ARITY_0_1_const.combin.I')).
fof(f1261,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X1),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X1)))))))) )),
  inference(subsumption_resolution,[],[f1253,f121])).
fof(f121,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.PROD_SET_THM|split|1',aHL_TRUTH)).
fof(f1253,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X1),s(fun('type.num.num',bool),sK1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun('type.num.num','type.num.num'),X0),s('type.num.num',X1))),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun('type.num.num','type.num.num'),X0),s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',X1)))))))) )),
  inference(superposition,[],[f173,f218])).
fof(f218,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),sK1)))),
  inference(unit_resulting_resolution,[],[f119,f128])).
fof(f128,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.PROD_SET_THM|split|1','thm.bool.EQ_CLAUSES|split|0')).
fof(f119,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),sK1))))),
  inference(cnf_transformation,[],[f85])).
fof(f85,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),sK1)))) & s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',sK0)))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1])],[f73])).
fof(f73,plain,(
  ? [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) & s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X1),s('type.num.num',X0))))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X1)))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X0),s(fun('type.num.num',bool),X1))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X1),s('type.num.num',X0))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X7)))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X6),s(fun('type.num.num',bool),X7))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X7),s('type.num.num',X6))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun('type.num.num',bool),X7)))) => s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',X6),s(fun('type.num.num',bool),X7))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',X6),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),X7),s('type.num.num',X6))))))))),
  file('thm.pred_set.PROD_SET_THM|split|1','thm.pred_set.PROD_SET_THM|split|1')).
fof(f173,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) | s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2)))))))) )),
  inference(cnf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) | s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(ennf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) => s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X3)))) => s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X0,'type.num.num'),X1),s(X0,X2))),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X0,'type.num.num'),X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X3),s(X0,X2))))))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X12] : ! [X19,X20,X7] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X12,bool),X7)))) => s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X12,'type.num.num'),X19),s(fun(X12,bool),'const.pred_set.INSERT_2'(s(X12,X20),s(fun(X12,bool),X7))))) = s('type.num.num','const.arithmetic.*_2'(s('type.num.num',happ(s(fun(X12,'type.num.num'),X19),s(X12,X20))),s('type.num.num','const.pred_set.PROD_IMAGE_2'(s(fun(X12,'type.num.num'),X19),s(fun(X12,bool),'const.pred_set.DELETE_2'(s(fun(X12,bool),X7),s(X12,X20))))))))),
  file('thm.pred_set.PROD_SET_THM|split|1','thm.pred_set.PROD_IMAGE_THM|split|1')).
fof(f120,plain,(
  s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.INSERT_2'(s('type.num.num',sK0),s(fun('type.num.num',bool),sK1))))) != s('type.num.num','const.arithmetic.*_2'(s('type.num.num',sK0),s('type.num.num','const.pred_set.PROD_SET_1'(s(fun('type.num.num',bool),'const.pred_set.DELETE_2'(s(fun('type.num.num',bool),sK1),s('type.num.num',sK0)))))))),
  inference(cnf_transformation,[],[f85])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 130
% Passive clauses: 399
% Generated clauses: 8528
% Final active clauses: 73
% Final passive clauses: 143
% Input formulas: 32
% Initial clauses: 73
% 
% Duplicate literals: 567
% Trivial inequalities: 163
% Fw subsumption resolutions: 504
% Fw demodulations: 1036
% Bw demodulations: 103
% Condensations: 732
% Global subsumptions: 17
% 
% Simple tautologies: 222
% Equational tautologies: 183
% Forward subsumptions: 4595
% Fw demodulations to eq. taut.: 28
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 705
% Forward superposition: 2258
% Backward superposition: 2014
% Self superposition: 116
% Equality factoring: 68
% Equality resolution: 21
% 
% Split clauses: 38
% Split components: 82
% SAT solver clauses: 513
% SAT solver unit clauses: 29
% SAT solver binary clauses: 128
% 
% Sat splits: 38
% Sat splitting refutations: 19
% 
% Memory used [KB]: 3709
% Time elapsed: 0.286 s
% ------------------------------
% ------------------------------
% Success in time 1.613 s
