Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.measure
% (82401)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 133
% Passive clauses: 2621
% Generated clauses: 12354
% Final active clauses: 120
% Final passive clauses: 2247
% Input formulas: 31
% Initial clauses: 108
% 
% Duplicate literals: 568
% Trivial inequalities: 199
% Fw subsumption resolutions: 2329
% Fw demodulations: 624
% Bw demodulations: 11
% 
% Simple tautologies: 317
% Equational tautologies: 82
% Forward subsumptions: 4522
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 719
% Factoring: 2
% Forward superposition: 3290
% Backward superposition: 3906
% Self superposition: 242
% Equality factoring: 210
% Equality resolution: 23
% 
% Split clauses: 427
% Split components: 894
% SAT solver clauses: 1272
% SAT solver unit clauses: 52
% SAT solver binary clauses: 410
% 
% Sat splits: 427
% 
% TWLsolver clauses: 1271
% TWLsolver calls for satisfiability: 202
% 
% Memory used [KB]: 5500
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.measure
% (82424)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 201
% Passive clauses: 860
% Generated clauses: 18867
% Final active clauses: 167
% Final passive clauses: 569
% Input formulas: 31
% Initial clauses: 108
% 
% Duplicate literals: 1024
% Trivial inequalities: 198
% Fw subsumption resolutions: 1997
% Fw demodulations: 3157
% Bw demodulations: 101
% Condensations: 1068
% Global subsumptions: 13
% 
% Simple tautologies: 309
% Equational tautologies: 15
% Forward subsumptions: 8765
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 783
% Factoring: 1
% Forward superposition: 5882
% Backward superposition: 4189
% Self superposition: 95
% Equality factoring: 28
% Equality resolution: 40
% 
% Split clauses: 75
% Split components: 152
% SAT solver clauses: 1274
% SAT solver unit clauses: 119
% SAT solver binary clauses: 48
% SAT solver learnt clauses: 26
% 
% Sat splits: 75
% Sat splitting refutations: 6
% 
% TWLsolver clauses: 1222
% TWLsolver calls for satisfiability: 1338
% 
% Memory used [KB]: 11641
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.measure
% (82442)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 469
% Passive clauses: 3923
% Generated clauses: 131115
% Final active clauses: 140
% Final passive clauses: 546
% Input formulas: 31
% Initial clauses: 108
% Discarded non-redundant clauses: 5949
% 
% Duplicate literals: 8911
% Trivial inequalities: 1044
% Fw subsumption resolutions: 14379
% Fw demodulations: 15280
% Bw demodulations: 2306
% Condensations: 16912
% Global subsumptions: 158
% 
% Simple tautologies: 1255
% Equational tautologies: 964
% Forward subsumptions: 67348
% Fw demodulations to eq. taut.: 10
% Bw demodulations to eq. taut.: 64
% 
% Unit resulting resolution: 3440
% Factoring: 13
% Forward superposition: 36017
% Backward superposition: 29781
% Self superposition: 693
% Equality factoring: 540
% Equality resolution: 56
% 
% Split clauses: 465
% Split components: 970
% SAT solver clauses: 5067
% SAT solver unit clauses: 21
% SAT solver binary clauses: 190
% 
% Sat splits: 465
% Sat splitting refutations: 54
% 
% Memory used [KB]: 58591
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.measure
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.measure
% SZS output start Proof for thm.measure
fof(f12854,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f164,f163,f267,f268,f269,f270,f271,f272,f273,f274,f275,f277,f291,f312,f391,f393,f394,f396,f398,f412,f432,f433,f434,f435,f436,f437,f540,f487,f481,f564,f565,f566,f627,f626,f631,f633,f637,f638,f639,f640,f675,f669,f751,f752,f753,f806,f841,f884,f883,f887,f888,f896,f898,f899,f900,f903,f915,f939,f941,f942,f944,f946,f960,f979,f980,f981,f982,f983,f984,f1206,f1207,f1620,f1621,f1622,f1623,f1624,f1625,f1584,f1586,f1587,f1588,f1591,f1603,f1642,f1450,f1643,f1438,f1647,f1648,f1649,f1650,f1651,f1652,f1653,f1654,f1717,f1703,f1691,f1690,f1688,f1687,f1686,f1685,f1684,f1818,f1819,f1820,f1821,f1823,f1824,f1825,f1828,f1840,f1994,f1827,f1822,f2006,f2007,f2008,f2009,f2010,f2011,f2012,f2031,f2032,f2033,f2034,f2035,f2037,f315,f314,f2038,f2050,f1723,f2674,f2678,f2633,f2635,f2636,f2637,f2640,f2652,f2723,f2722,f2786,f2787,f2788,f2789,f2790,f2791,f2798,f2800,f2801,f2803,f2805,f2819,f2843,f2909,f2910,f2911,f2912,f2913,f2914,f2870,f2871,f2873,f2875,f2876,f2877,f2880,f2892,f2935,f2937,f2938,f2939,f2942,f2954,f3001,f3002,f3003,f3004,f3005,f3006,f3042,f3041,f3103,f3104,f3105,f3106,f3107,f3108,f3151,f3152,f3153,f3203,f3205,f3206,f3208,f3210,f3224,f3263,f3507,f3508,f3509,f3510,f3511,f3512,f3471,f3473,f3474,f3475,f3478,f3490,f3570,f3572,f3574,f3576,f3578,f3580,f3533,f3535,f3536,f3537,f3540,f3552,f3395,f311,f3394,f3393,f308,f3612,f3613,f3614,f3615,f3616,f3617,f3625,f3626,f3627,f3628,f3629,f3630,f3678,f3679,f3608,f3606,f3692,f3694,f3695,f3696,f3699,f3711,f3755,f3756,f3759,f4032,f310,f3991,f3993,f3994,f3995,f3998,f4010,f2681,f4218,f4220,f4221,f4222,f4225,f4237,f4281,f4283,f4284,f4286,f4288,f4302,f4394,f4396,f4398,f4400,f4402,f4404,f4357,f4359,f4360,f4361,f4364,f4376,f4422,f4424,f4425,f4426,f4429,f4441,f4471,f4624,f4627,f4630,f4633,f4636,f4639,f4720,f4722,f4723,f4725,f4727,f4741,f4329,f4817,f4834,f4837,f4840,f4841,f4839,f4765,f4842,f4799,f4785,f4783,f4781,f4780,f4778,f4843,f3618,f3619,f3620,f3621,f3622,f3623,f259,f4850,f4851,f4852,f4853,f4854,f4855,f4897,f1995,f4898,f4899,f4948,f4950,f4951,f4953,f4955,f4969,f426,f5042,f5043,f5044,f5045,f5046,f5047,f5006,f5008,f5009,f5010,f5013,f5025,f1629,f5124,f5126,f5128,f5130,f5132,f444,f5134,f5087,f5089,f5090,f5091,f5094,f5106,f3581,f5321,f5323,f5324,f5325,f5328,f5340,f5379,f5594,f5596,f5597,f5599,f5601,f5615,f5644,f5679,f5680,f5681,f5698,f5700,f5701,f5702,f5705,f5717,f5745,f5752,f5754,f5755,f5756,f5759,f5771,f5803,f5805,f5806,f5807,f5810,f5822,f5842,f6006,f6867,f6979,f6994,f7001,f7004,f6625,f6624,f6623,f6009,f6008,f6007,f7011,f7013,f7016,f5799,f7006,f5798,f7007,f5797,f447,f5642,f4836,f5317,f7018,f5316,f7019,f5315,f7020,f7024,f7028,f7032,f5251,f7034,f7039,f7042,f5002,f5001,f5000,f7053,f7058,f7061,f446,f7073,f7078,f7080,f7081,f7082,f7085,f7097,f931,f7132,f7134,f7135,f7136,f7139,f7151,f4270,f7368,f7370,f7371,f7373,f7375,f7389,f7407,f7419,f7421,f7422,f7423,f7426,f7438,f5737,f7472,f7474,f7475,f7476,f7479,f7491,f5791,f8192,f8197,f8160,f8146,f8144,f8142,f8141,f8139,f8320,f8321,f8322,f8578,f8583,f8552,f8538,f8536,f8534,f8533,f8531,f8593,f8596,f8598,f8599,f8601,f8603,f8617,f8635,f8876,f8877,f8878,f8879,f8927,f8928,f8931,f8933,f8989,f8994,f7868,f7418,f7301,f8209,f8207,f9052,f2675,f7076,f309,f8460,f2680,f8456,f2676,f3758,f3757,f3754,f3751,f3750,f8957,f8943,f8941,f8939,f8938,f8936,f9056,f9063,f9064,f1660,f9065,f9066,f9067,f9068,f442,f440,f9110,f9109,f9113,f2001,f9118,f9119,f9120,f9121,f9122,f1992,f9123,f9124,f9127,f9129,f9130,f9131,f9134,f9146,f9171,f9179,f9181,f9182,f9183,f9186,f9198,f9231,f9286,f9287,f9170,f9241,f9243,f9244,f9246,f9248,f9262,f9288,f9296,f9298,f9299,f9301,f9303,f9317,f9340,f9636,f9638,f9639,f9640,f9643,f9655,f9719,f9718,f9724,f9723,f9728,f9729,f9730,f9731,f441,f9732,f9734,f9879,f9878,f9884,f9883,f9886,f9887,f9888,f9889,f439,f9890,f630,f9892,f10124,f10126,f10127,f10129,f10131,f10145,f8211,f10467,f10466,f10432,f10431,f10503,f10504,f10505,f10506,f10507,f10508,f10512,f10513,f10795,f10797,f10798,f10799,f10802,f10814,f10840,f11531,f11546,f11564,f11568,f11572,f11576,f11580,f11584,f11588,f11592,f9601,f9587,f9585,f9583,f9582,f9580,f9619,f11553,f11552,f1655,f11499,f11485,f11483,f11481,f11480,f11478,f11594,f11596,f11598,f11600,f11602,f11604,f11606,f11608,f11610,f11409,f11407,f10793,f11611,f10792,f11612,f10059,f10057,f9752,f9751,f11620,f11624,f11628,f11632,f1663,f7866,f7864,f7417,f7416,f11642,f11645,f11648,f11651,f11654,f11657,f11660,f11663,f1665,f11667,f1662,f11671,f9167,f1658,f1656,f11673,f9168,f11713,f11715,f1664,f11734,f11736,f11737,f11738,f11741,f11753,f11773,f11785,f11787,f11788,f11789,f11792,f11804,f2974,f11835,f11837,f11838,f11839,f11842,f11854,f11874,f11888,f11890,f11891,f11892,f11895,f11907,f11933,f12153,f12155,f12156,f313,f12157,f12160,f316,f12172,f9681,f12654,f12542,f12341,f12152,f11887,f11834,f11784,f10061,f9753,f260,f12538,f12536,f12534,f12337,f12335,f12333,f12149,f12757,f12148,f12758,f12147,f12746,f12080,f12078,f11884,f12761,f11883,f12762,f11882,f12748,f11831,f12753,f11830,f12754,f11829,f12742,f11781,f12755,f11780,f12756,f11779,f12744,f11730,f12759,f11729,f12760,f12764,f12768,f12770,f12774,f12780,f12784,f12786,f12790,f10055,f10053,f10051,f12796,f12800,f12802,f12806,f9750,f2919,f9749,f2918,f9748,f7009,f12812,f12816,f12818,f448,f12822,f8193,f12827,f8990,f7862,f7860,f7858,f12832,f12836,f7415,f11638,f7414,f11639,f7413,f9048,f6995,f12843,f452,f12847,f6996,f10177,f257,f306,f8919,f8918,f8930,f8929,f450,f8924,f8923,f318,f4901,f4894,f4893,f4900,f4902,f3156,f3155,f3154,f3148,f3147,f1989,f1988,f1993,f1998,f2000,f317,f2002,f1999,f632,f258,f12852,f12853])).
fof(f12853,plain,(
  $false | ($spl2 | $spl55)),
  inference(subsumption_resolution,[],[f264,f1658])).
fof(f264,plain,(
  sP0(sK9,sK10) | $spl2),
  inference(resolution,[],[f259,f150])).
fof(f150,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X9),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X9))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | (s('type.num.num',sK3(X6,X1,X0)) != s('type.num.num',sK4(X6,X1,X0)) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK3(X6,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK4(X6,X1,X0)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK5(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP1(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3,sK4,sK5,sK6])],[f96])).
fof(f96,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X9] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X9),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X9))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X6] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X7,X8] : (s('type.num.num',X7) != s('type.num.num',X8) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',X7))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',X8))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP1(X0,X1)))),
  inference(rectify,[],[f95])).
fof(f95,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP1(X0,X1)))),
  inference(flattening,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ~sP0(X0,X1) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | sP1(X0,X1)))),
  inference(nnf_transformation,[],[f90])).
fof(f90,plain,(
  ! [X0,X1] : (sP1(X0,X1) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & sP0(X0,X1) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f12852,plain,(
  $false | ($spl2 | $spl55)),
  inference(subsumption_resolution,[],[f4847,f1658])).
fof(f4847,plain,(
  sP0(sK9,sK10) | $spl2),
  inference(resolution,[],[f259,f150])).
fof(f258,plain,(
  ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl1),
  inference(cnf_transformation,[],[f258_D])).
fof(f258_D,plain,(
  ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f632,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) | $spl44),
  inference(cnf_transformation,[],[f632_D])).
fof(f632_D,plain,(
  ( ! [X21] : (s(bool,'T_0') != s(bool,X21)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f1999,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1984,f189])).
fof(f189,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.EQ_CLAUSES|split|0')).
fof(f1984,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f192])).
fof(f192,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f26])).
fof(f26,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.EQ_CLAUSES|split|1')).
fof(f2002,plain,(
  s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl70),
  inference(cnf_transformation,[],[f2002_D])).
fof(f2002_D,plain,(
  s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f317,plain,(
  ( ! [X19] : (~p(s(bool,X19))) ) | $spl22),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  ( ! [X19] : (~p(s(bool,X19))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2000,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl12),
  inference(subsumption_resolution,[],[f1985,f190])).
fof(f190,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f112])).
fof(f1985,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,X12) | ~p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl12),
  inference(superposition,[],[f312,f193])).
fof(f193,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT',aHL_BOOL_CASES)).
fof(f1998,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f1983,f189])).
fof(f1983,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f190])).
fof(f1993,plain,(
  s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl67),
  inference(cnf_transformation,[],[f1993_D])).
fof(f1993_D,plain,(
  s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f1988,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X11))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f1974])).
fof(f1974,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X11)) | ~p(s(bool,X11))) ) | $spl12),
  inference(superposition,[],[f312,f192])).
fof(f1989,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X10))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f1973])).
fof(f1973,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl12),
  inference(superposition,[],[f312,f190])).
fof(f3147,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X11))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f3133])).
fof(f3133,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X11)) | ~p(s(bool,X11))) ) | $spl12),
  inference(superposition,[],[f312,f192])).
fof(f3148,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X10))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f3132])).
fof(f3132,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl12),
  inference(superposition,[],[f312,f190])).
fof(f3154,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f3142,f189])).
fof(f3142,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f190])).
fof(f3155,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f3143,f189])).
fof(f3143,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f192])).
fof(f3156,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl12),
  inference(subsumption_resolution,[],[f3144,f190])).
fof(f3144,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,X12) | ~p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl12),
  inference(superposition,[],[f312,f193])).
fof(f4902,plain,(
  ( ! [X12] : (~p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl12),
  inference(subsumption_resolution,[],[f4890,f190])).
fof(f4890,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,X12) | ~p(s(bool,X12)) | s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl12),
  inference(superposition,[],[f312,f193])).
fof(f4900,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f4888,f189])).
fof(f4888,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f190])).
fof(f4893,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X11))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f4879])).
fof(f4879,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X11)) | ~p(s(bool,X11))) ) | $spl12),
  inference(superposition,[],[f312,f192])).
fof(f4894,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X10))) ) | $spl12),
  inference(duplicate_literal_removal,[],[f4878])).
fof(f4878,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X10)) | ~p(s(bool,X10))) ) | $spl12),
  inference(superposition,[],[f312,f190])).
fof(f4901,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(subsumption_resolution,[],[f4889,f189])).
fof(f4889,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f192])).
fof(f318,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) | $spl24),
  inference(cnf_transformation,[],[f318_D])).
fof(f318_D,plain,(
  ( ! [X17] : (p(s(bool,X17))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f8923,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl135),
  inference(cnf_transformation,[],[f8923_D])).
fof(f8923_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) <=> ~$spl135),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl135])])).
fof(f8924,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X13))) ) | $spl108),
  inference(subsumption_resolution,[],[f8900,f189])).
fof(f8900,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X13)) ) | $spl108),
  inference(superposition,[],[f5791,f193])).
fof(f450,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl38),
  inference(cnf_transformation,[],[f450_D])).
fof(f450_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f8929,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl108),
  inference(subsumption_resolution,[],[f8911,f186])).
fof(f186,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f110])).
fof(f110,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.EQ_CLAUSES|split|2')).
fof(f8911,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl108),
  inference(superposition,[],[f5791,f186])).
fof(f8930,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl108),
  inference(subsumption_resolution,[],[f8912,f186])).
fof(f8912,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X9)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl108),
  inference(superposition,[],[f5791,f188])).
fof(f188,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.EQ_CLAUSES|split|3')).
fof(f8918,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9))) ) | $spl108),
  inference(duplicate_literal_removal,[],[f8896])).
fof(f8896,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9)) | p(s(bool,X9))) ) | $spl108),
  inference(superposition,[],[f5791,f188])).
fof(f8919,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X8))) ) | $spl108),
  inference(duplicate_literal_removal,[],[f8895])).
fof(f8895,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X8)) | p(s(bool,X8))) ) | $spl108),
  inference(superposition,[],[f5791,f186])).
fof(f306,plain,(
  ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl5),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f257,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl0),
  inference(cnf_transformation,[],[f257_D])).
fof(f257_D,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f10177,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl75 | $spl77)),
  inference(subsumption_resolution,[],[f10123,f2676])).
fof(f10123,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl77),
  inference(resolution,[],[f2680,f169])).
fof(f169,plain,(
  ( ! [X4,X0,X3,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s('type.num.num',X4)))))) | s('type.num.num',X3) = s('type.num.num',X4) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | (s('type.num.num',sK11(X5,X1,X0)) != s('type.num.num',sK12(X5,X1,X0)) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',sK11(X5,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',sK12(X5,X1,X0)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK11,sK12,sK13])],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X5] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X6,X7] : (s('type.num.num',X6) != s('type.num.num',X7) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X6))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',X7))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(rectify,[],[f104])).
fof(f104,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(flattening,[],[f103])).
fof(f103,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & ((~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ? [X2] : (p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) = s('type.num.num',X4) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f72])).
fof(f72,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X2] : ((~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(ennf_transformation,[],[f36])).
fof(f36,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X2] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f35])).
fof(f35,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) & ! [X2] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (~s('type.num.num',X3) = s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) <=> (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) & ! [X10] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X6,bool)),X10),s(fun(fun('type.num.num',fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))) & ! [X11,X12] : (~s('type.num.num',X11) = s('type.num.num',X12) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X10),s('type.num.num',X11))),s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X10),s('type.num.num',X12)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X6,bool)),X10),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.measure.SIGMA_ALGEBRA_ALT_DISJOINT')).
fof(f6996,plain,(
  sP2(sK9,sK10) | $spl112),
  inference(cnf_transformation,[],[f6996_D])).
fof(f6996_D,plain,(
  sP2(sK9,sK10) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f12847,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12846,f219])).
fof(f219,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f126])).
fof(f126,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.IMP_CLAUSES|split|3')).
fof(f12846,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12845,f9892])).
fof(f12845,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12844,f220])).
fof(f220,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f127])).
fof(f12844,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(forward_demodulation,[],[f7366,f9734])).
fof(f7366,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl33),
  inference(resolution,[],[f443,f200])).
fof(f200,plain,(
  ( ! [X0,X1] : (sP2(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK16(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK14(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK15(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK16(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK15(X1,X0)),s(fun(X0,bool),sK16(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP2(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16])],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP2(X0,X1)))),
  inference(rectify,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & (~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP2(X0,X1)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0,X1] : ((~sP2(X0,X1) | (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))) & ((~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | ? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | sP2(X0,X1)))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0,X1] : (sP2(X0,X1) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f443,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl33),
  inference(cnf_transformation,[],[f443_D])).
fof(f443_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f452,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl41),
  inference(cnf_transformation,[],[f452_D])).
fof(f452_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f12843,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12842,f220])).
fof(f12842,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12841,f9892])).
fof(f12841,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12840,f220])).
fof(f12840,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(forward_demodulation,[],[f7367,f9734])).
fof(f7367,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl33),
  inference(resolution,[],[f443,f202])).
fof(f202,plain,(
  ( ! [X0,X1] : (sP2(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK15(X1,X0)),s(fun(X0,bool),sK16(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f6995,plain,(
  ~sP2(sK9,sK10) | $spl113),
  inference(cnf_transformation,[],[f6995_D])).
fof(f6995_D,plain,(
  ~sP2(sK9,sK10) <=> ~$spl113),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl113])])).
fof(f9048,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl138),
  inference(cnf_transformation,[],[f9048_D])).
fof(f9048_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f7413,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl34),
  inference(resolution,[],[f446,f196])).
fof(f196,plain,(
  ( ! [X0,X7,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f117])).
fof(f11639,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl176),
  inference(cnf_transformation,[],[f11639_D])).
fof(f11639_D,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f7414,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl34),
  inference(resolution,[],[f446,f197])).
fof(f197,plain,(
  ( ! [X6,X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f117])).
fof(f11638,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl174),
  inference(cnf_transformation,[],[f11638_D])).
fof(f11638_D,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl174),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl174])])).
fof(f7415,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl34),
  inference(resolution,[],[f446,f197])).
fof(f12836,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12835,f219])).
fof(f12835,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12834,f9892])).
fof(f12834,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12833,f220])).
fof(f12833,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(forward_demodulation,[],[f7635,f9734])).
fof(f7635,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl33),
  inference(subsumption_resolution,[],[f7533,f235])).
fof(f235,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f136])).
fof(f136,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : ((p(s(bool,X9)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.AND_CLAUSES|split|3')).
fof(f7533,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl33),
  inference(superposition,[],[f200,f7378])).
fof(f7378,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl33),
  inference(resolution,[],[f443,f186])).
fof(f12832,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12831,f220])).
fof(f12831,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12830,f9892])).
fof(f12830,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12829,f220])).
fof(f12829,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(forward_demodulation,[],[f12828,f9734])).
fof(f12828,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl33),
  inference(subsumption_resolution,[],[f7534,f217])).
fof(f217,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f124])).
fof(f124,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X9)))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.IMP_CLAUSES|split|4')).
fof(f7534,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl33),
  inference(superposition,[],[f202,f7378])).
fof(f7858,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl34),
  inference(subsumption_resolution,[],[f7754,f219])).
fof(f7754,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl34),
  inference(superposition,[],[f196,f7431])).
fof(f7431,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl34),
  inference(resolution,[],[f446,f190])).
fof(f7860,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl34),
  inference(subsumption_resolution,[],[f7755,f220])).
fof(f7755,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl34),
  inference(superposition,[],[f197,f7431])).
fof(f7862,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl34),
  inference(subsumption_resolution,[],[f7756,f220])).
fof(f7756,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl34),
  inference(superposition,[],[f197,f7431])).
fof(f8990,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl97 | $spl99)),
  inference(subsumption_resolution,[],[f8935,f4839])).
fof(f8935,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl97),
  inference(resolution,[],[f4836,f169])).
fof(f12827,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl75 | $spl77)),
  inference(subsumption_resolution,[],[f8138,f2676])).
fof(f8138,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl77),
  inference(resolution,[],[f2680,f169])).
fof(f8193,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl75 | $spl77)),
  inference(subsumption_resolution,[],[f8138,f2676])).
fof(f12822,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33 | $spl164)),
  inference(subsumption_resolution,[],[f12821,f219])).
fof(f12821,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33 | $spl164)),
  inference(forward_demodulation,[],[f12820,f9892])).
fof(f12820,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33)),
  inference(subsumption_resolution,[],[f12819,f235])).
fof(f12819,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33)),
  inference(forward_demodulation,[],[f9744,f7378])).
fof(f9744,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl30),
  inference(resolution,[],[f441,f200])).
fof(f448,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl37),
  inference(cnf_transformation,[],[f448_D])).
fof(f448_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f12818,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl164)),
  inference(subsumption_resolution,[],[f12817,f220])).
fof(f12817,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl164)),
  inference(forward_demodulation,[],[f9745,f9892])).
fof(f9745,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl30),
  inference(resolution,[],[f441,f201])).
fof(f201,plain,(
  ( ! [X0,X1] : (sP2(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK16(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK14(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f12816,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33 | $spl164)),
  inference(subsumption_resolution,[],[f12815,f220])).
fof(f12815,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33 | $spl164)),
  inference(forward_demodulation,[],[f12814,f9892])).
fof(f12814,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33)),
  inference(subsumption_resolution,[],[f12813,f217])).
fof(f12813,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl33)),
  inference(forward_demodulation,[],[f9746,f7378])).
fof(f9746,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl30),
  inference(resolution,[],[f441,f202])).
fof(f12812,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl164)),
  inference(subsumption_resolution,[],[f12811,f220])).
fof(f12811,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl30 | $spl164)),
  inference(forward_demodulation,[],[f9747,f9892])).
fof(f9747,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl30),
  inference(resolution,[],[f441,f203])).
fof(f203,plain,(
  ( ! [X0,X1] : (sP2(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK15(X1,X0)),s(fun(X0,bool),sK16(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK14(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f7009,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl118),
  inference(cnf_transformation,[],[f7009_D])).
fof(f7009_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f9748,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl30),
  inference(resolution,[],[f441,f196])).
fof(f2918,plain,(
  ( ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X6))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl82),
  inference(cnf_transformation,[],[f2918_D])).
fof(f2918_D,plain,(
  ( ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X6))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f9749,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl30),
  inference(resolution,[],[f441,f197])).
fof(f2919,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X7),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X7),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl84),
  inference(cnf_transformation,[],[f2919_D])).
fof(f2919_D,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X7),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X7),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9750,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X5),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl30),
  inference(resolution,[],[f441,f197])).
fof(f12806,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12805,f219])).
fof(f12805,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12804,f9892])).
fof(f12804,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12803,f235])).
fof(f12803,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(forward_demodulation,[],[f10047,f7378])).
fof(f10047,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(subsumption_resolution,[],[f9934,f220])).
fof(f9934,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(superposition,[],[f200,f9734])).
fof(f12802,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12801,f220])).
fof(f12801,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(forward_demodulation,[],[f10048,f9892])).
fof(f10048,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(subsumption_resolution,[],[f9935,f220])).
fof(f9935,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(superposition,[],[f201,f9734])).
fof(f12800,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12799,f220])).
fof(f12799,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12798,f9892])).
fof(f12798,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12797,f217])).
fof(f12797,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160)),
  inference(forward_demodulation,[],[f10049,f7378])).
fof(f10049,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(subsumption_resolution,[],[f9936,f220])).
fof(f9936,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(superposition,[],[f202,f9734])).
fof(f12796,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12795,f220])).
fof(f12795,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(forward_demodulation,[],[f10050,f9892])).
fof(f10050,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(subsumption_resolution,[],[f9937,f220])).
fof(f9937,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(superposition,[],[f203,f9734])).
fof(f10051,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl160),
  inference(subsumption_resolution,[],[f9938,f219])).
fof(f9938,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl160),
  inference(superposition,[],[f196,f9734])).
fof(f10053,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl160),
  inference(subsumption_resolution,[],[f9939,f220])).
fof(f9939,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl160),
  inference(superposition,[],[f197,f9734])).
fof(f10055,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X5),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl160),
  inference(subsumption_resolution,[],[f9940,f220])).
fof(f9940,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X5),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl160),
  inference(superposition,[],[f197,f9734])).
fof(f12790,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12789,f219])).
fof(f12789,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33 | $spl160)),
  inference(forward_demodulation,[],[f12788,f9734])).
fof(f12788,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33)),
  inference(subsumption_resolution,[],[f12787,f235])).
fof(f12787,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33)),
  inference(forward_demodulation,[],[f10742,f7378])).
fof(f10742,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl28),
  inference(resolution,[],[f439,f200])).
fof(f12786,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl160)),
  inference(subsumption_resolution,[],[f12785,f220])).
fof(f12785,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl160)),
  inference(forward_demodulation,[],[f10743,f9734])).
fof(f10743,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl28),
  inference(resolution,[],[f439,f201])).
fof(f12784,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33 | $spl160)),
  inference(subsumption_resolution,[],[f12783,f220])).
fof(f12783,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33 | $spl160)),
  inference(forward_demodulation,[],[f12782,f9734])).
fof(f12782,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33)),
  inference(subsumption_resolution,[],[f12781,f217])).
fof(f12781,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl33)),
  inference(forward_demodulation,[],[f10744,f7378])).
fof(f10744,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl28),
  inference(resolution,[],[f439,f202])).
fof(f12780,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl160)),
  inference(subsumption_resolution,[],[f12779,f220])).
fof(f12779,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl28 | $spl160)),
  inference(forward_demodulation,[],[f10745,f9734])).
fof(f10745,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl28),
  inference(resolution,[],[f439,f203])).
fof(f12774,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12773,f219])).
fof(f12773,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12772,f9734])).
fof(f12772,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl164)),
  inference(subsumption_resolution,[],[f12771,f235])).
fof(f12771,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl164)),
  inference(forward_demodulation,[],[f11191,f7378])).
fof(f11191,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(subsumption_resolution,[],[f11076,f220])).
fof(f11076,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(superposition,[],[f200,f9892])).
fof(f12770,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12769,f220])).
fof(f12769,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(forward_demodulation,[],[f11192,f9734])).
fof(f11192,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(subsumption_resolution,[],[f11077,f220])).
fof(f11077,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(superposition,[],[f201,f9892])).
fof(f12768,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12767,f220])).
fof(f12767,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f12766,f9734])).
fof(f12766,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl164)),
  inference(subsumption_resolution,[],[f12765,f217])).
fof(f12765,plain,(
  p(s(bool,'F_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl33 | $spl164)),
  inference(forward_demodulation,[],[f11193,f7378])).
fof(f11193,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(subsumption_resolution,[],[f11078,f220])).
fof(f11078,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(superposition,[],[f202,f9892])).
fof(f12764,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(subsumption_resolution,[],[f12763,f220])).
fof(f12763,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl160 | $spl164)),
  inference(forward_demodulation,[],[f11194,f9734])).
fof(f11194,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(subsumption_resolution,[],[f11079,f220])).
fof(f11079,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(superposition,[],[f203,f9892])).
fof(f12760,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl208),
  inference(cnf_transformation,[],[f12760_D])).
fof(f12760_D,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f11729,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl52),
  inference(resolution,[],[f1656,f197])).
fof(f12759,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl206),
  inference(cnf_transformation,[],[f12759_D])).
fof(f12759_D,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f11730,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl52),
  inference(resolution,[],[f1656,f197])).
fof(f12744,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl188),
  inference(cnf_transformation,[],[f12744_D])).
fof(f12744_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f11779,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl55),
  inference(resolution,[],[f11674,f196])).
fof(f11674,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl55),
  inference(resolution,[],[f1658,f160])).
fof(f160,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK7(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK7(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK8(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK7(X1,X0)),s(fun(X0,bool),sK8(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK7,sK8])],[f99])).
fof(f99,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X4,X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & (? [X2,X3] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(rectify,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & (? [X5,X6] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f12756,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl200),
  inference(cnf_transformation,[],[f12756_D])).
fof(f12756_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f11780,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(resolution,[],[f11674,f197])).
fof(f12755,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl198),
  inference(cnf_transformation,[],[f12755_D])).
fof(f12755_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f11781,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(resolution,[],[f11674,f197])).
fof(f12742,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl186),
  inference(cnf_transformation,[],[f12742_D])).
fof(f12742_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f11829,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl55),
  inference(resolution,[],[f11675,f196])).
fof(f11675,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl55),
  inference(resolution,[],[f1658,f161])).
fof(f161,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK8(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f12754,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl196),
  inference(cnf_transformation,[],[f12754_D])).
fof(f12754_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f11830,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(resolution,[],[f11675,f197])).
fof(f12753,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl194),
  inference(cnf_transformation,[],[f12753_D])).
fof(f12753_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f11831,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(resolution,[],[f11675,f197])).
fof(f12748,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl192),
  inference(cnf_transformation,[],[f12748_D])).
fof(f12748_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f11882,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl62),
  inference(resolution,[],[f1664,f196])).
fof(f12762,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl212),
  inference(cnf_transformation,[],[f12762_D])).
fof(f12762_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f11883,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl62),
  inference(resolution,[],[f1664,f197])).
fof(f12761,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl210),
  inference(cnf_transformation,[],[f12761_D])).
fof(f12761_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f11884,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl62),
  inference(resolution,[],[f1664,f197])).
fof(f12078,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl52),
  inference(subsumption_resolution,[],[f11973,f220])).
fof(f11973,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl52),
  inference(superposition,[],[f197,f11746])).
fof(f11746,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl52),
  inference(resolution,[],[f1656,f190])).
fof(f12080,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl52),
  inference(subsumption_resolution,[],[f11974,f220])).
fof(f11974,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl52),
  inference(superposition,[],[f197,f11746])).
fof(f12746,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl190),
  inference(cnf_transformation,[],[f12746_D])).
fof(f12746_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f12147,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl58),
  inference(resolution,[],[f1661,f196])).
fof(f1661,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl58),
  inference(cnf_transformation,[],[f1661_D])).
fof(f1661_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f12758,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl204),
  inference(cnf_transformation,[],[f12758_D])).
fof(f12758_D,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl204),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl204])])).
fof(f12148,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl58),
  inference(resolution,[],[f1661,f197])).
fof(f12757,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl202),
  inference(cnf_transformation,[],[f12757_D])).
fof(f12757_D,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f12149,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl58),
  inference(resolution,[],[f1661,f197])).
fof(f12333,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl55),
  inference(subsumption_resolution,[],[f12232,f219])).
fof(f12232,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl55),
  inference(superposition,[],[f196,f11797])).
fof(f11797,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl55),
  inference(resolution,[],[f11674,f190])).
fof(f12335,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(subsumption_resolution,[],[f12233,f220])).
fof(f12233,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(superposition,[],[f197,f11797])).
fof(f12337,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(subsumption_resolution,[],[f12234,f220])).
fof(f12234,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(superposition,[],[f197,f11797])).
fof(f12534,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl55),
  inference(subsumption_resolution,[],[f12433,f219])).
fof(f12433,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl55),
  inference(superposition,[],[f196,f11847])).
fof(f11847,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl55),
  inference(resolution,[],[f11675,f190])).
fof(f12536,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(subsumption_resolution,[],[f12434,f220])).
fof(f12434,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(superposition,[],[f197,f11847])).
fof(f12538,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(subsumption_resolution,[],[f12435,f220])).
fof(f12435,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl55),
  inference(superposition,[],[f197,f11847])).
fof(f260,plain,(
  ~sP1(sK9,sK10) | $spl3),
  inference(cnf_transformation,[],[f260_D])).
fof(f260_D,plain,(
  ~sP1(sK9,sK10) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f9753,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl30),
  inference(resolution,[],[f441,f149])).
fof(f149,plain,(
  ( ! [X0,X1,X9] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X9))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X9),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f10061,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl160),
  inference(subsumption_resolution,[],[f9943,f219])).
fof(f9943,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl160),
  inference(superposition,[],[f149,f9734])).
fof(f11784,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl55),
  inference(resolution,[],[f11674,f149])).
fof(f11834,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl55),
  inference(resolution,[],[f11675,f149])).
fof(f11887,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl62),
  inference(resolution,[],[f1664,f149])).
fof(f12152,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl58),
  inference(resolution,[],[f1661,f149])).
fof(f12341,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl55),
  inference(subsumption_resolution,[],[f12237,f219])).
fof(f12237,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK7(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl55),
  inference(superposition,[],[f149,f11797])).
fof(f12542,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl55),
  inference(subsumption_resolution,[],[f12438,f219])).
fof(f12438,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl55),
  inference(superposition,[],[f149,f11847])).
fof(f12654,plain,(
  $false | ($spl4 | $spl55)),
  inference(subsumption_resolution,[],[f12653,f253])).
fof(f253,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.TRUTH')).
fof(f12653,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl55)),
  inference(forward_demodulation,[],[f12652,f11797])).
fof(f12652,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl4 | $spl55)),
  inference(subsumption_resolution,[],[f12651,f223])).
fof(f223,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f128])).
fof(f128,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X9] : ((p(s(bool,'F_0')) => p(s(bool,X9))) <=> p(s(bool,'T_0')))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.IMP_CLAUSES|split|2')).
fof(f12651,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl4 | $spl55)),
  inference(forward_demodulation,[],[f12650,f11847])).
fof(f12650,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl4 | $spl55)),
  inference(subsumption_resolution,[],[f12605,f892])).
fof(f892,plain,(
  sP2(sK9,sK10) | $spl4),
  inference(resolution,[],[f305,f204])).
fof(f204,plain,(
  ( ! [X0,X1] : (sP2(X0,X1) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | sP2(X0,X1)) & (~sP2(X0,X1) | p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> sP2(X0,X1))),
  inference(definition_folding,[],[f79,f92])).
fof(f79,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X4))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2,X3] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X2),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X3),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & ! [X8] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(X6,bool),X8))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) & ! [X8,X9] : ((p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.measure.algebra_def')).
fof(f305,plain,(
  p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl4),
  inference(cnf_transformation,[],[f305_D])).
fof(f305_D,plain,(
  p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f12605,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl55),
  inference(resolution,[],[f11676,f197])).
fof(f11676,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(sK9,bool),sK8(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl55),
  inference(resolution,[],[f1658,f162])).
fof(f162,plain,(
  ( ! [X0,X1] : (sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),sK7(X1,X0)),s(fun(X0,bool),sK8(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f100])).
fof(f9681,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl156),
  inference(cnf_transformation,[],[f9681_D])).
fof(f9681_D,plain,(
  ( ! [X7] : (p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f12172,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f255])).
fof(f255,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(duplicate_literal_removal,[],[f175])).
fof(f175,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | ~p(s(bool,X1)) | p(s(bool,X3))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | ~p(s(bool,X0)) | p(s(bool,X2)))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | (((p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X1)) | p(s(bool,X3)))) & ((p(s(bool,X1)) & ~p(s(bool,X3))) | (~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0,X1,X2,X3] : (s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3)) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(flattening,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) != s(bool,X1) | (p(s(bool,X1)) & s(bool,X2) != s(bool,X3))) | ((~p(s(bool,X0)) | p(s(bool,X2))) <=> (~p(s(bool,X1)) | p(s(bool,X3)))))),
  inference(ennf_transformation,[],[f37])).
fof(f37,plain,(
  ! [X0,X1,X2,X3] : ((s(bool,X0) = s(bool,X1) & (p(s(bool,X1)) => s(bool,X2) = s(bool,X3))) => ((p(s(bool,X0)) => p(s(bool,X2))) <=> (p(s(bool,X1)) => p(s(bool,X3)))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X15,X18,X16,X19] : ((s(bool,X15) = s(bool,X18) & (p(s(bool,X18)) => s(bool,X16) = s(bool,X19))) => ((p(s(bool,X15)) => p(s(bool,X16))) <=> (p(s(bool,X18)) => p(s(bool,X19)))))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.IMP_CONG')).
fof(f316,plain,(
  ( ! [X10,X8] : (s(bool,X8) != s(bool,X10) | ~p(s(bool,X8))) ) | $spl20),
  inference(cnf_transformation,[],[f316_D])).
fof(f316_D,plain,(
  ( ! [X10,X8] : (s(bool,X8) != s(bool,X10) | ~p(s(bool,X8))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f12160,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X28) != s(bool,X29)) ) | $spl58),
  inference(resolution,[],[f1661,f178])).
fof(f178,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f12157,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f177])).
fof(f177,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X3)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f313,plain,(
  ( ! [X4,X2] : (s(bool,X2) != s(bool,X4) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f313_D])).
fof(f313_D,plain,(
  ( ! [X4,X2] : (s(bool,X2) != s(bool,X4) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f12156,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f176])).
fof(f176,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X2)) | ~p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f12155,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X13) != s(bool,X14)) ) | $spl58),
  inference(resolution,[],[f1661,f174])).
fof(f174,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f12153,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f172])).
fof(f172,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,X3)) | ~p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X2) != s(bool,X3) | s(bool,X0) != s(bool,X1)) )),
  inference(cnf_transformation,[],[f108])).
fof(f11933,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl184),
  inference(cnf_transformation,[],[f11933_D])).
fof(f11933_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f11907,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl62),
  inference(resolution,[],[f1664,f255])).
fof(f11895,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl62),
  inference(resolution,[],[f1664,f178])).
fof(f11892,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl62),
  inference(resolution,[],[f1664,f177])).
fof(f11891,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl62),
  inference(resolution,[],[f1664,f176])).
fof(f11890,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl62),
  inference(resolution,[],[f1664,f174])).
fof(f11888,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl62),
  inference(resolution,[],[f1664,f172])).
fof(f11874,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl182),
  inference(cnf_transformation,[],[f11874_D])).
fof(f11874_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f11854,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11675,f255])).
fof(f11842,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl55),
  inference(resolution,[],[f11675,f178])).
fof(f11839,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11675,f177])).
fof(f11838,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11675,f176])).
fof(f11837,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl55),
  inference(resolution,[],[f11675,f174])).
fof(f11835,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK8(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11675,f172])).
fof(f2974,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl86),
  inference(cnf_transformation,[],[f2974_D])).
fof(f2974_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f11804,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11674,f255])).
fof(f11792,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl55),
  inference(resolution,[],[f11674,f178])).
fof(f11789,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11674,f177])).
fof(f11788,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11674,f176])).
fof(f11787,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl55),
  inference(resolution,[],[f11674,f174])).
fof(f11785,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f11674,f172])).
fof(f11773,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl180),
  inference(cnf_transformation,[],[f11773_D])).
fof(f11773_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f11753,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl52),
  inference(resolution,[],[f1656,f255])).
fof(f11741,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl52),
  inference(resolution,[],[f1656,f178])).
fof(f11738,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl52),
  inference(resolution,[],[f1656,f177])).
fof(f11737,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl52),
  inference(resolution,[],[f1656,f176])).
fof(f11736,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl52),
  inference(resolution,[],[f1656,f174])).
fof(f11734,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl52),
  inference(resolution,[],[f1656,f172])).
fof(f1664,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl62),
  inference(cnf_transformation,[],[f1664_D])).
fof(f1664_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f11715,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9))))))) | $spl179),
  inference(cnf_transformation,[],[f11715_D])).
fof(f11715_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9))))))) <=> ~$spl179),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl179])])).
fof(f11713,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl56 | $spl66 | $spl142)),
  inference(subsumption_resolution,[],[f11712,f220])).
fof(f11712,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl56 | $spl66 | $spl142)),
  inference(forward_demodulation,[],[f11711,f1992])).
fof(f11711,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl56 | $spl142)),
  inference(subsumption_resolution,[],[f11710,f220])).
fof(f11710,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl56 | $spl142)),
  inference(forward_demodulation,[],[f11686,f9139])).
fof(f9139,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | $spl56),
  inference(resolution,[],[f1660,f190])).
fof(f11686,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl142),
  inference(superposition,[],[f167,f9167])).
fof(f167,plain,(
  ( ! [X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',sK11(X5,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X5),s('type.num.num',sK12(X5,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f9168,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | $spl143),
  inference(cnf_transformation,[],[f9168_D])).
fof(f9168_D,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f11673,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ($spl56 | $spl66)),
  inference(subsumption_resolution,[],[f11672,f219])).
fof(f11672,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ($spl56 | $spl66)),
  inference(forward_demodulation,[],[f10985,f1992])).
fof(f10985,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl56),
  inference(subsumption_resolution,[],[f10882,f220])).
fof(f10882,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl56),
  inference(superposition,[],[f166,f9139])).
fof(f166,plain,(
  ( ! [X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | s('type.num.num',sK11(X5,X1,X0)) != s('type.num.num',sK12(X5,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X5),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f1656,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl52),
  inference(cnf_transformation,[],[f1656_D])).
fof(f1656_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1658,plain,(
  ~sP0(sK9,sK10) | $spl55),
  inference(cnf_transformation,[],[f1658_D])).
fof(f1658_D,plain,(
  ~sP0(sK9,sK10) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f9167,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | $spl142),
  inference(cnf_transformation,[],[f9167_D])).
fof(f9167_D,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f11671,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11670,f219])).
fof(f11670,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11669,f9892])).
fof(f11669,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160)),
  inference(subsumption_resolution,[],[f11668,f220])).
fof(f11668,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160)),
  inference(forward_demodulation,[],[f11532,f9734])).
fof(f11532,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145)),
  inference(subsumption_resolution,[],[f11476,f260])).
fof(f11476,plain,(
  sP1(sK9,sK10) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl145),
  inference(resolution,[],[f9170,f155])).
fof(f155,plain,(
  ( ! [X4,X0,X3,X1] : (sP1(X0,X1) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s('type.num.num',X4)))))) | s('type.num.num',X3) = s('type.num.num',X4) | ~sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f1662,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl59),
  inference(cnf_transformation,[],[f1662_D])).
fof(f1662_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f11667,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11666,f220])).
fof(f11666,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11665,f9892])).
fof(f11665,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160)),
  inference(subsumption_resolution,[],[f11664,f220])).
fof(f11664,plain,(
  ~p(s(bool,'T_0')) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145 | $spl160)),
  inference(forward_demodulation,[],[f11539,f9734])).
fof(f11539,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl145)),
  inference(subsumption_resolution,[],[f11477,f260])).
fof(f11477,plain,(
  sP1(sK9,sK10) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl145),
  inference(resolution,[],[f9170,f156])).
fof(f156,plain,(
  ( ! [X4,X0,X3,X1] : (sP1(X0,X1) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s('type.num.num',X4)))))) | s('type.num.num',X3) = s('type.num.num',X4) | ~sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK5(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f1665,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl63),
  inference(cnf_transformation,[],[f1665_D])).
fof(f1665_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f11663,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl30 | $spl164)),
  inference(subsumption_resolution,[],[f11662,f220])).
fof(f11662,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl30 | $spl164)),
  inference(forward_demodulation,[],[f11661,f9892])).
fof(f11661,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f9740,f260])).
fof(f9740,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl30),
  inference(resolution,[],[f441,f157])).
fof(f157,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f11660,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl30 | $spl164)),
  inference(subsumption_resolution,[],[f11659,f220])).
fof(f11659,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl30 | $spl164)),
  inference(forward_demodulation,[],[f11658,f9892])).
fof(f11658,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f9741,f260])).
fof(f9741,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl30),
  inference(resolution,[],[f441,f158])).
fof(f158,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK5(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f11657,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11656,f220])).
fof(f11656,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11655,f9892])).
fof(f11655,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f10043,f260])).
fof(f10043,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(subsumption_resolution,[],[f9929,f220])).
fof(f9929,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(superposition,[],[f157,f9734])).
fof(f11654,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11653,f220])).
fof(f11653,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11652,f9892])).
fof(f11652,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f10044,f260])).
fof(f10044,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(subsumption_resolution,[],[f9930,f220])).
fof(f9930,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl160),
  inference(superposition,[],[f158,f9734])).
fof(f11651,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl28 | $spl160)),
  inference(subsumption_resolution,[],[f11650,f220])).
fof(f11650,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl28 | $spl160)),
  inference(forward_demodulation,[],[f11649,f9734])).
fof(f11649,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl28)),
  inference(subsumption_resolution,[],[f10738,f260])).
fof(f10738,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl28),
  inference(resolution,[],[f439,f157])).
fof(f11648,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl28 | $spl160)),
  inference(subsumption_resolution,[],[f11647,f220])).
fof(f11647,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl28 | $spl160)),
  inference(forward_demodulation,[],[f11646,f9734])).
fof(f11646,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl28)),
  inference(subsumption_resolution,[],[f10739,f260])).
fof(f10739,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl28),
  inference(resolution,[],[f439,f158])).
fof(f11645,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11644,f220])).
fof(f11644,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11643,f9734])).
fof(f11643,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl164)),
  inference(subsumption_resolution,[],[f11187,f260])).
fof(f11187,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(subsumption_resolution,[],[f11071,f220])).
fof(f11071,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(superposition,[],[f157,f9892])).
fof(f11642,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11641,f220])).
fof(f11641,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11640,f9734])).
fof(f11640,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl164)),
  inference(subsumption_resolution,[],[f11188,f260])).
fof(f11188,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(subsumption_resolution,[],[f11072,f220])).
fof(f11072,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl164),
  inference(superposition,[],[f158,f9892])).
fof(f7416,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl34),
  inference(resolution,[],[f446,f159])).
fof(f159,plain,(
  ( ! [X4,X0,X5,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X4),s(fun(X0,bool),X5))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X4),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f100])).
fof(f7417,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl34),
  inference(resolution,[],[f446,f159])).
fof(f7864,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl34),
  inference(subsumption_resolution,[],[f7757,f220])).
fof(f7757,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl34),
  inference(superposition,[],[f159,f7431])).
fof(f7866,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl34),
  inference(subsumption_resolution,[],[f7758,f220])).
fof(f7758,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl34),
  inference(superposition,[],[f159,f7431])).
fof(f1663,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl60),
  inference(cnf_transformation,[],[f1663_D])).
fof(f1663_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f11632,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11631,f219])).
fof(f11631,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11630,f9892])).
fof(f11630,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f11629,f220])).
fof(f11629,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160)),
  inference(forward_demodulation,[],[f9059,f9734])).
fof(f9059,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f155])).
fof(f11628,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11627,f220])).
fof(f11627,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11626,f9892])).
fof(f11626,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f11625,f220])).
fof(f11625,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160)),
  inference(forward_demodulation,[],[f9060,f9734])).
fof(f9060,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f156])).
fof(f11624,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11623,f220])).
fof(f11623,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11622,f9892])).
fof(f11622,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f11621,f220])).
fof(f11621,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160)),
  inference(forward_demodulation,[],[f9061,f9734])).
fof(f9061,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f157])).
fof(f11620,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11619,f220])).
fof(f11619,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11618,f9892])).
fof(f11618,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f11617,f220])).
fof(f11617,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl160)),
  inference(forward_demodulation,[],[f9062,f9734])).
fof(f9062,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f158])).
fof(f9751,plain,(
  ( ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X6))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl30),
  inference(resolution,[],[f441,f159])).
fof(f9752,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X7),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X7),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl30),
  inference(resolution,[],[f441,f159])).
fof(f10057,plain,(
  ( ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X6))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl160),
  inference(subsumption_resolution,[],[f9941,f220])).
fof(f9941,plain,(
  ( ! [X6] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X6))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl160),
  inference(superposition,[],[f159,f9734])).
fof(f10059,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X7),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X7),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl160),
  inference(subsumption_resolution,[],[f9942,f220])).
fof(f9942,plain,(
  ( ! [X7] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X7),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X7),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl160),
  inference(superposition,[],[f159,f9734])).
fof(f11612,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl172),
  inference(cnf_transformation,[],[f11612_D])).
fof(f11612_D,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl172),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl172])])).
fof(f10792,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl40),
  inference(resolution,[],[f451,f159])).
fof(f451,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl40),
  inference(cnf_transformation,[],[f451_D])).
fof(f451_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f11611,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl170),
  inference(cnf_transformation,[],[f11611_D])).
fof(f11611_D,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl170),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl170])])).
fof(f10793,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl40),
  inference(resolution,[],[f451,f159])).
fof(f11407,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl40),
  inference(subsumption_resolution,[],[f11299,f220])).
fof(f11299,plain,(
  ( ! [X2] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(sK9,bool),X2))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl40),
  inference(superposition,[],[f159,f10807])).
fof(f10807,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl40),
  inference(resolution,[],[f451,f190])).
fof(f11409,plain,(
  ( ! [X3] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl40),
  inference(subsumption_resolution,[],[f11300,f220])).
fof(f11300,plain,(
  ( ! [X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X3),s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl40),
  inference(superposition,[],[f159,f10807])).
fof(f11610,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11609,f220])).
fof(f11609,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11561,f9734])).
fof(f11561,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl164)),
  inference(subsumption_resolution,[],[f11186,f260])).
fof(f11186,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl164),
  inference(subsumption_resolution,[],[f11070,f220])).
fof(f11070,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl164),
  inference(superposition,[],[f156,f9892])).
fof(f11608,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11607,f219])).
fof(f11607,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11565,f9734])).
fof(f11565,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl164)),
  inference(subsumption_resolution,[],[f11185,f260])).
fof(f11185,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl164),
  inference(subsumption_resolution,[],[f11069,f220])).
fof(f11069,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl164),
  inference(superposition,[],[f155,f9892])).
fof(f11606,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl160)),
  inference(subsumption_resolution,[],[f11605,f220])).
fof(f11605,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl160)),
  inference(forward_demodulation,[],[f11569,f9734])).
fof(f11569,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28)),
  inference(subsumption_resolution,[],[f10737,f260])).
fof(f10737,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl28),
  inference(resolution,[],[f439,f156])).
fof(f11604,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl160)),
  inference(subsumption_resolution,[],[f11603,f219])).
fof(f11603,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl160)),
  inference(forward_demodulation,[],[f11573,f9734])).
fof(f11573,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28)),
  inference(subsumption_resolution,[],[f10736,f260])).
fof(f10736,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl28),
  inference(resolution,[],[f439,f155])).
fof(f11602,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11601,f220])).
fof(f11601,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11577,f9892])).
fof(f11577,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f10042,f260])).
fof(f10042,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl160),
  inference(subsumption_resolution,[],[f9928,f220])).
fof(f9928,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl160),
  inference(superposition,[],[f156,f9734])).
fof(f11600,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11599,f219])).
fof(f11599,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11581,f9892])).
fof(f11581,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl160)),
  inference(subsumption_resolution,[],[f10041,f260])).
fof(f10041,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl160),
  inference(subsumption_resolution,[],[f9927,f220])).
fof(f9927,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl160),
  inference(superposition,[],[f155,f9734])).
fof(f11598,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl164)),
  inference(subsumption_resolution,[],[f11597,f220])).
fof(f11597,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl164)),
  inference(forward_demodulation,[],[f11585,f9892])).
fof(f11585,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f9739,f260])).
fof(f9739,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl30),
  inference(resolution,[],[f441,f156])).
fof(f11596,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl164)),
  inference(subsumption_resolution,[],[f11595,f219])).
fof(f11595,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl164)),
  inference(forward_demodulation,[],[f11589,f9892])).
fof(f11589,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f9738,f260])).
fof(f9738,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl30),
  inference(resolution,[],[f441,f155])).
fof(f11594,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) ) | $spl168),
  inference(cnf_transformation,[],[f11594_D])).
fof(f11594_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) ) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f11478,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) ) | $spl145),
  inference(resolution,[],[f9170,f172])).
fof(f11480,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)))))) | s(bool,X6) != s(bool,X8)) ) | $spl145),
  inference(resolution,[],[f9170,f174])).
fof(f11481,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) ) | $spl145),
  inference(resolution,[],[f9170,f176])).
fof(f11483,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) ) | $spl145),
  inference(resolution,[],[f9170,f177])).
fof(f11485,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)))))) | s(bool,X21) != s(bool,X23)) ) | $spl145),
  inference(resolution,[],[f9170,f178])).
fof(f11499,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) ) | $spl145),
  inference(resolution,[],[f9170,f255])).
fof(f1655,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl53),
  inference(cnf_transformation,[],[f1655_D])).
fof(f1655_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f11552,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl4 | $spl59)),
  inference(duplicate_literal_removal,[],[f9609])).
fof(f9609,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,X2))) ) | ($spl4 | $spl59)),
  inference(superposition,[],[f9339,f182])).
fof(f182,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f76])).
fof(f76,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X13,X14] : ((p(s(bool,X13)) => p(s(bool,X14))) => ((p(s(bool,X14)) => p(s(bool,X13))) => s(bool,X13) = s(bool,X14)))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.IMP_ANTISYM_AX')).
fof(f9339,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl4 | $spl59)),
  inference(subsumption_resolution,[],[f9294,f892])).
fof(f9294,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl59),
  inference(resolution,[],[f1662,f196])).
fof(f11553,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl4 | $spl59)),
  inference(duplicate_literal_removal,[],[f9610])).
fof(f9610,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,X3)) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl4 | $spl59)),
  inference(superposition,[],[f9339,f182])).
fof(f9619,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl154),
  inference(cnf_transformation,[],[f9619_D])).
fof(f9619_D,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f9580,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl4 | $spl59)),
  inference(resolution,[],[f9339,f172])).
fof(f9582,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X8) != s(bool,X10)) ) | ($spl4 | $spl59)),
  inference(resolution,[],[f9339,f174])).
fof(f9583,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl4 | $spl59)),
  inference(resolution,[],[f9339,f176])).
fof(f9585,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl4 | $spl59)),
  inference(resolution,[],[f9339,f177])).
fof(f9587,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X23) != s(bool,X25)) ) | ($spl4 | $spl59)),
  inference(resolution,[],[f9339,f178])).
fof(f9601,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl4 | $spl59)),
  inference(resolution,[],[f9339,f255])).
fof(f11592,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl54 | $spl164)),
  inference(subsumption_resolution,[],[f11591,f219])).
fof(f11591,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl54 | $spl164)),
  inference(forward_demodulation,[],[f11590,f9892])).
fof(f11590,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl54)),
  inference(subsumption_resolution,[],[f11589,f1657])).
fof(f1657,plain,(
  sP0(sK9,sK10) | $spl54),
  inference(cnf_transformation,[],[f1657_D])).
fof(f1657_D,plain,(
  sP0(sK9,sK10) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f11588,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl54 | $spl164)),
  inference(subsumption_resolution,[],[f11587,f220])).
fof(f11587,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl54 | $spl164)),
  inference(forward_demodulation,[],[f11586,f9892])).
fof(f11586,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl30 | $spl54)),
  inference(subsumption_resolution,[],[f11585,f1657])).
fof(f11584,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11583,f219])).
fof(f11583,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11582,f9892])).
fof(f11582,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160)),
  inference(subsumption_resolution,[],[f11581,f1657])).
fof(f11580,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11579,f220])).
fof(f11579,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11578,f9892])).
fof(f11578,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160)),
  inference(subsumption_resolution,[],[f11577,f1657])).
fof(f11576,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl54 | $spl160)),
  inference(subsumption_resolution,[],[f11575,f219])).
fof(f11575,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl54 | $spl160)),
  inference(forward_demodulation,[],[f11574,f9734])).
fof(f11574,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl54)),
  inference(subsumption_resolution,[],[f11573,f1657])).
fof(f11572,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl54 | $spl160)),
  inference(subsumption_resolution,[],[f11571,f220])).
fof(f11571,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl54 | $spl160)),
  inference(forward_demodulation,[],[f11570,f9734])).
fof(f11570,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl28 | $spl54)),
  inference(subsumption_resolution,[],[f11569,f1657])).
fof(f11568,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11567,f219])).
fof(f11567,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11566,f9734])).
fof(f11566,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl164)),
  inference(subsumption_resolution,[],[f11565,f1657])).
fof(f11564,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11563,f220])).
fof(f11563,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11562,f9734])).
fof(f11562,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54 | $spl164)),
  inference(subsumption_resolution,[],[f11561,f1657])).
fof(f11546,plain,(
  $false | ($spl3 | $spl54 | $spl58 | $spl143 | $spl145 | $spl160 | $spl164)),
  inference(subsumption_resolution,[],[f11545,f253])).
fof(f11545,plain,(
  ~p(s(bool,'T_0')) | ($spl3 | $spl54 | $spl58 | $spl143 | $spl145 | $spl160 | $spl164)),
  inference(forward_demodulation,[],[f11544,f9892])).
fof(f11544,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54 | $spl58 | $spl143 | $spl145 | $spl160)),
  inference(subsumption_resolution,[],[f11543,f223])).
fof(f11543,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54 | $spl58 | $spl143 | $spl145 | $spl160)),
  inference(forward_demodulation,[],[f11542,f9734])).
fof(f11542,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54 | $spl58 | $spl143 | $spl145)),
  inference(subsumption_resolution,[],[f11541,f1661])).
fof(f11541,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54 | $spl143 | $spl145)),
  inference(subsumption_resolution,[],[f11540,f1657])).
fof(f11540,plain,(
  ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl143 | $spl145)),
  inference(subsumption_resolution,[],[f11539,f9168])).
fof(f11531,plain,(
  $false | ($spl60 | $spl143 | $spl145)),
  inference(subsumption_resolution,[],[f11475,f9168])).
fof(f11475,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ($spl60 | $spl145)),
  inference(resolution,[],[f9170,f1663])).
fof(f10840,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl166),
  inference(cnf_transformation,[],[f10840_D])).
fof(f10840_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f10814,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl40),
  inference(resolution,[],[f451,f255])).
fof(f10802,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl40),
  inference(resolution,[],[f451,f178])).
fof(f10799,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl40),
  inference(resolution,[],[f451,f177])).
fof(f10798,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl40),
  inference(resolution,[],[f451,f176])).
fof(f10797,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl40),
  inference(resolution,[],[f451,f174])).
fof(f10795,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl40),
  inference(resolution,[],[f451,f172])).
fof(f10513,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl135),
  inference(subsumption_resolution,[],[f10189,f186])).
fof(f10189,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl135),
  inference(superposition,[],[f8923,f179])).
fof(f179,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f77])).
fof(f10512,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X1))) ) | $spl135),
  inference(subsumption_resolution,[],[f10190,f186])).
fof(f10190,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X1))) ) | $spl135),
  inference(superposition,[],[f8923,f179])).
fof(f10508,plain,(
  $false | ($spl132 | $spl136)),
  inference(subsumption_resolution,[],[f10430,f10397])).
fof(f10397,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl132 | $spl136)),
  inference(subsumption_resolution,[],[f10396,f192])).
fof(f10396,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl132 | $spl136)),
  inference(forward_demodulation,[],[f8635,f8933])).
fof(f10430,plain,(
  p(s(bool,'F_0')) | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f250])).
fof(f250,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f145])).
fof(f145,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f23])).
fof(f23,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.NOT_CLAUSES|split|2')).
fof(f10507,plain,(
  $false | ($spl132 | $spl136)),
  inference(subsumption_resolution,[],[f10429,f10397])).
fof(f10429,plain,(
  p(s(bool,'F_0')) | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f248])).
fof(f248,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f22])).
fof(f22,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.NOT_CLAUSES|split|1')).
fof(f10506,plain,(
  $false | ($spl132 | $spl136)),
  inference(subsumption_resolution,[],[f10426,f10397])).
fof(f10426,plain,(
  ( ! [X48] : (p(s(bool,X48))) ) | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f228])).
fof(f228,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : ((p(s(bool,'T_0')) => p(s(bool,X9))) <=> p(s(bool,X9)))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.IMP_CLAUSES|split|0')).
fof(f10505,plain,(
  $false | ($spl132 | $spl136)),
  inference(subsumption_resolution,[],[f10425,f10397])).
fof(f10425,plain,(
  ( ! [X47] : (p(s(bool,X47))) ) | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f225])).
fof(f225,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X9] : ((p(s(bool,X9)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.bool.IMP_CLAUSES|split|1')).
fof(f10504,plain,(
  $false | ($spl132 | $spl136)),
  inference(subsumption_resolution,[],[f10423,f10397])).
fof(f10423,plain,(
  p(s(bool,'F_0')) | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f222])).
fof(f222,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f129])).
fof(f10503,plain,(
  $false | ($spl132 | $spl136)),
  inference(subsumption_resolution,[],[f10421,f10397])).
fof(f10421,plain,(
  ( ! [X44] : (p(s(bool,X44))) ) | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f219])).
fof(f10431,plain,(
  $false | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f252])).
fof(f252,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT',aHL_TRUTH)).
fof(f10432,plain,(
  $false | ($spl132 | $spl136)),
  inference(resolution,[],[f10397,f253])).
fof(f10466,plain,(
  $false | ($spl56 | $spl132 | $spl136)),
  inference(resolution,[],[f10397,f1660])).
fof(f10467,plain,(
  $false | ($spl58 | $spl132 | $spl136)),
  inference(resolution,[],[f10397,f1661])).
fof(f8211,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl128),
  inference(cnf_transformation,[],[f8211_D])).
fof(f8211_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f10145,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f255])).
fof(f10131,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)))))) | s(bool,X21) != s(bool,X23)) ) | $spl77),
  inference(resolution,[],[f2680,f178])).
fof(f10129,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f177])).
fof(f10127,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f176])).
fof(f10126,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)))))) | s(bool,X6) != s(bool,X8)) ) | $spl77),
  inference(resolution,[],[f2680,f174])).
fof(f10124,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f172])).
fof(f9892,plain,(
  s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl164),
  inference(cnf_transformation,[],[f9892_D])).
fof(f9892_D,plain,(
  s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f630,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13)) ) | $spl42),
  inference(cnf_transformation,[],[f630_D])).
fof(f630_D,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f9890,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl88),
  inference(subsumption_resolution,[],[f9876,f185])).
fof(f185,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f110])).
fof(f9876,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl88),
  inference(superposition,[],[f3581,f193])).
fof(f439,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl28),
  inference(cnf_transformation,[],[f439_D])).
fof(f439_D,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f9889,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl88),
  inference(subsumption_resolution,[],[f9872,f186])).
fof(f9872,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X9)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl88),
  inference(superposition,[],[f3581,f188])).
fof(f9888,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl88),
  inference(subsumption_resolution,[],[f9871,f186])).
fof(f9871,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl88),
  inference(superposition,[],[f3581,f186])).
fof(f9887,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl88),
  inference(subsumption_resolution,[],[f9868,f255])).
fof(f9868,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X2))) ) | $spl88),
  inference(superposition,[],[f3581,f179])).
fof(f9886,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl88),
  inference(subsumption_resolution,[],[f9867,f255])).
fof(f9867,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl88),
  inference(superposition,[],[f3581,f179])).
fof(f9883,plain,(
  s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl163),
  inference(cnf_transformation,[],[f9883_D])).
fof(f9883_D,plain,(
  s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) <=> ~$spl163),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl163])])).
fof(f9884,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X13))) ) | $spl88),
  inference(subsumption_resolution,[],[f9861,f189])).
fof(f9861,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X13)) ) | $spl88),
  inference(superposition,[],[f3581,f193])).
fof(f9878,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9))) ) | $spl88),
  inference(duplicate_literal_removal,[],[f9857])).
fof(f9857,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9)) | p(s(bool,X9))) ) | $spl88),
  inference(superposition,[],[f3581,f188])).
fof(f9879,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X8))) ) | $spl88),
  inference(duplicate_literal_removal,[],[f9856])).
fof(f9856,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X8)) | p(s(bool,X8))) ) | $spl88),
  inference(superposition,[],[f3581,f186])).
fof(f9734,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl160),
  inference(cnf_transformation,[],[f9734_D])).
fof(f9734_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f9732,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl50),
  inference(subsumption_resolution,[],[f9716,f185])).
fof(f9716,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl50),
  inference(superposition,[],[f1629,f193])).
fof(f441,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl30),
  inference(cnf_transformation,[],[f441_D])).
fof(f441_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f9731,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl50),
  inference(subsumption_resolution,[],[f9712,f186])).
fof(f9712,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X9)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl50),
  inference(superposition,[],[f1629,f188])).
fof(f9730,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl50),
  inference(subsumption_resolution,[],[f9711,f186])).
fof(f9711,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl50),
  inference(superposition,[],[f1629,f186])).
fof(f9729,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl50),
  inference(subsumption_resolution,[],[f9708,f255])).
fof(f9708,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X2))) ) | $spl50),
  inference(superposition,[],[f1629,f179])).
fof(f9728,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl50),
  inference(subsumption_resolution,[],[f9707,f255])).
fof(f9707,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl50),
  inference(superposition,[],[f1629,f179])).
fof(f9723,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl159),
  inference(cnf_transformation,[],[f9723_D])).
fof(f9723_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) <=> ~$spl159),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl159])])).
fof(f9724,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X13))) ) | $spl50),
  inference(subsumption_resolution,[],[f9699,f189])).
fof(f9699,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X13)) ) | $spl50),
  inference(superposition,[],[f1629,f193])).
fof(f9718,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9))) ) | $spl50),
  inference(duplicate_literal_removal,[],[f9695])).
fof(f9695,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9)) | p(s(bool,X9))) ) | $spl50),
  inference(superposition,[],[f1629,f188])).
fof(f9719,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X8))) ) | $spl50),
  inference(duplicate_literal_removal,[],[f9694])).
fof(f9694,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X8)) | p(s(bool,X8))) ) | $spl50),
  inference(superposition,[],[f1629,f186])).
fof(f9655,plain,(
  ( ! [X33,X32] : (p(s(bool,X32)) | p(s(bool,X33)) | s(bool,X33) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f255])).
fof(f9643,plain,(
  ( ! [X28,X29,X27] : (p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X28) != s(bool,X29)) ) | $spl58),
  inference(resolution,[],[f1661,f178])).
fof(f9640,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,X18)) | ~p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X20) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f177])).
fof(f9639,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X15) != s(bool,X17) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f176])).
fof(f9638,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | s(bool,X12) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X13) != s(bool,X14)) ) | $spl58),
  inference(resolution,[],[f1661,f174])).
fof(f9636,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | p(s(bool,X7)) | s(bool,X6) != s(bool,X8) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl58),
  inference(resolution,[],[f1661,f172])).
fof(f9340,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl152),
  inference(cnf_transformation,[],[f9340_D])).
fof(f9340_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f9317,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl59),
  inference(resolution,[],[f1662,f255])).
fof(f9303,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl59),
  inference(resolution,[],[f1662,f178])).
fof(f9301,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl59),
  inference(resolution,[],[f1662,f177])).
fof(f9299,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl59),
  inference(resolution,[],[f1662,f176])).
fof(f9298,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl59),
  inference(resolution,[],[f1662,f174])).
fof(f9296,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl59),
  inference(resolution,[],[f1662,f172])).
fof(f9288,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl150),
  inference(cnf_transformation,[],[f9288_D])).
fof(f9288_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f9262,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl63),
  inference(resolution,[],[f1665,f255])).
fof(f9248,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl63),
  inference(resolution,[],[f1665,f178])).
fof(f9246,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl63),
  inference(resolution,[],[f1665,f177])).
fof(f9244,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl63),
  inference(resolution,[],[f1665,f176])).
fof(f9243,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl63),
  inference(resolution,[],[f1665,f174])).
fof(f9241,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl63),
  inference(resolution,[],[f1665,f172])).
fof(f9170,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) | $spl145),
  inference(cnf_transformation,[],[f9170_D])).
fof(f9170_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f9287,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl56 | $spl63)),
  inference(subsumption_resolution,[],[f9238,f1660])).
fof(f9238,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl63),
  inference(resolution,[],[f1665,f167])).
fof(f9286,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl56 | $spl63)),
  inference(subsumption_resolution,[],[f9237,f1660])).
fof(f9237,plain,(
  s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl63),
  inference(resolution,[],[f1665,f166])).
fof(f9231,plain,(
  ( ! [X2,X5,X3] : (s('type.num.num',X2) = s('type.num.num',X3) | s(bool,X5) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3))))) | p(s(bool,X5))) ) | $spl148),
  inference(cnf_transformation,[],[f9231_D])).
fof(f9231_D,plain,(
  ( ! [X2,X5,X3] : (s('type.num.num',X2) = s('type.num.num',X3) | s(bool,X5) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3))))) | p(s(bool,X5))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f9198,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.num.num',X66) = s('type.num.num',X67) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X69) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X66))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X67)))))) ) | $spl60),
  inference(resolution,[],[f1663,f255])).
fof(f9186,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.num.num',X37) = s('type.num.num',X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X37))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X38))))) | s(bool,X40) != s(bool,X41)) ) | $spl60),
  inference(resolution,[],[f1663,f178])).
fof(f9183,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.num.num',X22) = s('type.num.num',X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X22))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X23)))))) ) | $spl60),
  inference(resolution,[],[f1663,f177])).
fof(f9182,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.num.num',X17) = s('type.num.num',X18) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X17))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X18)))))) ) | $spl60),
  inference(resolution,[],[f1663,f176])).
fof(f9181,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.num.num',X12) = s('type.num.num',X13) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X12))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X13))))) | s(bool,X15) != s(bool,X16)) ) | $spl60),
  inference(resolution,[],[f1663,f174])).
fof(f9179,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.num.num',X2) = s('type.num.num',X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) ) | $spl60),
  inference(resolution,[],[f1663,f172])).
fof(f9171,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl146),
  inference(cnf_transformation,[],[f9171_D])).
fof(f9171_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f9146,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl56),
  inference(resolution,[],[f1660,f255])).
fof(f9134,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X22) != s(bool,X23)) ) | $spl56),
  inference(resolution,[],[f1660,f178])).
fof(f9131,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl56),
  inference(resolution,[],[f1660,f177])).
fof(f9130,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl56),
  inference(resolution,[],[f1660,f176])).
fof(f9129,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X7) != s(bool,X8)) ) | $spl56),
  inference(resolution,[],[f1660,f174])).
fof(f9127,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl56),
  inference(resolution,[],[f1660,f172])).
fof(f9124,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl56),
  inference(resolution,[],[f1660,f167])).
fof(f9123,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK6(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK6(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl56),
  inference(resolution,[],[f1660,f166])).
fof(f1992,plain,(
  s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl66),
  inference(cnf_transformation,[],[f1992_D])).
fof(f1992_D,plain,(
  s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f9122,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl64),
  inference(subsumption_resolution,[],[f9107,f185])).
fof(f9107,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl64),
  inference(superposition,[],[f1723,f193])).
fof(f9121,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f9103,f186])).
fof(f9103,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X9)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(superposition,[],[f1723,f188])).
fof(f9120,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f9102,f186])).
fof(f9102,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(superposition,[],[f1723,f186])).
fof(f9119,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f9099,f255])).
fof(f9099,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | p(s(bool,X2))) ) | $spl64),
  inference(superposition,[],[f1723,f179])).
fof(f9118,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f9098,f255])).
fof(f9098,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(superposition,[],[f1723,f179])).
fof(f2001,plain,(
  s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl71),
  inference(cnf_transformation,[],[f2001_D])).
fof(f2001_D,plain,(
  s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f9113,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X13))) ) | $spl64),
  inference(subsumption_resolution,[],[f9093,f189])).
fof(f9093,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X13)) ) | $spl64),
  inference(superposition,[],[f1723,f193])).
fof(f9109,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X9))) ) | $spl64),
  inference(duplicate_literal_removal,[],[f9089])).
fof(f9089,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X9)) | p(s(bool,X9))) ) | $spl64),
  inference(superposition,[],[f1723,f188])).
fof(f9110,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X8))) ) | $spl64),
  inference(duplicate_literal_removal,[],[f9088])).
fof(f9088,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X8)) | p(s(bool,X8))) ) | $spl64),
  inference(superposition,[],[f1723,f186])).
fof(f440,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl29),
  inference(cnf_transformation,[],[f440_D])).
fof(f440_D,plain,(
  ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f442,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl31),
  inference(cnf_transformation,[],[f442_D])).
fof(f442_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f9068,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f9062,f1657])).
fof(f9067,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f9061,f1657])).
fof(f9066,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f9060,f1657])).
fof(f9065,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f9059,f1657])).
fof(f1660,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | $spl56),
  inference(cnf_transformation,[],[f1660_D])).
fof(f1660_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f9064,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f9058,f1657])).
fof(f9058,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f154])).
fof(f154,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~sP0(X0,X1) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK5(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f9063,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f9057,f1657])).
fof(f9057,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f153])).
fof(f153,plain,(
  ( ! [X0,X1] : (sP1(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),sK6(X1,X0)),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~sP0(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK5(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f97])).
fof(f9056,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl140),
  inference(cnf_transformation,[],[f9056_D])).
fof(f9056_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) ) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f8936,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl97),
  inference(resolution,[],[f4836,f172])).
fof(f8938,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)))))) | s(bool,X6) != s(bool,X8)) ) | $spl97),
  inference(resolution,[],[f4836,f174])).
fof(f8939,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl97),
  inference(resolution,[],[f4836,f176])).
fof(f8941,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl97),
  inference(resolution,[],[f4836,f177])).
fof(f8943,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)))))) | s(bool,X21) != s(bool,X23)) ) | $spl97),
  inference(resolution,[],[f4836,f178])).
fof(f8957,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl97),
  inference(resolution,[],[f4836,f255])).
fof(f3750,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X9))) ) | $spl64),
  inference(duplicate_literal_removal,[],[f3734])).
fof(f3734,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X9)) | p(s(bool,X9))) ) | $spl64),
  inference(superposition,[],[f1723,f188])).
fof(f3751,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X8))) ) | $spl64),
  inference(duplicate_literal_removal,[],[f3733])).
fof(f3733,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X8)) | p(s(bool,X8))) ) | $spl64),
  inference(superposition,[],[f1723,f186])).
fof(f3754,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X13))) ) | $spl64),
  inference(subsumption_resolution,[],[f3738,f189])).
fof(f3738,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,X13)) ) | $spl64),
  inference(superposition,[],[f1723,f193])).
fof(f3757,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f3743,f186])).
fof(f3743,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(superposition,[],[f1723,f186])).
fof(f3758,plain,(
  ( ! [X9] : (p(s(bool,X9)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f3744,f186])).
fof(f3744,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,X9)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(superposition,[],[f1723,f188])).
fof(f2676,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | $spl75),
  inference(cnf_transformation,[],[f2676_D])).
fof(f2676_D,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) <=> ~$spl75),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl75])])).
fof(f8456,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f8455,f311])).
fof(f8455,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(subsumption_resolution,[],[f8354,f220])).
fof(f8354,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(superposition,[],[f166,f4003])).
fof(f4003,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | $spl6),
  inference(resolution,[],[f308,f190])).
fof(f2680,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | $spl77),
  inference(cnf_transformation,[],[f2680_D])).
fof(f2680_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f8460,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f8459,f311])).
fof(f8459,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(subsumption_resolution,[],[f8355,f220])).
fof(f8355,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(superposition,[],[f167,f4003])).
fof(f309,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl8),
  inference(cnf_transformation,[],[f309_D])).
fof(f309_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f7076,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(resolution,[],[f3631,f169])).
fof(f3631,plain,(
  p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl4),
  inference(resolution,[],[f892,f205])).
fof(f205,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ~sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f118])).
fof(f2675,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | $spl74),
  inference(cnf_transformation,[],[f2675_D])).
fof(f2675_D,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f9052,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ($spl4 | $spl77)),
  inference(subsumption_resolution,[],[f9051,f219])).
fof(f9051,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ($spl4 | $spl77)),
  inference(forward_demodulation,[],[f8138,f7090])).
fof(f7090,plain,(
  s(bool,'T_0') = s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl4),
  inference(resolution,[],[f3631,f190])).
fof(f8207,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4716,f308])).
fof(f4716,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl11),
  inference(resolution,[],[f311,f166])).
fof(f8209,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f3987,f311])).
fof(f3987,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(resolution,[],[f308,f166])).
fof(f7301,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(subsumption_resolution,[],[f7202,f219])).
fof(f7202,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(superposition,[],[f169,f7090])).
fof(f7418,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl34),
  inference(resolution,[],[f446,f149])).
fof(f7868,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl34),
  inference(subsumption_resolution,[],[f7759,f219])).
fof(f7759,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK15(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP1(sK9,sK10) | $spl34),
  inference(superposition,[],[f149,f7431])).
fof(f8994,plain,(
  $false | ($spl4 | $spl70 | $spl97 | $spl99)),
  inference(subsumption_resolution,[],[f8993,f253])).
fof(f8993,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl70 | $spl97 | $spl99)),
  inference(forward_demodulation,[],[f8992,f7090])).
fof(f8992,plain,(
  ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl70 | $spl97 | $spl99)),
  inference(subsumption_resolution,[],[f8991,f217])).
fof(f8991,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl70 | $spl97 | $spl99)),
  inference(forward_demodulation,[],[f8990,f2002])).
fof(f8989,plain,(
  $false | ($spl8 | $spl97 | $spl99)),
  inference(subsumption_resolution,[],[f8934,f4839])).
fof(f8934,plain,(
  s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) | ($spl8 | $spl97)),
  inference(resolution,[],[f4836,f309])).
fof(f8933,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl136),
  inference(cnf_transformation,[],[f8933_D])).
fof(f8933_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) <=> ~$spl136),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl136])])).
fof(f8931,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl108),
  inference(subsumption_resolution,[],[f8916,f185])).
fof(f8916,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl108),
  inference(superposition,[],[f5791,f193])).
fof(f8928,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl108),
  inference(subsumption_resolution,[],[f8908,f255])).
fof(f8908,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X2))) ) | $spl108),
  inference(superposition,[],[f5791,f179])).
fof(f8927,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl108),
  inference(subsumption_resolution,[],[f8907,f255])).
fof(f8907,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl108),
  inference(superposition,[],[f5791,f179])).
fof(f8879,plain,(
  $false | ($spl39 | $spl70 | $spl108)),
  inference(trivial_inequality_removal,[],[f8868])).
fof(f8868,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl39 | $spl70 | $spl108)),
  inference(superposition,[],[f8841,f2002])).
fof(f8841,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5)) ) | ($spl39 | $spl108)),
  inference(subsumption_resolution,[],[f8840,f187])).
fof(f187,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f111])).
fof(f8840,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5))) ) | ($spl39 | $spl108)),
  inference(forward_demodulation,[],[f5791,f8606])).
fof(f8606,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl39),
  inference(resolution,[],[f449,f186])).
fof(f449,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl39),
  inference(cnf_transformation,[],[f449_D])).
fof(f449_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f8878,plain,(
  $false | ($spl33 | $spl39 | $spl108)),
  inference(trivial_inequality_removal,[],[f8871])).
fof(f8871,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl33 | $spl39 | $spl108)),
  inference(superposition,[],[f8841,f7378])).
fof(f8877,plain,(
  $false | ($spl39 | $spl108)),
  inference(trivial_inequality_removal,[],[f8873])).
fof(f8873,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl39 | $spl108)),
  inference(superposition,[],[f8841,f8606])).
fof(f8876,plain,(
  $false | ($spl39 | $spl108)),
  inference(equality_resolution,[],[f8841])).
fof(f8635,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl132),
  inference(cnf_transformation,[],[f8635_D])).
fof(f8635_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f8617,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl39),
  inference(resolution,[],[f449,f255])).
fof(f8603,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl39),
  inference(resolution,[],[f449,f178])).
fof(f8601,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl39),
  inference(resolution,[],[f449,f177])).
fof(f8599,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl39),
  inference(resolution,[],[f449,f176])).
fof(f8598,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl39),
  inference(resolution,[],[f449,f174])).
fof(f8596,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl39),
  inference(resolution,[],[f449,f172])).
fof(f8593,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl130),
  inference(cnf_transformation,[],[f8593_D])).
fof(f8593_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f8531,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl41),
  inference(resolution,[],[f452,f172])).
fof(f8533,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl41),
  inference(resolution,[],[f452,f174])).
fof(f8534,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl41),
  inference(resolution,[],[f452,f176])).
fof(f8536,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl41),
  inference(resolution,[],[f452,f177])).
fof(f8538,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl41),
  inference(resolution,[],[f452,f178])).
fof(f8552,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl41),
  inference(resolution,[],[f452,f255])).
fof(f8583,plain,(
  $false | ($spl34 | $spl38 | $spl41 | $spl54)),
  inference(subsumption_resolution,[],[f8582,f253])).
fof(f8582,plain,(
  ~p(s(bool,'T_0')) | ($spl34 | $spl38 | $spl41 | $spl54)),
  inference(forward_demodulation,[],[f8581,f7431])).
fof(f8581,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl38 | $spl41 | $spl54)),
  inference(subsumption_resolution,[],[f8580,f223])).
fof(f8580,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl38 | $spl41 | $spl54)),
  inference(forward_demodulation,[],[f8579,f7484])).
fof(f7484,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl38),
  inference(resolution,[],[f450,f190])).
fof(f8579,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl41 | $spl54)),
  inference(subsumption_resolution,[],[f8530,f1657])).
fof(f8530,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | $spl41),
  inference(resolution,[],[f452,f159])).
fof(f8578,plain,(
  $false | ($spl4 | $spl34 | $spl38 | $spl41)),
  inference(subsumption_resolution,[],[f8577,f253])).
fof(f8577,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl34 | $spl38 | $spl41)),
  inference(forward_demodulation,[],[f8576,f7431])).
fof(f8576,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl4 | $spl38 | $spl41)),
  inference(subsumption_resolution,[],[f8575,f223])).
fof(f8575,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl4 | $spl38 | $spl41)),
  inference(forward_demodulation,[],[f8574,f7484])).
fof(f8574,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl4 | $spl41)),
  inference(subsumption_resolution,[],[f8529,f892])).
fof(f8529,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl41),
  inference(resolution,[],[f452,f197])).
fof(f8322,plain,(
  $false | ($spl33 | $spl70 | $spl110)),
  inference(trivial_inequality_removal,[],[f8313])).
fof(f8313,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl33 | $spl70 | $spl110)),
  inference(superposition,[],[f8286,f2002])).
fof(f8286,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5)) ) | ($spl33 | $spl110)),
  inference(subsumption_resolution,[],[f8285,f187])).
fof(f8285,plain,(
  ( ! [X5] : (s(bool,'F_0') != s(bool,X5) | p(s(bool,X5))) ) | ($spl33 | $spl110)),
  inference(forward_demodulation,[],[f5842,f7378])).
fof(f8321,plain,(
  $false | ($spl33 | $spl110)),
  inference(trivial_inequality_removal,[],[f8316])).
fof(f8316,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl33 | $spl110)),
  inference(superposition,[],[f8286,f7378])).
fof(f8320,plain,(
  $false | ($spl33 | $spl110)),
  inference(equality_resolution,[],[f8286])).
fof(f8139,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f172])).
fof(f8141,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)))))) | s(bool,X6) != s(bool,X8)) ) | $spl77),
  inference(resolution,[],[f2680,f174])).
fof(f8142,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f176])).
fof(f8144,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f177])).
fof(f8146,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)))))) | s(bool,X21) != s(bool,X23)) ) | $spl77),
  inference(resolution,[],[f2680,f178])).
fof(f8160,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) ) | $spl77),
  inference(resolution,[],[f2680,f255])).
fof(f8197,plain,(
  $false | ($spl4 | $spl70 | $spl75 | $spl77)),
  inference(subsumption_resolution,[],[f8196,f253])).
fof(f8196,plain,(
  ~p(s(bool,'T_0')) | ($spl4 | $spl70 | $spl75 | $spl77)),
  inference(forward_demodulation,[],[f8195,f7090])).
fof(f8195,plain,(
  ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl70 | $spl75 | $spl77)),
  inference(subsumption_resolution,[],[f8194,f217])).
fof(f8194,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl70 | $spl75 | $spl77)),
  inference(forward_demodulation,[],[f8193,f2002])).
fof(f8192,plain,(
  $false | ($spl8 | $spl75 | $spl77)),
  inference(subsumption_resolution,[],[f8137,f2676])).
fof(f8137,plain,(
  s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) = s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ($spl8 | $spl77)),
  inference(resolution,[],[f2680,f309])).
fof(f5791,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X5))) ) | $spl108),
  inference(cnf_transformation,[],[f5791_D])).
fof(f5791_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X5))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f7491,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl38),
  inference(resolution,[],[f450,f255])).
fof(f7479,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl38),
  inference(resolution,[],[f450,f178])).
fof(f7476,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl38),
  inference(resolution,[],[f450,f177])).
fof(f7475,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl38),
  inference(resolution,[],[f450,f176])).
fof(f7474,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl38),
  inference(resolution,[],[f450,f174])).
fof(f7472,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl38),
  inference(resolution,[],[f450,f172])).
fof(f5737,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl106),
  inference(cnf_transformation,[],[f5737_D])).
fof(f5737_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f7438,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl34),
  inference(resolution,[],[f446,f255])).
fof(f7426,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl34),
  inference(resolution,[],[f446,f178])).
fof(f7423,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl34),
  inference(resolution,[],[f446,f177])).
fof(f7422,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl34),
  inference(resolution,[],[f446,f176])).
fof(f7421,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl34),
  inference(resolution,[],[f446,f174])).
fof(f7419,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl34),
  inference(resolution,[],[f446,f172])).
fof(f7407,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl126),
  inference(cnf_transformation,[],[f7407_D])).
fof(f7407_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl126),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl126])])).
fof(f7389,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl33),
  inference(resolution,[],[f443,f255])).
fof(f7375,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl33),
  inference(resolution,[],[f443,f178])).
fof(f7373,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl33),
  inference(resolution,[],[f443,f177])).
fof(f7371,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl33),
  inference(resolution,[],[f443,f176])).
fof(f7370,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl33),
  inference(resolution,[],[f443,f174])).
fof(f7368,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl33),
  inference(resolution,[],[f443,f172])).
fof(f4270,plain,(
  ( ! [X2,X5,X3] : (s('type.num.num',X2) = s('type.num.num',X3) | s(bool,X5) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3))))) | p(s(bool,X5))) ) | $spl90),
  inference(cnf_transformation,[],[f4270_D])).
fof(f4270_D,plain,(
  ( ! [X2,X5,X3] : (s('type.num.num',X2) = s('type.num.num',X3) | s(bool,X5) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3))))) | p(s(bool,X5))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f7151,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.num.num',X66) = s('type.num.num',X67) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X69) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X66))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X67)))))) ) | $spl8),
  inference(resolution,[],[f309,f255])).
fof(f7139,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.num.num',X37) = s('type.num.num',X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X37))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X38))))) | s(bool,X40) != s(bool,X41)) ) | $spl8),
  inference(resolution,[],[f309,f178])).
fof(f7136,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.num.num',X22) = s('type.num.num',X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X22))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X23)))))) ) | $spl8),
  inference(resolution,[],[f309,f177])).
fof(f7135,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.num.num',X17) = s('type.num.num',X18) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X17))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X18)))))) ) | $spl8),
  inference(resolution,[],[f309,f176])).
fof(f7134,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.num.num',X12) = s('type.num.num',X13) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X12))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X13))))) | s(bool,X15) != s(bool,X16)) ) | $spl8),
  inference(resolution,[],[f309,f174])).
fof(f7132,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.num.num',X2) = s('type.num.num',X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3)))))) ) | $spl8),
  inference(resolution,[],[f309,f172])).
fof(f931,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl48),
  inference(cnf_transformation,[],[f931_D])).
fof(f931_D,plain,(
  ( ! [X3] : (p(s(bool,X3)) | s(bool,X3) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f7097,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f3631,f255])).
fof(f7085,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X24) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f3631,f178])).
fof(f7082,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f3631,f177])).
fof(f7081,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f3631,f176])).
fof(f7080,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f3631,f174])).
fof(f7078,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f3631,f172])).
fof(f7073,plain,(
  $false | ($spl4 | $spl113)),
  inference(subsumption_resolution,[],[f6995,f892])).
fof(f446,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl34),
  inference(cnf_transformation,[],[f446_D])).
fof(f446_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f7061,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7060,f220])).
fof(f7060,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f4995,f5099])).
fof(f5099,plain,(
  s(bool,'T_0') = s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl2),
  inference(resolution,[],[f4844,f190])).
fof(f4844,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f259,f147])).
fof(f147,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f4995,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4845,f199])).
fof(f199,plain,(
  ( ! [X0,X1] : (sP2(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK15(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),sK14(X1,X0)))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f4845,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f259,f148])).
fof(f148,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f7058,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7057,f220])).
fof(f7057,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f4997,f5099])).
fof(f4997,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4845,f201])).
fof(f7053,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7052,f220])).
fof(f7052,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f4999,f5099])).
fof(f4999,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4845,f203])).
fof(f5000,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl2),
  inference(resolution,[],[f4845,f196])).
fof(f5001,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(resolution,[],[f4845,f197])).
fof(f5002,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X5),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(resolution,[],[f4845,f197])).
fof(f7042,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7041,f220])).
fof(f7041,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f5082,f5018])).
fof(f5018,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl2),
  inference(resolution,[],[f4845,f190])).
fof(f5082,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4844,f199])).
fof(f7039,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7038,f220])).
fof(f7038,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f5084,f5018])).
fof(f5084,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4844,f201])).
fof(f7034,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7033,f220])).
fof(f7033,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f5086,f5018])).
fof(f5086,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4844,f203])).
fof(f5251,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl70),
  inference(subsumption_resolution,[],[f5154,f235])).
fof(f5154,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl70),
  inference(superposition,[],[f169,f2002])).
fof(f7032,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7031,f219])).
fof(f7031,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f7030,f5099])).
fof(f7030,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7029,f220])).
fof(f7029,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f5364,f5018])).
fof(f5364,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f5312,f198])).
fof(f198,plain,(
  ( ! [X0,X1] : (sP2(X0,X1) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK15(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),sK14(X1,X0)),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f5312,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4846,f199])).
fof(f4846,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl2),
  inference(resolution,[],[f259,f149])).
fof(f7028,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7027,f219])).
fof(f7027,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f7026,f5099])).
fof(f7026,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7025,f220])).
fof(f7025,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f5368,f5018])).
fof(f5368,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f5313,f200])).
fof(f5313,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4846,f201])).
fof(f7024,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7023,f220])).
fof(f7023,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f7022,f5099])).
fof(f7022,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7021,f220])).
fof(f7021,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f5372,f5018])).
fof(f5372,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f5314,f202])).
fof(f5314,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4846,f203])).
fof(f7020,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X2))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl124),
  inference(cnf_transformation,[],[f7020_D])).
fof(f7020_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X2))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f5315,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X2),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X2))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(resolution,[],[f4846,f196])).
fof(f7019,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X3))),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl122),
  inference(cnf_transformation,[],[f7019_D])).
fof(f7019_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X3))),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f5316,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X3),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X3))),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(resolution,[],[f4846,f197])).
fof(f7018,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X6),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X5))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl120),
  inference(cnf_transformation,[],[f7018_D])).
fof(f7018_D,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X6),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X5))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f5317,plain,(
  ( ! [X6,X5] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X6),s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X5))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(resolution,[],[f4846,f197])).
fof(f4836,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | $spl97),
  inference(cnf_transformation,[],[f4836_D])).
fof(f4836_D,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f5642,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl2 | $spl70)),
  inference(subsumption_resolution,[],[f5641,f217])).
fof(f5641,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl2 | $spl70)),
  inference(forward_demodulation,[],[f5640,f2002])).
fof(f5640,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl2),
  inference(subsumption_resolution,[],[f5593,f170])).
fof(f170,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f5593,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl2),
  inference(resolution,[],[f4849,f168])).
fof(f168,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),sK13(X1,X0)),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f4849,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X2),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X2),s('type.num.num',sK3(X2,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X2),s('type.num.num',sK4(X2,sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl2),
  inference(resolution,[],[f259,f152])).
fof(f152,plain,(
  ( ! [X6,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK3(X6,X1,X0)))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X6),s('type.num.num',sK4(X6,X1,X0))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f447,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl36),
  inference(cnf_transformation,[],[f447_D])).
fof(f447_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f5797,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl32),
  inference(resolution,[],[f444,f196])).
fof(f7007,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl116),
  inference(cnf_transformation,[],[f7007_D])).
fof(f7007_D,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl116),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl116])])).
fof(f5798,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl32),
  inference(resolution,[],[f444,f197])).
fof(f7006,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl114),
  inference(cnf_transformation,[],[f7006_D])).
fof(f7006_D,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f5799,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl32),
  inference(resolution,[],[f444,f197])).
fof(f7016,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7015,f220])).
fof(f7015,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f6000,f5099])).
fof(f6000,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f5887,f220])).
fof(f5887,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(superposition,[],[f199,f5018])).
fof(f7013,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7012,f220])).
fof(f7012,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f6002,f5099])).
fof(f6002,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f5889,f220])).
fof(f5889,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(superposition,[],[f201,f5018])).
fof(f7011,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7010,f220])).
fof(f7010,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f6004,f5099])).
fof(f6004,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f5891,f220])).
fof(f5891,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(superposition,[],[f203,f5018])).
fof(f6007,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl2),
  inference(subsumption_resolution,[],[f5892,f219])).
fof(f5892,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl2),
  inference(superposition,[],[f196,f5018])).
fof(f6008,plain,(
  ( ! [X4] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(subsumption_resolution,[],[f5893,f220])).
fof(f5893,plain,(
  ( ! [X4] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X4))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X4),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(superposition,[],[f197,f5018])).
fof(f6009,plain,(
  ( ! [X5] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X5),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(subsumption_resolution,[],[f5894,f220])).
fof(f5894,plain,(
  ( ! [X5] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X5),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X5),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl2),
  inference(superposition,[],[f197,f5018])).
fof(f6623,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl32),
  inference(subsumption_resolution,[],[f6519,f219])).
fof(f6519,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10) | $spl32),
  inference(superposition,[],[f196,f5815])).
fof(f5815,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl32),
  inference(resolution,[],[f444,f190])).
fof(f6624,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl32),
  inference(subsumption_resolution,[],[f6520,f220])).
fof(f6520,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(sK9,bool),X0))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X0),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl32),
  inference(superposition,[],[f197,f5815])).
fof(f6625,plain,(
  ( ! [X1] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl32),
  inference(subsumption_resolution,[],[f6521,f220])).
fof(f6521,plain,(
  ( ! [X1] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X1),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X1),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP2(sK9,sK10)) ) | $spl32),
  inference(superposition,[],[f197,f5815])).
fof(f7004,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7003,f220])).
fof(f7003,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f6846,f5018])).
fof(f6846,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f6737,f220])).
fof(f6737,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(superposition,[],[f199,f5099])).
fof(f7001,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f7000,f220])).
fof(f7000,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f6854,f5018])).
fof(f6854,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f6739,f220])).
fof(f6739,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(superposition,[],[f201,f5099])).
fof(f6994,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f6993,f220])).
fof(f6993,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(forward_demodulation,[],[f6864,f5018])).
fof(f6864,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(subsumption_resolution,[],[f6741,f220])).
fof(f6741,plain,(
  ~p(s(bool,'T_0')) | sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(superposition,[],[f203,f5099])).
fof(f6979,plain,(
  $false | ($spl2 | $spl5 | $spl54)),
  inference(subsumption_resolution,[],[f6978,f253])).
fof(f6978,plain,(
  ~p(s(bool,'T_0')) | ($spl2 | $spl5 | $spl54)),
  inference(forward_demodulation,[],[f6977,f5710])).
fof(f5710,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | ($spl2 | $spl5)),
  inference(resolution,[],[f5367,f190])).
fof(f5367,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5366,f4844])).
fof(f5366,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5365,f4845])).
fof(f5365,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5364,f389])).
fof(f389,plain,(
  ~sP2(sK9,sK10) | $spl5),
  inference(resolution,[],[f306,f205])).
fof(f6977,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5 | $spl54)),
  inference(subsumption_resolution,[],[f6976,f223])).
fof(f6976,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5 | $spl54)),
  inference(forward_demodulation,[],[f6975,f5764])).
fof(f5764,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | ($spl2 | $spl5)),
  inference(resolution,[],[f5371,f190])).
fof(f5371,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5370,f4844])).
fof(f5370,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5369,f4845])).
fof(f5369,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5368,f389])).
fof(f6975,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5 | $spl54)),
  inference(subsumption_resolution,[],[f6931,f1657])).
fof(f6931,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ($spl2 | $spl5)),
  inference(resolution,[],[f5375,f159])).
fof(f5375,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5374,f4844])).
fof(f5374,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5373,f4845])).
fof(f5373,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5372,f389])).
fof(f6867,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f6866,f220])).
fof(f6866,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(forward_demodulation,[],[f6865,f5018])).
fof(f6865,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f6864,f389])).
fof(f6006,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f6005,f4844])).
fof(f6005,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f6004,f389])).
fof(f5842,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl110),
  inference(cnf_transformation,[],[f5842_D])).
fof(f5842_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f5822,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl32),
  inference(resolution,[],[f444,f255])).
fof(f5810,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl32),
  inference(resolution,[],[f444,f178])).
fof(f5807,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl32),
  inference(resolution,[],[f444,f177])).
fof(f5806,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl32),
  inference(resolution,[],[f444,f176])).
fof(f5805,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl32),
  inference(resolution,[],[f444,f174])).
fof(f5803,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl32),
  inference(resolution,[],[f444,f172])).
fof(f5771,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5371,f255])).
fof(f5759,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5371,f178])).
fof(f5756,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5371,f177])).
fof(f5755,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5371,f176])).
fof(f5754,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5371,f174])).
fof(f5752,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5371,f172])).
fof(f5745,plain,(
  $false | ($spl2 | $spl31)),
  inference(subsumption_resolution,[],[f442,f4845])).
fof(f5717,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5367,f255])).
fof(f5705,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5367,f178])).
fof(f5702,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5367,f177])).
fof(f5701,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5367,f176])).
fof(f5700,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5367,f174])).
fof(f5698,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | ($spl2 | $spl5)),
  inference(resolution,[],[f5367,f172])).
fof(f5681,plain,(
  $false | ($spl5 | $spl48 | $spl70)),
  inference(trivial_inequality_removal,[],[f5677])).
fof(f5677,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl5 | $spl48 | $spl70)),
  inference(superposition,[],[f5650,f2002])).
fof(f5650,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3)) ) | ($spl5 | $spl48)),
  inference(subsumption_resolution,[],[f5649,f187])).
fof(f5649,plain,(
  ( ! [X3] : (s(bool,'F_0') != s(bool,X3) | p(s(bool,X3))) ) | ($spl5 | $spl48)),
  inference(forward_demodulation,[],[f931,f4958])).
fof(f4958,plain,(
  s(bool,'F_0') = s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl5),
  inference(resolution,[],[f3007,f186])).
fof(f3007,plain,(
  ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl5),
  inference(resolution,[],[f389,f204])).
fof(f5680,plain,(
  $false | ($spl5 | $spl48)),
  inference(trivial_inequality_removal,[],[f5678])).
fof(f5678,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl5 | $spl48)),
  inference(superposition,[],[f5650,f4958])).
fof(f5679,plain,(
  $false | ($spl5 | $spl48)),
  inference(equality_resolution,[],[f5650])).
fof(f5644,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X0),s('type.num.num',sK3(X0,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X0),s('type.num.num',sK4(X0,sK10,sK9))))))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X0),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X0),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl104),
  inference(cnf_transformation,[],[f5644_D])).
fof(f5644_D,plain,(
  ( ! [X2,X0] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X0),s('type.num.num',sK3(X0,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X0),s('type.num.num',sK4(X0,sK10,sK9))))))) | ~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X0),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X0),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f5615,plain,(
  ( ! [X47,X48,X49] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X47),s('type.num.num',sK3(X47,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X47),s('type.num.num',sK4(X47,sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X47),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X48)) | ~p(s(bool,X49)) | s(bool,X49) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X47),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl2),
  inference(resolution,[],[f4849,f255])).
fof(f5601,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X28),s('type.num.num',sK3(X28,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X28),s('type.num.num',sK4(X28,sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X28),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X30) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X28),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X29) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f4849,f178])).
fof(f5599,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X20),s('type.num.num',sK3(X20,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X20),s('type.num.num',sK4(X20,sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X20),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | ~p(s(bool,X23)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X20),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl2),
  inference(resolution,[],[f4849,f177])).
fof(f5597,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X12),s('type.num.num',sK3(X12,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X12),s('type.num.num',sK4(X12,sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X12),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X13)) | ~p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X12),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl2),
  inference(resolution,[],[f4849,f176])).
fof(f5596,plain,(
  ( ! [X10,X8,X11,X9] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X8),s('type.num.num',sK3(X8,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X8),s('type.num.num',sK4(X8,sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X8),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X8),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X9) != s(bool,X11)) ) | $spl2),
  inference(resolution,[],[f4849,f174])).
fof(f5594,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X0),s('type.num.num',sK3(X0,sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),X0),s('type.num.num',sK4(X0,sK10,sK9))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),X0),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X1)) | ~p(s(bool,X2)) | s(bool,X1) != s(bool,X3) | s(bool,X2) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),X0),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl2),
  inference(resolution,[],[f4849,f172])).
fof(f5379,plain,(
  ( ! [X14,X12] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X12),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X12))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl102),
  inference(cnf_transformation,[],[f5379_D])).
fof(f5379_D,plain,(
  ( ! [X14,X12] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X12),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X12))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f5340,plain,(
  ( ! [X59,X57,X58] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X57),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X58)) | p(s(bool,X59)) | s(bool,X59) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X57))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4846,f255])).
fof(f5328,plain,(
  ( ! [X43,X41,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X40),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X41)) | ~p(s(bool,X42)) | s(bool,X41) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X40))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X42) != s(bool,X43)) ) | $spl2),
  inference(resolution,[],[f4846,f178])).
fof(f5325,plain,(
  ( ! [X30,X28,X31,X29] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X28),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X29)) | ~p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X28))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4846,f177])).
fof(f5324,plain,(
  ( ! [X26,X24,X27,X25] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X24),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X25) != s(bool,X27) | s(bool,X26) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X24))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4846,f176])).
fof(f5323,plain,(
  ( ! [X23,X21,X22,X20] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X20),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X20))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X22) != s(bool,X23)) ) | $spl2),
  inference(resolution,[],[f4846,f174])).
fof(f5321,plain,(
  ( ! [X14,X12,X15,X13] : (~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X12),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),X12))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4846,f172])).
fof(f3581,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X5))) ) | $spl88),
  inference(cnf_transformation,[],[f3581_D])).
fof(f3581_D,plain,(
  ( ! [X5] : (s(bool,X5) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X5))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f5106,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4844,f255])).
fof(f5094,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl2),
  inference(resolution,[],[f4844,f178])).
fof(f5091,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4844,f177])).
fof(f5090,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4844,f176])).
fof(f5089,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl2),
  inference(resolution,[],[f4844,f174])).
fof(f5087,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4844,f172])).
fof(f5134,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5133,f4845])).
fof(f5133,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5086,f389])).
fof(f444,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl32),
  inference(cnf_transformation,[],[f444_D])).
fof(f444_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f5132,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5131,f4845])).
fof(f5131,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5085,f389])).
fof(f5085,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4844,f202])).
fof(f5130,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5129,f4845])).
fof(f5129,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5084,f389])).
fof(f5128,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5127,f4845])).
fof(f5127,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5083,f389])).
fof(f5083,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4844,f200])).
fof(f5126,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5125,f4845])).
fof(f5125,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5082,f389])).
fof(f5124,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5123,f4845])).
fof(f5123,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f5081,f389])).
fof(f5081,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4844,f198])).
fof(f1629,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9))) ) | $spl50),
  inference(cnf_transformation,[],[f1629_D])).
fof(f1629_D,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | p(s(bool,X9))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5025,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4845,f255])).
fof(f5013,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f4845,f178])).
fof(f5010,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4845,f177])).
fof(f5009,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4845,f176])).
fof(f5008,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X15) != s(bool,X16)) ) | $spl2),
  inference(resolution,[],[f4845,f174])).
fof(f5006,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f4845,f172])).
fof(f5047,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4999,f389])).
fof(f5046,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4998,f389])).
fof(f4998,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4845,f202])).
fof(f5045,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4997,f389])).
fof(f5044,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4996,f389])).
fof(f4996,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4845,f200])).
fof(f5043,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4995,f389])).
fof(f5042,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f4994,f389])).
fof(f4994,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f4845,f198])).
fof(f426,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl26),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4969,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f255])).
fof(f4955,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f3007,f178])).
fof(f4953,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f177])).
fof(f4951,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f176])).
fof(f4950,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f3007,f174])).
fof(f4948,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f172])).
fof(f4899,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | ($spl12 | $spl18)),
  inference(subsumption_resolution,[],[f4885,f315])).
fof(f4885,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,X6)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f182])).
fof(f4898,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | ($spl12 | $spl18)),
  inference(subsumption_resolution,[],[f4884,f315])).
fof(f4884,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,X4))) ) | $spl12),
  inference(superposition,[],[f312,f182])).
fof(f1995,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) | $spl68),
  inference(cnf_transformation,[],[f1995_D])).
fof(f1995_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,X12)) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4897,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,'F_0') = s(bool,X12)) ) | $spl12),
  inference(subsumption_resolution,[],[f4880,f186])).
fof(f4880,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X12)) | s(bool,'F_0') = s(bool,X12)) ) | $spl12),
  inference(superposition,[],[f312,f193])).
fof(f4855,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f203])).
fof(f4854,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f202])).
fof(f4853,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f201])).
fof(f4852,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f200])).
fof(f4851,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f199])).
fof(f4850,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f198])).
fof(f259,plain,(
  sP1(sK9,sK10) | $spl2),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  sP1(sK9,sK10) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f3623,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl54),
  inference(resolution,[],[f1657,f158])).
fof(f3622,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl54),
  inference(resolution,[],[f1657,f157])).
fof(f3621,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl54),
  inference(resolution,[],[f1657,f156])).
fof(f3620,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl54),
  inference(resolution,[],[f1657,f155])).
fof(f3619,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl54),
  inference(resolution,[],[f1657,f154])).
fof(f3618,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl54),
  inference(resolution,[],[f1657,f153])).
fof(f4843,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl100),
  inference(cnf_transformation,[],[f4843_D])).
fof(f4843_D,plain,(
  ( ! [X27] : (~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f4778,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl29),
  inference(resolution,[],[f440,f172])).
fof(f4780,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl29),
  inference(resolution,[],[f440,f174])).
fof(f4781,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl29),
  inference(resolution,[],[f440,f176])).
fof(f4783,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl29),
  inference(resolution,[],[f440,f177])).
fof(f4785,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl29),
  inference(resolution,[],[f440,f178])).
fof(f4799,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl29),
  inference(resolution,[],[f440,f255])).
fof(f4842,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f3988,f311])).
fof(f3988,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(resolution,[],[f308,f167])).
fof(f4765,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4717,f308])).
fof(f4717,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl11),
  inference(resolution,[],[f311,f167])).
fof(f4839,plain,(
  s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) | $spl99),
  inference(cnf_transformation,[],[f4839_D])).
fof(f4839_D,plain,(
  s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) <=> ~$spl99),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl99])])).
fof(f4841,plain,(
  s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) | ~sP1(sK9,sK10) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f3989,f311])).
fof(f3989,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) | ~sP1(sK9,sK10) | $spl6),
  inference(resolution,[],[f308,f151])).
fof(f151,plain,(
  ( ! [X6,X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | s('type.num.num',sK3(X6,X1,X0)) != s('type.num.num',sK4(X6,X1,X0)) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X6),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ~sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f97])).
fof(f4840,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | ~sP1(sK9,sK10) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f3990,f311])).
fof(f3990,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | ~sP1(sK9,sK10) | $spl6),
  inference(resolution,[],[f308,f152])).
fof(f4837,plain,(
  s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) | ~sP1(sK9,sK10) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4718,f308])).
fof(f4718,plain,(
  s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP1(sK9,sK10) | $spl11),
  inference(resolution,[],[f311,f151])).
fof(f4834,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | ~sP1(sK9,sK10) | ($spl6 | $spl11)),
  inference(subsumption_resolution,[],[f4719,f308])).
fof(f4719,plain,(
  ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK3(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK4(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP1(sK9,sK10) | $spl11),
  inference(resolution,[],[f311,f152])).
fof(f4817,plain,(
  $false | ($spl4 | $spl29)),
  inference(subsumption_resolution,[],[f4776,f892])).
fof(f4776,plain,(
  ~sP2(sK9,sK10) | $spl29),
  inference(resolution,[],[f440,f194])).
fof(f194,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f117])).
fof(f4329,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl92),
  inference(cnf_transformation,[],[f4329_D])).
fof(f4329_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f4741,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f255])).
fof(f4727,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl11),
  inference(resolution,[],[f311,f178])).
fof(f4725,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f177])).
fof(f4723,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f176])).
fof(f4722,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl11),
  inference(resolution,[],[f311,f174])).
fof(f4720,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f172])).
fof(f4639,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4638,f1657])).
fof(f4638,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4637,f260])).
fof(f4637,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(subsumption_resolution,[],[f4513,f220])).
fof(f4513,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(superposition,[],[f158,f4369])).
fof(f4369,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | $spl4),
  inference(resolution,[],[f3633,f190])).
fof(f3633,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f892,f195])).
fof(f195,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~sP2(X0,X1)) )),
  inference(cnf_transformation,[],[f117])).
fof(f4636,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4635,f1657])).
fof(f4635,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4634,f260])).
fof(f4634,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(subsumption_resolution,[],[f4512,f220])).
fof(f4512,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(superposition,[],[f157,f4369])).
fof(f4633,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4632,f1657])).
fof(f4632,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4631,f260])).
fof(f4631,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f4511,f220])).
fof(f4511,plain,(
  ( ! [X2,X3] : (~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(superposition,[],[f156,f4369])).
fof(f4630,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4629,f1657])).
fof(f4629,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4628,f260])).
fof(f4628,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(subsumption_resolution,[],[f4510,f220])).
fof(f4510,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(superposition,[],[f155,f4369])).
fof(f4627,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4626,f1657])).
fof(f4626,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4625,f260])).
fof(f4625,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(subsumption_resolution,[],[f4509,f220])).
fof(f4509,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(superposition,[],[f154,f4369])).
fof(f4624,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4623,f1657])).
fof(f4623,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4622,f260])).
fof(f4622,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(subsumption_resolution,[],[f4508,f220])).
fof(f4508,plain,(
  ~p(s(bool,'T_0')) | sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(superposition,[],[f153,f4369])).
fof(f4471,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl94),
  inference(cnf_transformation,[],[f4471_D])).
fof(f4471_D,plain,(
  ( ! [X5] : (p(s(bool,X5)) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f4441,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl10),
  inference(resolution,[],[f310,f255])).
fof(f4429,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl10),
  inference(resolution,[],[f310,f178])).
fof(f4426,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl10),
  inference(resolution,[],[f310,f177])).
fof(f4425,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl10),
  inference(resolution,[],[f310,f176])).
fof(f4424,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl10),
  inference(resolution,[],[f310,f174])).
fof(f4422,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl10),
  inference(resolution,[],[f310,f172])).
fof(f4376,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f3633,f255])).
fof(f4364,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X30) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f3633,f178])).
fof(f4361,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f3633,f177])).
fof(f4360,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f3633,f176])).
fof(f4359,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X15) != s(bool,X16)) ) | $spl4),
  inference(resolution,[],[f3633,f174])).
fof(f4357,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f3633,f172])).
fof(f4404,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4403,f1657])).
fof(f4403,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4344,f260])).
fof(f4344,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f3633,f158])).
fof(f4402,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4401,f1657])).
fof(f4401,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4343,f260])).
fof(f4343,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f3633,f157])).
fof(f4400,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4399,f1657])).
fof(f4399,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4342,f260])).
fof(f4342,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(resolution,[],[f3633,f156])).
fof(f4398,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4397,f1657])).
fof(f4397,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4341,f260])).
fof(f4341,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(resolution,[],[f3633,f155])).
fof(f4396,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4395,f1657])).
fof(f4395,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4340,f260])).
fof(f4340,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f3633,f154])).
fof(f4394,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4 | $spl54)),
  inference(subsumption_resolution,[],[f4393,f1657])).
fof(f4393,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f4339,f260])).
fof(f4339,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f3633,f153])).
fof(f4302,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f255])).
fof(f4288,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X21) != s(bool,X23)) ) | $spl11),
  inference(resolution,[],[f311,f178])).
fof(f4286,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f177])).
fof(f4284,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f176])).
fof(f4283,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X6) != s(bool,X8)) ) | $spl11),
  inference(resolution,[],[f311,f174])).
fof(f4281,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl11),
  inference(resolution,[],[f311,f172])).
fof(f4237,plain,(
  ( ! [X68,X66,X69,X67] : (s('type.num.num',X66) = s('type.num.num',X67) | p(s(bool,X68)) | p(s(bool,X69)) | s(bool,X69) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X66))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X67)))))) ) | $spl8),
  inference(resolution,[],[f309,f255])).
fof(f4225,plain,(
  ( ! [X39,X37,X41,X38,X40] : (s('type.num.num',X37) = s('type.num.num',X38) | p(s(bool,X39)) | ~p(s(bool,X40)) | s(bool,X39) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X37))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X38))))) | s(bool,X40) != s(bool,X41)) ) | $spl8),
  inference(resolution,[],[f309,f178])).
fof(f4222,plain,(
  ( ! [X26,X24,X23,X25,X22] : (s('type.num.num',X22) = s('type.num.num',X23) | p(s(bool,X24)) | ~p(s(bool,X25)) | p(s(bool,X26)) | s(bool,X26) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X22))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X23)))))) ) | $spl8),
  inference(resolution,[],[f309,f177])).
fof(f4221,plain,(
  ( ! [X21,X19,X17,X20,X18] : (s('type.num.num',X17) = s('type.num.num',X18) | p(s(bool,X19)) | p(s(bool,X20)) | s(bool,X19) != s(bool,X21) | s(bool,X20) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X17))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X18)))))) ) | $spl8),
  inference(resolution,[],[f309,f176])).
fof(f4220,plain,(
  ( ! [X14,X12,X15,X13,X16] : (s('type.num.num',X12) = s('type.num.num',X13) | p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X12))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X13))))) | s(bool,X15) != s(bool,X16)) ) | $spl8),
  inference(resolution,[],[f309,f174])).
fof(f4218,plain,(
  ( ! [X6,X4,X2,X5,X3] : (s('type.num.num',X2) = s('type.num.num',X3) | p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3)))))) ) | $spl8),
  inference(resolution,[],[f309,f172])).
fof(f2681,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl78),
  inference(cnf_transformation,[],[f2681_D])).
fof(f2681_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f4010,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f255])).
fof(f3998,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X22) != s(bool,X23)) ) | $spl6),
  inference(resolution,[],[f308,f178])).
fof(f3995,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f177])).
fof(f3994,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f176])).
fof(f3993,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X7) != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f308,f174])).
fof(f3991,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f172])).
fof(f310,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl10),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4032,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ($spl6 | $spl66)),
  inference(subsumption_resolution,[],[f4031,f220])).
fof(f4031,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ($spl6 | $spl66)),
  inference(forward_demodulation,[],[f3988,f1992])).
fof(f3759,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl64),
  inference(subsumption_resolution,[],[f3748,f185])).
fof(f3748,plain,(
  ( ! [X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl64),
  inference(superposition,[],[f1723,f193])).
fof(f3756,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f3740,f255])).
fof(f3740,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | p(s(bool,X2))) ) | $spl64),
  inference(superposition,[],[f1723,f179])).
fof(f3755,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(subsumption_resolution,[],[f3739,f255])).
fof(f3739,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl64),
  inference(superposition,[],[f1723,f179])).
fof(f3711,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f1681,f255])).
fof(f1681,plain,(
  p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl0),
  inference(resolution,[],[f257,f165])).
fof(f165,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))) )),
  inference(cnf_transformation,[],[f106])).
fof(f3699,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X24) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f1681,f178])).
fof(f3696,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f1681,f177])).
fof(f3695,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f1681,f176])).
fof(f3694,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X9) != s(bool,X10)) ) | $spl0),
  inference(resolution,[],[f1681,f174])).
fof(f3692,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f1681,f172])).
fof(f3606,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl67),
  inference(subsumption_resolution,[],[f3273,f189])).
fof(f3273,plain,(
  ( ! [X3] : (s(bool,'T_0') != s(bool,X3) | ~p(s(bool,X3)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl67),
  inference(superposition,[],[f1993,f182])).
fof(f3608,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl67),
  inference(subsumption_resolution,[],[f3272,f189])).
fof(f3272,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,X2) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,X2))) ) | $spl67),
  inference(superposition,[],[f1993,f182])).
fof(f3679,plain,(
  $false | ($spl64 | $spl70)),
  inference(trivial_inequality_removal,[],[f3677])).
fof(f3677,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl64 | $spl70)),
  inference(superposition,[],[f3611,f2002])).
fof(f3611,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29)) ) | ($spl64 | $spl70)),
  inference(subsumption_resolution,[],[f3610,f187])).
fof(f3610,plain,(
  ( ! [X29] : (s(bool,'F_0') != s(bool,X29) | p(s(bool,X29))) ) | ($spl64 | $spl70)),
  inference(forward_demodulation,[],[f1723,f2002])).
fof(f3678,plain,(
  $false | ($spl64 | $spl70)),
  inference(equality_resolution,[],[f3611])).
fof(f3630,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f3623,f260])).
fof(f3629,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f3622,f260])).
fof(f3628,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f3621,f260])).
fof(f3627,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f3620,f260])).
fof(f3626,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f3619,f260])).
fof(f3625,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl54)),
  inference(subsumption_resolution,[],[f3618,f260])).
fof(f3617,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f158])).
fof(f3616,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f157])).
fof(f3615,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f156])).
fof(f3614,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f155])).
fof(f3613,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f154])).
fof(f3612,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f153])).
fof(f308,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | $spl6),
  inference(cnf_transformation,[],[f308_D])).
fof(f308_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f3393,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl70),
  inference(subsumption_resolution,[],[f3296,f235])).
fof(f3296,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl70),
  inference(superposition,[],[f168,f2002])).
fof(f3394,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl70),
  inference(subsumption_resolution,[],[f3297,f235])).
fof(f3297,plain,(
  ( ! [X2,X3] : (p(s(bool,'F_0')) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl70),
  inference(superposition,[],[f169,f2002])).
fof(f311,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl11),
  inference(cnf_transformation,[],[f311_D])).
fof(f311_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f3395,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl70),
  inference(subsumption_resolution,[],[f3298,f217])).
fof(f3298,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl70),
  inference(superposition,[],[f170,f2002])).
fof(f3552,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2988,f255])).
fof(f2988,plain,(
  p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f259,f147])).
fof(f3540,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl2),
  inference(resolution,[],[f2988,f178])).
fof(f3537,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2988,f177])).
fof(f3536,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2988,f176])).
fof(f3535,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl2),
  inference(resolution,[],[f2988,f174])).
fof(f3533,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2988,f172])).
fof(f3580,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3579,f2989])).
fof(f2989,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f259,f148])).
fof(f3579,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3532,f389])).
fof(f3532,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2988,f203])).
fof(f3578,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3577,f2989])).
fof(f3577,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3531,f389])).
fof(f3531,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2988,f202])).
fof(f3576,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3575,f2989])).
fof(f3575,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3530,f389])).
fof(f3530,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2988,f201])).
fof(f3574,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3573,f2989])).
fof(f3573,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3529,f389])).
fof(f3529,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2988,f200])).
fof(f3572,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3571,f2989])).
fof(f3571,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3528,f389])).
fof(f3528,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2988,f199])).
fof(f3570,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3569,f2989])).
fof(f3569,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3527,f389])).
fof(f3527,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2988,f198])).
fof(f3490,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2989,f255])).
fof(f3478,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f2989,f178])).
fof(f3475,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2989,f177])).
fof(f3474,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2989,f176])).
fof(f3473,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X15) != s(bool,X16)) ) | $spl2),
  inference(resolution,[],[f2989,f174])).
fof(f3471,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f2989,f172])).
fof(f3512,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3464,f389])).
fof(f3464,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2989,f203])).
fof(f3511,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3463,f389])).
fof(f3463,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2989,f202])).
fof(f3510,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3462,f389])).
fof(f3462,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2989,f201])).
fof(f3509,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3461,f389])).
fof(f3461,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2989,f200])).
fof(f3508,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3460,f389])).
fof(f3460,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2989,f199])).
fof(f3507,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f3459,f389])).
fof(f3459,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f2989,f198])).
fof(f3263,plain,(
  $false | ($spl12 | $spl64)),
  inference(equality_resolution,[],[f3200])).
fof(f3200,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | ($spl12 | $spl64)),
  inference(subsumption_resolution,[],[f1723,f312])).
fof(f3224,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f255])).
fof(f3210,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f3007,f178])).
fof(f3208,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f177])).
fof(f3206,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f176])).
fof(f3205,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f3007,f174])).
fof(f3203,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f3007,f172])).
fof(f3153,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | ($spl12 | $spl18)),
  inference(subsumption_resolution,[],[f3139,f315])).
fof(f3139,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,X6)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl12),
  inference(superposition,[],[f312,f182])).
fof(f3152,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | ($spl12 | $spl18)),
  inference(subsumption_resolution,[],[f3138,f315])).
fof(f3138,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,X4))) ) | $spl12),
  inference(superposition,[],[f312,f182])).
fof(f3151,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,'F_0') = s(bool,X12)) ) | $spl12),
  inference(subsumption_resolution,[],[f3134,f186])).
fof(f3134,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X12)) | s(bool,'F_0') = s(bool,X12)) ) | $spl12),
  inference(superposition,[],[f312,f193])).
fof(f3108,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3040,f317])).
fof(f3040,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(resolution,[],[f317,f250])).
fof(f3107,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3039,f317])).
fof(f3039,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(resolution,[],[f317,f248])).
fof(f3106,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3036,f317])).
fof(f3036,plain,(
  ( ! [X48] : (p(s(bool,X48))) ) | $spl22),
  inference(resolution,[],[f317,f228])).
fof(f3105,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3035,f317])).
fof(f3035,plain,(
  ( ! [X47] : (p(s(bool,X47))) ) | $spl22),
  inference(resolution,[],[f317,f225])).
fof(f3104,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3033,f317])).
fof(f3033,plain,(
  p(s(bool,'F_0')) | $spl22),
  inference(resolution,[],[f317,f222])).
fof(f3103,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f3031,f317])).
fof(f3031,plain,(
  ( ! [X44] : (p(s(bool,X44))) ) | $spl22),
  inference(resolution,[],[f317,f219])).
fof(f3041,plain,(
  $false | $spl22),
  inference(resolution,[],[f317,f252])).
fof(f3042,plain,(
  $false | $spl22),
  inference(resolution,[],[f317,f253])).
fof(f3006,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f203])).
fof(f3005,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f202])).
fof(f3004,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f201])).
fof(f3003,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f200])).
fof(f3002,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f199])).
fof(f3001,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f198])).
fof(f2954,plain,(
  ( ! [X30,X31] : (p(s(bool,X30)) | p(s(bool,X31)) | s(bool,X31) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f2925,f255])).
fof(f2925,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl55),
  inference(resolution,[],[f1658,f160])).
fof(f2942,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X25)) | ~p(s(bool,X26)) | s(bool,X25) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X26) != s(bool,X27)) ) | $spl55),
  inference(resolution,[],[f2925,f178])).
fof(f2939,plain,(
  ( ! [X17,X18,X16] : (p(s(bool,X16)) | ~p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f2925,f177])).
fof(f2938,plain,(
  ( ! [X14,X15,X13] : (p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X13) != s(bool,X15) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f2925,f176])).
fof(f2937,plain,(
  ( ! [X12,X10,X11] : (p(s(bool,X10)) | ~p(s(bool,X11)) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X11) != s(bool,X12)) ) | $spl55),
  inference(resolution,[],[f2925,f174])).
fof(f2935,plain,(
  ( ! [X6,X4,X5] : (p(s(bool,X4)) | p(s(bool,X5)) | s(bool,X4) != s(bool,X6) | s(bool,X5) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK7(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl55),
  inference(resolution,[],[f2925,f172])).
fof(f2892,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f1742,f255])).
fof(f1742,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f892,f195])).
fof(f2880,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X30) != s(bool,X31)) ) | $spl4),
  inference(resolution,[],[f1742,f178])).
fof(f2877,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f1742,f177])).
fof(f2876,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f1742,f176])).
fof(f2875,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X15) != s(bool,X16)) ) | $spl4),
  inference(resolution,[],[f1742,f174])).
fof(f2873,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl4),
  inference(resolution,[],[f1742,f172])).
fof(f2871,plain,(
  ( ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),X7),s(fun(sK9,bool),'const.pred_set.EMPTY_0'))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X7),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl4),
  inference(resolution,[],[f1742,f159])).
fof(f2870,plain,(
  ( ! [X6] : (p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(sK9,bool),X6))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),X6),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10)) ) | $spl4),
  inference(resolution,[],[f1742,f159])).
fof(f2914,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f2860,f260])).
fof(f2860,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f1742,f158])).
fof(f2913,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f2859,f260])).
fof(f2859,plain,(
  sP1(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f1742,f157])).
fof(f2912,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f2858,f260])).
fof(f2858,plain,(
  ( ! [X2,X3] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(resolution,[],[f1742,f156])).
fof(f2911,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f2857,f260])).
fof(f2857,plain,(
  ( ! [X0,X1] : (sP1(sK9,sK10) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl4),
  inference(resolution,[],[f1742,f155])).
fof(f2910,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f2856,f260])).
fof(f2856,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f1742,f154])).
fof(f2909,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl3 | $spl4)),
  inference(subsumption_resolution,[],[f2855,f260])).
fof(f2855,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f1742,f153])).
fof(f2843,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl80),
  inference(cnf_transformation,[],[f2843_D])).
fof(f2843_D,plain,(
  ( ! [X1] : (~p(s(bool,X1)) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2819,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl7),
  inference(resolution,[],[f307,f255])).
fof(f307,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | $spl7),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f2805,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X21) != s(bool,X23)) ) | $spl7),
  inference(resolution,[],[f307,f178])).
fof(f2803,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl7),
  inference(resolution,[],[f307,f177])).
fof(f2801,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl7),
  inference(resolution,[],[f307,f176])).
fof(f2800,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X6) != s(bool,X8)) ) | $spl7),
  inference(resolution,[],[f307,f174])).
fof(f2798,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl7),
  inference(resolution,[],[f307,f172])).
fof(f2791,plain,(
  $false | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f2721,f2688])).
fof(f2688,plain,(
  ( ! [X1] : (~p(s(bool,X1))) ) | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f2687,f192])).
fof(f2687,plain,(
  ( ! [X1] : (s(bool,'T_0') != s(bool,X1) | ~p(s(bool,X1))) ) | ($spl4 | $spl26)),
  inference(forward_demodulation,[],[f426,f1833])).
fof(f1833,plain,(
  s(bool,'T_0') = s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl4),
  inference(resolution,[],[f305,f190])).
fof(f2721,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f250])).
fof(f2790,plain,(
  $false | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f2720,f2688])).
fof(f2720,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f248])).
fof(f2789,plain,(
  $false | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f2717,f2688])).
fof(f2717,plain,(
  ( ! [X48] : (p(s(bool,X48))) ) | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f228])).
fof(f2788,plain,(
  $false | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f2716,f2688])).
fof(f2716,plain,(
  ( ! [X47] : (p(s(bool,X47))) ) | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f225])).
fof(f2787,plain,(
  $false | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f2714,f2688])).
fof(f2714,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f222])).
fof(f2786,plain,(
  $false | ($spl4 | $spl26)),
  inference(subsumption_resolution,[],[f2712,f2688])).
fof(f2712,plain,(
  ( ! [X44] : (p(s(bool,X44))) ) | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f219])).
fof(f2722,plain,(
  $false | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f252])).
fof(f2723,plain,(
  $false | ($spl4 | $spl26)),
  inference(resolution,[],[f2688,f253])).
fof(f2652,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | p(s(bool,X27)) | s(bool,X27) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f255])).
fof(f2640,plain,(
  ( ! [X23,X21,X22] : (p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X21) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X22) != s(bool,X23)) ) | $spl6),
  inference(resolution,[],[f308,f178])).
fof(f2637,plain,(
  ( ! [X14,X12,X13] : (p(s(bool,X12)) | ~p(s(bool,X13)) | p(s(bool,X14)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f177])).
fof(f2636,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f176])).
fof(f2635,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X6) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) | s(bool,X7) != s(bool,X8)) ) | $spl6),
  inference(resolution,[],[f308,f174])).
fof(f2633,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) ) | $spl6),
  inference(resolution,[],[f308,f172])).
fof(f2678,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f2677,f220])).
fof(f2677,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f2630,f2043])).
fof(f2043,plain,(
  s(bool,'T_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl0),
  inference(resolution,[],[f257,f190])).
fof(f2630,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9))))))) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(resolution,[],[f308,f167])).
fof(f2674,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ($spl0 | $spl6)),
  inference(subsumption_resolution,[],[f2673,f219])).
fof(f2673,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ($spl0 | $spl6)),
  inference(forward_demodulation,[],[f2629,f2043])).
fof(f2629,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | s('type.num.num',sK11(sK13(sK10,sK9),sK10,sK9)) != s('type.num.num',sK12(sK13(sK10,sK9),sK10,sK9)) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl6),
  inference(resolution,[],[f308,f166])).
fof(f1723,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X29))) ) | $spl64),
  inference(cnf_transformation,[],[f1723_D])).
fof(f1723_D,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | p(s(bool,X29))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2050,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f255])).
fof(f2038,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X24) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f257,f178])).
fof(f314,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl16),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  ( ! [X7] : (s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f315,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X5)) | p(s(bool,X6))) ) | $spl18),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  ( ! [X6,X5] : (s(bool,X5) != s(bool,X6) | ~p(s(bool,X5)) | p(s(bool,X6))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2037,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f178])).
fof(f2035,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f177])).
fof(f2034,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f176])).
fof(f2033,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X9) != s(bool,X10)) ) | $spl0),
  inference(resolution,[],[f257,f174])).
fof(f2032,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f174])).
fof(f2031,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f172])).
fof(f2012,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f158])).
fof(f2011,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f157])).
fof(f2010,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f156])).
fof(f2009,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f155])).
fof(f2008,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f154])).
fof(f2007,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f153])).
fof(f2006,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl72),
  inference(cnf_transformation,[],[f2006_D])).
fof(f2006_D,plain,(
  ( ! [X22] : (s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1822,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f174])).
fof(f1827,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f178])).
fof(f1994,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,'F_0') = s(bool,X12)) ) | $spl12),
  inference(subsumption_resolution,[],[f1975,f186])).
fof(f1975,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X12)) | s(bool,'F_0') = s(bool,X12)) ) | $spl12),
  inference(superposition,[],[f312,f193])).
fof(f1840,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f255])).
fof(f1828,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X24) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f305,f178])).
fof(f1825,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f177])).
fof(f1824,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f176])).
fof(f1823,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f305,f174])).
fof(f1821,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f172])).
fof(f1820,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl4),
  inference(resolution,[],[f305,f170])).
fof(f1819,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1)) ) | $spl4),
  inference(resolution,[],[f305,f169])).
fof(f1818,plain,(
  p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | $spl4),
  inference(resolution,[],[f305,f168])).
fof(f1684,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f172])).
fof(f1685,plain,(
  ( ! [X6,X7,X5] : (p(s(bool,X5)) | ~p(s(bool,X6)) | s(bool,X5) != s(bool,X6) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f174])).
fof(f1686,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X9) != s(bool,X10)) ) | $spl0),
  inference(resolution,[],[f257,f174])).
fof(f1687,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f176])).
fof(f1688,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f177])).
fof(f1690,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | s(bool,X20) != s(bool,X21) | s(bool,X22) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f178])).
fof(f1691,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X24) != s(bool,X25)) ) | $spl0),
  inference(resolution,[],[f257,f178])).
fof(f1703,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl0),
  inference(resolution,[],[f257,f255])).
fof(f1717,plain,(
  $false | ($spl0 | $spl5)),
  inference(subsumption_resolution,[],[f1716,f247])).
fof(f247,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f65])).
fof(f65,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT',aHL_FALSITY)).
fof(f1716,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl5)),
  inference(forward_demodulation,[],[f1681,f949])).
fof(f949,plain,(
  s(bool,'F_0') = s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl5),
  inference(resolution,[],[f306,f186])).
fof(f1654,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f158])).
fof(f1653,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f157])).
fof(f1652,plain,(
  ( ! [X2,X3] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X2))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X3)))))) | s('type.num.num',X2) = s('type.num.num',X3) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f156])).
fof(f1651,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))))) ) | $spl3),
  inference(resolution,[],[f260,f155])).
fof(f1650,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK5(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f154])).
fof(f1649,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK6(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~sP0(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK5(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl3),
  inference(resolution,[],[f260,f153])).
fof(f1648,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(forward_demodulation,[],[f940,f949])).
fof(f940,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | s(bool,X5) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X3) != s(bool,X4)) ) | $spl5),
  inference(resolution,[],[f306,f172])).
fof(f1647,plain,(
  ( ! [X14,X12,X13] : (s(bool,'F_0') != s(bool,X14) | ~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(forward_demodulation,[],[f943,f949])).
fof(f943,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | p(s(bool,X13)) | s(bool,X14) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X12) != s(bool,X13)) ) | $spl5),
  inference(resolution,[],[f306,f176])).
fof(f1438,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,X8) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X10)) ) | $spl5),
  inference(subsumption_resolution,[],[f1340,f185])).
fof(f1340,plain,(
  ( ! [X10,X8,X9] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | ~p(s(bool,X9)) | p(s(bool,X10)) | s(bool,X9) != s(bool,X10)) ) | $spl5),
  inference(superposition,[],[f172,f949])).
fof(f1643,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl5),
  inference(subsumption_resolution,[],[f1346,f255])).
fof(f1346,plain,(
  ( ! [X28,X26,X27] : (s(bool,'F_0') != s(bool,X26) | p(s(bool,X27)) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,X28)) | s(bool,X27) != s(bool,X28)) ) | $spl5),
  inference(superposition,[],[f174,f949])).
fof(f1450,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') != s(bool,X53) | ~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X54) != s(bool,X55)) ) | $spl5),
  inference(subsumption_resolution,[],[f1355,f185])).
fof(f1355,plain,(
  ( ! [X54,X55,X53] : (s(bool,'F_0') != s(bool,X53) | p(s(bool,X53)) | ~p(s(bool,X54)) | p(s(bool,X55)) | s(bool,X54) != s(bool,X55)) ) | $spl5),
  inference(superposition,[],[f176,f949])).
fof(f1642,plain,(
  ( ! [X85,X83,X84] : (s(bool,'F_0') != s(bool,X83) | p(s(bool,X84)) | ~p(s(bool,X85)) | s(bool,X84) != s(bool,X85)) ) | $spl5),
  inference(subsumption_resolution,[],[f1365,f255])).
fof(f1365,plain,(
  ( ! [X85,X83,X84] : (s(bool,'F_0') != s(bool,X83) | p(s(bool,X84)) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~p(s(bool,X85)) | s(bool,X84) != s(bool,X85)) ) | $spl5),
  inference(superposition,[],[f178,f949])).
fof(f1603,plain,(
  ( ! [X35,X34] : (p(s(bool,X34)) | p(s(bool,X35)) | s(bool,X35) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f262,f255])).
fof(f262,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f259,f148])).
fof(f1591,plain,(
  ( ! [X30,X31,X29] : (p(s(bool,X29)) | ~p(s(bool,X30)) | s(bool,X29) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X30) != s(bool,X31)) ) | $spl2),
  inference(resolution,[],[f262,f178])).
fof(f1588,plain,(
  ( ! [X21,X22,X20] : (p(s(bool,X20)) | ~p(s(bool,X21)) | p(s(bool,X22)) | s(bool,X22) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f262,f177])).
fof(f1587,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | p(s(bool,X18)) | s(bool,X17) != s(bool,X19) | s(bool,X18) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f262,f176])).
fof(f1586,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | s(bool,X14) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) | s(bool,X15) != s(bool,X16)) ) | $spl2),
  inference(resolution,[],[f262,f174])).
fof(f1584,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | p(s(bool,X9)) | s(bool,X8) != s(bool,X10) | s(bool,X9) != s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) ) | $spl2),
  inference(resolution,[],[f262,f172])).
fof(f1625,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1577,f389])).
fof(f1577,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f262,f203])).
fof(f1624,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1576,f389])).
fof(f1576,plain,(
  sP2(sK9,sK10) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f262,f202])).
fof(f1623,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1575,f389])).
fof(f1575,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f262,f201])).
fof(f1622,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1574,f389])).
fof(f1574,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f262,f200])).
fof(f1621,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1573,f389])).
fof(f1573,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f262,f199])).
fof(f1620,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ($spl2 | $spl5)),
  inference(subsumption_resolution,[],[f1572,f389])).
fof(f1572,plain,(
  sP2(sK9,sK10) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl2),
  inference(resolution,[],[f262,f198])).
fof(f1207,plain,(
  $false | ($spl1 | $spl16)),
  inference(trivial_inequality_removal,[],[f1194])).
fof(f1194,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl1 | $spl16)),
  inference(superposition,[],[f314,f280])).
fof(f280,plain,(
  s(bool,'F_0') = s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | $spl1),
  inference(resolution,[],[f258,f186])).
fof(f1206,plain,(
  $false | $spl16),
  inference(equality_resolution,[],[f314])).
fof(f984,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f203])).
fof(f983,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f202])).
fof(f982,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f201])).
fof(f981,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f200])).
fof(f980,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f199])).
fof(f979,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f198])).
fof(f960,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f255])).
fof(f946,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f306,f178])).
fof(f944,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f177])).
fof(f942,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f176])).
fof(f941,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f306,f174])).
fof(f939,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f172])).
fof(f915,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | p(s(bool,X29)) | s(bool,X29) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f255])).
fof(f903,plain,(
  ( ! [X24,X23,X25] : (p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X23) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X24) != s(bool,X25)) ) | $spl4),
  inference(resolution,[],[f305,f178])).
fof(f900,plain,(
  ( ! [X14,X15,X16] : (p(s(bool,X14)) | ~p(s(bool,X15)) | p(s(bool,X16)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f177])).
fof(f899,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f176])).
fof(f898,plain,(
  ( ! [X10,X8,X9] : (p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X8) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X9) != s(bool,X10)) ) | $spl4),
  inference(resolution,[],[f305,f174])).
fof(f896,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl4),
  inference(resolution,[],[f305,f172])).
fof(f888,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X22))) ) | $spl20),
  inference(subsumption_resolution,[],[f859,f316])).
fof(f859,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | ~p(s(bool,X22)) | s(bool,'T_0') = s(bool,X22)) ) | $spl20),
  inference(superposition,[],[f316,f193])).
fof(f887,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X20))) ) | $spl20),
  inference(subsumption_resolution,[],[f858,f316])).
fof(f858,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,X20)) | s(bool,'F_0') = s(bool,X20)) ) | $spl20),
  inference(superposition,[],[f316,f193])).
fof(f883,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,X18))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f857])).
fof(f857,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') != s(bool,X19) | ~p(s(bool,X18)) | ~p(s(bool,X18))) ) | $spl20),
  inference(superposition,[],[f316,f192])).
fof(f884,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X16))) ) | $spl20),
  inference(duplicate_literal_removal,[],[f856])).
fof(f856,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,X17) | ~p(s(bool,X16)) | ~p(s(bool,X16))) ) | $spl20),
  inference(superposition,[],[f316,f190])).
fof(f841,plain,(
  $false | $spl44),
  inference(equality_resolution,[],[f632])).
fof(f806,plain,(
  $false | $spl42),
  inference(equality_resolution,[],[f630])).
fof(f753,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f671,f625])).
fof(f625,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl14),
  inference(equality_resolution,[],[f313])).
fof(f671,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(resolution,[],[f625,f251])).
fof(f251,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f145])).
fof(f752,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f670,f625])).
fof(f670,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(resolution,[],[f625,f249])).
fof(f249,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f144])).
fof(f751,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f664,f625])).
fof(f664,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) | $spl14),
  inference(resolution,[],[f625,f217])).
fof(f669,plain,(
  $false | $spl14),
  inference(resolution,[],[f625,f247])).
fof(f675,plain,(
  $false | ($spl1 | $spl14)),
  inference(resolution,[],[f625,f258])).
fof(f640,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22)) ) | $spl46),
  inference(cnf_transformation,[],[f640_D])).
fof(f640_D,plain,(
  ( ! [X22] : (s(bool,'T_0') = s(bool,X22)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f639,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | s(bool,'T_0') = s(bool,X22)) ) | $spl14),
  inference(subsumption_resolution,[],[f618,f185])).
fof(f618,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X23)) | s(bool,'T_0') = s(bool,X22)) ) | $spl14),
  inference(superposition,[],[f313,f193])).
fof(f638,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl14),
  inference(subsumption_resolution,[],[f614,f185])).
fof(f614,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X15)) | p(s(bool,X14))) ) | $spl14),
  inference(superposition,[],[f313,f188])).
fof(f637,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12))) ) | $spl14),
  inference(subsumption_resolution,[],[f613,f185])).
fof(f613,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X13)) | p(s(bool,X12))) ) | $spl14),
  inference(superposition,[],[f313,f186])).
fof(f633,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22))) ) | $spl14),
  inference(subsumption_resolution,[],[f602,f313])).
fof(f602,plain,(
  ( ! [X23,X22] : (s(bool,'F_0') != s(bool,X23) | p(s(bool,X22)) | s(bool,'T_0') = s(bool,X22)) ) | $spl14),
  inference(superposition,[],[f313,f193])).
fof(f631,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X20))) ) | $spl14),
  inference(subsumption_resolution,[],[f601,f313])).
fof(f601,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | p(s(bool,X20)) | s(bool,'F_0') = s(bool,X20)) ) | $spl14),
  inference(superposition,[],[f313,f193])).
fof(f626,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f598])).
fof(f598,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') != s(bool,X15) | p(s(bool,X14)) | p(s(bool,X14))) ) | $spl14),
  inference(superposition,[],[f313,f188])).
fof(f627,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12))) ) | $spl14),
  inference(duplicate_literal_removal,[],[f597])).
fof(f597,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,X13) | p(s(bool,X12)) | p(s(bool,X12))) ) | $spl14),
  inference(superposition,[],[f313,f186])).
fof(f566,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f483,f318])).
fof(f483,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(resolution,[],[f318,f251])).
fof(f565,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f482,f318])).
fof(f482,plain,(
  ~p(s(bool,'T_0')) | $spl24),
  inference(resolution,[],[f318,f249])).
fof(f564,plain,(
  $false | $spl24),
  inference(subsumption_resolution,[],[f476,f318])).
fof(f476,plain,(
  ( ! [X43] : (~p(s(bool,X43))) ) | $spl24),
  inference(resolution,[],[f318,f217])).
fof(f481,plain,(
  $false | $spl24),
  inference(resolution,[],[f318,f247])).
fof(f487,plain,(
  $false | ($spl1 | $spl24)),
  inference(resolution,[],[f318,f258])).
fof(f540,plain,(
  $false | ($spl5 | $spl24)),
  inference(resolution,[],[f318,f306])).
fof(f437,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f203])).
fof(f436,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.UNION_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(sK9,bool),sK16(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f202])).
fof(f435,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f201])).
fof(f434,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK16(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f200])).
fof(f433,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.DIFF_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(sK9,bool),sK14(sK10,sK9)))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f199])).
fof(f432,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK15(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),sK14(sK10,sK9)),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.EMPTY_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.subset_class_2'(s(fun(sK9,bool),'const.measure.space_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | $spl5),
  inference(resolution,[],[f389,f198])).
fof(f412,plain,(
  ( ! [X26,X27] : (p(s(bool,X26)) | ~p(s(bool,X27)) | s(bool,X27) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f255])).
fof(f398,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,X21)) | ~p(s(bool,X22)) | s(bool,X22) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X21) != s(bool,X23)) ) | $spl5),
  inference(resolution,[],[f306,f178])).
fof(f396,plain,(
  ( ! [X17,X15,X16] : (p(s(bool,X15)) | ~p(s(bool,X16)) | ~p(s(bool,X17)) | s(bool,X16) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f177])).
fof(f394,plain,(
  ( ! [X10,X11,X9] : (p(s(bool,X9)) | ~p(s(bool,X10)) | s(bool,X9) != s(bool,X11) | s(bool,X10) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f176])).
fof(f393,plain,(
  ( ! [X6,X8,X7] : (~p(s(bool,X6)) | ~p(s(bool,X7)) | s(bool,X7) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X6) != s(bool,X8)) ) | $spl5),
  inference(resolution,[],[f306,f174])).
fof(f391,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X0)) | ~p(s(bool,X1)) | s(bool,X0) != s(bool,X2) | s(bool,X1) != s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl5),
  inference(resolution,[],[f306,f172])).
fof(f312,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X3))) ) | $spl12),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | ~p(s(bool,X3))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f291,plain,(
  ( ! [X28,X29] : (p(s(bool,X28)) | ~p(s(bool,X29)) | s(bool,X29) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl1),
  inference(resolution,[],[f258,f255])).
fof(f277,plain,(
  ( ! [X24,X23,X25] : (~p(s(bool,X23)) | ~p(s(bool,X24)) | s(bool,X24) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X23) != s(bool,X25)) ) | $spl1),
  inference(resolution,[],[f258,f178])).
fof(f275,plain,(
  ( ! [X19,X17,X18] : (p(s(bool,X17)) | ~p(s(bool,X18)) | ~p(s(bool,X19)) | s(bool,X18) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl1),
  inference(resolution,[],[f258,f177])).
fof(f274,plain,(
  ( ! [X14,X15,X16] : (~p(s(bool,X14)) | p(s(bool,X15)) | s(bool,X16) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X14) != s(bool,X15)) ) | $spl1),
  inference(resolution,[],[f258,f176])).
fof(f273,plain,(
  ( ! [X12,X13,X11] : (p(s(bool,X11)) | ~p(s(bool,X12)) | s(bool,X11) != s(bool,X13) | s(bool,X12) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl1),
  inference(resolution,[],[f258,f176])).
fof(f272,plain,(
  ( ! [X10,X8,X9] : (~p(s(bool,X8)) | ~p(s(bool,X9)) | s(bool,X9) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X8) != s(bool,X10)) ) | $spl1),
  inference(resolution,[],[f258,f174])).
fof(f271,plain,(
  ( ! [X6,X7,X5] : (~p(s(bool,X5)) | p(s(bool,X6)) | s(bool,X7) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))) | s(bool,X5) != s(bool,X6)) ) | $spl1),
  inference(resolution,[],[f258,f172])).
fof(f270,plain,(
  ( ! [X4,X2,X3] : (p(s(bool,X2)) | ~p(s(bool,X3)) | s(bool,X2) != s(bool,X4) | s(bool,X3) != s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) ) | $spl1),
  inference(resolution,[],[f258,f172])).
fof(f269,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(fun(sK9,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(sK9,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl1),
  inference(resolution,[],[f258,f170])).
fof(f268,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X0))),s(fun(sK9,bool),happ(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s('type.num.num',X1)))))) | s('type.num.num',X0) = s('type.num.num',X1) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))) ) | $spl1),
  inference(resolution,[],[f258,f169])).
fof(f267,plain,(
  p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(sK9,bool)),sK13(sK10,sK9)),s(fun(fun('type.num.num',fun(sK9,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(sK9,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))))))) | ~p(s(bool,'const.measure.algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | $spl1),
  inference(resolution,[],[f258,f168])).
fof(f163,plain,(
  sP1(sK9,sK10) | p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | sP1(sK9,sK10)) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10)))) | ~sP1(sK9,sK10))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK9,sK10])],[f101])).
fof(f101,plain,(
  ? [X0,X1] : ((p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | sP1(X0,X1)) & (~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) | ~sP1(X0,X1)))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ? [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <~> sP1(X0,X1))),
  inference(definition_folding,[],[f71,f90,f89])).
fof(f71,plain,(
  ? [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <~> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f70])).
fof(f70,plain,(
  ? [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <~> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : ((~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) | ~p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((~p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) | ? [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) & ~p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) | p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(ennf_transformation,[],[f34])).
fof(f34,plain,(
  ~! [X0,X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (s('type.num.num',X3) != s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(flattening,[],[f33])).
fof(f33,plain,(
  ~! [X0] : ! [X1] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & ! [X7] : (p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X7),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.DIFF_2'(s(fun(X0,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))),s(fun(X0,bool),X7))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X5,X6] : ((p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X5),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),X6),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.UNION_2'(s(fun(X0,bool),X5),s(fun(X0,bool),X6))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1))))))) & ! [X2] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun(fun('type.num.num',fun(X0,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))) & ! [X3,X4] : (~s('type.num.num',X3) = s('type.num.num',X4) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X3))),s(fun(X0,bool),happ(s(fun('type.num.num',fun(X0,bool)),X2),s('type.num.num',X4)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X0,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X0,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X0,bool)),X2),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X0,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X0,bool),fun(fun(X0,bool),bool)),X1)))))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & ! [X8] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(X6,bool),X8))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) & ! [X8,X9] : ((p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) & ! [X10] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X6,bool)),X10),s(fun(fun('type.num.num',fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))) & ! [X11,X12] : (~s('type.num.num',X11) = s('type.num.num',X12) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X10),s('type.num.num',X11))),s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X10),s('type.num.num',X12)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X6,bool)),X10),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7] : (p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))) <=> (p(s(bool,'const.measure.subset_class_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.EMPTY_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & ! [X8] : (p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.DIFF_2'(s(fun(X6,bool),'const.measure.space_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))),s(fun(X6,bool),X8))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) & ! [X8,X9] : ((p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X8),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))) & p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),X9),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.UNION_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7))))))) & ! [X10] : ((p(s(bool,'const.bool.IN_2'(s(fun('type.num.num',fun(X6,bool)),X10),s(fun(fun('type.num.num',fun(X6,bool)),bool),'const.util_prob.FUNSET_2'(s(fun('type.num.num',bool),'const.pred_set.UNIV_0'),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))) & ! [X11,X12] : (~s('type.num.num',X11) = s('type.num.num',X12) => p(s(bool,'const.pred_set.DISJOINT_2'(s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X10),s('type.num.num',X11))),s(fun(X6,bool),happ(s(fun('type.num.num',fun(X6,bool)),X10),s('type.num.num',X12)))))))) => p(s(bool,'const.bool.IN_2'(s(fun(X6,bool),'const.pred_set.BIGUNION_1'(s(fun(fun(X6,bool),bool),'const.pred_set.IMAGE_2'(s(fun('type.num.num',fun(X6,bool)),X10),s(fun('type.num.num',bool),'const.pred_set.UNIV_0'))))),s(fun(fun(X6,bool),bool),'const.measure.subsets_1'(s('type.pair.prod'(fun(X6,bool),fun(fun(X6,bool),bool)),X7)))))))))),
  file('thm.measure.SIGMA_ALGEBRA_FN_DISJOINT','thm.measure.SIGMA_ALGEBRA_FN_DISJOINT')).
fof(f164,plain,(
  ~sP1(sK9,sK10) | ~p(s(bool,'const.measure.sigma_algebra_1'(s('type.pair.prod'(fun(sK9,bool),fun(fun(sK9,bool),bool)),sK10))))),
  inference(cnf_transformation,[],[f102])).
% SZS output end Proof for thm.measure
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 265
% Passive clauses: 1645
% Generated clauses: 13644
% Final active clauses: 96
% Final passive clauses: 428
% Input formulas: 31
% Initial clauses: 108
% 
% Selected by SInE selection: 31
% SInE iterations: 1
% 
% Duplicate literals: 373
% Trivial inequalities: 244
% Fw subsumption resolutions: 2608
% Fw demodulations: 1324
% Bw demodulations: 60
% 
% Simple tautologies: 685
% Forward subsumptions: 5885
% 
% Binary resolution: 2219
% Forward superposition: 2725
% Backward superposition: 2828
% Equality resolution: 29
% 
% Split clauses: 759
% Split components: 1803
% SAT solver clauses: 830
% SAT solver unit clauses: 19
% SAT solver binary clauses: 38
% 
% Sat splits: 759
% Sat splitting refutations: 70
% 
% Memory used [KB]: 4349
% Time elapsed: 0.227 s
% ------------------------------
% ------------------------------
% Success in time 11.924 s
