Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.list
% (113570)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 639
% Passive clauses: 8498
% Generated clauses: 18314
% Final active clauses: 634
% Final passive clauses: 7573
% Input formulas: 69
% Initial clauses: 224
% 
% Duplicate literals: 247
% Trivial inequalities: 131
% Fw subsumption resolutions: 1181
% Fw demodulations: 1135
% Bw demodulations: 1
% 
% Simple tautologies: 261
% Equational tautologies: 59
% Forward subsumptions: 6937
% 
% Binary resolution: 8253
% Factoring: 1162
% Forward superposition: 3597
% Backward superposition: 2293
% Self superposition: 43
% Equality factoring: 2
% Equality resolution: 46
% 
% Split clauses: 122
% Split components: 245
% SAT solver clauses: 133
% SAT solver unit clauses: 5
% SAT solver binary clauses: 26
% SAT solver learnt clauses: 1
% SAT solver learnt literals: 1
% 
% Sat splits: 122
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 132
% TWLsolver calls for satisfiability: 62
% 
% Memory used [KB]: 14967
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (113587)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 329
% Passive clauses: 2310
% Generated clauses: 10362
% Final active clauses: 326
% Final passive clauses: 1980
% Input formulas: 69
% Initial clauses: 224
% 
% Duplicate literals: 319
% Trivial inequalities: 13
% Fw subsumption resolutions: 554
% Fw demodulations: 566
% Bw demodulations: 1
% Condensations: 332
% Global subsumptions: 17
% 
% Simple tautologies: 188
% Equational tautologies: 9
% Forward subsumptions: 3521
% 
% Binary resolution: 2519
% Factoring: 610
% Forward superposition: 1884
% Backward superposition: 954
% Self superposition: 46
% Equality factoring: 2
% Equality resolution: 10
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 2771
% SAT solver unit clauses: 237
% SAT solver binary clauses: 112
% SAT solver learnt clauses: 35
% SAT solver learnt literals: 1
% 
% Sat splits: 16
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 2702
% TWLsolver calls for satisfiability: 1887
% 
% Memory used [KB]: 19701
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f4314,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f318,f554,f552,f317,f553,f551,f560,f4208,f4212,f4297,f4298,f4300,f4301])).
fof(f4301,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4222,f588])).
fof(f588,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f555,f332])).
fof(f332,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17] : (s(bool,'T_0') = s(bool,X17) <=> p(s(bool,X17)))),
  file('thm.list.EVERY2_THM|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f555,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f322,f319])).
fof(f319,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.EVERY2_THM|split|2',aHL_TRUTH)).
fof(f322,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f46])).
fof(f46,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.EVERY2_THM|split|2','thm.bool.NOT_CLAUSES|split|1')).
fof(f4222,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3)),
  inference(superposition,[],[f597,f4099])).
fof(f4099,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(X0,bool)),X1),s('type.list.list'(sK11),sK13),s('type.list.list'(X0),'const.list.NIL_0')))) ) | ($spl0 | $spl3)),
  inference(superposition,[],[f439,f2094])).
fof(f2094,plain,(
  s('type.list.list'(sK11),sK13) = s('type.list.list'(sK11),'const.list.CONS_2'(s(sK11,sK28(sK13,'const.list.NIL_0',sK12,sK10,sK11)),s('type.list.list'(sK11),sK30(sK13,'const.list.NIL_0',sK12,sK10,sK11)))) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f1472,f554,f444])).
fof(f444,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP1(X0,X1,X2,X3,X4) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) | s('type.list.list'(X0),X4) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK28(X4,X3,X2,X1,X0)),s('type.list.list'(X0),sK30(X4,X3,X2,X1,X0))))) )),
  inference(cnf_transformation,[],[f281])).
fof(f281,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP1(X0,X1,X2,X3,X4) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3)) | (s('type.list.list'(X0),X4) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,sK28(X4,X3,X2,X1,X0)),s('type.list.list'(X0),sK30(X4,X3,X2,X1,X0)))) & s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,sK29(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK31(X4,X3,X2,X1,X0)))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,sK28(X4,X3,X2,X1,X0)))),s(X1,sK29(X4,X3,X2,X1,X0))))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),sK30(X4,X3,X2,X1,X0)),s('type.list.list'(X1),sK31(X4,X3,X2,X1,X0))))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X4) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP1(X0,X1,X2,X3,X4)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK28,sK29,sK30,sK31])],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2,X3,X4] : ((~sP1(X0,X1,X2,X3,X4) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X4) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3)) | ? [X9,X10,X11,X12] : (s('type.list.list'(X0),X4) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X9),s('type.list.list'(X0),X11))) & s('type.list.list'(X1),X3) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X10),s('type.list.list'(X1),X12))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X9))),s(X1,X10)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X11),s('type.list.list'(X1),X12)))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X4) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X4) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X3) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP1(X0,X1,X2,X3,X4)))),
  inference(rectify,[],[f279])).
fof(f279,plain,(
  ! [X0,X1,X2,X4,X3] : ((~sP1(X0,X1,X2,X4,X3) | (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X3) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X4)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X4) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP1(X0,X1,X2,X4,X3)))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X0,X1,X2,X4,X3] : ((~sP1(X0,X1,X2,X4,X3) | ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8))))))) & (((s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X3) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X4)) & ! [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) != s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) | s('type.list.list'(X1),X4) != s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))) | sP1(X0,X1,X2,X4,X3)))),
  inference(nnf_transformation,[],[f200])).
fof(f200,plain,(
  ! [X0,X1,X2,X4,X3] : (sP1(X0,X1,X2,X4,X3) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f1472,plain,(
  sP1(sK11,sK10,sK12,'const.list.NIL_0',sK13) | $spl0),
  inference(subsumption_resolution,[],[f1469,f319])).
fof(f1469,plain,(
  ~p(s(bool,'T_0')) | sP1(sK11,sK10,sK12,'const.list.NIL_0',sK13) | $spl0),
  inference(superposition,[],[f454,f597])).
fof(f454,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | sP1(X0,X1,X2,X4,X3)) )),
  inference(cnf_transformation,[],[f282])).
fof(f282,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) | sP1(X0,X1,X2,X4,X3)) & (~sP1(X0,X1,X2,X4,X3) | p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4))))))),
  inference(nnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> sP1(X0,X1,X2,X4,X3))),
  inference(definition_folding,[],[f157,f200])).
fof(f157,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X3),s('type.list.list'(X1),X4)))) <=> ((s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X3) & s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X4)) | ? [X5,X6,X7,X8] : (s('type.list.list'(X0),X3) = s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X5),s('type.list.list'(X0),X7))) & s('type.list.list'(X1),X4) = s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X6),s('type.list.list'(X1),X8))) & p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X5))),s(X1,X6)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),X7),s('type.list.list'(X1),X8)))))))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10,X11] : ! [X12,X23,X24] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X10,fun(X11,bool)),X12),s('type.list.list'(X10),X23),s('type.list.list'(X11),X24)))) <=> ((s('type.list.list'(X10),'const.list.NIL_0') = s('type.list.list'(X10),X23) & s('type.list.list'(X11),'const.list.NIL_0') = s('type.list.list'(X11),X24)) | ? [X25,X26,X15,X16] : (s('type.list.list'(X10),X23) = s('type.list.list'(X10),'const.list.CONS_2'(s(X10,X25),s('type.list.list'(X10),X15))) & s('type.list.list'(X11),X24) = s('type.list.list'(X11),'const.list.CONS_2'(s(X11,X26),s('type.list.list'(X11),X16))) & p(s(bool,happ(s(fun(X11,bool),happ(s(fun(X10,fun(X11,bool)),X12),s(X10,X25))),s(X11,X26)))) & p(s(bool,'const.list.LIST_REL_3'(s(fun(X10,fun(X11,bool)),X12),s('type.list.list'(X10),X15),s('type.list.list'(X11),X16)))))))),
  file('thm.list.EVERY2_THM|split|2','thm.list.LIST_REL_cases')).
fof(f439,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2,X3,X4] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))),s('type.list.list'(X1),'const.list.NIL_0')))),
  inference(flattening,[],[f150])).
fof(f150,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X4),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X3),s('type.list.list'(X0),X2))),s('type.list.list'(X1),'const.list.NIL_0')))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X10,X11] : ! [X13,X14,X12] : s(bool,'F_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X10,fun(X11,bool)),X12),s('type.list.list'(X10),'const.list.CONS_2'(s(X10,X14),s('type.list.list'(X10),X13))),s('type.list.list'(X11),'const.list.NIL_0')))),
  file('thm.list.EVERY2_THM|split|2','thm.list.LIST_REL_def|split|1')).
fof(f597,plain,(
  s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f551,f333])).
fof(f333,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f217])).
fof(f4300,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4221,f588])).
fof(f4221,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl0 | $spl3)),
  inference(superposition,[],[f4099,f597])).
fof(f4298,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4207,f555])).
fof(f4207,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f1323,f4099,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f299])).
fof(f299,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f206])).
fof(f206,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f163,f205,f204,f203,f202])).
fof(f202,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f203,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f204,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f205,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f163,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f70])).
fof(f70,axiom,(
  ! [X40,X58,X59] : ((p(s(bool,X59)) <=> s(bool,X40) = s(bool,X58)) <=> ((p(s(bool,X59)) | p(s(bool,X58)) | p(s(bool,X40))) & (p(s(bool,X59)) | ~p(s(bool,X40)) | ~p(s(bool,X58))) & (p(s(bool,X58)) | ~p(s(bool,X40)) | ~p(s(bool,X59))) & (p(s(bool,X40)) | ~p(s(bool,X58)) | ~p(s(bool,X59)))))),
  file('thm.list.EVERY2_THM|split|2','thm.sat.sat_2|49|_prove_ltb')).
fof(f1323,plain,(
  sP5('F_0','const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')),'F_0') | $spl0),
  inference(unit_resulting_resolution,[],[f579,f576,f565,f555,f469])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f290])).
fof(f290,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f289])).
fof(f289,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f288])).
fof(f288,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f205])).
fof(f565,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f555,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f296])).
fof(f296,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f295])).
fof(f295,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f294])).
fof(f294,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f203])).
fof(f576,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')))) ) | $spl0),
  inference(unit_resulting_resolution,[],[f551,f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f293])).
fof(f293,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f292])).
fof(f292,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f291])).
fof(f291,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f204])).
fof(f579,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')),X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f551,f480])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f298])).
fof(f298,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f297])).
fof(f297,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f202])).
fof(f4297,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f4211,f555])).
fof(f4211,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f1266,f4099,f485])).
fof(f1266,plain,(
  sP5('F_0','F_0','const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0'))) | $spl0),
  inference(unit_resulting_resolution,[],[f580,f563,f566,f555,f468])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f290])).
fof(f566,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f555,f477])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f296])).
fof(f563,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f555,f473])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f293])).
fof(f580,plain,(
  ( ! [X0,X1] : (sP2('const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')),X0,X1)) ) | $spl0),
  inference(unit_resulting_resolution,[],[f551,f481])).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f298])).
fof(f4212,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f555,f1266,f4099,f485])).
fof(f4208,plain,(
  $false | ($spl0 | $spl3)),
  inference(unit_resulting_resolution,[],[f555,f1323,f4099,f485])).
fof(f560,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f559,f319])).
fof(f559,plain,(
  ~p(s(bool,'T_0')) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f558,f421])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0')))) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0')))),
  inference(flattening,[],[f130])).
fof(f130,plain,(
  ! [X0,X1] : ! [X2] : s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X0,fun(X1,bool)),X2),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X1),'const.list.NIL_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ! [X12] : s(bool,'T_0') = s(bool,'const.list.LIST_REL_3'(s(fun(X10,fun(X11,bool)),X12),s('type.list.list'(X10),'const.list.NIL_0'),s('type.list.list'(X11),'const.list.NIL_0')))),
  file('thm.list.EVERY2_THM|split|2','thm.list.LIST_REL_def|split|0')).
fof(f558,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),'const.list.NIL_0'),s('type.list.list'(sK10),'const.list.NIL_0')))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f552,f553])).
fof(f551,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')))) | $spl0),
  inference(cnf_transformation,[],[f551_D])).
fof(f551_D,plain,(
  p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f553,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),sK13) | $spl2),
  inference(cnf_transformation,[],[f553_D])).
fof(f553_D,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),sK13) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f317,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),sK13) | p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0'))))),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  (p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')))) | s('type.list.list'(sK11),'const.list.NIL_0') = s('type.list.list'(sK11),sK13)) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')))) | s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f213])).
fof(f213,plain,(
  ? [X0,X1,X2,X3] : ((p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) | s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3)) & (~p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) | s('type.list.list'(X1),'const.list.NIL_0') != s('type.list.list'(X1),X3)))),
  inference(nnf_transformation,[],[f176])).
fof(f176,plain,(
  ? [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) <~> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X1,fun(X0,bool)),X2),s('type.list.list'(X1),X3),s('type.list.list'(X0),'const.list.NIL_0')))) <=> s('type.list.list'(X1),'const.list.NIL_0') = s('type.list.list'(X1),X3))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X7,fun(X6,bool)),X8),s('type.list.list'(X7),X9),s('type.list.list'(X6),'const.list.NIL_0')))) <=> s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X9))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (p(s(bool,'const.list.LIST_REL_3'(s(fun(X7,fun(X6,bool)),X8),s('type.list.list'(X7),X9),s('type.list.list'(X6),'const.list.NIL_0')))) <=> s('type.list.list'(X7),'const.list.NIL_0') = s('type.list.list'(X7),X9))),
  file('thm.list.EVERY2_THM|split|2','thm.list.EVERY2_THM|split|2')).
fof(f552,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')))) | $spl1),
  inference(cnf_transformation,[],[f552_D])).
fof(f552_D,plain,(
  ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0')))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f554,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),sK13) | $spl3),
  inference(cnf_transformation,[],[f554_D])).
fof(f554_D,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),sK13) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f318,plain,(
  s('type.list.list'(sK11),'const.list.NIL_0') != s('type.list.list'(sK11),sK13) | ~p(s(bool,'const.list.LIST_REL_3'(s(fun(sK11,fun(sK10,bool)),sK12),s('type.list.list'(sK11),sK13),s('type.list.list'(sK10),'const.list.NIL_0'))))),
  inference(cnf_transformation,[],[f214])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 303
% Passive clauses: 911
% Generated clauses: 3995
% Final active clauses: 301
% Final passive clauses: 605
% Input formulas: 69
% Initial clauses: 224
% 
% Duplicate literals: 5
% Trivial inequalities: 2
% Fw subsumption resolutions: 49
% Fw demodulations: 485
% Bw demodulations: 4
% Condensations: 7
% 
% Simple tautologies: 59
% Equational tautologies: 20
% Forward subsumptions: 2440
% Fw demodulations to eq. taut.: 8
% 
% Unit resulting resolution: 936
% Factoring: 1
% Forward superposition: 1209
% Backward superposition: 1027
% Self superposition: 34
% Equality factoring: 2
% Equality resolution: 6
% 
% Split clauses: 2
% Split components: 4
% SAT solver clauses: 924
% SAT solver unit clauses: 184
% SAT solver binary clauses: 463
% 
% Sat splits: 2
% Sat splitting refutations: 7
% 
% Memory used [KB]: 8571
% Time elapsed: 0.206 s
% ------------------------------
% ------------------------------
% Success in time 1.526 s
