Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 300 next slice time: 5
dis+11_7_3 on thm.pred_set
% (83791)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 155
% Passive clauses: 2198
% Generated clauses: 18987
% Final active clauses: 146
% Final passive clauses: 1877
% Input formulas: 22
% Initial clauses: 53
% 
% Duplicate literals: 963
% Trivial inequalities: 312
% Fw subsumption resolutions: 3185
% Fw demodulations: 1820
% Bw demodulations: 6
% 
% Simple tautologies: 467
% Equational tautologies: 55
% Forward subsumptions: 7959
% 
% Binary resolution: 645
% Forward superposition: 4695
% Backward superposition: 6707
% Self superposition: 190
% Equality factoring: 157
% Equality resolution: 30
% 
% Split clauses: 1300
% Split components: 2704
% SAT solver clauses: 2020
% SAT solver unit clauses: 39
% SAT solver binary clauses: 212
% 
% Sat splits: 1300
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2016
% TWLsolver calls for satisfiability: 220
% 
% Memory used [KB]: 9338
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f5295,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f5271,f186,f190,f90,f164,f98,f170,f249,f252,f250,f96,f169,f97,f171,f271,f95,f168,f305,f205,f210,f92,f160,f102,f172,f1318,f1329,f124,f182,f1617,f1619,f93,f166,f188,f192,f91,f162,f704,f711,f121,f181,f1772,f1774,f884,f894,f120,f178,f122,f180])).
fof(f180,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl24),
  inference(cnf_transformation,[],[f180_D])).
fof(f180_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f122,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & (~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0,X1] : ((~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))))) & ((~p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) | p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(nnf_transformation,[],[f42])).
fof(f42,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(flattening,[],[f41])).
fof(f41,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) <=> (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X17] : ! [X18] : (p(s(bool,'const.relation.StrongOrder_1'(s(fun(X17,fun(X17,bool)),X18)))) <=> (p(s(bool,'const.relation.irreflexive_1'(s(fun(X17,fun(X17,bool)),X18)))) & p(s(bool,'const.relation.transitive_1'(s(fun(X17,fun(X17,bool)),X18))))))),
  file('thm.pred_set.FINITE_StrongOrder_WF','thm.relation.StrongOrder')).
fof(f178,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))) ) | $spl20),
  inference(cnf_transformation,[],[f178_D])).
fof(f178_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f120,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))) | ~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (~p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ! [X0,X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'const.relation.transitive_1'(s(fun(X0,fun(X0,bool)),X1)))) => s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X6] : ! [X7] : (p(s(bool,'const.relation.transitive_1'(s(fun(X6,fun(X6,bool)),X7)))) => s(fun(X6,fun(X6,bool)),X7) = s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),X7))))),
  file('thm.pred_set.FINITE_StrongOrder_WF','thm.relation.transitive_TC_identity')).
fof(f894,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))) ) | $spl114),
  inference(cnf_transformation,[],[f894_D])).
fof(f894_D,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1)))) | s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f884,plain,(
  ( ! [X0,X1] : (s(fun(X0,fun(X0,bool)),X1) = s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),X1))) | ~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | ($spl20 | $spl24)),
  inference(resolution,[],[f178,f180])).
fof(f1774,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,fun(sK0,bool)),'const.relation.TC_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | $spl208),
  inference(cnf_transformation,[],[f1774_D])).
fof(f1774_D,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,fun(sK0,bool)),'const.relation.TC_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f1772,plain,(
  s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,fun(sK0,bool)),'const.relation.TC_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | ($spl6 | $spl38 | $spl114)),
  inference(subsumption_resolution,[],[f1761,f166])).
fof(f1761,plain,(
  ~p(s(bool,'T_0')) | s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))) = s(fun(sK0,fun(sK0,bool)),'const.relation.TC_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | ($spl38 | $spl114)),
  inference(superposition,[],[f894,f192])).
fof(f181,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) ) | $spl26),
  inference(cnf_transformation,[],[f181_D])).
fof(f181_D,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f121,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),X1)))) | ~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),X1))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f711,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X3,fun(X3,bool)),X4)))) | s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(X3,fun(X3,bool)),X4)))) ) | $spl86),
  inference(cnf_transformation,[],[f711_D])).
fof(f711_D,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X3,fun(X3,bool)),X4)))) | s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(X3,fun(X3,bool)),X4)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f704,plain,(
  ( ! [X4,X3] : (~p(s(bool,'const.relation.StrongOrder_1'(s(fun(X3,fun(X3,bool)),X4)))) | s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(X3,fun(X3,bool)),X4)))) ) | ($spl12 | $spl26)),
  inference(resolution,[],[f181,f170])).
fof(f162,plain,(
  p(s(bool,'const.relation.StrongOrder_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2)))))) | $spl2),
  inference(cnf_transformation,[],[f162_D])).
fof(f162_D,plain,(
  p(s(bool,'const.relation.StrongOrder_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f91,plain,(
  p(s(bool,'const.relation.StrongOrder_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))))),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2)))) & p(s(bool,'const.relation.StrongOrder_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2)))))) & ~p(s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f54])).
fof(f54,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2)))))) & ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ? [X0,X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))))) & ~p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2)))))))),
  inference(ennf_transformation,[],[f25])).
fof(f25,plain,(
  ~! [X0,X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))))) => p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2)))))))),
  inference(flattening,[],[f24])).
fof(f24,plain,(
  ~! [X0] : ! [X1,X2] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) & p(s(bool,'const.relation.StrongOrder_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))))) => p(s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2)))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) & p(s(bool,'const.relation.StrongOrder_1'(s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),X8))))))) => p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),X8)))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : ((p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) & p(s(bool,'const.relation.StrongOrder_1'(s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),X8))))))) => p(s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),X8)))))))),
  file('thm.pred_set.FINITE_StrongOrder_WF','thm.pred_set.FINITE_StrongOrder_WF')).
fof(f192,plain,(
  s(bool,'T_0') = s(bool,'const.relation.StrongOrder_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | $spl38),
  inference(cnf_transformation,[],[f192_D])).
fof(f192_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.StrongOrder_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f188,plain,(
  s(bool,'T_0') = s(bool,'const.relation.StrongOrder_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | ($spl2 | $spl12)),
  inference(resolution,[],[f170,f162])).
fof(f166,plain,(
  p(s(bool,'T_0')) | $spl6),
  inference(cnf_transformation,[],[f166_D])).
fof(f166_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f93,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.FINITE_StrongOrder_WF',aHL_TRUTH)).
fof(f1619,plain,(
  s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | $spl194),
  inference(cnf_transformation,[],[f1619_D])).
fof(f1619_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f1617,plain,(
  s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | ($spl6 | $spl38 | $spl86)),
  inference(subsumption_resolution,[],[f1606,f166])).
fof(f1606,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,'const.relation.irreflexive_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | ($spl38 | $spl86)),
  inference(superposition,[],[f711,f192])).
fof(f182,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2)))))))) ) | $spl28),
  inference(cnf_transformation,[],[f182_D])).
fof(f182_D,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2)))))))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f124,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))))) | ~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) | s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))))))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))))))),
  inference(flattening,[],[f43])).
fof(f43,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X0,bool),X2)))) => s(bool,'const.relation.WF_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X0,fun(X0,bool)),'const.relation.TC_1'(s(fun(X0,fun(X0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X0,fun(X0,bool)),X1),s(fun(X0,bool),X2))))))))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.pred_set.FINITE_1'(s(fun(X6,bool),X8)))) => s(bool,'const.relation.WF_1'(s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),X8))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X6,fun(X6,bool)),'const.relation.TC_1'(s(fun(X6,fun(X6,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X6,fun(X6,bool)),X7),s(fun(X6,bool),X8))))))))),
  file('thm.pred_set.FINITE_StrongOrder_WF','thm.pred_set.FINITE_WF_noloops')).
fof(f1329,plain,(
  ( ! [X6,X4,X5] : (s(bool,'const.relation.WF_1'(s(fun(X4,fun(X4,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X4,fun(X4,bool)),X5),s(fun(X4,bool),X6))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X4,fun(X4,bool)),'const.relation.TC_1'(s(fun(X4,fun(X4,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X4,fun(X4,bool)),X5),s(fun(X4,bool),X6))))))) | s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(X4,bool),X6)))) ) | $spl162),
  inference(cnf_transformation,[],[f1329_D])).
fof(f1329_D,plain,(
  ( ! [X6,X4,X5] : (s(bool,'const.relation.WF_1'(s(fun(X4,fun(X4,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X4,fun(X4,bool)),X5),s(fun(X4,bool),X6))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X4,fun(X4,bool)),'const.relation.TC_1'(s(fun(X4,fun(X4,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X4,fun(X4,bool)),X5),s(fun(X4,bool),X6))))))) | s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(X4,bool),X6)))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f1318,plain,(
  ( ! [X6,X4,X5] : (s(bool,'const.relation.WF_1'(s(fun(X4,fun(X4,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X4,fun(X4,bool)),X5),s(fun(X4,bool),X6))))) = s(bool,'const.relation.irreflexive_1'(s(fun(X4,fun(X4,bool)),'const.relation.TC_1'(s(fun(X4,fun(X4,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(X4,fun(X4,bool)),X5),s(fun(X4,bool),X6))))))) | s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(X4,bool),X6)))) ) | ($spl16 | $spl28)),
  inference(resolution,[],[f182,f172])).
fof(f172,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl16),
  inference(cnf_transformation,[],[f172_D])).
fof(f172_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f102,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f30])).
fof(f30,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f29])).
fof(f29,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.pred_set.FINITE_StrongOrder_WF','thm.bool.EQ_CLAUSES|split|2')).
fof(f160,plain,(
  ~p(s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2)))))) | $spl1),
  inference(cnf_transformation,[],[f160_D])).
fof(f160_D,plain,(
  ~p(s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2)))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f92,plain,(
  ~p(s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))))),
  inference(cnf_transformation,[],[f67])).
fof(f210,plain,(
  s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | $spl40),
  inference(cnf_transformation,[],[f210_D])).
fof(f210_D,plain,(
  s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f205,plain,(
  s(bool,'F_0') = s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | ($spl1 | $spl16)),
  inference(resolution,[],[f172,f160])).
fof(f305,plain,(
  $false | ($spl40 | $spl49 | $spl60)),
  inference(subsumption_resolution,[],[f304,f252])).
fof(f304,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl40 | $spl60)),
  inference(backward_demodulation,[],[f303,f210])).
fof(f303,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | $spl60),
  inference(cnf_transformation,[],[f303_D])).
fof(f303_D,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f168,plain,(
  ~p(s(bool,'F_0')) | $spl9),
  inference(cnf_transformation,[],[f168_D])).
fof(f168_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f95,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f26])).
fof(f26,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.FINITE_StrongOrder_WF',aHL_FALSITY)).
fof(f271,plain,(
  $false | ($spl9 | $spl46)),
  inference(resolution,[],[f250,f168])).
fof(f171,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) | $spl14),
  inference(cnf_transformation,[],[f171_D])).
fof(f171_D,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f97,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f27])).
fof(f27,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.pred_set.FINITE_StrongOrder_WF','thm.bool.EQ_CLAUSES|split|0')).
fof(f169,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) | $spl10),
  inference(cnf_transformation,[],[f169_D])).
fof(f169_D,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f96,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.FINITE_StrongOrder_WF',aHL_BOOL_CASES)).
fof(f250,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) | $spl46),
  inference(cnf_transformation,[],[f250_D])).
fof(f250_D,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f252,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl49),
  inference(cnf_transformation,[],[f252_D])).
fof(f252_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f249,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X2))) ) | ($spl10 | $spl14)),
  inference(subsumption_resolution,[],[f236,f171])).
fof(f236,plain,(
  ( ! [X2] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X2)) | s(bool,'T_0') = s(bool,X2)) ) | ($spl10 | $spl14)),
  inference(superposition,[],[f171,f169])).
fof(f170,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl12),
  inference(cnf_transformation,[],[f170_D])).
fof(f170_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f98,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f68])).
fof(f164,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2)))) | $spl4),
  inference(cnf_transformation,[],[f164_D])).
fof(f164_D,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2)))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f90,plain,(
  p(s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f67])).
fof(f190,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))) | $spl36),
  inference(cnf_transformation,[],[f190_D])).
fof(f190_D,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f186,plain,(
  s(bool,'T_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))) | ($spl4 | $spl12)),
  inference(resolution,[],[f170,f164])).
fof(f5271,plain,(
  $false | ($spl36 | $spl49 | $spl61 | $spl162 | $spl194 | $spl208)),
  inference(subsumption_resolution,[],[f5270,f252])).
fof(f5270,plain,(
  s(bool,'T_0') = s(bool,'F_0') | ($spl36 | $spl61 | $spl162 | $spl194 | $spl208)),
  inference(forward_demodulation,[],[f5269,f190])).
fof(f5269,plain,(
  s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))) | ($spl61 | $spl162 | $spl194 | $spl208)),
  inference(subsumption_resolution,[],[f5268,f302])).
fof(f302,plain,(
  s(bool,'T_0') != s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | $spl61),
  inference(cnf_transformation,[],[f302_D])).
fof(f302_D,plain,(
  s(bool,'T_0') != s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f5268,plain,(
  s(bool,'T_0') = s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))) | ($spl162 | $spl194 | $spl208)),
  inference(forward_demodulation,[],[f5249,f1619])).
fof(f5249,plain,(
  s(bool,'const.relation.WF_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) = s(bool,'const.relation.irreflexive_1'(s(fun(sK0,fun(sK0,bool)),'const.pred_set.REL_RESTRICT_2'(s(fun(sK0,fun(sK0,bool)),sK1),s(fun(sK0,bool),sK2))))) | s(bool,'F_0') = s(bool,'const.pred_set.FINITE_1'(s(fun(sK0,bool),sK2))) | ($spl162 | $spl208)),
  inference(superposition,[],[f1329,f1774])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 80
% Passive clauses: 301
% Generated clauses: 5150
% Final active clauses: 69
% Final passive clauses: 200
% Input formulas: 22
% Initial clauses: 53
% 
% Duplicate literals: 405
% Trivial inequalities: 39
% Fw subsumption resolutions: 422
% Fw demodulations: 707
% Bw demodulations: 26
% Condensations: 338
% Global subsumptions: 4
% 
% Simple tautologies: 104
% Equational tautologies: 19
% Forward subsumptions: 2359
% 
% Binary resolution: 102
% Forward superposition: 1659
% Backward superposition: 991
% Self superposition: 51
% Equality factoring: 32
% Equality resolution: 7
% 
% Split clauses: 29
% Split components: 67
% SAT solver clauses: 501
% SAT solver unit clauses: 59
% SAT solver binary clauses: 28
% SAT solver learnt clauses: 9
% 
% Sat splits: 29
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 484
% TWLsolver calls for satisfiability: 259
% 
% Memory used [KB]: 4093
% Time elapsed: 0.151 s
% ------------------------------
% ------------------------------
% Success in time 0.765 s
