Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.gcd
% (82712)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 666
% Passive clauses: 10230
% Generated clauses: 18704
% Final active clauses: 613
% Final passive clauses: 9012
% Input formulas: 44
% Initial clauses: 200
% 
% Duplicate literals: 286
% Trivial inequalities: 69
% Fw subsumption resolutions: 1605
% Fw demodulations: 145
% Bw demodulations: 52
% 
% Simple tautologies: 228
% Equational tautologies: 59
% Forward subsumptions: 5855
% 
% Binary resolution: 11820
% Factoring: 1012
% Forward superposition: 1921
% Backward superposition: 1332
% Self superposition: 24
% Equality factoring: 26
% Equality resolution: 27
% 
% Split clauses: 323
% Split components: 652
% SAT solver clauses: 696
% SAT solver unit clauses: 16
% SAT solver binary clauses: 199
% 
% Sat splits: 323
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 694
% TWLsolver calls for satisfiability: 421
% 
% Memory used [KB]: 9210
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.gcd
% (82731)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 392
% Passive clauses: 2860
% Generated clauses: 11749
% Final active clauses: 382
% Final passive clauses: 2397
% Input formulas: 44
% Initial clauses: 200
% 
% Duplicate literals: 196
% Trivial inequalities: 18
% Fw subsumption resolutions: 888
% Fw demodulations: 378
% Bw demodulations: 72
% Condensations: 355
% Global subsumptions: 13
% 
% Simple tautologies: 258
% Equational tautologies: 17
% Forward subsumptions: 3302
% Bw demodulations to eq. taut.: 6
% 
% Binary resolution: 4005
% Factoring: 214
% Forward superposition: 1074
% Backward superposition: 1231
% Self superposition: 28
% Equality factoring: 28
% Equality resolution: 7
% 
% Split clauses: 54
% Split components: 109
% SAT solver clauses: 3652
% SAT solver unit clauses: 184
% SAT solver binary clauses: 85
% SAT solver learnt clauses: 26
% SAT solver learnt literals: 10
% 
% Sat splits: 54
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 3600
% TWLsolver calls for satisfiability: 1859
% 
% Memory used [KB]: 13560
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.gcd
% (82746)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 1616
% Passive clauses: 6329
% Generated clauses: 111538
% Final active clauses: 342
% Final passive clauses: 637
% Input formulas: 44
% Initial clauses: 200
% 
% Duplicate literals: 8267
% Trivial inequalities: 218
% Fw subsumption resolutions: 5773
% Fw demodulations: 11305
% Bw demodulations: 2613
% Condensations: 13838
% Global subsumptions: 287
% 
% Simple tautologies: 328
% Equational tautologies: 1205
% Forward subsumptions: 62655
% Fw demodulations to eq. taut.: 4
% Bw demodulations to eq. taut.: 397
% 
% Unit resulting resolution: 12679
% Factoring: 13
% Forward superposition: 31602
% Backward superposition: 19469
% Self superposition: 331
% Equality factoring: 228
% Equality resolution: 10
% 
% Split clauses: 517
% Split components: 1047
% SAT solver clauses: 8547
% SAT solver unit clauses: 416
% SAT solver binary clauses: 1424
% 
% Sat splits: 517
% Sat splitting refutations: 894
% 
% Memory used [KB]: 47845
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.gcd
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.gcd
% SZS output start Proof for thm.gcd
fof(f48876,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f407,f408,f409,f410,f411,f412,f413,f414,f415,f416,f429,f430,f431,f701,f705,f694,f695,f842,f1181,f1183,f1185,f1187,f1189,f1191,f1416,f1440,f1441,f419,f1442,f1443,f420,f1444,f1445,f418,f1446,f1447,f417,f1448,f424,f426,f1449,f1640,f1641,f703,f1691,f1693,f702,f1694,f1943,f1485,f1484,f1483,f533,f2325,f2556,f535,f2558,f537,f2560,f539,f531,f2801,f3009,f3008,f2993,f2992,f3072,f3073,f3074,f3075,f3076,f3077,f3089,f3100,f3135,f3203,f5435,f5446,f5449,f5439,f5455,f5459,f5475,f5486,f5476,f5488,f7341,f7346,f7343,f7347,f7430,f7435,f7432,f7436,f20679,f20712,f20716,f21031,f422,f428,f22426,f22427,f22428,f22429,f22430,f22539,f22545,f22685,f22691,f22793,f22810,f22794,f22795,f22859,f22861,f22862,f22860,f4077,f4078,f4079,f22903,f22904,f1692,f22905,f22906,f707,f22718,f22621,f22516,f23041,f23034,f23105,f23106,f23107,f23108,f23213,f23212,f23200,f23199,f23286,f23287,f23288,f23289,f23290,f23291,f23334,f23386,f704,f23424,f23453,f23490,f3090,f24036,f24059,f22813,f22812,f26865,f26872,f27020,f27027,f3136,f48867,f48872])).
fof(f48872,plain,(
  $false | ($spl0 | $spl30 | $spl32 | $spl35)),
  inference(subsumption_resolution,[],[f48871,f392])).
fof(f392,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.gcd.IS_GCD_SYM',aHL_FALSITY)).
fof(f48871,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl30 | $spl32 | $spl35)),
  inference(forward_demodulation,[],[f48870,f21743])).
fof(f21743,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | $spl35),
  inference(resolution,[],[f1191,f251])).
fof(f251,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f52])).
fof(f52,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.EQ_CLAUSES|split|2')).
fof(f48870,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | ($spl0 | $spl30 | $spl32)),
  inference(subsumption_resolution,[],[f48869,f397])).
fof(f397,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.gcd.IS_GCD_SYM',aHL_TRUTH)).
fof(f48869,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | ($spl0 | $spl30 | $spl32)),
  inference(forward_demodulation,[],[f48868,f21484])).
fof(f21484,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12))) | $spl30),
  inference(resolution,[],[f1183,f255])).
fof(f255,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.EQ_CLAUSES|split|0')).
fof(f48868,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | ($spl0 | $spl32)),
  inference(subsumption_resolution,[],[f48836,f373])).
fof(f373,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f211])).
fof(f211,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f210])).
fof(f210,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f112])).
fof(f112,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X12] : ((p(s(bool,'F_0')) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.IMP_CLAUSES|split|2')).
fof(f48836,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | ($spl0 | $spl32)),
  inference(superposition,[],[f20733,f21629])).
fof(f21629,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13))) | $spl32),
  inference(resolution,[],[f1187,f255])).
fof(f20733,plain,(
  ( ! [X0] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK12)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK14))))) ) | $spl0),
  inference(resolution,[],[f418,f261])).
fof(f261,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X2)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X1)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X0)))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f162,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) | (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ! [X4] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X1)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X2))))))) & (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | (p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(X2,X1,X0)),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(X2,X1,X0)),s('type.num.num',X1)))) & ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(X2,X1,X0)),s('type.num.num',X2))))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) | (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ! [X4] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X1)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X4),s('type.num.num',X2))))))) & (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ? [X3] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1)))) & ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2))))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))))))),
  inference(rectify,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) | (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ! [X3] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2))))))) & (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ? [X3] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1)))) & ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2))))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))))))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) | (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ! [X3] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2))))))) & ((~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ? [X3] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1)))) & ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2)))))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))))))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) <=> (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ! [X3] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2)))))))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) <=> (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ! [X3] : ((~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1))))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2)))))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) <=> (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) & ! [X3] : ((p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X0)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X1))))) => p(s(bool,'const.divides.divides_2'(s('type.num.num',X3),s('type.num.num',X2)))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X6),s('type.num.num',X7),s('type.num.num',X8)))) <=> (p(s(bool,'const.divides.divides_2'(s('type.num.num',X8),s('type.num.num',X6)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X8),s('type.num.num',X7)))) & ! [X21] : ((p(s(bool,'const.divides.divides_2'(s('type.num.num',X21),s('type.num.num',X6)))) & p(s(bool,'const.divides.divides_2'(s('type.num.num',X21),s('type.num.num',X7))))) => p(s(bool,'const.divides.divides_2'(s('type.num.num',X21),s('type.num.num',X8)))))))),
  file('thm.gcd.IS_GCD_SYM','thm.gcd.is_gcd_def')).
fof(f48867,plain,(
  $false | ($spl0 | $spl18 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f48866,f392])).
fof(f48866,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl18 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f48865,f3471])).
fof(f3471,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) | $spl23),
  inference(resolution,[],[f539,f251])).
fof(f48865,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ($spl0 | $spl18 | $spl20)),
  inference(subsumption_resolution,[],[f48864,f397])).
fof(f48864,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ($spl0 | $spl18 | $spl20)),
  inference(forward_demodulation,[],[f48863,f3362])).
fof(f3362,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12))) | $spl20),
  inference(resolution,[],[f537,f255])).
fof(f48863,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ($spl0 | $spl18)),
  inference(subsumption_resolution,[],[f48835,f373])).
fof(f48835,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ($spl0 | $spl18)),
  inference(superposition,[],[f20733,f3246])).
fof(f3246,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13))) | $spl18),
  inference(resolution,[],[f535,f255])).
fof(f3136,plain,(
  ( ! [X4] : (~p(s(bool,X4))) ) | $spl40),
  inference(cnf_transformation,[],[f3136_D])).
fof(f3136_D,plain,(
  ( ! [X4] : (~p(s(bool,X4))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f27027,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X14),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK13,sK12))))) ) | $spl76),
  inference(cnf_transformation,[],[f27027_D])).
fof(f27027_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X14),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK13,sK12))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f27020,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X14),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK13,sK12)))) | ~p(s(bool,X15))) ) | ($spl6 | $spl12 | $spl35)),
  inference(resolution,[],[f25981,f23457])).
fof(f23457,plain,(
  ( ! [X6,X7] : (~sP2(X6,X7,'F_0') | s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X6))) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f23456,f237])).
fof(f237,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f151])).
fof(f151,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f23456,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | ~sP2(X6,X7,'F_0') | ~sP0(X7,'F_0',X6)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f23455,f233])).
fof(f233,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f149])).
fof(f149,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f148])).
fof(f148,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f23455,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | ~sP2(X6,X7,'F_0') | ~sP1('F_0',X7,X6) | ~sP0(X7,'F_0',X6)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f23439,f254])).
fof(f254,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f157])).
fof(f23439,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | s(bool,'T_0') != s(bool,X7) | ~p(s(bool,X7)) | ~sP2(X6,X7,'F_0') | ~sP1('F_0',X7,X6) | ~sP0(X7,'F_0',X6)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f22930,f225])).
fof(f225,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f143])).
fof(f143,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f22930,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21928,f242])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f48,f126,f125,f124,f123])).
fof(f125,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f48,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X22,X23,X24] : ((p(s(bool,X24)) <=> s(bool,X22) = s(bool,X23)) <=> ((p(s(bool,X24)) | p(s(bool,X23)) | p(s(bool,X22))) & (p(s(bool,X24)) | ~p(s(bool,X22)) | ~p(s(bool,X23))) & (p(s(bool,X23)) | ~p(s(bool,X22)) | ~p(s(bool,X24))) & (p(s(bool,X22)) | ~p(s(bool,X23)) | ~p(s(bool,X24)))))),
  file('thm.gcd.IS_GCD_SYM','thm.sat.sat_2|49|_prove_ltb')).
fof(f21928,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl6 | $spl12)),
  inference(backward_demodulation,[],[f428,f962])).
fof(f962,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | $spl6),
  inference(backward_demodulation,[],[f421,f220])).
fof(f220,plain,(
  s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12,sK13,sK14])],[f99])).
fof(f99,plain,(
  ? [X0,X1,X2] : s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',X2)))),
  inference(ennf_transformation,[],[f46])).
fof(f46,plain,(
  ~! [X0,X1,X2] : s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))) = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X1),s('type.num.num',X0),s('type.num.num',X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X6),s('type.num.num',X7),s('type.num.num',X8))) = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X7),s('type.num.num',X6),s('type.num.num',X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X6),s('type.num.num',X7),s('type.num.num',X8))) = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X7),s('type.num.num',X6),s('type.num.num',X8)))),
  file('thm.gcd.IS_GCD_SYM','thm.gcd.IS_GCD_SYM')).
fof(f421,plain,(
  s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | $spl6),
  inference(cnf_transformation,[],[f421_D])).
fof(f421_D,plain,(
  s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f25981,plain,(
  ( ! [X17,X15,X16] : (sP2(X15,'const.gcd.is_gcd_3'(s('type.num.num',X16),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK13,sK12))),X17)) ) | $spl35),
  inference(resolution,[],[f21725,f230])).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f146])).
fof(f146,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f125])).
fof(f21725,plain,(
  ( ! [X1] : (~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK13,sK12)))))) ) | $spl35),
  inference(resolution,[],[f1191,f260])).
fof(f260,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f26872,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X14),s('type.num.num',sK15(sK14,sK13,sK12))))) ) | $spl74),
  inference(cnf_transformation,[],[f26872_D])).
fof(f26872_D,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X14),s('type.num.num',sK15(sK14,sK13,sK12))))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f26865,plain,(
  ( ! [X14,X15] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X14),s('type.num.num',sK15(sK14,sK13,sK12)))) | ~p(s(bool,X15))) ) | ($spl6 | $spl12 | $spl35)),
  inference(resolution,[],[f25821,f23457])).
fof(f25821,plain,(
  ( ! [X17,X15,X16] : (sP2(X15,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X16),s('type.num.num',sK15(sK14,sK13,sK12))),X17)) ) | $spl35),
  inference(resolution,[],[f21724,f230])).
fof(f21724,plain,(
  ( ! [X0] : (~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',sK15(sK14,sK13,sK12)))))) ) | $spl35),
  inference(resolution,[],[f1191,f259])).
fof(f259,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0)))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f22812,plain,(
  s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | $spl67),
  inference(cnf_transformation,[],[f22812_D])).
fof(f22812_D,plain,(
  s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f22813,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) ) | $spl68),
  inference(cnf_transformation,[],[f22813_D])).
fof(f22813_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f24059,plain,(
  ( ! [X14] : (sP3(X14,'F_0','const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) ) | ($spl6 | $spl12 | $spl35)),
  inference(resolution,[],[f23529,f21731])).
fof(f21731,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)),X13)) ) | $spl35),
  inference(resolution,[],[f1191,f230])).
fof(f23529,plain,(
  ( ! [X4,X5] : (~sP2(X5,X4,'F_0') | sP3(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f23528,f392])).
fof(f23528,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0')) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f23524,f235])).
fof(f235,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f150])).
fof(f23524,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0') | ~sP1('F_0',X4,X5)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f23387,f226])).
fof(f226,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f144])).
fof(f23387,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | ($spl6 | $spl12)),
  inference(subsumption_resolution,[],[f23363,f237])).
fof(f23363,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f22928,f221])).
fof(f221,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f144])).
fof(f22928,plain,(
  ( ! [X6,X5] : (sP3(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21928,f240])).
fof(f240,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f153])).
fof(f24036,plain,(
  ( ! [X14] : (s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | ~p(s(bool,X14))) ) | ($spl6 | $spl12 | $spl35)),
  inference(resolution,[],[f23457,f21731])).
fof(f3090,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl38),
  inference(cnf_transformation,[],[f3090_D])).
fof(f3090_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f23490,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f23488,f428])).
fof(f23488,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f23473])).
fof(f23473,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,X4))) ) | $spl6),
  inference(resolution,[],[f20841,f241])).
fof(f241,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f153])).
fof(f20841,plain,(
  ( ! [X6,X7] : (~sP3(X7,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)),X6) | ~p(s(bool,X7)) | s(bool,'F_0') != s(bool,X6)) ) | $spl6),
  inference(superposition,[],[f962,f242])).
fof(f23453,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f23438])).
fof(f23438,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f22930,f241])).
fof(f23424,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f23409])).
fof(f23409,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4))) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f22929,f241])).
fof(f22929,plain,(
  ( ! [X8,X7] : (~sP3(X8,X7,'F_0') | ~p(s(bool,X8)) | s(bool,'T_0') != s(bool,X7)) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21928,f242])).
fof(f704,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | $spl26),
  inference(cnf_transformation,[],[f704_D])).
fof(f704_D,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f23386,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f23362])).
fof(f23362,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f22928,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f23334,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(duplicate_literal_removal,[],[f23310])).
fof(f23310,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | ($spl6 | $spl12)),
  inference(resolution,[],[f22927,f243])).
fof(f22927,plain,(
  ( ! [X4,X3] : (sP3(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | ($spl6 | $spl12)),
  inference(superposition,[],[f21928,f240])).
fof(f23291,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f23198,f3136])).
fof(f23198,plain,(
  p(s(bool,'F_0')) | $spl40),
  inference(resolution,[],[f3136,f395])).
fof(f395,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f98])).
fof(f98,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.NOT_CLAUSES|split|2')).
fof(f23290,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f23197,f3136])).
fof(f23197,plain,(
  p(s(bool,'F_0')) | $spl40),
  inference(resolution,[],[f3136,f393])).
fof(f393,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.NOT_CLAUSES|split|1')).
fof(f23289,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f23196,f3136])).
fof(f23196,plain,(
  ( ! [X216] : (p(s(bool,X216))) ) | $spl40),
  inference(resolution,[],[f3136,f381])).
fof(f381,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f217])).
fof(f217,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f216])).
fof(f216,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X12] : ((p(s(bool,'T_0')) => p(s(bool,X12))) <=> p(s(bool,X12)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.IMP_CLAUSES|split|0')).
fof(f23288,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f23195,f3136])).
fof(f23195,plain,(
  ( ! [X215] : (p(s(bool,X215))) ) | $spl40),
  inference(resolution,[],[f3136,f378])).
fof(f378,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f215])).
fof(f215,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f214])).
fof(f214,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.IMP_CLAUSES|split|1')).
fof(f23287,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f23193,f3136])).
fof(f23193,plain,(
  ( ! [X213] : (p(s(bool,X213))) ) | $spl40),
  inference(resolution,[],[f3136,f375])).
fof(f375,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f213])).
fof(f213,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f212])).
fof(f212,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,X12))) <=> p(s(bool,'T_0')))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.IMP_CLAUSES|split|3')).
fof(f23286,plain,(
  $false | $spl40),
  inference(subsumption_resolution,[],[f23191,f3136])).
fof(f23191,plain,(
  p(s(bool,'F_0')) | $spl40),
  inference(resolution,[],[f3136,f372])).
fof(f372,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f211])).
fof(f23199,plain,(
  $false | $spl40),
  inference(resolution,[],[f3136,f397])).
fof(f23200,plain,(
  $false | $spl40),
  inference(resolution,[],[f3136,f398])).
fof(f398,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.TRUTH')).
fof(f23212,plain,(
  $false | ($spl30 | $spl40)),
  inference(resolution,[],[f3136,f1183])).
fof(f23213,plain,(
  $false | ($spl32 | $spl40)),
  inference(resolution,[],[f3136,f1187])).
fof(f23108,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f23036,f704])).
fof(f23036,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f704,f396])).
fof(f396,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f219])).
fof(f23107,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f23035,f704])).
fof(f23035,plain,(
  ~p(s(bool,'T_0')) | $spl26),
  inference(resolution,[],[f704,f394])).
fof(f394,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f218])).
fof(f23106,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f23033,f704])).
fof(f23033,plain,(
  ( ! [X222] : (~p(s(bool,X222))) ) | $spl26),
  inference(resolution,[],[f704,f390])).
fof(f390,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) => ~p(s(bool,X12)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.IMP_F')).
fof(f23105,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f23031,f704])).
fof(f23031,plain,(
  ( ! [X220] : (~p(s(bool,X220))) ) | $spl26),
  inference(resolution,[],[f704,f370])).
fof(f370,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f209])).
fof(f209,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f208])).
fof(f208,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f16])).
fof(f16,axiom,(
  ! [X12] : ((p(s(bool,X12)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X12)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.IMP_CLAUSES|split|4')).
fof(f23034,plain,(
  $false | $spl26),
  inference(resolution,[],[f704,f392])).
fof(f23041,plain,(
  $false | ($spl26 | $spl35)),
  inference(resolution,[],[f704,f1191])).
fof(f22516,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f22515,f428])).
fof(f22515,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f22482])).
fof(f22482,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f20838,f243])).
fof(f20838,plain,(
  ( ! [X0,X1] : (sP3(X1,X0,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | s(bool,'F_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl6),
  inference(superposition,[],[f962,f240])).
fof(f22621,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f22620,f428])).
fof(f22620,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f22596])).
fof(f22596,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl6),
  inference(resolution,[],[f20839,f243])).
fof(f20839,plain,(
  ( ! [X2,X3] : (sP3(X3,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)),X2) | s(bool,'F_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl6),
  inference(superposition,[],[f962,f240])).
fof(f22718,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5)) ) | ($spl6 | $spl12)),
  inference(forward_demodulation,[],[f22716,f428])).
fof(f22716,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl6),
  inference(duplicate_literal_removal,[],[f22694])).
fof(f22694,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'F_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,X4))) ) | $spl6),
  inference(resolution,[],[f20840,f241])).
fof(f20840,plain,(
  ( ! [X4,X5] : (~sP3(X5,X4,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,X5)) | s(bool,'F_0') != s(bool,X4)) ) | $spl6),
  inference(superposition,[],[f962,f242])).
fof(f707,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl29),
  inference(cnf_transformation,[],[f707_D])).
fof(f707_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f22906,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X22))) ) | $spl38),
  inference(subsumption_resolution,[],[f4065,f255])).
fof(f4065,plain,(
  ( ! [X22] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X22) | ~p(s(bool,X22))) ) | $spl38),
  inference(superposition,[],[f3090,f255])).
fof(f22905,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X23))) ) | $spl38),
  inference(subsumption_resolution,[],[f4066,f255])).
fof(f4066,plain,(
  ( ! [X23] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23))) ) | $spl38),
  inference(superposition,[],[f3090,f257])).
fof(f257,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f57])).
fof(f57,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X12] : (s(bool,'T_0') = s(bool,X12) <=> p(s(bool,X12)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.EQ_CLAUSES|split|1')).
fof(f1692,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10)) ) | $spl36),
  inference(cnf_transformation,[],[f1692_D])).
fof(f1692_D,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f22904,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24)) ) | $spl38),
  inference(subsumption_resolution,[],[f4067,f3090])).
fof(f4067,plain,(
  ( ! [X24] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X24) | s(bool,'F_0') = s(bool,X24)) ) | $spl38),
  inference(superposition,[],[f3090,f258])).
fof(f258,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.gcd.IS_GCD_SYM',aHL_BOOL_CASES)).
fof(f22903,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) | $spl72),
  inference(cnf_transformation,[],[f22903_D])).
fof(f22903_D,plain,(
  ( ! [X29] : (s(bool,'T_0') = s(bool,X29)) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4079,plain,(
  ( ! [X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl38),
  inference(trivial_inequality_removal,[],[f4069])).
fof(f4069,plain,(
  ( ! [X29] : (s(bool,X29) != s(bool,X29) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X29)) ) | $spl38),
  inference(superposition,[],[f3090,f258])).
fof(f4078,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl38),
  inference(trivial_inequality_removal,[],[f4070])).
fof(f4070,plain,(
  ( ! [X30] : (s(bool,X30) != s(bool,X30) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X30))) ) | $spl38),
  inference(superposition,[],[f3090,f253])).
fof(f253,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X12] : (s(bool,'F_0') = s(bool,X12) <=> ~p(s(bool,X12)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.EQ_CLAUSES|split|3')).
fof(f4077,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl38),
  inference(trivial_inequality_removal,[],[f4071])).
fof(f4071,plain,(
  ( ! [X31] : (s(bool,X31) != s(bool,X31) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X31))) ) | $spl38),
  inference(superposition,[],[f3090,f251])).
fof(f22860,plain,(
  ( ! [X39,X38] : (s(bool,X38) != s(bool,X39) | p(s(bool,X39))) ) | $spl70),
  inference(cnf_transformation,[],[f22860_D])).
fof(f22860_D,plain,(
  ( ! [X39,X38] : (s(bool,X38) != s(bool,X39) | p(s(bool,X39))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f22862,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X28) != s(bool,X29) | p(s(bool,X29))) ) | ($spl16 | $spl38)),
  inference(subsumption_resolution,[],[f12127,f7854])).
fof(f7854,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl16),
  inference(subsumption_resolution,[],[f7853,f373])).
fof(f7853,plain,(
  ( ! [X10,X11] : (~p(s(bool,'T_0')) | s(bool,X10) != s(bool,X11) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl16),
  inference(forward_demodulation,[],[f7846,f2219])).
fof(f2219,plain,(
  s(bool,'T_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12))) | $spl16),
  inference(resolution,[],[f532,f255])).
fof(f532,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl16),
  inference(cnf_transformation,[],[f532_D])).
fof(f532_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f7846,plain,(
  ( ! [X10,X11] : (s(bool,X10) != s(bool,X11) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,X10)) | p(s(bool,X11))) ) | $spl16),
  inference(resolution,[],[f2213,f224])).
fof(f224,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f144])).
fof(f2213,plain,(
  ( ! [X35,X34] : (sP3('const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)),X34,X35) | s(bool,X34) != s(bool,X35)) ) | $spl16),
  inference(resolution,[],[f532,f241])).
fof(f12127,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X28) != s(bool,X29) | p(s(bool,X29)) | p(s(bool,X28))) ) | $spl38),
  inference(superposition,[],[f4039,f251])).
fof(f4039,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X3) | s(bool,X2) != s(bool,X3) | p(s(bool,X2))) ) | $spl38),
  inference(superposition,[],[f3090,f253])).
fof(f22861,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | p(s(bool,X31))) ) | ($spl16 | $spl38)),
  inference(subsumption_resolution,[],[f12128,f7854])).
fof(f12128,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X30) != s(bool,X31) | p(s(bool,X31)) | p(s(bool,X30))) ) | $spl38),
  inference(superposition,[],[f4039,f253])).
fof(f22859,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | p(s(bool,X39))) ) | $spl38),
  inference(subsumption_resolution,[],[f12132,f4039])).
fof(f12132,plain,(
  ( ! [X39,X38] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,X38) != s(bool,X39) | p(s(bool,X39)) | s(bool,'T_0') = s(bool,X38)) ) | $spl38),
  inference(superposition,[],[f4039,f258])).
fof(f22795,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | ~p(s(bool,X10))) ) | ($spl29 | $spl35)),
  inference(resolution,[],[f21731,f5426])).
fof(f5426,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl29),
  inference(subsumption_resolution,[],[f5425,f254])).
fof(f5425,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0')) ) | $spl29),
  inference(subsumption_resolution,[],[f5424,f5332])).
fof(f5332,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl29),
  inference(subsumption_resolution,[],[f5326,f237])).
fof(f5326,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | $spl29),
  inference(resolution,[],[f1421,f221])).
fof(f1421,plain,(
  ( ! [X6,X5] : (sP3(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl29),
  inference(superposition,[],[f707,f240])).
fof(f5424,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP0(X13,'F_0',X12)) ) | $spl29),
  inference(subsumption_resolution,[],[f5414,f5333])).
fof(f5333,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl29),
  inference(subsumption_resolution,[],[f5327,f233])).
fof(f5327,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl29),
  inference(resolution,[],[f1421,f222])).
fof(f222,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f144])).
fof(f5414,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | $spl29),
  inference(resolution,[],[f1423,f225])).
fof(f1423,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl29),
  inference(superposition,[],[f707,f242])).
fof(f22794,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) ) | ($spl29 | $spl35)),
  inference(resolution,[],[f21731,f5340])).
fof(f5340,plain,(
  ( ! [X4,X5] : (~sP2(X5,X4,'F_0') | sP3(X5,'F_0',X4) | s(bool,'T_0') != s(bool,X4)) ) | $spl29),
  inference(subsumption_resolution,[],[f5339,f392])).
fof(f5339,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0')) ) | $spl29),
  inference(subsumption_resolution,[],[f5336,f235])).
fof(f5336,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | sP3(X5,'F_0',X4) | p(s(bool,'F_0')) | ~sP2(X5,X4,'F_0') | ~sP1('F_0',X4,X5)) ) | $spl29),
  inference(resolution,[],[f5332,f226])).
fof(f22810,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | $spl64),
  inference(cnf_transformation,[],[f22810_D])).
fof(f22810_D,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f22793,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14))) | ~p(s(bool,X8))) ) | ($spl29 | $spl35)),
  inference(resolution,[],[f21731,f5356])).
fof(f5356,plain,(
  ( ! [X0,X1] : (~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | $spl29),
  inference(subsumption_resolution,[],[f5349,f251])).
fof(f5349,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | $spl29),
  inference(resolution,[],[f5346,f242])).
fof(f5346,plain,(
  ( ! [X2,X3] : (sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl29),
  inference(subsumption_resolution,[],[f5345,f255])).
fof(f5345,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl29),
  inference(subsumption_resolution,[],[f5342,f5332])).
fof(f5342,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) ) | $spl29),
  inference(resolution,[],[f5333,f225])).
fof(f22691,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13))) | $spl63),
  inference(cnf_transformation,[],[f22691_D])).
fof(f22691_D,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f22685,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13)))) ) | ($spl29 | $spl32)),
  inference(resolution,[],[f21622,f5382])).
fof(f5382,plain,(
  ( ! [X2,X3] : (~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl29),
  inference(subsumption_resolution,[],[f5375,f231])).
fof(f231,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f147])).
fof(f5375,plain,(
  ( ! [X2,X3] : (~sP2(X2,X3,'F_0') | ~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl29),
  inference(resolution,[],[f5372,f243])).
fof(f5372,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | $spl29),
  inference(subsumption_resolution,[],[f5370,f392])).
fof(f5370,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | $spl29),
  inference(resolution,[],[f5367,f226])).
fof(f5367,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5)) ) | $spl29),
  inference(subsumption_resolution,[],[f5361,f233])).
fof(f5361,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5) | p(s(bool,X5))) ) | $spl29),
  inference(resolution,[],[f5358,f231])).
fof(f5358,plain,(
  ( ! [X6,X7] : (~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | $spl29),
  inference(subsumption_resolution,[],[f5352,f234])).
fof(f234,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f150])).
fof(f5352,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | $spl29),
  inference(resolution,[],[f5346,f222])).
fof(f21622,plain,(
  ( ! [X33,X32] : (sP0('const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13)),X32,X33)) ) | $spl32),
  inference(resolution,[],[f1187,f239])).
fof(f239,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f152])).
fof(f22545,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12))) | $spl61),
  inference(cnf_transformation,[],[f22545_D])).
fof(f22545_D,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f22539,plain,(
  ( ! [X8] : (p(s(bool,X8)) | s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) ) | ($spl29 | $spl30)),
  inference(resolution,[],[f21477,f5382])).
fof(f21477,plain,(
  ( ! [X33,X32] : (sP0('const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)),X32,X33)) ) | $spl30),
  inference(resolution,[],[f1183,f239])).
fof(f22430,plain,(
  $false | ($spl2 | $spl6 | $spl23)),
  inference(trivial_inequality_removal,[],[f22417])).
fof(f22417,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6 | $spl23)),
  inference(superposition,[],[f22382,f6547])).
fof(f6547,plain,(
  ( ! [X49] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X49),s('type.num.num',sK15(sK14,sK12,sK13))))) ) | $spl23),
  inference(resolution,[],[f3452,f251])).
fof(f3452,plain,(
  ( ! [X0] : (~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X0),s('type.num.num',sK15(sK14,sK12,sK13)))))) ) | $spl23),
  inference(resolution,[],[f539,f259])).
fof(f22382,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8)) ) | ($spl2 | $spl6)),
  inference(subsumption_resolution,[],[f22381,f252])).
fof(f252,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f156])).
fof(f22381,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8))) ) | ($spl2 | $spl6)),
  inference(forward_demodulation,[],[f419,f421])).
fof(f22429,plain,(
  $false | ($spl2 | $spl6 | $spl23)),
  inference(trivial_inequality_removal,[],[f22418])).
fof(f22418,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6 | $spl23)),
  inference(superposition,[],[f22382,f6657])).
fof(f6657,plain,(
  ( ! [X49] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X49),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13))))) ) | $spl23),
  inference(resolution,[],[f3453,f251])).
fof(f3453,plain,(
  ( ! [X1] : (~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X1),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13)))))) ) | $spl23),
  inference(resolution,[],[f539,f260])).
fof(f22428,plain,(
  $false | ($spl2 | $spl6)),
  inference(trivial_inequality_removal,[],[f22419])).
fof(f22419,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6)),
  inference(superposition,[],[f22382,f421])).
fof(f22427,plain,(
  $false | ($spl2 | $spl6 | $spl23)),
  inference(trivial_inequality_removal,[],[f22425])).
fof(f22425,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl2 | $spl6 | $spl23)),
  inference(superposition,[],[f22382,f3471])).
fof(f22426,plain,(
  $false | ($spl2 | $spl6)),
  inference(equality_resolution,[],[f22382])).
fof(f428,plain,(
  s(bool,'T_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | $spl12),
  inference(cnf_transformation,[],[f428_D])).
fof(f428_D,plain,(
  s(bool,'T_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f422,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | $spl7),
  inference(cnf_transformation,[],[f422_D])).
fof(f422_D,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f21031,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))),
  inference(superposition,[],[f220,f258])).
fof(f20716,plain,(
  $false | ($spl0 | $spl10)),
  inference(subsumption_resolution,[],[f20715,f392])).
fof(f20715,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl10)),
  inference(forward_demodulation,[],[f418,f426])).
fof(f20712,plain,(
  $false | ($spl6 | $spl10)),
  inference(subsumption_resolution,[],[f962,f426])).
fof(f20679,plain,(
  $false | ($spl8 | $spl18 | $spl20 | $spl23)),
  inference(subsumption_resolution,[],[f20678,f398])).
fof(f20678,plain,(
  ~p(s(bool,'T_0')) | ($spl8 | $spl18 | $spl20 | $spl23)),
  inference(forward_demodulation,[],[f20677,f3362])).
fof(f20677,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ($spl8 | $spl18 | $spl23)),
  inference(subsumption_resolution,[],[f20676,f390])).
fof(f20676,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ($spl8 | $spl18 | $spl23)),
  inference(forward_demodulation,[],[f20675,f3471])).
fof(f20675,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ($spl8 | $spl18)),
  inference(subsumption_resolution,[],[f20648,f373])).
fof(f20648,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ($spl8 | $spl18)),
  inference(superposition,[],[f1944,f3246])).
fof(f1944,plain,(
  ( ! [X0] : (~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK13)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK12))))) ) | $spl8),
  inference(subsumption_resolution,[],[f1741,f373])).
fof(f1741,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X0),s('type.num.num',sK12))))) ) | $spl8),
  inference(superposition,[],[f261,f423])).
fof(f423,plain,(
  s(bool,'T_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | $spl8),
  inference(cnf_transformation,[],[f423_D])).
fof(f423_D,plain,(
  s(bool,'T_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f7436,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X16),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13))))) ) | $spl58),
  inference(cnf_transformation,[],[f7436_D])).
fof(f7436_D,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X16),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13))))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f7432,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X16),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13)))) | ~p(s(bool,X17))) ) | ($spl23 | $spl29)),
  inference(resolution,[],[f6645,f5426])).
fof(f6645,plain,(
  ( ! [X17,X15,X16] : (sP2(X15,'const.gcd.is_gcd_3'(s('type.num.num',X16),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13))),X17)) ) | $spl23),
  inference(resolution,[],[f3453,f230])).
fof(f7435,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X12),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13))))) ) | $spl56),
  inference(cnf_transformation,[],[f7435_D])).
fof(f7435_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X12),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f7430,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X12),s('type.num.num',sK14),s('type.num.num',sK15(sK14,sK12,sK13)))) | ~p(s(bool,X13))) ) | ($spl23 | $spl29)),
  inference(resolution,[],[f6645,f5356])).
fof(f7347,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X16),s('type.num.num',sK15(sK14,sK12,sK13))))) ) | $spl54),
  inference(cnf_transformation,[],[f7347_D])).
fof(f7347_D,plain,(
  ( ! [X16] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X16),s('type.num.num',sK15(sK14,sK12,sK13))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f7343,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X16),s('type.num.num',sK15(sK14,sK12,sK13)))) | ~p(s(bool,X17))) ) | ($spl23 | $spl29)),
  inference(resolution,[],[f6535,f5426])).
fof(f6535,plain,(
  ( ! [X17,X15,X16] : (sP2(X15,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X16),s('type.num.num',sK15(sK14,sK12,sK13))),X17)) ) | $spl23),
  inference(resolution,[],[f3452,f230])).
fof(f7346,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num',sK15(sK14,sK12,sK13))))) ) | $spl52),
  inference(cnf_transformation,[],[f7346_D])).
fof(f7346_D,plain,(
  ( ! [X12] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num',sK15(sK14,sK12,sK13))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f7341,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK14),s('type.num.num',X12),s('type.num.num',sK15(sK14,sK12,sK13)))) | ~p(s(bool,X13))) ) | ($spl23 | $spl29)),
  inference(resolution,[],[f6535,f5356])).
fof(f5488,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12))) | $spl51),
  inference(cnf_transformation,[],[f5488_D])).
fof(f5488_D,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f5476,plain,(
  ( ! [X20] : (p(s(bool,X20)) | s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) ) | ($spl20 | $spl29)),
  inference(resolution,[],[f5382,f3355])).
fof(f3355,plain,(
  ( ! [X33,X32] : (sP0('const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)),X32,X33)) ) | $spl20),
  inference(resolution,[],[f537,f239])).
fof(f5486,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13))) | $spl49),
  inference(cnf_transformation,[],[f5486_D])).
fof(f5486_D,plain,(
  s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13))) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f5475,plain,(
  ( ! [X19] : (p(s(bool,X19)) | s(bool,'F_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) ) | ($spl18 | $spl29)),
  inference(resolution,[],[f5382,f3239])).
fof(f3239,plain,(
  ( ! [X33,X32] : (sP0('const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)),X32,X33)) ) | $spl18),
  inference(resolution,[],[f535,f239])).
fof(f5459,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) | $spl46),
  inference(cnf_transformation,[],[f5459_D])).
fof(f5459_D,plain,(
  s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5455,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | ($spl23 | $spl29)),
  inference(resolution,[],[f5356,f3459])).
fof(f3459,plain,(
  ( ! [X12,X13] : (sP2(X12,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)),X13)) ) | $spl23),
  inference(resolution,[],[f539,f230])).
fof(f5439,plain,(
  s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) | $spl43),
  inference(cnf_transformation,[],[f5439_D])).
fof(f5439_D,plain,(
  s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f5449,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) ) | $spl44),
  inference(cnf_transformation,[],[f5449_D])).
fof(f5449_D,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f5446,plain,(
  ( ! [X9] : (sP3(X9,'F_0','const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) ) | ($spl23 | $spl29)),
  inference(resolution,[],[f5340,f3459])).
fof(f5435,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14))) | ~p(s(bool,X9))) ) | ($spl23 | $spl29)),
  inference(resolution,[],[f5426,f3459])).
fof(f3203,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5)) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f3201,f426])).
fof(f3201,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3195])).
fof(f3195,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,X4))) ) | $spl8),
  inference(resolution,[],[f2311,f241])).
fof(f2311,plain,(
  ( ! [X6,X7] : (~sP3(X7,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)),X6) | ~p(s(bool,X7)) | s(bool,'T_0') != s(bool,X6)) ) | $spl8),
  inference(superposition,[],[f1727,f242])).
fof(f1727,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | $spl8),
  inference(backward_demodulation,[],[f423,f220])).
fof(f3135,plain,(
  ( ! [X4,X5] : (s(bool,'F_0') != s(bool,X5) | ~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5)) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f3133,f426])).
fof(f3133,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3127])).
fof(f3127,plain,(
  ( ! [X4,X5] : (~p(s(bool,X4)) | s(bool,'T_0') != s(bool,X5) | s(bool,X5) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,X4))) ) | $spl8),
  inference(resolution,[],[f2310,f241])).
fof(f2310,plain,(
  ( ! [X4,X5] : (~sP3(X5,X4,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | ~p(s(bool,X5)) | s(bool,'T_0') != s(bool,X4)) ) | $spl8),
  inference(superposition,[],[f1727,f242])).
fof(f3100,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f3099,f426])).
fof(f3099,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3094])).
fof(f3094,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl8),
  inference(resolution,[],[f2309,f243])).
fof(f2309,plain,(
  ( ! [X2,X3] : (sP3(X3,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)),X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | $spl8),
  inference(superposition,[],[f1727,f240])).
fof(f3089,plain,(
  ( ! [X2,X3] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2) | p(s(bool,X3))) ) | ($spl8 | $spl10)),
  inference(forward_demodulation,[],[f3088,f426])).
fof(f3088,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl8),
  inference(duplicate_literal_removal,[],[f3083])).
fof(f3083,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,X2) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) ) | $spl8),
  inference(resolution,[],[f2308,f243])).
fof(f2308,plain,(
  ( ! [X0,X1] : (sP3(X1,X0,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | s(bool,'T_0') != s(bool,X0) | p(s(bool,X1))) ) | $spl8),
  inference(superposition,[],[f1727,f240])).
fof(f3077,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2991,f2901])).
fof(f2901,plain,(
  ( ! [X10] : (~p(s(bool,X10))) ) | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2900,f257])).
fof(f2900,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,X10))) ) | ($spl4 | $spl8)),
  inference(forward_demodulation,[],[f420,f423])).
fof(f2991,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f395])).
fof(f3076,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2990,f2901])).
fof(f2990,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f393])).
fof(f3075,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2989,f2901])).
fof(f2989,plain,(
  ( ! [X216] : (p(s(bool,X216))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f381])).
fof(f3074,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2988,f2901])).
fof(f2988,plain,(
  ( ! [X215] : (p(s(bool,X215))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f378])).
fof(f3073,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2986,f2901])).
fof(f2986,plain,(
  ( ! [X213] : (p(s(bool,X213))) ) | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f375])).
fof(f3072,plain,(
  $false | ($spl4 | $spl8)),
  inference(subsumption_resolution,[],[f2984,f2901])).
fof(f2984,plain,(
  p(s(bool,'F_0')) | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f372])).
fof(f2992,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f397])).
fof(f2993,plain,(
  $false | ($spl4 | $spl8)),
  inference(resolution,[],[f2901,f398])).
fof(f3008,plain,(
  $false | ($spl4 | $spl8 | $spl16)),
  inference(resolution,[],[f2901,f532])).
fof(f3009,plain,(
  $false | ($spl4 | $spl8 | $spl14)),
  inference(resolution,[],[f2901,f530])).
fof(f530,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl14),
  inference(cnf_transformation,[],[f530_D])).
fof(f530_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2801,plain,(
  $false | ($spl8 | $spl12)),
  inference(subsumption_resolution,[],[f428,f1727])).
fof(f531,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl15),
  inference(cnf_transformation,[],[f531_D])).
fof(f531_D,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f539,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | $spl23),
  inference(cnf_transformation,[],[f539_D])).
fof(f539_D,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f2560,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ($spl10 | $spl16)),
  inference(subsumption_resolution,[],[f2559,f532])).
fof(f2559,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl10),
  inference(subsumption_resolution,[],[f2352,f370])).
fof(f2352,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl10),
  inference(superposition,[],[f264,f426])).
fof(f264,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(X2,X1,X0)),s('type.num.num',X2)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f537,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | $spl20),
  inference(cnf_transformation,[],[f537_D])).
fof(f537_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2558,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ($spl10 | $spl16)),
  inference(subsumption_resolution,[],[f2557,f532])).
fof(f2557,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl10),
  inference(subsumption_resolution,[],[f2351,f370])).
fof(f2351,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl10),
  inference(superposition,[],[f263,f426])).
fof(f263,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(X2,X1,X0)),s('type.num.num',X1)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f535,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) | $spl18),
  inference(cnf_transformation,[],[f535_D])).
fof(f535_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2556,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ($spl10 | $spl16)),
  inference(subsumption_resolution,[],[f2555,f532])).
fof(f2555,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl10),
  inference(subsumption_resolution,[],[f2350,f370])).
fof(f2350,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl10),
  inference(superposition,[],[f262,f426])).
fof(f262,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',X0),s('type.num.num',X1),s('type.num.num',X2)))) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(X2,X1,X0)),s('type.num.num',X0)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X1)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',X2),s('type.num.num',X0))))) )),
  inference(cnf_transformation,[],[f162])).
fof(f2325,plain,(
  $false | ($spl8 | $spl15)),
  inference(subsumption_resolution,[],[f2082,f531])).
fof(f2082,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl8),
  inference(subsumption_resolution,[],[f1740,f397])).
fof(f1740,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl8),
  inference(superposition,[],[f260,f423])).
fof(f533,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl17),
  inference(cnf_transformation,[],[f533_D])).
fof(f533_D,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f1483,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f417,f262])).
fof(f1484,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f417,f263])).
fof(f1485,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f417,f264])).
fof(f1943,plain,(
  $false | ($spl8 | $spl17)),
  inference(subsumption_resolution,[],[f1942,f533])).
fof(f1942,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl8),
  inference(subsumption_resolution,[],[f1739,f397])).
fof(f1739,plain,(
  ~p(s(bool,'T_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl8),
  inference(superposition,[],[f259,f423])).
fof(f1694,plain,(
  $false | ($spl7 | $spl9)),
  inference(subsumption_resolution,[],[f1688,f422])).
fof(f1688,plain,(
  s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | $spl9),
  inference(trivial_inequality_removal,[],[f1686])).
fof(f1686,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | $spl9),
  inference(superposition,[],[f424,f258])).
fof(f702,plain,(
  ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl25),
  inference(cnf_transformation,[],[f702_D])).
fof(f702_D,plain,(
  ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f1693,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))))) ) | $spl9),
  inference(subsumption_resolution,[],[f1681,f254])).
fof(f1681,plain,(
  ( ! [X11] : (s(bool,'T_0') != s(bool,X11) | ~p(s(bool,X11)) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))))) ) | $spl9),
  inference(superposition,[],[f424,f247])).
fof(f247,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f100])).
fof(f100,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.gcd.IS_GCD_SYM','thm.bool.IMP_ANTISYM_AX')).
fof(f1691,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))))) ) | $spl9),
  inference(subsumption_resolution,[],[f1680,f254])).
fof(f1680,plain,(
  ( ! [X10] : (s(bool,'T_0') != s(bool,X10) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | ~p(s(bool,X10))) ) | $spl9),
  inference(superposition,[],[f424,f247])).
fof(f703,plain,(
  p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl24),
  inference(cnf_transformation,[],[f703_D])).
fof(f703_D,plain,(
  p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f1641,plain,(
  ( ! [X9] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl7),
  inference(subsumption_resolution,[],[f1628,f251])).
fof(f1628,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl7),
  inference(superposition,[],[f422,f244])).
fof(f244,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f101])).
fof(f1640,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))))) ) | $spl7),
  inference(subsumption_resolution,[],[f1627,f251])).
fof(f1627,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))))) ) | $spl7),
  inference(superposition,[],[f422,f244])).
fof(f1449,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))),
  inference(superposition,[],[f220,f258])).
fof(f426,plain,(
  s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) | $spl10),
  inference(cnf_transformation,[],[f426_D])).
fof(f426_D,plain,(
  s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f424,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | $spl9),
  inference(cnf_transformation,[],[f424_D])).
fof(f424_D,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1448,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))),
  inference(superposition,[],[f220,f258])).
fof(f417,plain,(
  ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl1),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f1447,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f257])).
fof(f1446,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f255])).
fof(f418,plain,(
  p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) | $spl0),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f1445,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f253])).
fof(f1444,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f251])).
fof(f420,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,X10))) ) | $spl4),
  inference(cnf_transformation,[],[f420_D])).
fof(f420_D,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,X10))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f1443,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,X11)) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))) )),
  inference(superposition,[],[f220,f247])).
fof(f1442,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f220,f247])).
fof(f419,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,X8))) ) | $spl2),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,X8))) ) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f1441,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) | p(s(bool,X9))) )),
  inference(superposition,[],[f220,f244])).
fof(f1440,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,X8)) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))) )),
  inference(superposition,[],[f220,f244])).
fof(f1416,plain,(
  $false | ($spl0 | $spl17)),
  inference(subsumption_resolution,[],[f741,f533])).
fof(f741,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl0),
  inference(resolution,[],[f418,f260])).
fof(f1191,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | $spl35),
  inference(cnf_transformation,[],[f1191_D])).
fof(f1191_D,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f1189,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f1188,f530])).
fof(f1188,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl6),
  inference(subsumption_resolution,[],[f979,f370])).
fof(f979,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl6),
  inference(superposition,[],[f264,f421])).
fof(f1187,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13)))) | $spl32),
  inference(cnf_transformation,[],[f1187_D])).
fof(f1187_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1185,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f1184,f530])).
fof(f1184,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl6),
  inference(subsumption_resolution,[],[f978,f370])).
fof(f978,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl6),
  inference(superposition,[],[f263,f421])).
fof(f1183,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) | $spl30),
  inference(cnf_transformation,[],[f1183_D])).
fof(f1183_D,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1181,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ($spl6 | $spl14)),
  inference(subsumption_resolution,[],[f1180,f530])).
fof(f1180,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl6),
  inference(subsumption_resolution,[],[f977,f370])).
fof(f977,plain,(
  p(s(bool,'F_0')) | p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK13,sK12)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | $spl6),
  inference(superposition,[],[f262,f421])).
fof(f842,plain,(
  $false | ($spl0 | $spl15)),
  inference(subsumption_resolution,[],[f740,f531])).
fof(f740,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl0),
  inference(resolution,[],[f418,f259])).
fof(f695,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl7),
  inference(superposition,[],[f422,f257])).
fof(f694,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | $spl7),
  inference(superposition,[],[f422,f255])).
fof(f705,plain,(
  ( ! [X9] : (p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl7),
  inference(subsumption_resolution,[],[f689,f251])).
fof(f689,plain,(
  ( ! [X9] : (s(bool,'F_0') != s(bool,X9) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14)))) | p(s(bool,X9))) ) | $spl7),
  inference(superposition,[],[f422,f244])).
fof(f701,plain,(
  ( ! [X8] : (p(s(bool,X8)) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))))) ) | $spl7),
  inference(subsumption_resolution,[],[f688,f251])).
fof(f688,plain,(
  ( ! [X8] : (s(bool,'F_0') != s(bool,X8) | p(s(bool,X8)) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))))) ) | $spl7),
  inference(superposition,[],[f422,f244])).
fof(f431,plain,(
  ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK14)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f417,f264])).
fof(f430,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f417,f263])).
fof(f429,plain,(
  p(s(bool,'const.divides.divides_2'(s('type.num.num',sK15(sK14,sK12,sK13)),s('type.num.num',sK13)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK12)))) | ~p(s(bool,'const.divides.divides_2'(s('type.num.num',sK14),s('type.num.num',sK13)))) | $spl1),
  inference(resolution,[],[f417,f262])).
fof(f416,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | s(bool,'T_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))),
  inference(superposition,[],[f220,f258])).
fof(f415,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | s(bool,'F_0') = s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))),
  inference(superposition,[],[f220,f258])).
fof(f414,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f257])).
fof(f413,plain,(
  s(bool,'T_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f255])).
fof(f412,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f253])).
fof(f411,plain,(
  s(bool,'F_0') != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))),
  inference(superposition,[],[f220,f251])).
fof(f410,plain,(
  ( ! [X11] : (s(bool,X11) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,X11)) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))) )),
  inference(superposition,[],[f220,f247])).
fof(f409,plain,(
  ( ! [X10] : (s(bool,X10) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | ~p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) | ~p(s(bool,X10))) )),
  inference(superposition,[],[f220,f247])).
fof(f408,plain,(
  ( ! [X9] : (s(bool,X9) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14)))) | p(s(bool,X9))) )),
  inference(superposition,[],[f220,f244])).
fof(f407,plain,(
  ( ! [X8] : (s(bool,X8) != s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK12),s('type.num.num',sK13),s('type.num.num',sK14))) | p(s(bool,X8)) | p(s(bool,'const.gcd.is_gcd_3'(s('type.num.num',sK13),s('type.num.num',sK12),s('type.num.num',sK14))))) )),
  inference(superposition,[],[f220,f244])).
% SZS output end Proof for thm.gcd
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 1203
% Passive clauses: 4737
% Generated clauses: 48970
% Final active clauses: 929
% Final passive clauses: 1433
% Input formulas: 44
% Initial clauses: 179
% 
% Selected by SInE selection: 37
% SInE iterations: 1
% 
% Duplicate literals: 2877
% Trivial inequalities: 219
% Fw subsumption resolutions: 10277
% Fw demodulations: 11290
% Bw demodulations: 26
% 
% Simple tautologies: 4249
% Equational tautologies: 44
% Forward subsumptions: 16166
% 
% Binary resolution: 19095
% Forward superposition: 1353
% Backward superposition: 3348
% Equality resolution: 9
% 
% Split clauses: 88
% Split components: 182
% SAT solver clauses: 131
% SAT solver unit clauses: 13
% SAT solver binary clauses: 42
% 
% Sat splits: 88
% Sat splitting refutations: 42
% 
% Memory used [KB]: 12153
% Time elapsed: 0.558 s
% ------------------------------
% ------------------------------
% Success in time 12.282 s
