Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.arithmetic
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.arithmetic
% SZS output start Proof for thm.arithmetic
fof(f4753,plain,(
  $false),
  inference(trivial_inequality_removal,[],[f4723])).
fof(f4723,plain,(
  s('type.num.num',sK12) != s('type.num.num',sK12)),
  inference(superposition,[],[f240,f4707])).
fof(f4707,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4706,f292])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))) )),
  inference(cnf_transformation,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(flattening,[],[f88])).
fof(f88,plain,(
  ! [X0] : ! [X1,X2] : s(X0,X2) = s(X0,'const.bool.COND_3'(s(bool,'F_0'),s(X0,X1),s(X0,X2)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : ! [X8,X9] : s(X11,X9) = s(X11,'const.bool.COND_3'(s(bool,'F_0'),s(X11,X8),s(X11,X9)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|0','thm.bool.COND_CLAUSES|split|1')).
fof(f4706,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'F_0'),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f4705,f438])).
fof(f438,plain,(
  ( ! [X6] : (s(bool,'F_0') = s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))) )),
  inference(resolution,[],[f264,f288])).
fof(f288,plain,(
  ( ! [X0] : (~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))) )),
  inference(cnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6] : ~p(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0'))))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|0','thm.prim_rec.NOT_LESS_|48|')).
fof(f264,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|0','thm.bool.EQ_CLAUSES|split|3')).
fof(f4705,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(forward_demodulation,[],[f4662,f249])).
fof(f249,plain,(
  ( ! [X0] : (s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X7] : s('type.num.num',X7) = s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|0','thm.arithmetic.SUB_|48||split|1')).
fof(f4662,plain,(
  ( ! [X0] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0'))),s('type.num.num','const.num.|48|_0'),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f299,f248])).
fof(f248,plain,(
  ( ! [X0] : (s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))) )),
  inference(cnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X0)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X7] : s('type.num.num','const.num.|48|_0') = s('type.num.num','const.arithmetic.-_2'(s('type.num.num','const.num.|48|_0'),s('type.num.num',X7)))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|0','thm.arithmetic.SUB|split|0')).
fof(f299,plain,(
  ( ! [X0,X1] : (s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0,X1] : s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num',X1))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X0),s('type.num.num',X1))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X1),s('type.num.num',X0))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X0),s('type.num.num',X1)))))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X6,X7] : s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X6),s('type.num.num',X7))) = s('type.num.num','const.bool.COND_3'(s(bool,'const.prim_rec.<_2'(s('type.num.num',X6),s('type.num.num',X7))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X7),s('type.num.num',X6))),s('type.num.num','const.arithmetic.-_2'(s('type.num.num',X6),s('type.num.num',X7)))))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|0','thm.arithmetic.ABS_DIFF_def')).
fof(f240,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  s('type.num.num',sK12) != s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',sK12),s('type.num.num','const.num.|48|_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f120])).
fof(f120,plain,(
  ? [X0] : s('type.num.num',X0) != s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(ennf_transformation,[],[f49])).
fof(f49,plain,(
  ~! [X0] : s('type.num.num',X0) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X0),s('type.num.num','const.num.|48|_0')))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : s('type.num.num',X6) = s('type.num.num','const.arithmetic.ABS_DIFF_2'(s('type.num.num',X6),s('type.num.num','const.num.|48|_0')))),
  file('thm.arithmetic.ABS_DIFF_ZERO|split|0','thm.arithmetic.ABS_DIFF_ZERO|split|0')).
% SZS output end Proof for thm.arithmetic
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 337
% Passive clauses: 1369
% Generated clauses: 4501
% Final active clauses: 335
% Final passive clauses: 886
% Input formulas: 47
% Initial clauses: 174
% 
% Duplicate literals: 143
% Trivial inequalities: 22
% Fw subsumption resolutions: 287
% Fw demodulations: 92
% Bw demodulations: 1
% 
% Simple tautologies: 144
% Equational tautologies: 21
% Forward subsumptions: 2391
% 
% Binary resolution: 2383
% Factoring: 726
% Forward superposition: 362
% Backward superposition: 264
% Self superposition: 11
% Equality factoring: 2
% Equality resolution: 9
% 
% Split clauses: 79
% Split components: 159
% SAT solver clauses: 154
% SAT solver unit clauses: 7
% SAT solver binary clauses: 29
% 
% Sat splits: 79
% 
% TWLsolver clauses: 154
% TWLsolver calls for satisfiability: 135
% 
% Memory used [KB]: 2302
% Time elapsed: 0.103 s
% ------------------------------
% ------------------------------
% Success in time 0.106 s
