Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.quotient
% (105204)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 836
% Passive clauses: 7830
% Generated clauses: 23036
% Final active clauses: 822
% Final passive clauses: 6692
% Input formulas: 69
% Initial clauses: 194
% 
% Duplicate literals: 352
% Trivial inequalities: 34
% Fw subsumption resolutions: 1646
% Fw demodulations: 2217
% Bw demodulations: 11
% 
% Simple tautologies: 204
% Equational tautologies: 36
% Forward subsumptions: 10061
% 
% Binary resolution: 7656
% Factoring: 1092
% Forward superposition: 3592
% Backward superposition: 6031
% Self superposition: 54
% Equality factoring: 26
% Equality resolution: 9
% 
% Split clauses: 420
% Split components: 841
% SAT solver clauses: 658
% SAT solver unit clauses: 19
% SAT solver binary clauses: 111
% 
% Sat splits: 420
% Sat splitting refutations: 1
% 
% TWLsolver clauses: 658
% TWLsolver calls for satisfiability: 392
% 
% Memory used [KB]: 14839
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.quotient
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.quotient
% SZS output start Proof for thm.quotient
fof(f4834,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f4832,f335,f550,f2070,f2072,f386,f573,f712,f718,f333,f546,f349,f558,f714,f720,f332,f547,f4770,f4771,f551,f553,f693,f695,f392,f578,f400,f582,f713,f719,f331,f548,f741,f747,f334,f544,f353,f560])).
fof(f560,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) | $spl18),
  inference(cnf_transformation,[],[f560_D])).
fof(f560_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f353,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f240])).
fof(f240,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR','thm.bool.EQ_CLAUSES|split|2')).
fof(f544,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK8,bool),sK10),s(fun(sK8,bool),sK11)))) | $spl1),
  inference(cnf_transformation,[],[f544_D])).
fof(f544_D,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK8,bool),sK10),s(fun(sK8,bool),sK11)))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f334,plain,(
  ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK8,bool),sK10),s(fun(sK8,bool),sK11))))),
  inference(cnf_transformation,[],[f235])).
fof(f235,plain,(
  ! [X4] : (p(s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X4)))) & (~p(s(bool,happ(s(fun(sK8,bool),sK9),s(sK8,X4)))) | p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X4)))))) & p(s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9)))) & ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK8,bool),sK10),s(fun(sK8,bool),sK11))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK8,sK9,sK10,sK11])],[f188])).
fof(f188,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) & p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3)))))),
  inference(flattening,[],[f187])).
fof(f187,plain,(
  ? [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) & (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) & ~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) => (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : (! [X4] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) & (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X4)))) => p(s(bool,happ(s(fun(X0,bool),X3),s(X0,X4)))))) => (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) => p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X2),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) & (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))))) => (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : (! [X10] : (p(s(bool,happ(s(fun(X6,bool),X8),s(X6,X10)))) & (p(s(bool,happ(s(fun(X6,bool),X7),s(X6,X10)))) => p(s(bool,happ(s(fun(X6,bool),X9),s(X6,X10)))))) => (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X7)))) => p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X8),s(fun(X6,bool),X9))))))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR','thm.quotient.RIGHT_RES_EXISTS_REGULAR')).
fof(f747,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK8,bool),sK10),s(fun(sK8,bool),sK11))) | $spl270),
  inference(cnf_transformation,[],[f747_D])).
fof(f747_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK8,bool),sK10),s(fun(sK8,bool),sK11))) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f741,plain,(
  s(bool,'F_0') = s(bool,'const.bool.RES_EXISTS_2'(s(fun(sK8,bool),sK10),s(fun(sK8,bool),sK11))) | ($spl1 | $spl18)),
  inference(resolution,[],[f560,f544])).
fof(f548,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X4))))) ) | $spl6),
  inference(cnf_transformation,[],[f548_D])).
fof(f548_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X4))))) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f331,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X4))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f719,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X0)))) ) | $spl262),
  inference(cnf_transformation,[],[f719_D])).
fof(f719_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X0)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f713,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X0)))) ) | ($spl6 | $spl14)),
  inference(resolution,[],[f558,f548])).
fof(f582,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) ) | $spl52),
  inference(cnf_transformation,[],[f582_D])).
fof(f582_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f400,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | (p(s(bool,'const.bool.IN_2'(s(X0,sK15(X2,X1,X0)),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK15(X2,X1,X0))))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X4] : (p(s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(rectify,[],[f265])).
fof(f265,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))))) & (! [X3] : (~p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2))))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(flattening,[],[f122])).
fof(f122,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) <=> ? [X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1)))) & p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X6] : ! [X7,X12] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X6,bool),X7),s(fun(X6,bool),X12)))) <=> ? [X10] : (p(s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))) & p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X10))))))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR','thm.bool.RES_EXISTS_THM')).
fof(f578,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) ) | $spl44),
  inference(cnf_transformation,[],[f578_D])).
fof(f578_D,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f392,plain,(
  ( ! [X2,X0,X1] : (s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0] : ! [X1,X2] : s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))) = s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X6] : ! [X7,X10] : s(bool,happ(s(fun(X6,bool),X7),s(X6,X10))) = s(bool,'const.bool.IN_2'(s(X6,X10),s(fun(X6,bool),X7)))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR','thm.pred_set.SPECIFICATION')).
fof(f695,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) ) | $spl248),
  inference(cnf_transformation,[],[f695_D])).
fof(f695_D,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) ) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f693,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))) | p(s(bool,'const.bool.RES_EXISTS_2'(s(fun(X0,bool),X1),s(fun(X0,bool),X2)))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3))))) ) | ($spl44 | $spl52)),
  inference(backward_demodulation,[],[f578,f582])).
fof(f553,plain,(
  ~p(s(bool,'F_0')) | $spl11),
  inference(cnf_transformation,[],[f553_D])).
fof(f553_D,plain,(
  ~p(s(bool,'F_0')) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f551,plain,(
  ~p(s(bool,'F_0'))),
  inference(global_subsumption,[],[f338,f335])).
fof(f338,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f236])).
fof(f236,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f73])).
fof(f73,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f67])).
fof(f67,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR','thm.bool.NOT_CLAUSES|split|2')).
fof(f4771,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X0))))) ) | $spl1810),
  inference(cnf_transformation,[],[f4771_D])).
fof(f4771_D,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X0))))) ) <=> ~$spl1810),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1810])])).
fof(f4770,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X0))))) ) | ($spl8 | $spl11 | $spl248 | $spl262 | $spl270)),
  inference(subsumption_resolution,[],[f4769,f550])).
fof(f4769,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | ~p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X0))))) ) | ($spl11 | $spl248 | $spl262 | $spl270)),
  inference(forward_demodulation,[],[f4768,f719])).
fof(f4768,plain,(
  ( ! [X0] : (~p(s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X0)))) | ~p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X0))))) ) | ($spl11 | $spl248 | $spl270)),
  inference(subsumption_resolution,[],[f4744,f553])).
fof(f4744,plain,(
  ( ! [X0] : (p(s(bool,'F_0')) | ~p(s(bool,happ(s(fun(sK8,bool),sK10),s(sK8,X0)))) | ~p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X0))))) ) | ($spl248 | $spl270)),
  inference(superposition,[],[f695,f747])).
fof(f547,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X4)))) | ~p(s(bool,happ(s(fun(sK8,bool),sK9),s(sK8,X4))))) ) | $spl4),
  inference(cnf_transformation,[],[f547_D])).
fof(f547_D,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X4)))) | ~p(s(bool,happ(s(fun(sK8,bool),sK9),s(sK8,X4))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f332,plain,(
  ( ! [X4] : (p(s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X4)))) | ~p(s(bool,happ(s(fun(sK8,bool),sK9),s(sK8,X4))))) )),
  inference(cnf_transformation,[],[f235])).
fof(f720,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK8,bool),sK9),s(sK8,X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X1)))) ) | $spl264),
  inference(cnf_transformation,[],[f720_D])).
fof(f720_D,plain,(
  ( ! [X1] : (~p(s(bool,happ(s(fun(sK8,bool),sK9),s(sK8,X1)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X1)))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f714,plain,(
  ( ! [X1] : (s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,X1))) | ~p(s(bool,happ(s(fun(sK8,bool),sK9),s(sK8,X1))))) ) | ($spl4 | $spl14)),
  inference(resolution,[],[f558,f547])).
fof(f558,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f558_D])).
fof(f558_D,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f349,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f238])).
fof(f238,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR','thm.bool.EQ_CLAUSES|split|0')).
fof(f546,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9)))) | $spl2),
  inference(cnf_transformation,[],[f546_D])).
fof(f546_D,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f333,plain,(
  p(s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9))))),
  inference(cnf_transformation,[],[f235])).
fof(f718,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9))) | $spl260),
  inference(cnf_transformation,[],[f718_D])).
fof(f718_D,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f712,plain,(
  s(bool,'T_0') = s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9))) | ($spl2 | $spl14)),
  inference(resolution,[],[f558,f546])).
fof(f573,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,X0))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) | $spl34),
  inference(cnf_transformation,[],[f573_D])).
fof(f573_D,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,X0))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f386,plain,(
  ( ! [X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,X0))))) | ~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK12(X1,X0)))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f259])).
fof(f259,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(rectify,[],[f258])).
fof(f258,plain,(
  ! [X0,X1] : ((~p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) | ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2))))) & (! [X2] : ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))) | p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1))))))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1] : (p(s(bool,'|dollar|exists_1'(s(fun(X0,bool),X1)))) <=> ? [X2] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X2)))))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6] : ! [X12] : (p(s(bool,'|dollar|exists_1'(s(fun(X6,bool),X12)))) <=> ? [X10] : p(s(bool,happ(s(fun(X6,bool),X12),s(X6,X10)))))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR','thm.bool.EXISTS_THM')).
fof(f2072,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,sK12(sK9,sK8)))) | $spl796),
  inference(cnf_transformation,[],[f2072_D])).
fof(f2072_D,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,sK12(sK9,sK8)))) <=> ~$spl796),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl796])])).
fof(f2070,plain,(
  s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,sK12(sK9,sK8)))) | ($spl8 | $spl34 | $spl260 | $spl264)),
  inference(subsumption_resolution,[],[f2069,f550])).
fof(f2069,plain,(
  ~p(s(bool,'T_0')) | s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,sK12(sK9,sK8)))) | ($spl34 | $spl260 | $spl264)),
  inference(forward_demodulation,[],[f2045,f718])).
fof(f2045,plain,(
  ~p(s(bool,'|dollar|exists_1'(s(fun(sK8,bool),sK9)))) | s(bool,'T_0') = s(bool,happ(s(fun(sK8,bool),sK11),s(sK8,sK12(sK9,sK8)))) | ($spl34 | $spl264)),
  inference(resolution,[],[f573,f720])).
fof(f550,plain,(
  p(s(bool,'T_0')) | $spl8),
  inference(cnf_transformation,[],[f550_D])).
fof(f550_D,plain,(
  p(s(bool,'T_0')) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f335,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.quotient.RIGHT_RES_EXISTS_REGULAR',aHL_TRUTH)).
fof(f4832,plain,(
  $false | ($spl8 | $spl796 | $spl1810)),
  inference(subsumption_resolution,[],[f4804,f550])).
fof(f4804,plain,(
  ~p(s(bool,'T_0')) | ($spl796 | $spl1810)),
  inference(superposition,[],[f4771,f2072])).
% SZS output end Proof for thm.quotient
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 213
% Passive clauses: 890
% Generated clauses: 4432
% Final active clauses: 207
% Final passive clauses: 663
% Input formulas: 69
% Initial clauses: 194
% 
% Duplicate literals: 44
% Trivial inequalities: 6
% Fw subsumption resolutions: 147
% Fw demodulations: 738
% Bw demodulations: 20
% Condensations: 18
% Global subsumptions: 14
% 
% Simple tautologies: 143
% Equational tautologies: 6
% Forward subsumptions: 1380
% 
% Binary resolution: 738
% Factoring: 7
% Forward superposition: 703
% Backward superposition: 885
% Self superposition: 16
% Equality factoring: 2
% Equality resolution: 1
% 
% Split clauses: 16
% Split components: 32
% SAT solver clauses: 1266
% SAT solver unit clauses: 214
% SAT solver binary clauses: 78
% SAT solver learnt clauses: 29
% SAT solver learnt literals: 10
% 
% Sat splits: 16
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 1208
% TWLsolver calls for satisfiability: 1443
% 
% Memory used [KB]: 6524
% Time elapsed: 0.167 s
% ------------------------------
% ------------------------------
% Success in time 0.774 s
