Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.rat
% (99381)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 659
% Passive clauses: 7016
% Generated clauses: 18716
% Final active clauses: 603
% Final passive clauses: 5824
% Input formulas: 61
% Initial clauses: 199
% 
% Duplicate literals: 328
% Trivial inequalities: 123
% Fw subsumption resolutions: 1567
% Fw demodulations: 1694
% Bw demodulations: 48
% 
% Simple tautologies: 224
% Equational tautologies: 60
% Forward subsumptions: 7345
% Fw demodulations to eq. taut.: 24
% 
% Binary resolution: 7847
% Factoring: 1167
% Forward superposition: 3010
% Backward superposition: 2397
% Self superposition: 77
% Equality factoring: 26
% Equality resolution: 36
% 
% Split clauses: 382
% Split components: 770
% SAT solver clauses: 677
% SAT solver unit clauses: 20
% SAT solver binary clauses: 182
% 
% Sat splits: 382
% Sat splitting refutations: 3
% 
% TWLsolver clauses: 675
% TWLsolver calls for satisfiability: 384
% 
% Memory used [KB]: 11257
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.rat
% (99396)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 356
% Passive clauses: 1700
% Generated clauses: 11542
% Final active clauses: 349
% Final passive clauses: 1332
% Input formulas: 61
% Initial clauses: 199
% 
% Duplicate literals: 216
% Trivial inequalities: 65
% Fw subsumption resolutions: 821
% Fw demodulations: 2023
% Bw demodulations: 16
% Condensations: 234
% Global subsumptions: 22
% 
% Simple tautologies: 295
% Equational tautologies: 20
% Forward subsumptions: 3941
% Fw demodulations to eq. taut.: 26
% 
% Binary resolution: 1987
% Factoring: 150
% Forward superposition: 1912
% Backward superposition: 2013
% Self superposition: 81
% Equality factoring: 32
% Equality resolution: 19
% 
% Split clauses: 76
% Split components: 169
% SAT solver clauses: 2370
% SAT solver unit clauses: 202
% SAT solver binary clauses: 73
% SAT solver learnt clauses: 44
% SAT solver learnt literals: 10
% 
% Sat splits: 76
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2282
% TWLsolver calls for satisfiability: 1414
% 
% Memory used [KB]: 16119
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.rat
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.rat
% SZS output start Proof for thm.rat
fof(f9166,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f798,f802,f800,f1438,f1464,f1515,f1518,f1517,f2196,f2200,f2203,f2207,f2210,f2213,f2216,f2219,f2222,f2225,f2228,f2231,f2234,f2237,f2240,f2243,f2269,f2197,f2271,f2201,f2273,f2204,f2276,f2208,f2278,f2211,f2280,f2214,f2282,f2217,f2284,f2220,f2286,f2223,f2288,f2226,f2290,f2229,f2292,f2232,f2294,f2235,f2296,f2238,f2298,f2241,f2300,f2198,f2244,f2319,f2320,f2321,f2883,f2888,f2927,f4122,f4123,f4130,f4127,f4124,f4125,f4133,f4126,f4135,f4128,f4132,f4146,f4147,f4159,f4157,f4673,f4674,f4518,f4682,f4680,f4686,f4676,f4687,f4944,f4948,f4921,f4958,f1462,f5073,f5085,f5147,f2925,f7647,f7728,f7653,f7729,f7674,f7730,f7680,f7731,f7727,f7999,f8004,f8005,f8006,f8007,f8008,f8011,f8010,f8009,f8014,f8019,f8020,f8077,f8078,f8091,f8096,f8099,f8100,f8101,f8172,f8173,f8346,f8407,f8408,f8409,f9165])).
fof(f9165,plain,(
  $false | $spl82),
  inference(subsumption_resolution,[],[f9164,f7728])).
fof(f9164,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))))),
  inference(unit_resulting_resolution,[],[f288,f794,f448])).
fof(f448,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f269,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f143,f190,f189,f188,f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f188,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f189,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f190,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f143,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X22)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X22)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X22))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X22)))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.sat.sat_2|49|_prove_ltb')).
fof(f794,plain,(
  ~sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)))),
  inference(unit_resulting_resolution,[],[f288,f287,f449])).
fof(f449,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f269])).
fof(f287,plain,(
  s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))))),
  inference(cnf_transformation,[],[f198])).
fof(f198,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) & s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12])],[f148])).
fof(f148,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1)))) & s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))) != s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X2))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1)))) => s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X2))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X7)))) => s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))),s('type.rat.rat',X8))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X7)))) => s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))),s('type.rat.rat',X8))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_LDIV_LES_POS')).
fof(f288,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.rat.RAT_LDIV_LES_POS',aHL_TRUTH)).
fof(f8409,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7990,f531])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (sP8('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f367,f478])).
fof(f478,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f283])).
fof(f283,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | ~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f282])).
fof(f282,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f281])).
fof(f281,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f367,plain,(
  ( ! [X0] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0] : ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0] : ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))))),
  inference(rectify,[],[f47])).
fof(f47,axiom,(
  ! [X6] : ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X6))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_LES_REF')).
fof(f7990,plain,(
  ~sP8('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0','T_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f7170])).
fof(f7170,plain,(
  ~sP8('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f288,f496,f7165,f474])).
fof(f474,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f280,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & ~p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f279])).
fof(f279,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f278])).
fof(f278,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & ~p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f196])).
fof(f196,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f7165,plain,(
  ~sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0','T_0')),
  inference(unit_resulting_resolution,[],[f288,f496,f1146])).
fof(f1146,plain,(
  ( ! [X0,X1] : (~sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1139,f288])).
fof(f1139,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0)) )),
  inference(superposition,[],[f482,f536])).
fof(f536,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f286,f311])).
fof(f311,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X9] : (s(bool,'T_0') = s(bool,X9) <=> p(s(bool,X9)))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.bool.EQ_CLAUSES|split|0')).
fof(f286,plain,(
  p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))))),
  inference(cnf_transformation,[],[f198])).
fof(f482,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f285,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | ~p(s(bool,X1)) | p(s(bool,X0))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f284])).
fof(f284,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (~p(s(bool,X1)) | p(s(bool,X0)))) & ((p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f182,f196,f195])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (~p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(ennf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) => p(s(bool,X0)))) <=> ((p(s(bool,X2)) | p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (~p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) => p(s(bool,X20)))) <=> ((p(s(bool,X22)) | p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (~p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.sat.sat_27_prove_ltb')).
fof(f496,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f291,f288])).
fof(f291,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f199])).
fof(f199,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f31])).
fof(f31,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.bool.NOT_CLAUSES|split|2')).
fof(f8408,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7989,f1092])).
fof(f1092,plain,(
  ( ! [X0] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f531,f288,f496,f474])).
fof(f7989,plain,(
  ~sP9('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0','T_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f7165])).
fof(f8407,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7988,f647])).
fof(f647,plain,(
  ( ! [X0,X1] : (sP9('T_0',X0,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1)))) )),
  inference(unit_resulting_resolution,[],[f530,f288,f485])).
fof(f485,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f471])).
fof(f471,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f530,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X2)))) )),
  inference(unit_resulting_resolution,[],[f367,f476])).
fof(f476,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f283])).
fof(f7988,plain,(
  ~sP9('T_0','F_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | $spl80),
  inference(backward_demodulation,[],[f7727,f7154])).
fof(f7154,plain,(
  ~sP9('T_0','F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f6639,f470])).
fof(f470,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f280])).
fof(f6639,plain,(
  ~sP8('T_0','F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f496,f288,f1113])).
fof(f1113,plain,(
  ( ! [X0,X1] : (~sP8(X1,X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))) | ~p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1105,f288])).
fof(f1105,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | ~p(s(bool,X1)) | ~sP8(X1,X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(superposition,[],[f475,f536])).
fof(f475,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f283])).
fof(f8346,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7956,f530])).
fof(f7956,plain,(
  ~sP8('T_0','F_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | $spl80),
  inference(backward_demodulation,[],[f7727,f6639])).
fof(f8173,plain,(
  $false | ($spl3 | $spl46 | $spl80)),
  inference(subsumption_resolution,[],[f8154,f8166])).
fof(f8166,plain,(
  ( ! [X24,X23] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK12)),X23,X24)) ) | ($spl3 | $spl46 | $spl80)),
  inference(subsumption_resolution,[],[f8130,f8013])).
fof(f8013,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl80),
  inference(condensation,[],[f8012])).
fof(f8012,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0))) ) | $spl80),
  inference(subsumption_resolution,[],[f7780,f529])).
fof(f529,plain,(
  ( ! [X2,X0,X1] : (sP6('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f367,f461])).
fof(f461,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f275])).
fof(f275,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(rectify,[],[f274])).
fof(f274,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(flattening,[],[f273])).
fof(f273,plain,(
  ! [X2,X0,X1] : ((~sP6(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP6(X2,X0,X1)))),
  inference(nnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X2,X0,X1] : (sP6(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f7780,plain,(
  ( ! [X0,X1] : (~sP6('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f946])).
fof(f946,plain,(
  ( ! [X0,X1] : (~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f940,f288])).
fof(f940,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(superposition,[],[f458,f536])).
fof(f458,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f275])).
fof(f8130,plain,(
  ( ! [X24,X23] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK12)),X23,X24) | ~p(s(bool,X23))) ) | ($spl3 | $spl46 | $spl80)),
  inference(backward_demodulation,[],[f7837,f1084])).
fof(f1084,plain,(
  ( ! [X24,X23] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X23,X24) | ~p(s(bool,X23))) ) | $spl3),
  inference(subsumption_resolution,[],[f1060,f288])).
fof(f1060,plain,(
  ( ! [X24,X23] : (~p(s(bool,'T_0')) | ~p(s(bool,X23)) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X23,X24)) ) | $spl3),
  inference(superposition,[],[f481,f803])).
fof(f803,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12))))) | $spl3),
  inference(unit_resulting_resolution,[],[f802,f297])).
fof(f297,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.rat.RAT_LDIV_LES_POS',aHL_BOOL_CASES)).
fof(f481,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f7837,plain,(
  ( ! [X1] : (s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))) ) | ($spl46 | $spl80)),
  inference(backward_demodulation,[],[f7727,f5157])).
fof(f5157,plain,(
  ( ! [X1] : (s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X1)))) ) | $spl46),
  inference(backward_demodulation,[],[f2925,f1308])).
fof(f1308,plain,(
  ( ! [X1] : (s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X1)))) )),
  inference(superposition,[],[f295,f372])).
fof(f372,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X17,X18] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X17)))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_MUL_COMM')).
fof(f295,plain,(
  ( ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(cnf_transformation,[],[f67])).
fof(f67,plain,(
  ! [X0] : s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X17] : s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_MUL_RID')).
fof(f8154,plain,(
  ~sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK12)),'F_0','T_0') | ($spl3 | $spl46 | $spl80)),
  inference(backward_demodulation,[],[f7837,f7609])).
fof(f7609,plain,(
  ~sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'F_0','T_0') | $spl3),
  inference(unit_resulting_resolution,[],[f288,f496,f1147])).
fof(f1147,plain,(
  ( ! [X4,X5] : (~sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X5,X4) | p(s(bool,X5)) | ~p(s(bool,X4))) ) | $spl3),
  inference(subsumption_resolution,[],[f1141,f288])).
fof(f1141,plain,(
  ( ! [X4,X5] : (~p(s(bool,'T_0')) | ~p(s(bool,X4)) | p(s(bool,X5)) | ~sP9('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X5,X4)) ) | $spl3),
  inference(superposition,[],[f482,f803])).
fof(f8172,plain,(
  $false | ($spl3 | $spl46 | $spl80)),
  inference(subsumption_resolution,[],[f8151,f8165])).
fof(f8165,plain,(
  ( ! [X15,X16] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK12)),X16,X15)) ) | ($spl3 | $spl46 | $spl80)),
  inference(subsumption_resolution,[],[f8126,f8013])).
fof(f8126,plain,(
  ( ! [X15,X16] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK12)),X16,X15) | ~p(s(bool,X15))) ) | ($spl3 | $spl46 | $spl80)),
  inference(backward_demodulation,[],[f7837,f1080])).
fof(f1080,plain,(
  ( ! [X15,X16] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X16,X15) | ~p(s(bool,X15))) ) | $spl3),
  inference(subsumption_resolution,[],[f1056,f288])).
fof(f1056,plain,(
  ( ! [X15,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X15)) | sP7('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X16,X15)) ) | $spl3),
  inference(superposition,[],[f463,f803])).
fof(f463,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f277,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f276])).
fof(f276,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f194])).
fof(f194,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f145,f193,f192])).
fof(f193,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f145,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X20,X21,X22] : ((p(s(bool,X22)) <=> (p(s(bool,X21)) | p(s(bool,X20)))) <=> ((p(s(bool,X22)) | ~p(s(bool,X21))) & (p(s(bool,X22)) | ~p(s(bool,X20))) & (p(s(bool,X21)) | p(s(bool,X20)) | ~p(s(bool,X22)))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.sat.sat_25_prove_ltb')).
fof(f8151,plain,(
  ~sP7('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK12)),'F_0','F_0') | ($spl3 | $spl46 | $spl80)),
  inference(backward_demodulation,[],[f7837,f7575])).
fof(f7575,plain,(
  ~sP7('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f7565,f453])).
fof(f453,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f272])).
fof(f272,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP6(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f271])).
fof(f271,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f270])).
fof(f270,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP6(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP6(X2,X0,X1)) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f193])).
fof(f7565,plain,(
  ~sP6('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'F_0','F_0') | $spl3),
  inference(unit_resulting_resolution,[],[f496,f496,f1079])).
fof(f1079,plain,(
  ( ! [X14,X13] : (~sP6('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X13,X14) | p(s(bool,X14)) | p(s(bool,X13))) ) | $spl3),
  inference(subsumption_resolution,[],[f1055,f288])).
fof(f1055,plain,(
  ( ! [X14,X13] : (~p(s(bool,'T_0')) | p(s(bool,X13)) | p(s(bool,X14)) | ~sP6('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),X13,X14)) ) | $spl3),
  inference(superposition,[],[f458,f803])).
fof(f8101,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7829,f927])).
fof(f927,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f529,f496,f496,f457])).
fof(f457,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f272])).
fof(f7829,plain,(
  ~sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0','F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f4139])).
fof(f4139,plain,(
  ~sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f4109,f453])).
fof(f4109,plain,(
  ~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f496,f946])).
fof(f8100,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7828,f529])).
fof(f7828,plain,(
  ~sP6('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0','F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f4109])).
fof(f8099,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7826,f1236])).
fof(f1236,plain,(
  ( ! [X0] : (sP5('T_0','F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f510,f506,f525,f496,f433])).
fof(f433,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X1)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f260])).
fof(f260,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f259])).
fof(f259,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f258])).
fof(f258,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f190])).
fof(f525,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f367,f437])).
fof(f437,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f263])).
fof(f263,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f262])).
fof(f262,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f261])).
fof(f261,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f189])).
fof(f506,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f440])).
fof(f440,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f266])).
fof(f266,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f265])).
fof(f265,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f264])).
fof(f264,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f188])).
fof(f510,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f444])).
fof(f444,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f268])).
fof(f268,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f267])).
fof(f267,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f187])).
fof(f7826,plain,(
  ~sP5('T_0','F_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | $spl80),
  inference(backward_demodulation,[],[f7727,f3943])).
fof(f3943,plain,(
  ~sP5('T_0','F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f3789,f430])).
fof(f430,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f260])).
fof(f3789,plain,(
  ~sP4('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f288,f496,f863])).
fof(f863,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f857,f288])).
fof(f857,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1)) )),
  inference(superposition,[],[f435,f536])).
fof(f435,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f263])).
fof(f8096,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7823,f525])).
fof(f7823,plain,(
  ~sP4('T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f3789])).
fof(f8091,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7816,f1569])).
fof(f1569,plain,(
  ( ! [X0] : (sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f288,f552,f448])).
fof(f552,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f367,f315])).
fof(f315,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f203,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f86])).
fof(f86,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X9] : (s(bool,'F_0') = s(bool,X9) <=> ~p(s(bool,X9)))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.bool.EQ_CLAUSES|split|2')).
fof(f7816,plain,(
  ~sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f3549])).
fof(f3549,plain,(
  ~sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f288,f496,f842])).
fof(f842,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f838,f288])).
fof(f838,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP5(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1)) )),
  inference(superposition,[],[f431,f536])).
fof(f431,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP5(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f260])).
fof(f8078,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f8028,f288])).
fof(f8028,plain,(
  ~p(s(bool,'T_0')) | $spl80),
  inference(backward_demodulation,[],[f8025,f496])).
fof(f8025,plain,(
  ( ! [X111] : (s(bool,'T_0') = s(bool,X111)) ) | $spl80),
  inference(global_subsumption,[],[f287,f286,f288,f297,f311,f310,f367,f372,f373,f377,f374,f388,f387,f497,f406,f403,f402,f401,f419,f420,f434,f433,f432,f431,f430,f429,f428,f438,f437,f436,f435,f442,f441,f440,f439,f446,f445,f444,f443,f450,f449,f448,f447,f457,f456,f455,f486,f453,f452,f451,f461,f460,f459,f458,f465,f464,f463,f462,f474,f473,f472,f485,f470,f469,f468,f478,f477,f476,f475,f482,f481,f480,f479,f499,f498,f501,f500,f503,f502,f504,f505,f507,f506,f508,f509,f511,f510,f513,f512,f515,f514,f517,f516,f519,f518,f520,f521,f523,f522,f524,f531,f530,f529,f528,f527,f526,f525,f533,f532,f536,f556,f557,f560,f562,f565,f568,f570,f579,f584,f585,f588,f591,f593,f596,f598,f605,f606,f608,f609,f617,f619,f623,f625,f628,f630,f649,f647,f651,f654,f665,f663,f667,f674,f724,f735,f759,f758,f767,f766,f779,f777,f788,f786,f797,f796,f795,f794,f798,f842,f856,f863,f879,f886,f890,f896,f905,f927,f926,f939,f946,f950,f1009,f1092,f1091,f1104,f1113,f1118,f1146,f1199,f1196,f1195,f1194,f1190,f1188,f1187,f1209,f1245,f1241,f1237,f1236,f1230,f1229,f1294,f1293,f1290,f1282,f1269,f1340,f1375,f1406,f1385,f1409,f1410,f1438,f1468,f1515,f1551,f1548,f1570,f1569,f1568,f1640,f1639,f1691,f1690,f1696,f1750,f1755,f1816,f1881,f1927,f2196,f2200,f2203,f2207,f2210,f2213,f2216,f2219,f2222,f2225,f2228,f2231,f2234,f2237,f2240,f2243,f2245,f2246,f2247,f2248,f2249,f2250,f2251,f2252,f2253,f2254,f2255,f2256,f2257,f2258,f2259,f2260,f2261,f2262,f2263,f2264,f2265,f2266,f1995,f2269,f2271,f2273,f2276,f2278,f2280,f2282,f2284,f2286,f2288,f2290,f2292,f2294,f2296,f2298,f2300,f421,f1497,f1509,f1522,f1523,f2267,f1882,f1996,f1423,f1422,f552,f496,f315,f314,f2883,f2888,f2959,f3098,f3319,f3322,f3389,f3332,f3397,f3431,f3511,f3513,f3449,f3450,f3517,f3518,f3529,f3531,f3549,f3555,f3556,f3557,f3561,f3562,f3565,f3749,f3753,f3754,f3755,f3619,f3620,f3627,f3765,f3773,f3671,f3782,f3728,f3789,f3839,f3840,f3912,f3915,f3917,f3936,f3943,f4072,f4076,f4109,f4122,f4123,f4124,f4125,f4126,f4139,f4146,f4147,f4161,f2890,f2882,f1466,f1440,f1436,f1434,f1308,f308,f295,f4802,f4850,f4852,f4853,f4941,f4945,f4946,f4866,f4951,f4954,f4921,f5073,f5083,f5085,f5135,f6239,f6254,f6264,f6279,f6285,f6294,f6321,f6327,f6366,f6373,f6388,f6415,f6419,f6216,f6446,f6472,f6545,f6551,f6572,f6591,f6595,f6509,f6615,f6515,f6516,f6622,f6627,f6639,f6729,f6742,f6750,f6763,f6769,f6799,f6805,f6827,f6842,f6868,f6874,f6912,f6946,f6962,f6989,f6993,f7034,f6697,f7094,f7096,f7121,f7150,f7154,f7165,f7170,f7354,f7352,f7626,f7625,f7624,f7623,f7644,f7647,f7648,f7650,f7651,f7652,f7671,f7674,f7675,f7677,f7678,f7679,f7683,f7685,f7698,f7699,f7700,f7701,f7702,f7703,f7704,f7705,f7706,f7707,f7708,f7709,f7710,f7711,f7712,f7713,f7714,f7715,f7716,f7717,f7727,f7742,f7743,f7744,f7745,f7746,f7747,f7748,f7749,f7750,f7751,f7752,f7754,f7755,f7756,f7757,f7758,f7759,f7760,f7761,f7762,f7763,f7764,f7765,f7766,f8001,f7769,f7770,f8003,f7777,f8008,f8011,f8013,f7786,f8018,f7790,f7791,f7797,f7799,f8022,f8023,f8024,f7805])).
fof(f7805,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,X111) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',X108))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X109),s('type.rat.rat',X110))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X108),s('type.rat.rat',X109))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X108),s('type.rat.rat',X110))))) | s(bool,'T_0') = s(bool,X111)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f2267])).
fof(f8024,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22)))))))))) ) | $spl80),
  inference(subsumption_resolution,[],[f7802,f8013])).
fof(f7802,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22)))))))))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f1522])).
fof(f8023,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))),s('type.rat.rat',X12))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X14))),s('type.rat.rat',X12)))))) ) | $spl80),
  inference(subsumption_resolution,[],[f7801,f8013])).
fof(f7801,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))),s('type.rat.rat',X12))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X14))),s('type.rat.rat',X12)))))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f1509])).
fof(f8022,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X3)))))))) ) | $spl80),
  inference(subsumption_resolution,[],[f7800,f8013])).
fof(f7800,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X3)))))))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f1497])).
fof(f7799,plain,(
  ( ! [X2,X1] : (s('type.rat.rat',sK11) = s('type.rat.rat',X1) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f1468])).
fof(f7797,plain,(
  ( ! [X0] : (sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f1294])).
fof(f7791,plain,(
  ( ! [X0] : (sP5('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f1237])).
fof(f7790,plain,(
  sP5('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0','F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f1230])).
fof(f8018,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(subsumption_resolution,[],[f8017,f527])).
fof(f8017,plain,(
  ( ! [X0,X1] : (~sP3(X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0) | sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(forward_demodulation,[],[f8016,f7727])).
fof(f8016,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | ~sP3(X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) ) | $spl80),
  inference(forward_demodulation,[],[f8015,f7727])).
fof(f8015,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))) | ~sP3(X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) ) | $spl80),
  inference(subsumption_resolution,[],[f7789,f525])).
fof(f7789,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1) | sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))) | ~sP3(X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f1209])).
fof(f7786,plain,(
  sP5('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'T_0','T_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f1188])).
fof(f7777,plain,(
  ( ! [X0,X1] : (~sP5(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f842])).
fof(f8003,plain,(
  ( ! [X0,X1] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1)) ) | $spl80),
  inference(subsumption_resolution,[],[f8002,f529])).
fof(f8002,plain,(
  ( ! [X0,X1] : (~sP6('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1) | sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1)) ) | $spl80),
  inference(forward_demodulation,[],[f7771,f7727])).
fof(f7771,plain,(
  ( ! [X0,X1] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1) | ~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f667])).
fof(f7770,plain,(
  ( ! [X0] : (sP7('T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f665])).
fof(f7769,plain,(
  ( ! [X0] : (sP7('T_0',X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f663])).
fof(f8001,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1)) ) | $spl80),
  inference(subsumption_resolution,[],[f8000,f531])).
fof(f8000,plain,(
  ( ! [X0,X1] : (~sP8('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1) | sP9('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1)) ) | $spl80),
  inference(forward_demodulation,[],[f7767,f7727])).
fof(f7767,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1) | ~sP8('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f651])).
fof(f7766,plain,(
  ( ! [X0] : (sP9('T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f649])).
fof(f7765,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1) | ~p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f630])).
fof(f7764,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f625])).
fof(f7763,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'T_0',X0)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f623])).
fof(f7762,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1,X0) | p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f619])).
fof(f7761,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f608])).
fof(f7760,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,'T_0')) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f606])).
fof(f7759,plain,(
  ( ! [X0,X1] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1) | ~p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f598])).
fof(f7758,plain,(
  ( ! [X0,X1] : (sP7('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1,X0) | ~p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f593])).
fof(f7757,plain,(
  sP1('F_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | $spl80),
  inference(backward_demodulation,[],[f7727,f584])).
fof(f7756,plain,(
  ( ! [X0] : (sP1(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f570])).
fof(f7755,plain,(
  sP1('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f565])).
fof(f7754,plain,(
  ( ! [X0] : (sP1('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0) | p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f562])).
fof(f7752,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f523])).
fof(f7751,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f519])).
fof(f7750,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f517])).
fof(f7749,plain,(
  ( ! [X0,X1] : (sP2('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0,X1)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f515])).
fof(f7748,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f513])).
fof(f7747,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f511])).
fof(f7746,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f507])).
fof(f7745,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f503])).
fof(f7744,plain,(
  ( ! [X0] : (sP0('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X0)) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f501])).
fof(f7743,plain,(
  ( ! [X0] : (sP0(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f499])).
fof(f7742,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',X2)))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f421])).
fof(f7717,plain,(
  ( ! [X94,X92,X93] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X92),s('type.rat.rat',X93)))) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X92))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X93)))),X94)) )),
  inference(superposition,[],[f598,f1523])).
fof(f7716,plain,(
  ( ! [X90,X91,X89] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X89),s('type.rat.rat',X90)))) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X91,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X89))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X90)))))) )),
  inference(superposition,[],[f593,f1523])).
fof(f7715,plain,(
  ( ! [X88,X87,X85,X86] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86)))) | ~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X85))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X86)))),X87,X88) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X85))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X86)))),X87,X88)) )),
  inference(superposition,[],[f486,f1523])).
fof(f7714,plain,(
  ( ! [X83,X81,X84,X82] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X81),s('type.rat.rat',X82)))) | ~sP8('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X81))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X82)))),X83,X84) | sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X81))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X82)))),X83,X84)) )),
  inference(superposition,[],[f485,f1523])).
fof(f7713,plain,(
  ( ! [X80,X78,X79,X77] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X77),s('type.rat.rat',X78)))) | ~p(s(bool,X79)) | p(s(bool,X80)) | ~sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X77))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X78)))),X80,X79)) )),
  inference(superposition,[],[f482,f1523])).
fof(f7712,plain,(
  ( ! [X76,X74,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X73),s('type.rat.rat',X74)))) | ~p(s(bool,X75)) | sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X73))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X74)))),X75,X76)) )),
  inference(superposition,[],[f481,f1523])).
fof(f7711,plain,(
  ( ! [X70,X72,X71,X69] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X69),s('type.rat.rat',X70)))) | p(s(bool,X71)) | sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X69))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X70)))),X72,X71)) )),
  inference(superposition,[],[f480,f1523])).
fof(f7710,plain,(
  ( ! [X68,X66,X67,X65] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66)))) | p(s(bool,X67)) | ~p(s(bool,X68)) | ~sP8(X68,X67,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X65))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X66)))))) )),
  inference(superposition,[],[f475,f1523])).
fof(f7709,plain,(
  ( ! [X61,X64,X62,X63] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62)))) | p(s(bool,X63)) | p(s(bool,X64)) | ~sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X61))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X62)))),X64,X63)) )),
  inference(superposition,[],[f465,f1523])).
fof(f7708,plain,(
  ( ! [X59,X57,X60,X58] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X57),s('type.rat.rat',X58)))) | ~p(s(bool,X59)) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X57))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X58)))),X59,X60)) )),
  inference(superposition,[],[f464,f1523])).
fof(f7707,plain,(
  ( ! [X54,X56,X55,X53] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X53),s('type.rat.rat',X54)))) | ~p(s(bool,X55)) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X53))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X54)))),X56,X55)) )),
  inference(superposition,[],[f463,f1523])).
fof(f7706,plain,(
  ( ! [X52,X50,X51,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X49),s('type.rat.rat',X50)))) | p(s(bool,X51)) | p(s(bool,X52)) | ~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X49))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X50)))),X51,X52)) )),
  inference(superposition,[],[f458,f1523])).
fof(f7705,plain,(
  ( ! [X47,X45,X48,X46] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46)))) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP3(X47,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X45))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X46)))),X48)) )),
  inference(superposition,[],[f439,f1523])).
fof(f7704,plain,(
  ( ! [X43,X41,X44,X42] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42)))) | ~p(s(bool,X43)) | p(s(bool,X44)) | ~sP4(X43,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X41))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X42)))),X44)) )),
  inference(superposition,[],[f435,f1523])).
fof(f7703,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X37),s('type.rat.rat',X38)))) | sP5(X39,X40,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X37))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X38))))) | ~sP4(X39,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X37))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X38)))),X40) | ~sP3(X40,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X37))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X38)))),X39) | ~sP2('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X37))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X38)))),X40,X39)) )),
  inference(superposition,[],[f432,f1523])).
fof(f7702,plain,(
  ( ! [X35,X33,X36,X34] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X33),s('type.rat.rat',X34)))) | ~p(s(bool,X35)) | p(s(bool,X36)) | ~sP5(X35,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X33))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X34)))),X36)) )),
  inference(superposition,[],[f431,f1523])).
fof(f7701,plain,(
  ( ! [X30,X31,X32] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X30),s('type.rat.rat',X31)))) | p(s(bool,X32)) | sP1(X32,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X30))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X31)))))) )),
  inference(superposition,[],[f406,f1523])).
fof(f7700,plain,(
  ( ! [X28,X29,X27] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X27),s('type.rat.rat',X28)))) | p(s(bool,X29)) | sP1('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X27))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X28)))),X29)) )),
  inference(superposition,[],[f403,f1523])).
fof(f7699,plain,(
  ( ! [X26,X24,X25] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X24),s('type.rat.rat',X25)))) | sP0('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X24))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X25)))),X26)) )),
  inference(superposition,[],[f388,f1523])).
fof(f7698,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22)))) | sP0(X23,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X22)))))) )),
  inference(superposition,[],[f387,f1523])).
fof(f7685,plain,(
  ( ! [X12,X13,X11] : (s(bool,X13) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X11))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X12))))) | s(bool,'T_0') = s(bool,X13)) )),
  inference(superposition,[],[f1523,f674])).
fof(f7683,plain,(
  ( ! [X6,X7,X5] : (s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X5),s('type.rat.rat',X6))) | s(bool,X7) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X5))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X6))))) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f1523,f674])).
fof(f7679,plain,(
  ( ! [X19,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X18))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X19))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK11),s('type.rat.rat',X18)))))) )),
  inference(superposition,[],[f1523,f373])).
fof(f7678,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X15)))))))) )),
  inference(superposition,[],[f1523,f1385])).
fof(f7677,plain,(
  ( ! [X14,X12,X13] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',X13))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X14))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',sK11)))))))) )),
  inference(superposition,[],[f1523,f1385])).
fof(f7675,plain,(
  ( ! [X8,X7,X9] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X9),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X9))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X7))),s('type.rat.rat',X8)))))) )),
  inference(superposition,[],[f1523,f1409])).
fof(f7671,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',sK11)))))) )),
  inference(superposition,[],[f1523,f372])).
fof(f7652,plain,(
  ( ! [X19,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X18))),s('type.rat.rat',X19))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK11),s('type.rat.rat',X18))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X19)))))) )),
  inference(superposition,[],[f1523,f373])).
fof(f7651,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16))),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X15))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X17)))))) )),
  inference(superposition,[],[f1523,f1385])).
fof(f7650,plain,(
  ( ! [X14,X12,X13] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',X13))),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',sK11))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X14)))))) )),
  inference(superposition,[],[f1523,f1385])).
fof(f7648,plain,(
  ( ! [X8,X7,X9] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))),s('type.rat.rat',X9))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X7))),s('type.rat.rat',X8))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X9)))))) )),
  inference(superposition,[],[f1523,f1409])).
fof(f7644,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',sK11))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))))) )),
  inference(superposition,[],[f1523,f372])).
fof(f7623,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))),'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(unit_resulting_resolution,[],[f496,f1523,f450])).
fof(f7624,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))),'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(unit_resulting_resolution,[],[f288,f1523,f448])).
fof(f7625,plain,(
  ( ! [X0,X1] : (~sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f496,f1523,f450])).
fof(f7626,plain,(
  ( ! [X0,X1] : (sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))))) )),
  inference(unit_resulting_resolution,[],[f288,f1523,f448])).
fof(f7352,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1)))) )),
  inference(unit_resulting_resolution,[],[f530,f651])).
fof(f7354,plain,(
  ( ! [X0] : (sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(unit_resulting_resolution,[],[f523,f651])).
fof(f7150,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X4))))))))) | s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(subsumption_resolution,[],[f6722,f288])).
fof(f6722,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X4))))))))) | s(bool,X5) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1522,f674])).
fof(f7121,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))))))))) )),
  inference(forward_demodulation,[],[f7120,f420])).
fof(f7120,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))))))))) )),
  inference(forward_demodulation,[],[f6710,f420])).
fof(f6710,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72)))))))))) )),
  inference(superposition,[],[f1522,f420])).
fof(f7096,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X40))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))))))))))) )),
  inference(forward_demodulation,[],[f7095,f420])).
fof(f7095,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat',X40))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))))))))))) )),
  inference(forward_demodulation,[],[f6704,f420])).
fof(f6704,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat',X40))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat',X41)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f7094,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X34))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X35))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X30))))))))))) )),
  inference(forward_demodulation,[],[f7093,f420])).
fof(f7093,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))),s('type.rat.rat',X34))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X35))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X30))))))))))) )),
  inference(forward_demodulation,[],[f6703,f420])).
fof(f6703,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X30)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))),s('type.rat.rat',X34))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))),s('type.rat.rat',X35)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f6697,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X3))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X4)))))))))) )),
  inference(superposition,[],[f1522,f372])).
fof(f7034,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X78))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))))))))))))) )),
  inference(forward_demodulation,[],[f7033,f420])).
fof(f7033,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X78))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))))))))))))) )),
  inference(forward_demodulation,[],[f7032,f420])).
fof(f7032,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))))))))))))) )),
  inference(forward_demodulation,[],[f7031,f420])).
fof(f7031,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X79))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))))))))))))) )),
  inference(forward_demodulation,[],[f6692,f420])).
fof(f6692,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat',X79)))))))))) )),
  inference(superposition,[],[f1522,f1375])).
fof(f6993,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat',X48))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))))))))))))) )),
  inference(forward_demodulation,[],[f6992,f420])).
fof(f6992,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))),s('type.rat.rat',X48))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))))))))))))) )),
  inference(forward_demodulation,[],[f6991,f420])).
fof(f6991,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))))))))))))) )),
  inference(forward_demodulation,[],[f6990,f420])).
fof(f6990,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))),s('type.rat.rat',X49))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))))))))))))) )),
  inference(forward_demodulation,[],[f6687,f420])).
fof(f6687,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat',X49)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f6989,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X41))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))))))))))))) )),
  inference(forward_demodulation,[],[f6988,f420])).
fof(f6988,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat',X41))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))))))))))))) )),
  inference(forward_demodulation,[],[f6987,f420])).
fof(f6987,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))))))))))))) )),
  inference(forward_demodulation,[],[f6986,f420])).
fof(f6986,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat',X42))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))))))))))))) )),
  inference(forward_demodulation,[],[f6686,f420])).
fof(f6686,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X42)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f6962,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X21))))))))),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6961,f1409])).
fof(f6961,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X21))))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6960,f1409])).
fof(f6960,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X21))))),s('type.rat.rat',X17))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6959,f1409])).
fof(f6959,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X21))),s('type.rat.rat',X17))))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6958,f1410])).
fof(f6958,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))),s('type.rat.rat',X21))))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6957,f420])).
fof(f6957,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6956,f1409])).
fof(f6956,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))))),s('type.rat.rat',X17))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6955,f1409])).
fof(f6955,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))))),s('type.rat.rat',X17))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6954,f1409])).
fof(f6954,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X22))),s('type.rat.rat',X17))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6953,f1410])).
fof(f6953,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))),s('type.rat.rat',X22))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6683,f420])).
fof(f6683,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X18)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat',X22))))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1522,f1468])).
fof(f6946,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X4))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X5))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))))))) )),
  inference(forward_demodulation,[],[f6945,f420])).
fof(f6945,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X5))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))))))) )),
  inference(forward_demodulation,[],[f6680,f420])).
fof(f6680,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X5)))))))))) )),
  inference(superposition,[],[f1522,f372])).
fof(f6912,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6911,f420])).
fof(f6911,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6910,f420])).
fof(f6910,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6909,f420])).
fof(f6909,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6908,f420])).
fof(f6908,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat',X77)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79)))))))))) )),
  inference(forward_demodulation,[],[f6907,f420])).
fof(f6907,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))),s('type.rat.rat',X77)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79)))))))))) )),
  inference(forward_demodulation,[],[f6675,f420])).
fof(f6675,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))))),s('type.rat.rat',X77)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79)))))))))) )),
  inference(superposition,[],[f1522,f1375])).
fof(f6874,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6873,f420])).
fof(f6873,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6872,f420])).
fof(f6872,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6871,f420])).
fof(f6871,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6870,f420])).
fof(f6870,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X47)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49)))))))))) )),
  inference(forward_demodulation,[],[f6869,f420])).
fof(f6869,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))),s('type.rat.rat',X47)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49)))))))))) )),
  inference(forward_demodulation,[],[f6670,f420])).
fof(f6670,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))))),s('type.rat.rat',X47)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f6868,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6867,f420])).
fof(f6867,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6866,f420])).
fof(f6866,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6865,f420])).
fof(f6865,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6864,f420])).
fof(f6864,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X40)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42)))))))))) )),
  inference(forward_demodulation,[],[f6863,f420])).
fof(f6863,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))),s('type.rat.rat',X40)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42)))))))))) )),
  inference(forward_demodulation,[],[f6669,f420])).
fof(f6669,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))))),s('type.rat.rat',X40)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f6842,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))))),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6841,f1409])).
fof(f6841,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6840,f1409])).
fof(f6840,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X17))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6839,f1410])).
fof(f6839,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6838,f420])).
fof(f6838,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6837,f1409])).
fof(f6837,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))),s('type.rat.rat',X17))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6836,f1409])).
fof(f6836,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X17))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6835,f1410])).
fof(f6835,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6666,f420])).
fof(f6666,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1522,f1468])).
fof(f6827,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3))))))))))) )),
  inference(forward_demodulation,[],[f6826,f420])).
fof(f6826,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3))))))))))) )),
  inference(forward_demodulation,[],[f6825,f420])).
fof(f6825,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5)))))))))) )),
  inference(forward_demodulation,[],[f6663,f420])).
fof(f6663,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat',X3)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5)))))))))) )),
  inference(superposition,[],[f1522,f372])).
fof(f6805,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X86))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat',X84))))))))))))) )),
  inference(forward_demodulation,[],[f6804,f420])).
fof(f6804,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X86))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat',X84))))))))))))) )),
  inference(forward_demodulation,[],[f6803,f420])).
fof(f6803,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X86))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat',X84))))))))))))) )),
  inference(forward_demodulation,[],[f6802,f420])).
fof(f6802,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X86))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat',X84))))))))))))) )),
  inference(forward_demodulation,[],[f6801,f420])).
fof(f6801,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X82),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X86)))))))))) )),
  inference(forward_demodulation,[],[f6800,f420])).
fof(f6800,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X86)))))))))) )),
  inference(forward_demodulation,[],[f6659,f420])).
fof(f6659,plain,(
  ( ! [X80,X85,X83,X81,X86,X84,X82] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X81),s('type.rat.rat',X82))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat',X84)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X85),s('type.rat.rat',X86))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X85))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X80),s('type.rat.rat',X81))),s('type.rat.rat',X82))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X83),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X84),s('type.rat.rat',X86)))))))))) )),
  inference(superposition,[],[f1522,f420])).
fof(f6799,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6798,f420])).
fof(f6798,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6797,f420])).
fof(f6797,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6796,f420])).
fof(f6796,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat',X77))))))))))))) )),
  inference(forward_demodulation,[],[f6795,f420])).
fof(f6795,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X75),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat',X77)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79)))))))))) )),
  inference(forward_demodulation,[],[f6794,f420])).
fof(f6794,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat',X77)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79)))))))))) )),
  inference(forward_demodulation,[],[f6658,f420])).
fof(f6658,plain,(
  ( ! [X78,X76,X74,X79,X77,X75,X73] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X74),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X75))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat',X77)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X78),s('type.rat.rat',X79))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X78))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X73),s('type.rat.rat',X74))),s('type.rat.rat',X75))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X76),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X77),s('type.rat.rat',X79)))))))))) )),
  inference(superposition,[],[f1522,f1375])).
fof(f6769,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6768,f420])).
fof(f6768,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6767,f420])).
fof(f6767,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6766,f420])).
fof(f6766,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat',X47))))))))))))) )),
  inference(forward_demodulation,[],[f6765,f420])).
fof(f6765,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat',X47)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49)))))))))) )),
  inference(forward_demodulation,[],[f6764,f420])).
fof(f6764,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat',X47)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49)))))))))) )),
  inference(forward_demodulation,[],[f6653,f420])).
fof(f6653,plain,(
  ( ! [X47,X45,X43,X48,X46,X44,X49] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X45),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat',X44))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat',X47)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X48),s('type.rat.rat',X49))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X48))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X43),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X46),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X47),s('type.rat.rat',X49)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f6763,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6762,f420])).
fof(f6762,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6761,f420])).
fof(f6761,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6760,f420])).
fof(f6760,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))))))))))) )),
  inference(forward_demodulation,[],[f6759,f420])).
fof(f6759,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40)))))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42)))))))))) )),
  inference(forward_demodulation,[],[f6758,f420])).
fof(f6758,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42)))))))))) )),
  inference(forward_demodulation,[],[f6652,f420])).
fof(f6652,plain,(
  ( ! [X39,X37,X41,X38,X36,X42,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X36))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X41),s('type.rat.rat',X42))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X40),s('type.rat.rat',X42)))))))))) )),
  inference(superposition,[],[f1522,f1385])).
fof(f6750,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))))),s('type.rat.rat',X25))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29))))))))),s('type.rat.rat',X25))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))),s('type.rat.rat',X25))))))) )),
  inference(forward_demodulation,[],[f6749,f1409])).
fof(f6749,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29))))))))),s('type.rat.rat',X25))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))),s('type.rat.rat',X25))))))) )),
  inference(forward_demodulation,[],[f6748,f1410])).
fof(f6748,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29))))))))),s('type.rat.rat',X25))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))),s('type.rat.rat',X25))))))) )),
  inference(forward_demodulation,[],[f6747,f420])).
fof(f6747,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29))))))))),s('type.rat.rat',X25))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))),s('type.rat.rat',X25))))))) )),
  inference(forward_demodulation,[],[f6746,f1409])).
fof(f6746,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29))))))),s('type.rat.rat',X25))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))),s('type.rat.rat',X25))))))) )),
  inference(forward_demodulation,[],[f6745,f1410])).
fof(f6745,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))),s('type.rat.rat',X25))))))) )),
  inference(forward_demodulation,[],[f6744,f420])).
fof(f6744,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))),s('type.rat.rat',X25)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29)))))))))) )),
  inference(forward_demodulation,[],[f6743,f420])).
fof(f6743,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))),s('type.rat.rat',X25)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29)))))))))) )),
  inference(forward_demodulation,[],[f6650,f1410])).
fof(f6650,plain,(
  ( ! [X28,X26,X24,X23,X29,X27,X25] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))),s('type.rat.rat',X25))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X28),s('type.rat.rat',X29))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X28))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat',X29)))))))))) )),
  inference(superposition,[],[f1522,f1409])).
fof(f6742,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))))),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6741,f1409])).
fof(f6741,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6740,f1410])).
fof(f6740,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6739,f420])).
fof(f6739,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))),s('type.rat.rat',X17))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6738,f1409])).
fof(f6738,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))),s('type.rat.rat',X17))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6737,f1410])).
fof(f6737,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6649,f420])).
fof(f6649,plain,(
  ( ! [X21,X19,X17,X22,X20,X18] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X18),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))) | s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1522,f1468])).
fof(f6729,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))))))))))) )),
  inference(forward_demodulation,[],[f6728,f420])).
fof(f6728,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))))))))))) )),
  inference(forward_demodulation,[],[f6727,f420])).
fof(f6727,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5)))))))))) )),
  inference(forward_demodulation,[],[f6646,f420])).
fof(f6646,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5)))))))))) )),
  inference(superposition,[],[f1522,f372])).
fof(f6627,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3))),s('type.rat.rat',X1))))) | s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6527,f288])).
fof(f6527,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))),s('type.rat.rat',X1))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3))),s('type.rat.rat',X1))))) | s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1509,f674])).
fof(f6622,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X65))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat',X63))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X66))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat',X63))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))))))))) )),
  inference(forward_demodulation,[],[f6521,f420])).
fof(f6521,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X65))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X66))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63)))))))) )),
  inference(superposition,[],[f1509,f1375])).
fof(f6516,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X41))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38)))))))))) )),
  inference(superposition,[],[f1509,f1385])).
fof(f6515,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X30)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat',X34))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat',X35))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32)))))))))) )),
  inference(superposition,[],[f1509,f1385])).
fof(f6615,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15))),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15))),s('type.rat.rat',X14))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6512,f1409])).
fof(f6512,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1509,f1468])).
fof(f6509,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X4))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))))))) )),
  inference(superposition,[],[f1509,f372])).
fof(f6595,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))))),s('type.rat.rat',X70))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f6594,f420])).
fof(f6594,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))),s('type.rat.rat',X70))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f6593,f420])).
fof(f6593,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))),s('type.rat.rat',X70))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f6592,f420])).
fof(f6592,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))),s('type.rat.rat',X70))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f6505,f420])).
fof(f6505,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))),s('type.rat.rat',X70)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))),s('type.rat.rat',X70))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X72))),s('type.rat.rat',X70)))))) )),
  inference(superposition,[],[f1509,f420])).
fof(f6591,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X65))))))),s('type.rat.rat',X64))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))))),s('type.rat.rat',X64))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))),s('type.rat.rat',X64))))))) )),
  inference(forward_demodulation,[],[f6590,f420])).
fof(f6590,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X65))))),s('type.rat.rat',X64))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))))),s('type.rat.rat',X64))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))),s('type.rat.rat',X64))))))) )),
  inference(forward_demodulation,[],[f6589,f420])).
fof(f6589,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))),s('type.rat.rat',X64))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))))),s('type.rat.rat',X64))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))),s('type.rat.rat',X64))))))) )),
  inference(forward_demodulation,[],[f6588,f420])).
fof(f6588,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))),s('type.rat.rat',X64))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))),s('type.rat.rat',X64))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))),s('type.rat.rat',X64))))))) )),
  inference(forward_demodulation,[],[f6504,f420])).
fof(f6504,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))),s('type.rat.rat',X64)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))),s('type.rat.rat',X64))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X66))),s('type.rat.rat',X64)))))) )),
  inference(superposition,[],[f1509,f1375])).
fof(f6572,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X45))),s('type.rat.rat',X43))),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))),s('type.rat.rat',X43))),s('type.rat.rat',X44))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44))))))) )),
  inference(forward_demodulation,[],[f6571,f1409])).
fof(f6571,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X45),s('type.rat.rat',X43))))),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))),s('type.rat.rat',X43))),s('type.rat.rat',X44))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44))))))) )),
  inference(forward_demodulation,[],[f6570,f1340])).
fof(f6570,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))),s('type.rat.rat',X45))))),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))),s('type.rat.rat',X43))),s('type.rat.rat',X44))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44))))))) )),
  inference(forward_demodulation,[],[f6569,f420])).
fof(f6569,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))),s('type.rat.rat',X43))),s('type.rat.rat',X44))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44))))))) )),
  inference(forward_demodulation,[],[f6568,f1409])).
fof(f6568,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X46),s('type.rat.rat',X43))))),s('type.rat.rat',X44))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44))))))) )),
  inference(forward_demodulation,[],[f6567,f1340])).
fof(f6567,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))),s('type.rat.rat',X46))))),s('type.rat.rat',X44))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44))))))) )),
  inference(forward_demodulation,[],[f6500,f420])).
fof(f6500,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))),s('type.rat.rat',X44))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X46))),s('type.rat.rat',X44)))))) )),
  inference(superposition,[],[f1509,f373])).
fof(f6551,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X23))))),s('type.rat.rat',X21))),s('type.rat.rat',X22))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))))) )),
  inference(forward_demodulation,[],[f6550,f1409])).
fof(f6550,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X23))),s('type.rat.rat',X21))))),s('type.rat.rat',X22))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))))) )),
  inference(forward_demodulation,[],[f6549,f1410])).
fof(f6549,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))),s('type.rat.rat',X23))))),s('type.rat.rat',X22))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))))) )),
  inference(forward_demodulation,[],[f6548,f420])).
fof(f6548,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))),s('type.rat.rat',X22))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))))) )),
  inference(forward_demodulation,[],[f6547,f1409])).
fof(f6547,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))),s('type.rat.rat',X22))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))),s('type.rat.rat',X21))))),s('type.rat.rat',X22))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))))) )),
  inference(forward_demodulation,[],[f6546,f1410])).
fof(f6546,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))),s('type.rat.rat',X22))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))),s('type.rat.rat',X24))))),s('type.rat.rat',X22))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22))))))) )),
  inference(forward_demodulation,[],[f6496,f420])).
fof(f6496,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))),s('type.rat.rat',X22))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X24))),s('type.rat.rat',X22)))))) )),
  inference(superposition,[],[f1509,f1409])).
fof(f6545,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))))),s('type.rat.rat',X14))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))),s('type.rat.rat',X14))),s('type.rat.rat',X16))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6544,f1409])).
fof(f6544,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X14))))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))),s('type.rat.rat',X14))),s('type.rat.rat',X16))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6543,f1410])).
fof(f6543,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))),s('type.rat.rat',X17))))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))),s('type.rat.rat',X14))),s('type.rat.rat',X16))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6542,f420])).
fof(f6542,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))),s('type.rat.rat',X14))),s('type.rat.rat',X16))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6541,f1409])).
fof(f6541,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))),s('type.rat.rat',X14))))),s('type.rat.rat',X16))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6540,f1410])).
fof(f6540,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))),s('type.rat.rat',X18))))),s('type.rat.rat',X16))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6495,f420])).
fof(f6495,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X18))),s('type.rat.rat',X16))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1509,f1468])).
fof(f6472,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3))))))) | s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(subsumption_resolution,[],[f6232,f288])).
fof(f6232,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'T_0')) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3))))))) | s(bool,X4) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f1497,f674])).
fof(f6446,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X12))))),s('type.rat.rat',X11))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X12))))),s('type.rat.rat',X11))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X12)))) | s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6445,f1409])).
fof(f6445,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X12))),s('type.rat.rat',X11))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X12))))),s('type.rat.rat',X11))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X12)))) | s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6444,f1409])).
fof(f6444,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X11))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X12))))),s('type.rat.rat',X11))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X12)))) | s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6443,f1385])).
fof(f6443,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X11))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X12))))),s('type.rat.rat',X11))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X12)))) | s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6442,f1409])).
fof(f6442,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X11))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X12))),s('type.rat.rat',X11))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X12)))) | s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6441,f1409])).
fof(f6441,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X11))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X11))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X12)))) | s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6219,f1385])).
fof(f6219,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X12)))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X11))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X11))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X14))))))) | s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1497,f1468])).
fof(f6216,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X3)))))))) )),
  inference(superposition,[],[f1497,f372])).
fof(f6419,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X71))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))))))))) )),
  inference(forward_demodulation,[],[f6418,f420])).
fof(f6418,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X71))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))))))))) )),
  inference(forward_demodulation,[],[f6417,f420])).
fof(f6417,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))))))))) )),
  inference(forward_demodulation,[],[f6416,f420])).
fof(f6416,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X72))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))))))))) )),
  inference(forward_demodulation,[],[f6212,f420])).
fof(f6212,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat',X72)))))))) )),
  inference(superposition,[],[f1497,f420])).
fof(f6415,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X65))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X66))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))))))))) )),
  inference(forward_demodulation,[],[f6414,f420])).
fof(f6414,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X65))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X66))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))))))))) )),
  inference(forward_demodulation,[],[f6413,f420])).
fof(f6413,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X66))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))))))))) )),
  inference(forward_demodulation,[],[f6412,f420])).
fof(f6412,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X66))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))))))))) )),
  inference(forward_demodulation,[],[f6211,f420])).
fof(f6211,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat',X66)))))))) )),
  inference(superposition,[],[f1497,f1375])).
fof(f6388,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))))))))))) )),
  inference(forward_demodulation,[],[f6387,f420])).
fof(f6387,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))))))))))) )),
  inference(forward_demodulation,[],[f6386,f420])).
fof(f6386,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))))))))))) )),
  inference(forward_demodulation,[],[f6385,f420])).
fof(f6385,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X41))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))))))))))) )),
  inference(forward_demodulation,[],[f6206,f420])).
fof(f6206,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat',X41)))))))) )),
  inference(superposition,[],[f1497,f1385])).
fof(f6373,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))))),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6372,f1409])).
fof(f6372,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6371,f1410])).
fof(f6371,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6370,f420])).
fof(f6370,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6369,f1409])).
fof(f6369,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24))))),s('type.rat.rat',X21))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6368,f1410])).
fof(f6368,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6367,f420])).
fof(f6367,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))),s('type.rat.rat',X21)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24)))))))) )),
  inference(forward_demodulation,[],[f6203,f1409])).
fof(f6203,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24)))))))) )),
  inference(superposition,[],[f1497,f1409])).
fof(f6366,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))))),s('type.rat.rat',X14))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6365,f1409])).
fof(f6365,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6364,f1410])).
fof(f6364,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6363,f420])).
fof(f6363,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6362,f1409])).
fof(f6362,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))))),s('type.rat.rat',X14))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6361,f1410])).
fof(f6361,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6202,f420])).
fof(f6202,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X15)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X18))))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1497,f1468])).
fof(f6327,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))))))))) )),
  inference(forward_demodulation,[],[f6326,f420])).
fof(f6326,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))))))))) )),
  inference(forward_demodulation,[],[f6325,f420])).
fof(f6325,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))))))))) )),
  inference(forward_demodulation,[],[f6324,f420])).
fof(f6324,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))))))))) )),
  inference(forward_demodulation,[],[f6323,f420])).
fof(f6323,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X70)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X72)))))))) )),
  inference(forward_demodulation,[],[f6322,f420])).
fof(f6322,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X72)))))))) )),
  inference(forward_demodulation,[],[f6195,f420])).
fof(f6195,plain,(
  ( ! [X70,X68,X72,X71,X69,X67] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))))),s('type.rat.rat',X70)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X70),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X67),s('type.rat.rat',X68))),s('type.rat.rat',X69))),s('type.rat.rat',X72)))))))) )),
  inference(superposition,[],[f1497,f420])).
fof(f6321,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X65))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X64))))))))))) )),
  inference(forward_demodulation,[],[f6320,f420])).
fof(f6320,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X65))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X64))))))))))) )),
  inference(forward_demodulation,[],[f6319,f420])).
fof(f6319,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X64))))))))))) )),
  inference(forward_demodulation,[],[f6318,f420])).
fof(f6318,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X66))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X64))))))))))) )),
  inference(forward_demodulation,[],[f6317,f420])).
fof(f6317,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X63),s('type.rat.rat',X64)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X66)))))))) )),
  inference(forward_demodulation,[],[f6316,f420])).
fof(f6316,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))),s('type.rat.rat',X64)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X66)))))))) )),
  inference(forward_demodulation,[],[f6194,f420])).
fof(f6194,plain,(
  ( ! [X61,X66,X64,X62,X65,X63] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X62),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X63))))),s('type.rat.rat',X64)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X65),s('type.rat.rat',X66))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X65))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X64),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))),s('type.rat.rat',X63))),s('type.rat.rat',X66)))))))) )),
  inference(superposition,[],[f1497,f1375])).
fof(f6294,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X45))))),s('type.rat.rat',X43))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))))),s('type.rat.rat',X43))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6293,f1409])).
fof(f6293,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X45))),s('type.rat.rat',X43))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))))),s('type.rat.rat',X43))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6292,f1409])).
fof(f6292,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X45),s('type.rat.rat',X43))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))))),s('type.rat.rat',X43))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6291,f1340])).
fof(f6291,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))),s('type.rat.rat',X45))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))))),s('type.rat.rat',X43))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6290,f420])).
fof(f6290,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))))),s('type.rat.rat',X43))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6289,f1409])).
fof(f6289,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X46))),s('type.rat.rat',X43))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6288,f1409])).
fof(f6288,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X46),s('type.rat.rat',X43))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6287,f1340])).
fof(f6287,plain,(
  ( ! [X45,X43,X46,X44,X42] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))),s('type.rat.rat',X46))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43))))))) )),
  inference(forward_demodulation,[],[f6286,f420])).
fof(f6286,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat',X44))),s('type.rat.rat',X43)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X46)))))))) )),
  inference(forward_demodulation,[],[f6190,f1410])).
fof(f6190,plain,(
  ( ! [X45,X43,X46,X44,X42] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X42),s('type.rat.rat',X43))),s('type.rat.rat',X44)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X45),s('type.rat.rat',X46))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X45))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X44),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X42),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X43))))),s('type.rat.rat',X46)))))))) )),
  inference(superposition,[],[f1497,f373])).
fof(f6285,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X40))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X39))))))))))) )),
  inference(forward_demodulation,[],[f6284,f420])).
fof(f6284,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat',X40))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X39))))))))))) )),
  inference(forward_demodulation,[],[f6283,f420])).
fof(f6283,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X41))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X39))))))))))) )),
  inference(forward_demodulation,[],[f6282,f420])).
fof(f6282,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))),s('type.rat.rat',X41))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X39))))))))))) )),
  inference(forward_demodulation,[],[f6281,f420])).
fof(f6281,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X39)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X41)))))))) )),
  inference(forward_demodulation,[],[f6280,f420])).
fof(f6280,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))),s('type.rat.rat',X39)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X41)))))))) )),
  inference(forward_demodulation,[],[f6189,f420])).
fof(f6189,plain,(
  ( ! [X39,X37,X41,X38,X36,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat',X37))))),s('type.rat.rat',X39)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X40),s('type.rat.rat',X41))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X40))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X39),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X36),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat',X38))))),s('type.rat.rat',X41)))))))) )),
  inference(superposition,[],[f1497,f1385])).
fof(f6279,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X34))))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X35))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X33))))))))))) )),
  inference(forward_demodulation,[],[f6278,f420])).
fof(f6278,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))),s('type.rat.rat',X34))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X35))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X33))))))))))) )),
  inference(forward_demodulation,[],[f6277,f420])).
fof(f6277,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X34))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X35))))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X33))))))))))) )),
  inference(forward_demodulation,[],[f6276,f420])).
fof(f6276,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X34))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))),s('type.rat.rat',X35))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X33))))))))))) )),
  inference(forward_demodulation,[],[f6275,f420])).
fof(f6275,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X33)))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X34))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X35)))))))) )),
  inference(forward_demodulation,[],[f6274,f420])).
fof(f6274,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X30))),s('type.rat.rat',X33)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X34))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X35)))))))) )),
  inference(forward_demodulation,[],[f6188,f420])).
fof(f6188,plain,(
  ( ! [X30,X35,X33,X31,X34,X32] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X30))))),s('type.rat.rat',X33)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X34),s('type.rat.rat',X35))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X34))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat',X35)))))))) )),
  inference(superposition,[],[f1497,f1385])).
fof(f6264,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X23))))))),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6263,f1409])).
fof(f6263,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X23))))),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6262,f1409])).
fof(f6262,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X23))),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6261,f1410])).
fof(f6261,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))),s('type.rat.rat',X23))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6260,f420])).
fof(f6260,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))))),s('type.rat.rat',X21))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6259,f1409])).
fof(f6259,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))))),s('type.rat.rat',X21))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6258,f1409])).
fof(f6258,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X24))),s('type.rat.rat',X21))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6257,f1410])).
fof(f6257,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))),s('type.rat.rat',X24))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21))))))) )),
  inference(forward_demodulation,[],[f6256,f420])).
fof(f6256,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))),s('type.rat.rat',X21)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X24)))))))) )),
  inference(forward_demodulation,[],[f6255,f420])).
fof(f6255,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X22))),s('type.rat.rat',X21)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X24)))))))) )),
  inference(forward_demodulation,[],[f6186,f1410])).
fof(f6186,plain,(
  ( ! [X24,X23,X21,X19,X22,X20] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat',X22)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X23))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat',X24)))))))) )),
  inference(superposition,[],[f1497,f1409])).
fof(f6254,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))))))),s('type.rat.rat',X14))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6253,f1409])).
fof(f6253,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))))),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6252,f1409])).
fof(f6252,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X14))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6251,f1410])).
fof(f6251,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))),s('type.rat.rat',X17))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6250,f420])).
fof(f6250,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))))),s('type.rat.rat',X14))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6249,f1409])).
fof(f6249,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))))),s('type.rat.rat',X14))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6248,f1409])).
fof(f6248,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X18))),s('type.rat.rat',X14))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6247,f1410])).
fof(f6247,plain,(
  ( ! [X14,X17,X15,X18,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))),s('type.rat.rat',X18))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f6185,f420])).
fof(f6185,plain,(
  ( ! [X14,X17,X15,X18,X16] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X17))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X14))))),s('type.rat.rat',X18))))))) | s('type.rat.rat',X14) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1497,f1468])).
fof(f6239,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X3))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X4))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))))))))) )),
  inference(forward_demodulation,[],[f6238,f420])).
fof(f6238,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X3))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X4))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2))))))))) )),
  inference(forward_demodulation,[],[f6237,f420])).
fof(f6237,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X2)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X3))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X4)))))))) )),
  inference(forward_demodulation,[],[f6182,f420])).
fof(f6182,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat',X2)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X3))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X4)))))))) )),
  inference(superposition,[],[f1497,f372])).
fof(f5135,plain,(
  ( ! [X116,X115] : (s('type.rat.rat',X116) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X115),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X116),s('type.rat.rat',X115))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X115)))) )),
  inference(superposition,[],[f1468,f1308])).
fof(f5083,plain,(
  ( ! [X6,X7] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X6))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7)))))) )),
  inference(superposition,[],[f1385,f1308])).
fof(f4954,plain,(
  ( ! [X10,X11] : (s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X10))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X10))))) | s('type.rat.rat',X10) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4872,f295])).
fof(f4872,plain,(
  ( ! [X10,X11] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X10))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X10))))) | s('type.rat.rat',X10) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1385,f308])).
fof(f4951,plain,(
  ( ! [X6,X7] : (s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))))) | s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f4870,f1308])).
fof(f4870,plain,(
  ( ! [X6,X7] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X7))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))))) | s('type.rat.rat',X6) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1375,f308])).
fof(f4866,plain,(
  ( ! [X11] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X11))))),s('type.rat.rat',X11))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X11)))) )),
  inference(superposition,[],[f308,f1340])).
fof(f4946,plain,(
  ( ! [X8,X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8)))) )),
  inference(forward_demodulation,[],[f4864,f373])).
fof(f4864,plain,(
  ( ! [X8,X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8)))) )),
  inference(superposition,[],[f308,f1375])).
fof(f4945,plain,(
  ( ! [X6,X5] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',X6))))),s('type.rat.rat',X6))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',X6)))) )),
  inference(forward_demodulation,[],[f4863,f373])).
fof(f4863,plain,(
  ( ! [X6,X5] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',X6))))))),s('type.rat.rat',X6))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',X6)))) )),
  inference(superposition,[],[f308,f1410])).
fof(f4941,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X25)))))))))) )),
  inference(forward_demodulation,[],[f4858,f420])).
fof(f4858,plain,(
  ( ! [X24,X23,X25] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X25))))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X25))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))),s('type.rat.rat',X25)))) )),
  inference(superposition,[],[f308,f1375])).
fof(f4853,plain,(
  ( ! [X14,X15,X13] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15)))))) )),
  inference(superposition,[],[f308,f1385])).
fof(f4852,plain,(
  ( ! [X12,X10,X11] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',X10))))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',X10))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))))) )),
  inference(superposition,[],[f308,f1385])).
fof(f4850,plain,(
  ( ! [X6,X5] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X5))))) | s('type.rat.rat',X5) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f308,f1468])).
fof(f4802,plain,(
  ( ! [X114,X113] : (s('type.rat.rat',X114) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X113),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X114),s('type.rat.rat',X113))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X113),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f1468,f295])).
fof(f308,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : (s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))),
  inference(ennf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s('type.rat.rat',X0) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))),
  inference(flattening,[],[f80])).
fof(f80,plain,(
  ! [X0] : (~s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X0))))))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X17] : (~s('type.rat.rat',X17) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) => s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X17))))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_MUL_RINV')).
fof(f1434,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(superposition,[],[f308,f372])).
fof(f1436,plain,(
  ( ! [X4] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X4))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) )),
  inference(superposition,[],[f308,f295])).
fof(f1440,plain,(
  ( ! [X0] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f308,f373])).
fof(f1466,plain,(
  ( ! [X4,X5] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5)))) )),
  inference(forward_demodulation,[],[f1441,f373])).
fof(f1441,plain,(
  ( ! [X4,X5] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5)))) )),
  inference(superposition,[],[f308,f420])).
fof(f2882,plain,(
  ( ! [X3] : (s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X3))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X3)))) )),
  inference(superposition,[],[f1308,f373])).
fof(f2890,plain,(
  ( ! [X6] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X6)))) )),
  inference(superposition,[],[f308,f1308])).
fof(f4161,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,g(s(X0,sK13(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,f(s(X0,sK13(X1,X2,bool,X0))))) )),
  inference(condensation,[],[f4160])).
fof(f4160,plain,(
  ( ! [X10,X8,X11,X9] : (s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK13(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK13(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(subsumption_resolution,[],[f4153,f297])).
fof(f4153,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'F_0') != s(bool,X11) | s(fun(X8,bool),X9) = s(fun(X8,bool),X10) | s(bool,'T_0') = s(bool,g(s(X8,sK13(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,f(s(X8,sK13(X9,X10,bool,X8)))) | s(bool,'T_0') = s(bool,X11)) )),
  inference(superposition,[],[f1423,f674])).
fof(f4076,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat',X23))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X24))),s('type.rat.rat',X23)))))) )),
  inference(forward_demodulation,[],[f4000,f1410])).
fof(f4000,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat',X25))))),s('type.rat.rat',X23))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X25),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))),s('type.rat.rat',X24)))))) )),
  inference(superposition,[],[f1410,f1385])).
fof(f4072,plain,(
  ( ! [X12,X13,X11] : (s('type.rat.rat',X13) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))))),s('type.rat.rat',X12))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))) )),
  inference(forward_demodulation,[],[f3997,f1409])).
fof(f3997,plain,(
  ( ! [X12,X13,X11] : (s('type.rat.rat',X13) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X13),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))))))),s('type.rat.rat',X12))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))) )),
  inference(superposition,[],[f1410,f1468])).
fof(f3936,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))))),s('type.rat.rat',X70))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f3935,f1409])).
fof(f3935,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X71))),s('type.rat.rat',X70))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f3934,f1410])).
fof(f3934,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))),s('type.rat.rat',X71))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f3933,f420])).
fof(f3933,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))),s('type.rat.rat',X71))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))))),s('type.rat.rat',X70))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f3932,f1409])).
fof(f3932,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))),s('type.rat.rat',X71))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X69),s('type.rat.rat',X72))),s('type.rat.rat',X70))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f3931,f1410])).
fof(f3931,plain,(
  ( ! [X70,X68,X72,X71,X69] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))),s('type.rat.rat',X71))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))),s('type.rat.rat',X72))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70))))))) )),
  inference(forward_demodulation,[],[f3864,f420])).
fof(f3864,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat',X69))),s('type.rat.rat',X70)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X71),s('type.rat.rat',X72))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))),s('type.rat.rat',X71))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X68),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X69),s('type.rat.rat',X70))))),s('type.rat.rat',X72)))))) )),
  inference(superposition,[],[f421,f1409])).
fof(f3917,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))))),s('type.rat.rat',X4))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5))),s('type.rat.rat',X4)))))) )),
  inference(forward_demodulation,[],[f3916,f1410])).
fof(f3916,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))))),s('type.rat.rat',X4))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))),s('type.rat.rat',X5)))))) )),
  inference(forward_demodulation,[],[f3847,f1409])).
fof(f3847,plain,(
  ( ! [X4,X2,X5,X3] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X3),s('type.rat.rat',X4))),s('type.rat.rat',X5))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))),s('type.rat.rat',X4)))))) )),
  inference(superposition,[],[f1385,f1409])).
fof(f3915,plain,(
  ( ! [X23,X21,X22] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X22),s('type.rat.rat',X21))),s('type.rat.rat',X23))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))),s('type.rat.rat',X21)))) )),
  inference(forward_demodulation,[],[f3845,f1340])).
fof(f3845,plain,(
  ( ! [X23,X21,X22] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))),s('type.rat.rat',X21))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X21))),s('type.rat.rat',X22))),s('type.rat.rat',X23)))) )),
  inference(superposition,[],[f1409,f1340])).
fof(f3912,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15))))),s('type.rat.rat',X16))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X15))))),s('type.rat.rat',X16)))) )),
  inference(forward_demodulation,[],[f3911,f420])).
fof(f3911,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X15))))),s('type.rat.rat',X16))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))),s('type.rat.rat',X15))),s('type.rat.rat',X16)))) )),
  inference(forward_demodulation,[],[f3910,f1409])).
fof(f3910,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))),s('type.rat.rat',X15))),s('type.rat.rat',X16))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X15))),s('type.rat.rat',X16)))))) )),
  inference(forward_demodulation,[],[f3843,f1409])).
fof(f3843,plain,(
  ( ! [X14,X15,X13,X16] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))),s('type.rat.rat',X15))),s('type.rat.rat',X16))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))))) )),
  inference(superposition,[],[f1409,f1375])).
fof(f3840,plain,(
  ( ! [X6,X7,X5] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X6))),s('type.rat.rat',X7))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))),s('type.rat.rat',X5)))) )),
  inference(superposition,[],[f1409,f372])).
fof(f3839,plain,(
  ( ! [X0,X1] : (s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X0))) | s('type.rat.rat',X0) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1409,f1468])).
fof(f3728,plain,(
  ( ! [X167,X165,X166,X164] : (s('type.rat.rat',X167) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X165),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X166),s('type.rat.rat',X164))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X167),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X165),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X166),s('type.rat.rat',X164))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X164),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X165),s('type.rat.rat',X166)))))) )),
  inference(superposition,[],[f1468,f1385])).
fof(f3782,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X62))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X58))))))))) )),
  inference(forward_demodulation,[],[f3781,f420])).
fof(f3781,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))),s('type.rat.rat',X61))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X62))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X58))))))))) )),
  inference(forward_demodulation,[],[f3780,f420])).
fof(f3780,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X61))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X62))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X58))))))))) )),
  inference(forward_demodulation,[],[f3779,f420])).
fof(f3779,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X61))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))),s('type.rat.rat',X62))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X58))))))))) )),
  inference(forward_demodulation,[],[f3701,f420])).
fof(f3701,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X58)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X61))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X62)))))) )),
  inference(superposition,[],[f421,f1385])).
fof(f3671,plain,(
  ( ! [X167,X165,X166,X164] : (s('type.rat.rat',X167) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X166),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X164),s('type.rat.rat',X165))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X167),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X166),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X164),s('type.rat.rat',X165))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X164),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X165),s('type.rat.rat',X166)))))) )),
  inference(superposition,[],[f1468,f1385])).
fof(f3773,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X61))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X62))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat',X59))))))))) )),
  inference(forward_demodulation,[],[f3772,f420])).
fof(f3772,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))),s('type.rat.rat',X61))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X62))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat',X59))))))))) )),
  inference(forward_demodulation,[],[f3771,f420])).
fof(f3771,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X61))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat',X62))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat',X59))))))))) )),
  inference(forward_demodulation,[],[f3770,f420])).
fof(f3770,plain,(
  ( ! [X61,X59,X62,X60,X58] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X61))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))),s('type.rat.rat',X62))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat',X59))))))))) )),
  inference(forward_demodulation,[],[f3644,f420])).
fof(f3644,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X60),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat',X59)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X61),s('type.rat.rat',X62))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X61))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X58),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X59),s('type.rat.rat',X60))))),s('type.rat.rat',X62)))))) )),
  inference(superposition,[],[f421,f1385])).
fof(f3765,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))),s('type.rat.rat',X7)))) )),
  inference(forward_demodulation,[],[f3630,f420])).
fof(f3630,plain,(
  ( ! [X6,X4,X7,X5] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat',X6))),s('type.rat.rat',X7))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))),s('type.rat.rat',X7)))) )),
  inference(superposition,[],[f420,f1385])).
fof(f3627,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X21),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19)))))))) )),
  inference(superposition,[],[f1385,f420])).
fof(f3620,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X31))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X33)))))) )),
  inference(superposition,[],[f1385,f420])).
fof(f3619,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X29),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X28))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X28),s('type.rat.rat',X29)))))) )),
  inference(superposition,[],[f1385,f1375])).
fof(f3755,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))))) )),
  inference(forward_demodulation,[],[f3615,f420])).
fof(f3615,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))),s('type.rat.rat',X18))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))))) )),
  inference(superposition,[],[f1385,f1385])).
fof(f3754,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X11)))))))) )),
  inference(forward_demodulation,[],[f3614,f420])).
fof(f3614,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',X13))),s('type.rat.rat',X14))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X11)))))))) )),
  inference(superposition,[],[f1385,f1385])).
fof(f3753,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X9))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X9),s('type.rat.rat',X10))),s('type.rat.rat',X8))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3613,f1410])).
fof(f3613,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X9))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X9),s('type.rat.rat',X8))),s('type.rat.rat',X10))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1385,f1468])).
fof(f3749,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16)))))))) )),
  inference(forward_demodulation,[],[f3568,f420])).
fof(f3568,plain,(
  ( ! [X14,X17,X15,X16] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15))),s('type.rat.rat',X16)))))) )),
  inference(superposition,[],[f1385,f1375])).
fof(f3565,plain,(
  ( ! [X6,X8,X7] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))),s('type.rat.rat',X6)))) )),
  inference(superposition,[],[f1385,f372])).
fof(f3562,plain,(
  ( ! [X30,X33,X31,X32] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat',X31))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X33),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X30),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32)))))))) )),
  inference(superposition,[],[f1385,f420])).
fof(f3561,plain,(
  ( ! [X28,X26,X29,X27] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X28),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X29),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X27))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X29),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X27),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X26),s('type.rat.rat',X28)))))))) )),
  inference(superposition,[],[f1385,f1375])).
fof(f3557,plain,(
  ( ! [X17,X15,X18,X16] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X15)))))) )),
  inference(superposition,[],[f1385,f1385])).
fof(f3556,plain,(
  ( ! [X14,X12,X13,X11] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X13),s('type.rat.rat',X11))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat',X13))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X11)))))) )),
  inference(superposition,[],[f1385,f1385])).
fof(f3555,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X9))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X9),s('type.rat.rat',X8))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X8))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1385,f1468])).
fof(f3531,plain,(
  ( ! [X146,X147,X145,X148] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X145),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X146),s('type.rat.rat',X147))))) | s('type.rat.rat',X148) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X146),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X145),s('type.rat.rat',X147))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X148),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X146),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X145),s('type.rat.rat',X147)))))))))) )),
  inference(forward_demodulation,[],[f3495,f420])).
fof(f3495,plain,(
  ( ! [X146,X147,X145,X148] : (s('type.rat.rat',X148) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X146),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X145),s('type.rat.rat',X147))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X148),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X146),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X145),s('type.rat.rat',X147))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X145),s('type.rat.rat',X146))),s('type.rat.rat',X147)))) )),
  inference(superposition,[],[f1468,f1375])).
fof(f3529,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X57),s('type.rat.rat',X58))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X56),s('type.rat.rat',X57))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X56),s('type.rat.rat',X58))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X56))))))))) )),
  inference(forward_demodulation,[],[f3528,f420])).
fof(f3528,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X57),s('type.rat.rat',X58))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X55))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X56),s('type.rat.rat',X57))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X56),s('type.rat.rat',X58))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X56))))))))) )),
  inference(forward_demodulation,[],[f3527,f420])).
fof(f3527,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X57),s('type.rat.rat',X58))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X55))),s('type.rat.rat',X56))),s('type.rat.rat',X57))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X56),s('type.rat.rat',X58))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X56))))))))) )),
  inference(forward_demodulation,[],[f3526,f420])).
fof(f3526,plain,(
  ( ! [X57,X54,X58,X56,X55] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X57),s('type.rat.rat',X58))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X55))),s('type.rat.rat',X56))),s('type.rat.rat',X57))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X55))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X56),s('type.rat.rat',X58))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X56))))))))) )),
  inference(forward_demodulation,[],[f3471,f420])).
fof(f3471,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X55),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X56)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X57),s('type.rat.rat',X58))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X55))),s('type.rat.rat',X56))),s('type.rat.rat',X57))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X54),s('type.rat.rat',X55))),s('type.rat.rat',X56))),s('type.rat.rat',X58)))))) )),
  inference(superposition,[],[f421,f1375])).
fof(f3518,plain,(
  ( ! [X19,X17,X18] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))),s('type.rat.rat',X19))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X17),s('type.rat.rat',X19)))))) )),
  inference(forward_demodulation,[],[f3454,f373])).
fof(f3454,plain,(
  ( ! [X19,X17,X18] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))),s('type.rat.rat',X19))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X19)))))))) )),
  inference(superposition,[],[f1375,f373])).
fof(f3517,plain,(
  ( ! [X14,X15,X16] : (s('type.rat.rat',X16) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15)))) )),
  inference(forward_demodulation,[],[f3453,f1409])).
fof(f3453,plain,(
  ( ! [X14,X15,X16] : (s('type.rat.rat',X16) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X16),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X14),s('type.rat.rat',X15)))) )),
  inference(superposition,[],[f1375,f1468])).
fof(f3450,plain,(
  ( ! [X6,X8,X7] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X8))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat',X7)))))) )),
  inference(superposition,[],[f1375,f372])).
fof(f3449,plain,(
  ( ! [X4,X5,X3] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X5)))))) )),
  inference(superposition,[],[f1375,f420])).
fof(f3513,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X25))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X25)))))))) )),
  inference(forward_demodulation,[],[f3512,f420])).
fof(f3512,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X25))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))),s('type.rat.rat',X25)))))) )),
  inference(forward_demodulation,[],[f3436,f3511])).
fof(f3436,plain,(
  ( ! [X24,X23,X25,X22] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X24),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat',X23))),s('type.rat.rat',X25))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X22),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X23),s('type.rat.rat',X24))))),s('type.rat.rat',X25)))) )),
  inference(superposition,[],[f1375,f420])).
fof(f3511,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X21))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20))))),s('type.rat.rat',X21)))) )),
  inference(forward_demodulation,[],[f3435,f420])).
fof(f3435,plain,(
  ( ! [X21,X19,X20,X18] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))),s('type.rat.rat',X21))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X20))))),s('type.rat.rat',X21)))) )),
  inference(superposition,[],[f1375,f1375])).
fof(f3431,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X9),s('type.rat.rat',X10))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X9),s('type.rat.rat',X8))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X10))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1375,f1468])).
fof(f3397,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X39))))),s('type.rat.rat',X37))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X40))))),s('type.rat.rat',X37))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X38)))) | s('type.rat.rat',X37) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3396,f1409])).
fof(f3396,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X39))),s('type.rat.rat',X37))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X40))))),s('type.rat.rat',X37))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X38)))) | s('type.rat.rat',X37) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3395,f1410])).
fof(f3395,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))),s('type.rat.rat',X39))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X40))))),s('type.rat.rat',X37))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X38)))) | s('type.rat.rat',X37) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3394,f420])).
fof(f3394,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))))),s('type.rat.rat',X39))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X40))))),s('type.rat.rat',X37))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X38)))) | s('type.rat.rat',X37) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3393,f1409])).
fof(f3393,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))))),s('type.rat.rat',X39))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X38),s('type.rat.rat',X40))),s('type.rat.rat',X37))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X38)))) | s('type.rat.rat',X37) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3392,f1410])).
fof(f3392,plain,(
  ( ! [X39,X37,X38,X40] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))))),s('type.rat.rat',X39))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))),s('type.rat.rat',X40))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X38)))) | s('type.rat.rat',X37) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f3344,f420])).
fof(f3344,plain,(
  ( ! [X39,X37,X38,X40] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X38)))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X39),s('type.rat.rat',X40))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))))),s('type.rat.rat',X39))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X37),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X38),s('type.rat.rat',X37))))),s('type.rat.rat',X40))))) | s('type.rat.rat',X37) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f421,f1468])).
fof(f3332,plain,(
  ( ! [X10,X11] : (s('type.rat.rat',X11) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X10))))),s('type.rat.rat',X10))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X10)))) )),
  inference(superposition,[],[f1468,f1340])).
fof(f3389,plain,(
  ( ! [X8,X7,X9] : (s('type.rat.rat',X9) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8)))) )),
  inference(forward_demodulation,[],[f3331,f1409])).
fof(f3331,plain,(
  ( ! [X8,X7,X9] : (s('type.rat.rat',X9) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X9),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat',X8)))) )),
  inference(superposition,[],[f1468,f420])).
fof(f3322,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat',X10) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X9),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X10),s('type.rat.rat',X9))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X9),s('type.rat.rat',X8))))) | s('type.rat.rat',X8) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1468,f1468])).
fof(f3319,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X2),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) )),
  inference(superposition,[],[f1468,f372])).
fof(f3098,plain,(
  ( ! [X6,X7] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X6))),s('type.rat.rat',X7))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X7))),s('type.rat.rat',X6)))) )),
  inference(superposition,[],[f1340,f373])).
fof(f2959,plain,(
  ( ! [X0] : (~sP2('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f496,f527,f525,f1570,f433])).
fof(f314,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f203])).
fof(f1422,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,X2,bool))) )),
  inference(superposition,[],[f419,f297])).
fof(f1423,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK13(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK13(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f419,f297])).
fof(f1996,plain,(
  ( ! [X123,X121,X124,X122] : (s(bool,'T_0') != s(bool,f(s(X121,sK13(X122,X123,bool,X121)))) | s(fun(X121,bool),X122) = s(fun(X121,bool),X123) | s(bool,X124) = s(bool,g(s(X121,sK13(X122,X123,bool,X121)))) | s(bool,'T_0') = s(bool,X124)) )),
  inference(superposition,[],[f419,f674])).
fof(f1882,plain,(
  ( ! [X127,X130,X128,X129] : (s(bool,X130) != s(bool,f(s(X127,sK13(X128,X129,bool,X127)))) | s(fun(X127,bool),X128) = s(fun(X127,bool),X129) | s(bool,'T_0') = s(bool,X130) | s(bool,'T_0') = s(bool,g(s(X127,sK13(X128,X129,bool,X127))))) )),
  inference(superposition,[],[f419,f674])).
fof(f2267,plain,(
  ( ! [X111,X109,X110,X108] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X109),s('type.rat.rat',X110))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X108),s('type.rat.rat',X109))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X108),s('type.rat.rat',X110))))) | s(bool,X111) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X108))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(subsumption_resolution,[],[f1992,f288])).
fof(f1992,plain,(
  ( ! [X111,X109,X110,X108] : (~p(s(bool,'T_0')) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X109),s('type.rat.rat',X110))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X108),s('type.rat.rat',X109))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X108),s('type.rat.rat',X110))))) | s(bool,X111) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X108))) | s(bool,'T_0') = s(bool,X111)) )),
  inference(superposition,[],[f421,f674])).
fof(f1523,plain,(
  ( ! [X0,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))))) )),
  inference(subsumption_resolution,[],[f1494,f288])).
fof(f1494,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X1)))))) )),
  inference(superposition,[],[f421,f536])).
fof(f1522,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22)))))))))) )),
  inference(forward_demodulation,[],[f1521,f420])).
fof(f1521,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X21))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))))))) )),
  inference(forward_demodulation,[],[f1520,f420])).
fof(f1520,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))))))) )),
  inference(forward_demodulation,[],[f1519,f420])).
fof(f1519,plain,(
  ( ! [X21,X19,X22,X20,X18] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X20),s('type.rat.rat',X22))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20))))))))) )),
  inference(forward_demodulation,[],[f1493,f420])).
fof(f1493,plain,(
  ( ! [X21,X19,X22,X20,X18] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X19),s('type.rat.rat',X20)))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X21),s('type.rat.rat',X22))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))),s('type.rat.rat',X20))),s('type.rat.rat',X21))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))),s('type.rat.rat',X20))),s('type.rat.rat',X22)))))) )),
  inference(superposition,[],[f421,f420])).
fof(f1509,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X13))),s('type.rat.rat',X12))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X14))),s('type.rat.rat',X12)))))) )),
  inference(forward_demodulation,[],[f1508,f1409])).
fof(f1508,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X13),s('type.rat.rat',X12))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X14))),s('type.rat.rat',X12))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))))))) )),
  inference(forward_demodulation,[],[f1507,f1340])).
fof(f1507,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X12))),s('type.rat.rat',X13))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X14))),s('type.rat.rat',X12))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))))))) )),
  inference(forward_demodulation,[],[f1506,f420])).
fof(f1506,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X12))))),s('type.rat.rat',X13))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X14))),s('type.rat.rat',X12))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))))))) )),
  inference(forward_demodulation,[],[f1505,f1409])).
fof(f1505,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X12))))),s('type.rat.rat',X13))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X14),s('type.rat.rat',X12))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))))))) )),
  inference(forward_demodulation,[],[f1504,f1340])).
fof(f1504,plain,(
  ( ! [X14,X12,X13,X11] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X12))))),s('type.rat.rat',X13))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X12))),s('type.rat.rat',X14))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))))))) )),
  inference(forward_demodulation,[],[f1491,f420])).
fof(f1491,plain,(
  ( ! [X14,X12,X13,X11] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat',X12)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X13),s('type.rat.rat',X14))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X12))))),s('type.rat.rat',X13))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X12))))),s('type.rat.rat',X14)))))) )),
  inference(superposition,[],[f421,f373])).
fof(f1497,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X3)))))))) )),
  inference(forward_demodulation,[],[f1496,f420])).
fof(f1496,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X2))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X3))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))) )),
  inference(forward_demodulation,[],[f1488,f420])).
fof(f1488,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0)))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X2))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X3)))))) )),
  inference(superposition,[],[f421,f372])).
fof(f421,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X2)))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1)))))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X2)))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))))),
  inference(ennf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X2)))) => s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))))),
  inference(rectify,[],[f48])).
fof(f48,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X8)))) => s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X6))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X7))))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_LES_LMUL_POS')).
fof(f1995,plain,(
  ( ! [X118,X120,X119,X117] : (s(X119,f(s(bool,'T_0'))) != s(X119,g(s(bool,'T_0'))) | s(fun(bool,X119),X117) = s(fun(bool,X119),X118) | s(bool,X120) = s(bool,sK13(X117,X118,X119,bool)) | s(bool,'T_0') = s(bool,X120)) )),
  inference(superposition,[],[f419,f674])).
fof(f2266,plain,(
  ( ! [X92,X90,X91] : (sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X90,X91) | s(bool,X90) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(subsumption_resolution,[],[f1983,f288])).
fof(f1983,plain,(
  ( ! [X92,X90,X91] : (~p(s(bool,'T_0')) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X90,X91) | s(bool,X90) = s(bool,X92) | s(bool,'T_0') = s(bool,X92)) )),
  inference(superposition,[],[f598,f674])).
fof(f2265,plain,(
  ( ! [X88,X87,X89] : (sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X88,X87) | s(bool,X87) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f1982,f288])).
fof(f1982,plain,(
  ( ! [X88,X87,X89] : (~p(s(bool,'T_0')) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X88,X87) | s(bool,X87) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f593,f674])).
fof(f2264,plain,(
  ( ! [X85,X83,X86,X84] : (~sP6(X83,X84,X85) | sP7(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(subsumption_resolution,[],[f1981,f288])).
fof(f1981,plain,(
  ( ! [X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~sP6(X83,X84,X85) | sP7(X83,X84,X85) | s(bool,X83) = s(bool,X86) | s(bool,'T_0') = s(bool,X86)) )),
  inference(superposition,[],[f486,f674])).
fof(f2263,plain,(
  ( ! [X80,X81,X79,X82] : (~sP8(X79,X80,X81) | sP9(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(subsumption_resolution,[],[f1980,f288])).
fof(f1980,plain,(
  ( ! [X80,X81,X79,X82] : (~p(s(bool,'T_0')) | ~sP8(X79,X80,X81) | sP9(X79,X80,X81) | s(bool,X79) = s(bool,X82) | s(bool,'T_0') = s(bool,X82)) )),
  inference(superposition,[],[f485,f674])).
fof(f2262,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,X76)) | p(s(bool,X77)) | ~sP9(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(subsumption_resolution,[],[f1979,f288])).
fof(f1979,plain,(
  ( ! [X78,X76,X77,X75] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | p(s(bool,X77)) | ~sP9(X75,X77,X76) | s(bool,X75) = s(bool,X78) | s(bool,'T_0') = s(bool,X78)) )),
  inference(superposition,[],[f482,f674])).
fof(f2261,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,X72)) | sP9(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(subsumption_resolution,[],[f1978,f288])).
fof(f1978,plain,(
  ( ! [X74,X72,X71,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X72)) | sP9(X71,X72,X73) | s(bool,X71) = s(bool,X74) | s(bool,'T_0') = s(bool,X74)) )),
  inference(superposition,[],[f481,f674])).
fof(f2260,plain,(
  ( ! [X70,X68,X69,X67] : (p(s(bool,X68)) | sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f1977,f288])).
fof(f1977,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | p(s(bool,X68)) | sP9(X67,X69,X68) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f480,f674])).
fof(f2259,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP8(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f1976,f288])).
fof(f1976,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | ~p(s(bool,X65)) | ~sP8(X65,X64,X63) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f475,f674])).
fof(f2258,plain,(
  ( ! [X61,X59,X62,X60] : (p(s(bool,X60)) | p(s(bool,X61)) | ~sP7(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f1975,f288])).
fof(f1975,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | p(s(bool,X60)) | p(s(bool,X61)) | ~sP7(X59,X61,X60) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f465,f674])).
fof(f2257,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP7(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f1974,f288])).
fof(f1974,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP7(X55,X56,X57) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f464,f674])).
fof(f2256,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,X52)) | sP7(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f1973,f288])).
fof(f1973,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | ~p(s(bool,X52)) | sP7(X51,X53,X52) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f463,f674])).
fof(f2255,plain,(
  ( ! [X47,X50,X48,X49] : (p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X47,X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f1972,f288])).
fof(f1972,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | p(s(bool,X48)) | p(s(bool,X49)) | ~sP6(X47,X48,X49) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f458,f674])).
fof(f2254,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f1971,f288])).
fof(f1971,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP3(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f439,f674])).
fof(f2253,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f1970,f288])).
fof(f1970,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP4(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f435,f674])).
fof(f2252,plain,(
  ( ! [X37,X35,X38,X36] : (sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f1969,f288])).
fof(f1969,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | sP5(X36,X37,X35) | ~sP4(X36,X35,X37) | ~sP3(X37,X35,X36) | ~sP2(X35,X37,X36) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f432,f674])).
fof(f2251,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1968,f288])).
fof(f1968,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | ~p(s(bool,X32)) | p(s(bool,X33)) | ~sP5(X32,X31,X33) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f431,f674])).
fof(f2250,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f1967,f288])).
fof(f1967,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f406,f674])).
fof(f2249,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f1966,f288])).
fof(f1966,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f403,f674])).
fof(f2248,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f1965,f288])).
fof(f1965,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f388,f674])).
fof(f2247,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f1964,f288])).
fof(f1964,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f387,f674])).
fof(f2246,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f1963,f288])).
fof(f1963,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f377,f674])).
fof(f2245,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f1959,f288])).
fof(f1959,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f314,f674])).
fof(f1927,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f310,f674])).
fof(f1881,plain,(
  ( ! [X125,X123,X126,X124] : (s(X125,f(s(bool,X126))) != s(X125,g(s(bool,X126))) | s(fun(bool,X125),X123) = s(fun(bool,X125),X124) | s(bool,'T_0') = s(bool,X126) | s(bool,'T_0') = s(bool,sK13(X123,X124,X125,bool))) )),
  inference(superposition,[],[f419,f674])).
fof(f1816,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f674,f674])).
fof(f1755,plain,(
  ( ! [X0] : (~sP2('F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f496,f508,f504,f1568,f434])).
fof(f1750,plain,(
  ( ! [X0,X1] : (s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(subsumption_resolution,[],[f1744,f496])).
fof(f1744,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | s(bool,X1) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0))) | s(bool,'T_0') = s(bool,X1)) )),
  inference(superposition,[],[f735,f552])).
fof(f1696,plain,(
  ( ! [X0] : (~sP4('T_0','T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f288,f510,f506,f1548,f432])).
fof(f1690,plain,(
  sP3('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)),'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'F_0')),
  inference(unit_resulting_resolution,[],[f797,f429])).
fof(f1691,plain,(
  sP2('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f797,f428])).
fof(f1639,plain,(
  sP3('const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)),'F_0')),
  inference(unit_resulting_resolution,[],[f795,f429])).
fof(f1640,plain,(
  sP2('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)),'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'F_0')),
  inference(unit_resulting_resolution,[],[f795,f428])).
fof(f1568,plain,(
  ( ! [X0] : (~sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f496,f552,f450])).
fof(f1570,plain,(
  ( ! [X0] : (~sP5('F_0','F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f496,f552,f450])).
fof(f1548,plain,(
  ( ! [X0] : (~sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f533,f449])).
fof(f1551,plain,(
  ( ! [X0] : (~sP5('T_0','T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f288,f533,f449])).
fof(f1468,plain,(
  ( ! [X2,X1] : (s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1467,f1308])).
fof(f1467,plain,(
  ( ! [X2,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X2),s('type.rat.rat',X1))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X2))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1443,f1340])).
fof(f1443,plain,(
  ( ! [X2,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat',X2))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X1))),s('type.rat.rat',X2))))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f420,f308])).
fof(f1410,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat',X10))),s('type.rat.rat',X9))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))),s('type.rat.rat',X10)))) )),
  inference(backward_demodulation,[],[f1409,f1406])).
fof(f1409,plain,(
  ( ! [X12,X13,X11] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))),s('type.rat.rat',X13))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X12),s('type.rat.rat',X13)))))) )),
  inference(forward_demodulation,[],[f1388,f373])).
fof(f1388,plain,(
  ( ! [X12,X13,X11] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat',X12))),s('type.rat.rat',X13))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X12),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X13)))))))) )),
  inference(superposition,[],[f420,f373])).
fof(f1385,plain,(
  ( ! [X4,X5,X3] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X3),s('type.rat.rat',X4)))))) )),
  inference(superposition,[],[f420,f372])).
fof(f1406,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X10),s('type.rat.rat',X9))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))),s('type.rat.rat',X10)))) )),
  inference(forward_demodulation,[],[f1378,f1340])).
fof(f1378,plain,(
  ( ! [X10,X8,X9] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X8),s('type.rat.rat',X9))),s('type.rat.rat',X10))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X8),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X9))),s('type.rat.rat',X10)))))) )),
  inference(superposition,[],[f420,f373])).
fof(f1375,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat',X2)))) )),
  inference(superposition,[],[f420,f372])).
fof(f1340,plain,(
  ( ! [X2,X3] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X2),s('type.rat.rat',X3))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X3))),s('type.rat.rat',X2)))) )),
  inference(superposition,[],[f373,f372])).
fof(f1269,plain,(
  ( ! [X0] : (sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f515,f528,f505,f496,f434])).
fof(f1282,plain,(
  ~sP3('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f515,f788,f502,f496,f434])).
fof(f1290,plain,(
  sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f513,f508,f504,f496,f434])).
fof(f1293,plain,(
  ( ! [X0] : (sP5('F_0','T_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f512,f527,f525,f496,f434])).
fof(f1294,plain,(
  ( ! [X0] : (sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f513,f527,f525,f496,f434])).
fof(f1229,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f510,f506,f504,f496,f433])).
fof(f1230,plain,(
  sP5('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0','F_0')),
  inference(unit_resulting_resolution,[],[f511,f507,f504,f496,f433])).
fof(f1237,plain,(
  ( ! [X0] : (sP5('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f511,f507,f525,f496,f433])).
fof(f1241,plain,(
  sP5('F_0','F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f515,f509,f505,f496,f433])).
fof(f1245,plain,(
  ( ! [X0] : (sP5('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f515,f509,f526,f496,f433])).
fof(f1209,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))) | ~sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1) | ~sP3(X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(subsumption_resolution,[],[f1208,f515])).
fof(f1208,plain,(
  ( ! [X0,X1] : (sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))) | ~sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1) | ~sP3(X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0) | ~sP2('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0)) )),
  inference(subsumption_resolution,[],[f1202,f288])).
fof(f1202,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | sP5(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))) | ~sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1) | ~sP3(X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0) | ~sP2('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0)) )),
  inference(superposition,[],[f432,f536])).
fof(f1187,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f514,f506,f502,f288,f432])).
fof(f1188,plain,(
  sP5('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'T_0','T_0')),
  inference(unit_resulting_resolution,[],[f514,f507,f502,f288,f432])).
fof(f1190,plain,(
  sP5('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'T_0')),
  inference(unit_resulting_resolution,[],[f514,f507,f503,f288,f432])).
fof(f1194,plain,(
  ( ! [X0] : (sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f514,f528,f505,f288,f432])).
fof(f1195,plain,(
  ( ! [X0] : (sP5('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f514,f509,f526,f288,f432])).
fof(f1196,plain,(
  ( ! [X0,X1] : (sP5('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f514,f528,f526,f288,f432])).
fof(f1199,plain,(
  ~sP3('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'T_0','F_0')),
  inference(unit_resulting_resolution,[],[f514,f505,f786,f288,f432])).
fof(f1118,plain,(
  ~sP9('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f1104,f470])).
fof(f1104,plain,(
  ~sP8('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f288,f496,f288,f475])).
fof(f1091,plain,(
  sP9('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f524,f288,f496,f474])).
fof(f1009,plain,(
  ( ! [X0,X1] : (~sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f1003,f288])).
fof(f1003,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0)) )),
  inference(superposition,[],[f465,f536])).
fof(f950,plain,(
  ~sP7('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f939,f453])).
fof(f939,plain,(
  ~sP6('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f496,f288,f458])).
fof(f926,plain,(
  sP7('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f520,f496,f496,f457])).
fof(f905,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f896,f428])).
fof(f896,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f496,f496,f443])).
fof(f890,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f879,f429])).
fof(f886,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f880,f288])).
fof(f880,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X1)) | ~sP3(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1)) )),
  inference(superposition,[],[f439,f536])).
fof(f879,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f288,f288,f439])).
fof(f856,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f288,f288,f435])).
fof(f795,plain,(
  sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)))),
  inference(unit_resulting_resolution,[],[f496,f287,f447])).
fof(f796,plain,(
  ~sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)),'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))))),
  inference(unit_resulting_resolution,[],[f288,f287,f449])).
fof(f797,plain,(
  sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)),'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12)))))),
  inference(unit_resulting_resolution,[],[f496,f287,f447])).
fof(f786,plain,(
  ~sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'T_0')),
  inference(unit_resulting_resolution,[],[f496,f536,f450])).
fof(f788,plain,(
  ~sP5('F_0','T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f496,f536,f450])).
fof(f777,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f288,f532,f449])).
fof(f779,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f288,f532,f449])).
fof(f766,plain,(
  sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'T_0')),
  inference(unit_resulting_resolution,[],[f288,f536,f448])).
fof(f767,plain,(
  sP5('T_0','T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f288,f536,f448])).
fof(f758,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f496,f532,f447])).
fof(f759,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f532,f447])).
fof(f735,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f731,f496])).
fof(f731,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f374,f297])).
fof(f724,plain,(
  ( ! [X0] : (sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f288,f630])).
fof(f674,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f297,f297])).
fof(f667,plain,(
  ( ! [X0,X1] : (sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1) | ~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(subsumption_resolution,[],[f666,f288])).
fof(f666,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP6('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(superposition,[],[f486,f536])).
fof(f663,plain,(
  ( ! [X0] : (sP7('T_0',X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f517,f288,f486])).
fof(f665,plain,(
  ( ! [X0] : (sP7('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(unit_resulting_resolution,[],[f519,f288,f486])).
fof(f654,plain,(
  ( ! [X0] : (sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f496,f619])).
fof(f651,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1) | ~sP8('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(subsumption_resolution,[],[f650,f288])).
fof(f650,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~sP8('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1) | sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(superposition,[],[f485,f536])).
fof(f649,plain,(
  ( ! [X0] : (sP9('T_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(unit_resulting_resolution,[],[f523,f288,f485])).
fof(f630,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1) | ~p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f629,f288])).
fof(f629,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(superposition,[],[f481,f536])).
fof(f628,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f288,f288,f481])).
fof(f625,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(subsumption_resolution,[],[f624,f288])).
fof(f624,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(superposition,[],[f598,f536])).
fof(f623,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f288,f598])).
fof(f619,plain,(
  ( ! [X0,X1] : (sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f618,f288])).
fof(f618,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP9('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0)) )),
  inference(superposition,[],[f480,f536])).
fof(f617,plain,(
  ( ! [X0] : (sP9('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f496,f288,f480])).
fof(f609,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f288,f496,f469])).
fof(f608,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(subsumption_resolution,[],[f607,f288])).
fof(f607,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(superposition,[],[f593,f536])).
fof(f606,plain,(
  ( ! [X0] : (sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f593])).
fof(f605,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f496,f496,f468])).
fof(f598,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(subsumption_resolution,[],[f597,f288])).
fof(f597,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(superposition,[],[f464,f536])).
fof(f596,plain,(
  ( ! [X0] : (sP7('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f288,f288,f464])).
fof(f593,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0)) )),
  inference(subsumption_resolution,[],[f592,f288])).
fof(f592,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP7('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1,X0)) )),
  inference(superposition,[],[f463,f536])).
fof(f591,plain,(
  ( ! [X0] : (sP7('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f288,f463])).
fof(f588,plain,(
  ( ! [X0] : (~sP7('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f288,f496,f452])).
fof(f585,plain,(
  ( ! [X0] : (~sP7('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f496,f451])).
fof(f584,plain,(
  sP1('F_0','const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))),
  inference(unit_resulting_resolution,[],[f496,f570])).
fof(f579,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f496,f497])).
fof(f570,plain,(
  ( ! [X0] : (sP1(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11))) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f569,f288])).
fof(f569,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(superposition,[],[f406,f536])).
fof(f568,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f496,f288,f406])).
fof(f565,plain,(
  sP1('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),'F_0')),
  inference(unit_resulting_resolution,[],[f496,f562])).
fof(f562,plain,(
  ( ! [X0] : (sP1('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f561,f288])).
fof(f561,plain,(
  ( ! [X0] : (~p(s(bool,'T_0')) | p(s(bool,X0)) | sP1('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(superposition,[],[f403,f536])).
fof(f560,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f288,f403])).
fof(f557,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f288,f288,f402])).
fof(f556,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f496,f496,f401])).
fof(f532,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f496,f310])).
fof(f533,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)))) )),
  inference(unit_resulting_resolution,[],[f367,f310])).
fof(f526,plain,(
  ( ! [X2,X0,X1] : (sP4('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f367,f438])).
fof(f527,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,'const.rat.rat_les_2'(s('type.rat.rat',X1),s('type.rat.rat',X1)),X2)) )),
  inference(unit_resulting_resolution,[],[f367,f441])).
fof(f528,plain,(
  ( ! [X2,X0,X1] : (sP3('const.rat.rat_les_2'(s('type.rat.rat',X0),s('type.rat.rat',X0)),X1,X2)) )),
  inference(unit_resulting_resolution,[],[f367,f442])).
fof(f524,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f496,f478])).
fof(f522,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f288,f477])).
fof(f523,plain,(
  ( ! [X0,X1] : (sP8(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1)) )),
  inference(unit_resulting_resolution,[],[f286,f477])).
fof(f521,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'F_0')) )),
  inference(unit_resulting_resolution,[],[f496,f476])).
fof(f520,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f496,f461])).
fof(f518,plain,(
  ( ! [X0,X1] : (sP6(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f288,f460])).
fof(f519,plain,(
  ( ! [X0,X1] : (sP6(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1)) )),
  inference(unit_resulting_resolution,[],[f286,f460])).
fof(f516,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f459])).
fof(f517,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f286,f459])).
fof(f514,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f288,f446])).
fof(f515,plain,(
  ( ! [X0,X1] : (sP2('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0,X1)) )),
  inference(unit_resulting_resolution,[],[f286,f446])).
fof(f512,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f288,f445])).
fof(f513,plain,(
  ( ! [X0,X1] : (sP2(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X1)) )),
  inference(unit_resulting_resolution,[],[f286,f445])).
fof(f511,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f286,f444])).
fof(f509,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f496,f442])).
fof(f508,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f496,f441])).
fof(f507,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f286,f440])).
fof(f505,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f496,f438])).
fof(f504,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f496,f437])).
fof(f502,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f436])).
fof(f503,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f286,f436])).
fof(f500,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f288,f388])).
fof(f501,plain,(
  ( ! [X0] : (sP0('const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)),X0)) )),
  inference(unit_resulting_resolution,[],[f286,f388])).
fof(f498,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f288,f387])).
fof(f499,plain,(
  ( ! [X0] : (sP0(X0,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',sK11)))) )),
  inference(unit_resulting_resolution,[],[f286,f387])).
fof(f479,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f285])).
fof(f480,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f285])).
fof(f477,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f283])).
fof(f468,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f469,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f280])).
fof(f472,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f280])).
fof(f473,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f280])).
fof(f462,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f277])).
fof(f464,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f465,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f277])).
fof(f459,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f275])).
fof(f460,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f275])).
fof(f451,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f272])).
fof(f452,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f272])).
fof(f486,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP6(X0,X1,X2) | sP7(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f454])).
fof(f454,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f272])).
fof(f455,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f272])).
fof(f456,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~sP6(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f272])).
fof(f447,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f450,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f269])).
fof(f443,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f268])).
fof(f445,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f268])).
fof(f446,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f268])).
fof(f439,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f266])).
fof(f441,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f266])).
fof(f442,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f266])).
fof(f436,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f263])).
fof(f438,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f263])).
fof(f428,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f260])).
fof(f429,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f260])).
fof(f432,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f260])).
fof(f434,plain,(
  ( ! [X2,X0,X1] : (~sP4(X0,X2,X1) | p(s(bool,X0)) | sP5(X0,X1,X2) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f260])).
fof(f420,plain,(
  ( ! [X2,X0,X1] : (s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X2)))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0,X1,X2] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X2))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))),s('type.rat.rat',X2)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X17,X18,X19] : s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X18),s('type.rat.rat',X19))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X17),s('type.rat.rat',X18))),s('type.rat.rat',X19)))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_MUL_ASSOC')).
fof(f419,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f255])).
fof(f255,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK13(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK13(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK13])],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.rat.RAT_LDIV_LES_POS',aHL_EXT)).
fof(f401,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f249])).
fof(f249,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f248])).
fof(f248,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f247])).
fof(f247,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f402,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f249])).
fof(f403,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f406,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f249])).
fof(f497,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f386,f496])).
fof(f386,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f242,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f241])).
fof(f241,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f183])).
fof(f183,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f387,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f388,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f242])).
fof(f374,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f164])).
fof(f164,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f118])).
fof(f118,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.bool.IMP_ANTISYM_AX')).
fof(f377,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f165])).
fof(f373,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X1)))))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X0),s('type.rat.rat',X1))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X1)))))),
  inference(rectify,[],[f46])).
fof(f46,axiom,(
  ! [X6,X7] : s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X6),s('type.rat.rat',X7))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X7)))))),
  file('thm.rat.RAT_LDIV_LES_POS','thm.rat.RAT_DIV_MULMINV')).
fof(f310,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f201])).
fof(f8077,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f8076,f8025])).
fof(f8076,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl80),
  inference(subsumption_resolution,[],[f8026,f310])).
fof(f8026,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f8025,f314])).
fof(f8020,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7796,f1568])).
fof(f7796,plain,(
  sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f1290])).
fof(f8019,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7795,f527])).
fof(f7795,plain,(
  ~sP3('T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f1282])).
fof(f8014,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7788,f528])).
fof(f7788,plain,(
  ~sP3('const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'T_0','F_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f1199])).
fof(f8009,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl90),
  inference(cnf_transformation,[],[f8009_D])).
fof(f8009_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f8010,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl92),
  inference(cnf_transformation,[],[f8010_D])).
fof(f8010_D,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f8011,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl80),
  inference(subsumption_resolution,[],[f7779,f527])).
fof(f7779,plain,(
  ( ! [X0,X1] : (~sP3(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f886])).
fof(f8008,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl80),
  inference(subsumption_resolution,[],[f7778,f525])).
fof(f7778,plain,(
  ( ! [X0,X1] : (~sP4(X0,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),X1) | p(s(bool,X1)) | ~p(s(bool,X0))) ) | $spl80),
  inference(backward_demodulation,[],[f7727,f863])).
fof(f8007,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7776,f1293])).
fof(f7776,plain,(
  ~sP5('F_0','T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | $spl80),
  inference(backward_demodulation,[],[f7727,f788])).
fof(f8006,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7775,f1194])).
fof(f7775,plain,(
  ~sP5('F_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'T_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f786])).
fof(f8005,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7774,f1551])).
fof(f7774,plain,(
  sP5('T_0','T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | $spl80),
  inference(backward_demodulation,[],[f7727,f767])).
fof(f8004,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7773,f1548])).
fof(f7773,plain,(
  sP5('T_0','const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11)),'T_0') | $spl80),
  inference(backward_demodulation,[],[f7727,f766])).
fof(f7999,plain,(
  $false | $spl80),
  inference(subsumption_resolution,[],[f7753,f533])).
fof(f7753,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK11))) | $spl80),
  inference(backward_demodulation,[],[f7727,f536])).
fof(f7727,plain,(
  s('type.rat.rat',sK11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl80),
  inference(cnf_transformation,[],[f7727_D])).
fof(f7727_D,plain,(
  s('type.rat.rat',sK11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f7731,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK11))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X20))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl88),
  inference(cnf_transformation,[],[f7731_D])).
fof(f7731_D,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK11))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X20))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f7680,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X20),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK11))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X20))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.rat.rat',sK11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(superposition,[],[f1523,f1462])).
fof(f7730,plain,(
  ( ! [X6,X5] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',sK11))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X6))),s('type.rat.rat',X5)))) ) | $spl86),
  inference(cnf_transformation,[],[f7730_D])).
fof(f7730_D,plain,(
  ( ! [X6,X5] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',sK11))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X6))),s('type.rat.rat',X5)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f7674,plain,(
  ( ! [X6,X5] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X6),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',sK11))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X6))),s('type.rat.rat',X5))) | s('type.rat.rat',sK11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1523,f1468])).
fof(f7729,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK11))),s('type.rat.rat',X20))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X20)))))) ) | $spl84),
  inference(cnf_transformation,[],[f7729_D])).
fof(f7729_D,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK11))),s('type.rat.rat',X20))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X20)))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f7653,plain,(
  ( ! [X20] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',sK11))),s('type.rat.rat',X20))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X20))))) | s('type.rat.rat',sK11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(superposition,[],[f1523,f1462])).
fof(f7728,plain,(
  ( ! [X6,X5] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X6))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',sK11))),s('type.rat.rat',X6)))) ) | $spl82),
  inference(cnf_transformation,[],[f7728_D])).
fof(f7728_D,plain,(
  ( ! [X6,X5] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X6))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',sK11))),s('type.rat.rat',X6)))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f7647,plain,(
  ( ! [X6,X5] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',X6))))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat',sK11))),s('type.rat.rat',X6))) | s('type.rat.rat',sK11) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1523,f1468])).
fof(f2925,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl46),
  inference(cnf_transformation,[],[f2925_D])).
fof(f2925_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f5147,plain,(
  ( ! [X3] : (s('type.rat.rat',X3) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) | $spl78),
  inference(cnf_transformation,[],[f5147_D])).
fof(f5147_D,plain,(
  ( ! [X3] : (s('type.rat.rat',X3) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5085,plain,(
  ( ! [X10] : (s('type.rat.rat',X10) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X10),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f1468,f1308])).
fof(f5073,plain,(
  ( ! [X3] : (s('type.rat.rat',X3) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X3),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))) )),
  inference(superposition,[],[f1308,f1468])).
fof(f1462,plain,(
  ( ! [X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X7))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(cnf_transformation,[],[f1462_D])).
fof(f1462_D,plain,(
  ( ! [X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X7))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f4958,plain,(
  ( ! [X117] : (s('type.rat.rat',X117) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X117),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) ) | $spl76),
  inference(cnf_transformation,[],[f4958_D])).
fof(f4958_D,plain,(
  ( ! [X117] : (s('type.rat.rat',X117) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X117),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f4921,plain,(
  ( ! [X116,X117] : (s('type.rat.rat',X117) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X117),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X116),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X116))))) | s('type.rat.rat',X116) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f1468,f308])).
fof(f4948,plain,(
  $false | ($spl4 | $spl47)),
  inference(subsumption_resolution,[],[f4935,f2924])).
fof(f2924,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl47),
  inference(cnf_transformation,[],[f2924_D])).
fof(f2924_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f4935,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl4),
  inference(condensation,[],[f4934])).
fof(f4934,plain,(
  ( ! [X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f4868])).
fof(f4868,plain,(
  ( ! [X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat',X1) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(superposition,[],[f1462,f308])).
fof(f4944,plain,(
  $false | ($spl4 | $spl47)),
  inference(subsumption_resolution,[],[f4937,f2924])).
fof(f4937,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | $spl4),
  inference(condensation,[],[f4936])).
fof(f4936,plain,(
  ( ! [X2] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(duplicate_literal_removal,[],[f4862])).
fof(f4862,plain,(
  ( ! [X2] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.rat.rat',X2) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl4),
  inference(superposition,[],[f308,f1462])).
fof(f4687,plain,(
  ( ! [X31,X32] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X31)))) ) | $spl74),
  inference(cnf_transformation,[],[f4687_D])).
fof(f4687_D,plain,(
  ( ! [X31,X32] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X31)))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4676,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl68),
  inference(cnf_transformation,[],[f4676_D])).
fof(f4676_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f4686,plain,(
  ( ! [X31,X32] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X31)))) ) | $spl46),
  inference(forward_demodulation,[],[f4685,f2940])).
fof(f2940,plain,(
  ( ! [X3] : (s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X3))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat',X3)))) ) | $spl46),
  inference(backward_demodulation,[],[f2925,f2882])).
fof(f4685,plain,(
  ( ! [X31,X32] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X31)))) ) | $spl46),
  inference(forward_demodulation,[],[f4523,f373])).
fof(f4523,plain,(
  ( ! [X31,X32] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X31),s('type.rat.rat',X32))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X32),s('type.rat.rat',X31)))) ) | $spl46),
  inference(superposition,[],[f2935,f2935])).
fof(f2935,plain,(
  ( ! [X0,X1] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X1),s('type.rat.rat',X0))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X0),s('type.rat.rat',X1)))) ) | $spl46),
  inference(backward_demodulation,[],[f2925,f1434])).
fof(f4680,plain,(
  ( ! [X20] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X20))),s('type.rat.rat',X20))) | s('type.rat.rat',X20) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl70),
  inference(cnf_transformation,[],[f4680_D])).
fof(f4680_D,plain,(
  ( ! [X20] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X20))),s('type.rat.rat',X20))) | s('type.rat.rat',X20) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4682,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | $spl72),
  inference(cnf_transformation,[],[f4682_D])).
fof(f4682_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f4518,plain,(
  ( ! [X20] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X20))),s('type.rat.rat',X20))) | s('type.rat.rat',X20) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | ($spl4 | $spl46)),
  inference(superposition,[],[f2935,f1462])).
fof(f4674,plain,(
  ( ! [X10,X11] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11))))),s('type.rat.rat',X10))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11)))) ) | $spl66),
  inference(cnf_transformation,[],[f4674_D])).
fof(f4674_D,plain,(
  ( ! [X10,X11] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11))))),s('type.rat.rat',X10))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11)))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f4673,plain,(
  ( ! [X10,X11] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11))))),s('type.rat.rat',X10))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11)))) ) | $spl46),
  inference(forward_demodulation,[],[f4672,f2940])).
fof(f4672,plain,(
  ( ! [X10,X11] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11))))),s('type.rat.rat',X10))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11)))) ) | $spl46),
  inference(forward_demodulation,[],[f4514,f373])).
fof(f4514,plain,(
  ( ! [X10,X11] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X11),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11))))),s('type.rat.rat',X10))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X10),s('type.rat.rat',X11)))) ) | $spl46),
  inference(superposition,[],[f2935,f2938])).
fof(f2938,plain,(
  ( ! [X4,X5] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X5),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X4),s('type.rat.rat',X5)))) ) | $spl46),
  inference(backward_demodulation,[],[f2925,f1466])).
fof(f4157,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK13(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X0,X1,bool,bool))))) ) | $spl62),
  inference(cnf_transformation,[],[f4157_D])).
fof(f4157_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK13(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X0,X1,bool,bool))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f4159,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl65),
  inference(cnf_transformation,[],[f4159_D])).
fof(f4159_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f4147,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK13(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1423,f674])).
fof(f4146,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK13(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK13(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1423,f674])).
fof(f4132,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl56),
  inference(cnf_transformation,[],[f4132_D])).
fof(f4132_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f4128,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool))) ) | $spl52),
  inference(cnf_transformation,[],[f4128_D])).
fof(f4128_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f4135,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl61),
  inference(cnf_transformation,[],[f4135_D])).
fof(f4135_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f4126,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK13(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1422,f297])).
fof(f4133,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl58),
  inference(cnf_transformation,[],[f4133_D])).
fof(f4133_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f4125,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK13(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1422,f674])).
fof(f4124,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK13(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1422,f674])).
fof(f4127,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl50),
  inference(cnf_transformation,[],[f4127_D])).
fof(f4127_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f4130,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl55),
  inference(cnf_transformation,[],[f4130_D])).
fof(f4130_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f4123,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK13(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1422,f674])).
fof(f4122,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK13(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1422,f674])).
fof(f2927,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | $spl48),
  inference(cnf_transformation,[],[f2927_D])).
fof(f2927_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2888,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f308,f1308])).
fof(f2883,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))),
  inference(superposition,[],[f1308,f308])).
fof(f2321,plain,(
  $false | ($spl9 | $spl14)),
  inference(subsumption_resolution,[],[f2316,f288])).
fof(f2316,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl14)),
  inference(backward_demodulation,[],[f2198,f1517])).
fof(f2320,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2303,f288])).
fof(f2303,plain,(
  ~p(s(bool,'T_0')) | $spl14),
  inference(backward_demodulation,[],[f2198,f496])).
fof(f2319,plain,(
  $false | $spl14),
  inference(subsumption_resolution,[],[f2318,f2198])).
fof(f2318,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0)) ) | $spl14),
  inference(subsumption_resolution,[],[f2301,f310])).
fof(f2301,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl14),
  inference(backward_demodulation,[],[f2198,f314])).
fof(f2244,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~sP6('T_0',X1,X2)) ) | $spl44),
  inference(cnf_transformation,[],[f2244_D])).
fof(f2244_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~sP6('T_0',X1,X2)) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2198,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl14),
  inference(cnf_transformation,[],[f2198_D])).
fof(f2198_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2300,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP6('T_0',X1,X2) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f2299])).
fof(f2299,plain,(
  ( ! [X92,X90,X93,X91] : (~sP6('T_0',X91,X92) | sP7('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f2023,f735])).
fof(f2023,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X90)) | ~sP6('T_0',X91,X92) | sP7('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f486,f674])).
fof(f2241,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl42),
  inference(cnf_transformation,[],[f2241_D])).
fof(f2241_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2298,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2297])).
fof(f2297,plain,(
  ( ! [X88,X87,X89,X86] : (~sP8('T_0',X87,X88) | sP9('T_0',X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f2022,f735])).
fof(f2022,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | ~sP8('T_0',X87,X88) | sP9('T_0',X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f485,f674])).
fof(f2238,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl40),
  inference(cnf_transformation,[],[f2238_D])).
fof(f2238_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2296,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2295])).
fof(f2295,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X83)) | p(s(bool,X84)) | ~sP9('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f2021,f735])).
fof(f2021,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP9('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f482,f674])).
fof(f2235,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl38),
  inference(cnf_transformation,[],[f2235_D])).
fof(f2235_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2294,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2293])).
fof(f2293,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X79)) | sP9('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f2020,f735])).
fof(f2020,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | ~p(s(bool,X79)) | sP9('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f481,f674])).
fof(f2232,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | p(s(bool,X1))) ) | $spl36),
  inference(cnf_transformation,[],[f2232_D])).
fof(f2232_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | p(s(bool,X1))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2292,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2291])).
fof(f2291,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | sP9('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f2019,f735])).
fof(f2019,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | sP9('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f480,f674])).
fof(f2229,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) | $spl34),
  inference(cnf_transformation,[],[f2229_D])).
fof(f2229_D,plain,(
  ( ! [X2,X1] : (~sP8(X2,X1,'T_0') | ~p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2290,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f2289])).
fof(f2289,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP8(X72,X71,'T_0') | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f2018,f735])).
fof(f2018,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP8(X72,X71,'T_0') | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f475,f674])).
fof(f2226,plain,(
  ( ! [X2,X1] : (~sP7('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl32),
  inference(cnf_transformation,[],[f2226_D])).
fof(f2226_D,plain,(
  ( ! [X2,X1] : (~sP7('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2288,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f2287])).
fof(f2287,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP7('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f2017,f735])).
fof(f2017,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP7('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f465,f674])).
fof(f2223,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl30),
  inference(cnf_transformation,[],[f2223_D])).
fof(f2223_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2286,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f2285])).
fof(f2285,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | sP7('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f2016,f735])).
fof(f2016,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | sP7('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f464,f674])).
fof(f2220,plain,(
  ( ! [X2,X1] : (sP7('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f2220_D])).
fof(f2220_D,plain,(
  ( ! [X2,X1] : (sP7('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2284,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f2283])).
fof(f2283,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X59)) | sP7('T_0',X60,X59) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f2015,f735])).
fof(f2015,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | ~p(s(bool,X59)) | sP7('T_0',X60,X59) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f463,f674])).
fof(f2217,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl26),
  inference(cnf_transformation,[],[f2217_D])).
fof(f2217_D,plain,(
  ( ! [X2,X1] : (~sP6('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2282,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f2281])).
fof(f2281,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X55)) | p(s(bool,X56)) | ~sP6('T_0',X55,X56) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f2014,f735])).
fof(f2014,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP6('T_0',X55,X56) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f458,f674])).
fof(f2214,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl24),
  inference(cnf_transformation,[],[f2214_D])).
fof(f2214_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2280,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2279])).
fof(f2279,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP3(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f2013,f735])).
fof(f2013,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP3(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f439,f674])).
fof(f2211,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f2211_D])).
fof(f2211_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2278,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2277])).
fof(f2277,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP4(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f2012,f735])).
fof(f2012,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP4(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f435,f674])).
fof(f2208,plain,(
  ( ! [X2,X1] : (~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2)) ) | $spl20),
  inference(cnf_transformation,[],[f2208_D])).
fof(f2208_D,plain,(
  ( ! [X2,X1] : (~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2276,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2275,f514])).
fof(f2275,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2274])).
fof(f2274,plain,(
  ( ! [X45,X43,X44,X42] : (sP5(X43,X44,'T_0') | ~sP4(X43,'T_0',X44) | ~sP3(X44,'T_0',X43) | ~sP2('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f2011,f735])).
fof(f2011,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | sP5(X43,X44,'T_0') | ~sP4(X43,'T_0',X44) | ~sP3(X44,'T_0',X43) | ~sP2('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f432,f674])).
fof(f2204,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f2204_D])).
fof(f2204_D,plain,(
  ( ! [X2,X1] : (~sP5(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2273,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2272])).
fof(f2272,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f2010,f735])).
fof(f2010,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f431,f674])).
fof(f2201,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl16),
  inference(cnf_transformation,[],[f2201_D])).
fof(f2201_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f2271,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2270])).
fof(f2270,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f2009,f735])).
fof(f2009,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f406,f674])).
fof(f2197,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl12),
  inference(cnf_transformation,[],[f2197_D])).
fof(f2197_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2269,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2268])).
fof(f2268,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f2008,f735])).
fof(f2008,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f403,f674])).
fof(f2243,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP6('T_0',X1,X2) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f2242])).
fof(f2242,plain,(
  ( ! [X92,X90,X93,X91] : (~sP6('T_0',X91,X92) | sP7('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(subsumption_resolution,[],[f1952,f735])).
fof(f1952,plain,(
  ( ! [X92,X90,X93,X91] : (~p(s(bool,X90)) | ~sP6('T_0',X91,X92) | sP7('T_0',X91,X92) | s(bool,X90) = s(bool,X93) | s(bool,'T_0') = s(bool,X93)) )),
  inference(superposition,[],[f486,f674])).
fof(f2240,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2239])).
fof(f2239,plain,(
  ( ! [X88,X87,X89,X86] : (~sP8('T_0',X87,X88) | sP9('T_0',X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(subsumption_resolution,[],[f1951,f735])).
fof(f1951,plain,(
  ( ! [X88,X87,X89,X86] : (~p(s(bool,X86)) | ~sP8('T_0',X87,X88) | sP9('T_0',X87,X88) | s(bool,X86) = s(bool,X89) | s(bool,'T_0') = s(bool,X89)) )),
  inference(superposition,[],[f485,f674])).
fof(f2237,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2236])).
fof(f2236,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X83)) | p(s(bool,X84)) | ~sP9('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(subsumption_resolution,[],[f1950,f735])).
fof(f1950,plain,(
  ( ! [X85,X83,X84,X82] : (~p(s(bool,X82)) | ~p(s(bool,X83)) | p(s(bool,X84)) | ~sP9('T_0',X84,X83) | s(bool,X82) = s(bool,X85) | s(bool,'T_0') = s(bool,X85)) )),
  inference(superposition,[],[f482,f674])).
fof(f2234,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f2233])).
fof(f2233,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X79)) | sP9('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(subsumption_resolution,[],[f1949,f735])).
fof(f1949,plain,(
  ( ! [X80,X78,X81,X79] : (~p(s(bool,X78)) | ~p(s(bool,X79)) | sP9('T_0',X79,X80) | s(bool,X78) = s(bool,X81) | s(bool,'T_0') = s(bool,X81)) )),
  inference(superposition,[],[f481,f674])).
fof(f2231,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f2230])).
fof(f2230,plain,(
  ( ! [X76,X74,X77,X75] : (p(s(bool,X75)) | sP9('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(subsumption_resolution,[],[f1948,f735])).
fof(f1948,plain,(
  ( ! [X76,X74,X77,X75] : (~p(s(bool,X74)) | p(s(bool,X75)) | sP9('T_0',X76,X75) | s(bool,X74) = s(bool,X77) | s(bool,'T_0') = s(bool,X77)) )),
  inference(superposition,[],[f480,f674])).
fof(f2228,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP8(X2,X1,'T_0')) )),
  inference(condensation,[],[f2227])).
fof(f2227,plain,(
  ( ! [X70,X72,X71,X73] : (p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP8(X72,X71,'T_0') | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(subsumption_resolution,[],[f1947,f735])).
fof(f1947,plain,(
  ( ! [X70,X72,X71,X73] : (~p(s(bool,X70)) | p(s(bool,X71)) | ~p(s(bool,X72)) | ~sP8(X72,X71,'T_0') | s(bool,X70) = s(bool,X73) | s(bool,'T_0') = s(bool,X73)) )),
  inference(superposition,[],[f475,f674])).
fof(f2225,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f2224])).
fof(f2224,plain,(
  ( ! [X68,X66,X69,X67] : (p(s(bool,X67)) | p(s(bool,X68)) | ~sP7('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(subsumption_resolution,[],[f1946,f735])).
fof(f1946,plain,(
  ( ! [X68,X66,X69,X67] : (~p(s(bool,X66)) | p(s(bool,X67)) | p(s(bool,X68)) | ~sP7('T_0',X68,X67) | s(bool,X66) = s(bool,X69) | s(bool,'T_0') = s(bool,X69)) )),
  inference(superposition,[],[f465,f674])).
fof(f2222,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X1,X2)) )),
  inference(condensation,[],[f2221])).
fof(f2221,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X63)) | sP7('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(subsumption_resolution,[],[f1945,f735])).
fof(f1945,plain,(
  ( ! [X64,X62,X65,X63] : (~p(s(bool,X62)) | ~p(s(bool,X63)) | sP7('T_0',X63,X64) | s(bool,X62) = s(bool,X65) | s(bool,'T_0') = s(bool,X65)) )),
  inference(superposition,[],[f464,f674])).
fof(f2219,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP7('T_0',X2,X1)) )),
  inference(condensation,[],[f2218])).
fof(f2218,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X59)) | sP7('T_0',X60,X59) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f1944,f735])).
fof(f1944,plain,(
  ( ! [X61,X59,X60,X58] : (~p(s(bool,X58)) | ~p(s(bool,X59)) | sP7('T_0',X60,X59) | s(bool,X58) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f463,f674])).
fof(f2216,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP6('T_0',X1,X2)) )),
  inference(condensation,[],[f2215])).
fof(f2215,plain,(
  ( ! [X57,X54,X56,X55] : (p(s(bool,X55)) | p(s(bool,X56)) | ~sP6('T_0',X55,X56) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(subsumption_resolution,[],[f1943,f735])).
fof(f1943,plain,(
  ( ! [X57,X54,X56,X55] : (~p(s(bool,X54)) | p(s(bool,X55)) | p(s(bool,X56)) | ~sP6('T_0',X55,X56) | s(bool,X54) = s(bool,X57) | s(bool,'T_0') = s(bool,X57)) )),
  inference(superposition,[],[f458,f674])).
fof(f2213,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f2212])).
fof(f2212,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP3(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(subsumption_resolution,[],[f1942,f735])).
fof(f1942,plain,(
  ( ! [X52,X50,X53,X51] : (~p(s(bool,X50)) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP3(X51,'T_0',X52) | s(bool,X50) = s(bool,X53) | s(bool,'T_0') = s(bool,X53)) )),
  inference(superposition,[],[f439,f674])).
fof(f2210,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f2209])).
fof(f2209,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X47)) | p(s(bool,X48)) | ~sP4(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(subsumption_resolution,[],[f1941,f735])).
fof(f1941,plain,(
  ( ! [X47,X48,X46,X49] : (~p(s(bool,X46)) | ~p(s(bool,X47)) | p(s(bool,X48)) | ~sP4(X47,'T_0',X48) | s(bool,X46) = s(bool,X49) | s(bool,'T_0') = s(bool,X49)) )),
  inference(superposition,[],[f435,f674])).
fof(f2207,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f2206,f514])).
fof(f2206,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f2205])).
fof(f2205,plain,(
  ( ! [X45,X43,X44,X42] : (sP5(X43,X44,'T_0') | ~sP4(X43,'T_0',X44) | ~sP3(X44,'T_0',X43) | ~sP2('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(subsumption_resolution,[],[f1940,f735])).
fof(f1940,plain,(
  ( ! [X45,X43,X44,X42] : (~p(s(bool,X42)) | sP5(X43,X44,'T_0') | ~sP4(X43,'T_0',X44) | ~sP3(X44,'T_0',X43) | ~sP2('T_0',X44,X43) | s(bool,X42) = s(bool,X45) | s(bool,'T_0') = s(bool,X45)) )),
  inference(superposition,[],[f432,f674])).
fof(f2203,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP5(X1,'T_0',X2)) )),
  inference(condensation,[],[f2202])).
fof(f2202,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(subsumption_resolution,[],[f1939,f735])).
fof(f1939,plain,(
  ( ! [X39,X41,X38,X40] : (~p(s(bool,X38)) | ~p(s(bool,X39)) | p(s(bool,X40)) | ~sP5(X39,'T_0',X40) | s(bool,X38) = s(bool,X41) | s(bool,'T_0') = s(bool,X41)) )),
  inference(superposition,[],[f431,f674])).
fof(f2200,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f2199])).
fof(f2199,plain,(
  ( ! [X37,X35,X36] : (p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(subsumption_resolution,[],[f1938,f735])).
fof(f1938,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X35)) | p(s(bool,X36)) | sP1(X36,'T_0') | s(bool,X35) = s(bool,X37) | s(bool,'T_0') = s(bool,X37)) )),
  inference(superposition,[],[f406,f674])).
fof(f2196,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f2195])).
fof(f2195,plain,(
  ( ! [X33,X34,X32] : (p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f1937,f735])).
fof(f1937,plain,(
  ( ! [X33,X34,X32] : (~p(s(bool,X32)) | p(s(bool,X33)) | sP1('T_0',X33) | s(bool,X32) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f403,f674])).
fof(f1517,plain,(
  ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | $spl9),
  inference(cnf_transformation,[],[f1517_D])).
fof(f1517_D,plain,(
  ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f1518,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16))),s('type.rat.rat',X15))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X15))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl10),
  inference(cnf_transformation,[],[f1518_D])).
fof(f1518_D,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16))),s('type.rat.rat',X15))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X15))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f1515,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X16))),s('type.rat.rat',X15))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X15))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1514,f1409])).
fof(f1514,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X16),s('type.rat.rat',X15))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X15))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1513,f1340])).
fof(f1513,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15))),s('type.rat.rat',X16))))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X15))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1512,f420])).
fof(f1512,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15))))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat',X17))),s('type.rat.rat',X15))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1511,f1409])).
fof(f1511,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15))))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',X17),s('type.rat.rat',X15))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1510,f1340])).
fof(f1510,plain,(
  ( ! [X17,X15,X16] : (s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15))))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15))),s('type.rat.rat',X17))))))) | ~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(forward_demodulation,[],[f1492,f420])).
fof(f1492,plain,(
  ( ! [X17,X15,X16] : (~p(s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0')))))))))) | s(bool,'const.rat.rat_les_2'(s('type.rat.rat',X16),s('type.rat.rat',X17))) = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15))))),s('type.rat.rat',X16))),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X15),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X15))))),s('type.rat.rat',X17))))) | s('type.rat.rat',X15) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f421,f308])).
fof(f1464,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | $spl6),
  inference(cnf_transformation,[],[f1464_D])).
fof(f1464_D,plain,(
  s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f1438,plain,(
  ( ! [X7] : (s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.arithmetic.NUMERAL_1'(s('type.num.num','const.arithmetic.BIT|49|_1'(s('type.num.num','const.arithmetic.ZERO_0'))))))))))) | s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',X7),s('type.rat.rat','const.rat.rat_minv_1'(s('type.rat.rat',X7))))) | s('type.rat.rat',X7) = s('type.rat.rat','const.rat.rat_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f308,f308])).
fof(f800,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12))) | $spl0),
  inference(cnf_transformation,[],[f800_D])).
fof(f800_D,plain,(
  s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f802,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12))))) | $spl3),
  inference(cnf_transformation,[],[f802_D])).
fof(f802_D,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f798,plain,(
  s(bool,'F_0') != s(bool,'const.rat.rat_les_2'(s('type.rat.rat',sK10),s('type.rat.rat','const.rat.rat_mul_2'(s('type.rat.rat',sK11),s('type.rat.rat',sK12))))) | s(bool,'T_0') = s(bool,'const.rat.rat_les_2'(s('type.rat.rat','const.rat.rat_div_2'(s('type.rat.rat',sK10),s('type.rat.rat',sK11))),s('type.rat.rat',sK12)))),
  inference(superposition,[],[f287,f297])).
% SZS output end Proof for thm.rat
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 404
% Passive clauses: 1349
% Generated clauses: 9170
% Final active clauses: 245
% Final passive clauses: 418
% Input formulas: 61
% Initial clauses: 199
% 
% Duplicate literals: 146
% Trivial inequalities: 23
% Fw subsumption resolutions: 285
% Fw demodulations: 2112
% Bw demodulations: 423
% Condensations: 135
% Global subsumptions: 1
% 
% Simple tautologies: 105
% Equational tautologies: 54
% Forward subsumptions: 4832
% Fw demodulations to eq. taut.: 45
% Bw demodulations to eq. taut.: 2
% 
% Unit resulting resolution: 1699
% Factoring: 1
% Forward superposition: 1474
% Backward superposition: 2113
% Self superposition: 87
% Equality factoring: 28
% Equality resolution: 1
% 
% Split clauses: 56
% Split components: 117
% SAT solver clauses: 1491
% SAT solver unit clauses: 286
% SAT solver binary clauses: 441
% 
% Sat splits: 56
% Sat splitting refutations: 27
% 
% Memory used [KB]: 10746
% Time elapsed: 0.571 s
% ------------------------------
% ------------------------------
% Success in time 1.854 s
