Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.relation
% (62852)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 636
% Passive clauses: 9441
% Generated clauses: 18573
% Final active clauses: 570
% Final passive clauses: 8295
% Input formulas: 37
% Initial clauses: 142
% 
% Duplicate literals: 307
% Trivial inequalities: 96
% Fw subsumption resolutions: 1421
% Fw demodulations: 464
% Bw demodulations: 65
% 
% Simple tautologies: 187
% Equational tautologies: 66
% Forward subsumptions: 6378
% 
% Binary resolution: 10129
% Factoring: 1304
% Forward superposition: 2171
% Backward superposition: 2185
% Self superposition: 36
% Equality factoring: 26
% Equality resolution: 39
% 
% Split clauses: 318
% Split components: 642
% SAT solver clauses: 665
% SAT solver unit clauses: 21
% SAT solver binary clauses: 151
% 
% Sat splits: 317
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 659
% TWLsolver calls for satisfiability: 399
% 
% Memory used [KB]: 8827
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.relation
% (62861)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 368
% Passive clauses: 2156
% Generated clauses: 14677
% Final active clauses: 306
% Final passive clauses: 1584
% Input formulas: 37
% Initial clauses: 142
% 
% Duplicate literals: 276
% Trivial inequalities: 138
% Fw subsumption resolutions: 1050
% Fw demodulations: 1504
% Bw demodulations: 216
% Condensations: 440
% Global subsumptions: 14
% 
% Simple tautologies: 329
% Equational tautologies: 26
% Forward subsumptions: 5195
% 
% Binary resolution: 2305
% Factoring: 65
% Forward superposition: 1976
% Backward superposition: 2485
% Self superposition: 36
% Equality factoring: 32
% Equality resolution: 31
% 
% Split clauses: 169
% Split components: 352
% SAT solver clauses: 3050
% SAT solver unit clauses: 160
% SAT solver binary clauses: 79
% SAT solver learnt clauses: 30
% SAT solver learnt literals: 15
% 
% Sat splits: 169
% Sat splitting refutations: 15
% 
% TWLsolver clauses: 2989
% TWLsolver calls for satisfiability: 1434
% 
% Memory used [KB]: 13176
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.relation
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.relation
% SZS output start Proof for thm.relation
fof(f27159,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f200,f352,f350,f348,f199,f351,f198,f349,f347,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f4553,f4490,f4555,f4494,f4558,f4498,f4560,f4501,f4562,f4504,f4564,f4507,f4566,f4510,f4568,f4513,f4570,f4516,f4572,f4519,f4574,f4522,f4576,f4525,f4729,f4730,f4735,f4737,f4739,f4741,f4743,f4745,f4747,f4749,f4751,f4753,f4755,f4757,f4759,f4761,f4763,f4765,f4772,f4774,f4776,f4778,f4780,f4782,f4784,f4786,f4788,f4790,f4792,f4794,f4796,f4798,f4800,f4802,f4804,f4806,f4808,f4810,f4812,f4814,f4816,f4818,f4820,f4822,f4824,f4826,f4828,f4830,f4832,f4834,f4836,f7343,f7346,f7344,f7404,f7407,f7405,f7435,f7438,f7436,f7468,f7471,f7469,f8154,f8312,f8156,f8330,f8331,f8338,f8335,f8332,f8333,f8341,f8334,f8343,f8336,f8340,f8771,f8812,f8813,f8825,f8823,f8901,f8904,f8902,f8953,f8956,f8954,f9005,f9008,f9006,f9058,f9061,f9059,f12358,f12361,f12359,f12403,f12406,f12404,f12941,f12944,f12942,f12988,f12991,f12989,f13048,f13051,f13049,f13106,f13109,f13107,f20999,f21001,f21000,f21002,f21004,f21003,f21011,f21014,f8772,f8155,f21887,f21889,f21888,f21898,f21900,f21899,f4491,f21907,f21012,f8313,f22651,f22652,f22654,f22653,f22655,f24622,f24751,f24752,f27150,f27154])).
fof(f27154,plain,(
  $false | ($spl1 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f27153,f353])).
fof(f353,plain,(
  ~p(s(bool,'F_0'))),
  inference(subsumption_resolution,[],[f204,f201])).
fof(f201,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.relation.IN_RDOM_DELETE',aHL_TRUTH)).
fof(f204,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f137])).
fof(f137,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f19])).
fof(f19,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.relation.IN_RDOM_DELETE','thm.bool.NOT_CLAUSES|split|2')).
fof(f27153,plain,(
  p(s(bool,'F_0')) | ($spl1 | $spl2 | $spl5)),
  inference(forward_demodulation,[],[f26768,f26018])).
fof(f26018,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13),s(sK11,sK12),s(sK10,X0)))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f386,f25215,f606])).
fof(f606,plain,(
  ( ! [X4,X3] : (p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(subsumption_resolution,[],[f602,f353])).
fof(f602,plain,(
  ( ! [X4,X3] : (p(s(bool,'F_0')) | p(s(bool,X4)) | s(bool,X3) = s(bool,X4) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f244,f208])).
fof(f208,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.relation.IN_RDOM_DELETE',aHL_BOOL_CASES)).
fof(f244,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X12,X13] : ((p(s(bool,X12)) => p(s(bool,X13))) => ((p(s(bool,X13)) => p(s(bool,X12))) => s(bool,X12) = s(bool,X13)))),
  file('thm.relation.IN_RDOM_DELETE','thm.bool.IMP_ANTISYM_AX')).
fof(f25215,plain,(
  ( ! [X0] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13),s(sK11,sK12),s(sK10,X0))))) ) | $spl1),
  inference(forward_demodulation,[],[f24986,f1518])).
fof(f1518,plain,(
  ( ! [X6,X10,X8,X7,X11,X9] : (s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X7,fun(X6,bool)),X8),s(X7,X9),s(X7,X10),s(X6,X11))) = s(bool,'const.bool.IN_2'(s(X6,X11),s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X7,fun(X6,bool)),X8),s(X7,X9))),s(X7,X10)))))) )),
  inference(superposition,[],[f279,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(cnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0] : ! [X1,X2] : s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2))) = s(bool,happ(s(fun(X0,bool),X2),s(X0,X1)))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X7] : ! [X5,X11] : s(bool,'const.bool.IN_2'(s(X7,X5),s(fun(X7,bool),X11))) = s(bool,happ(s(fun(X7,bool),X11),s(X7,X5)))),
  file('thm.relation.IN_RDOM_DELETE','thm.bool.IN_DEF')).
fof(f279,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(cnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5))) = s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X7,X6] : ! [X21,X22,X23,X24] : s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X7,fun(X6,bool)),X21),s(X7,X22),s(X7,X23),s(X6,X24))) = s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X7,fun(X6,bool)),X21),s(X7,X22))),s(X7,X23))),s(X6,X24)))),
  file('thm.relation.IN_RDOM_DELETE','ARITY_2_4_const.relation.RDOM_DELETE')).
fof(f24986,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))),s(sK11,sK12))))))) ) | $spl1),
  inference(unit_resulting_resolution,[],[f24753,f355])).
fof(f355,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))) )),
  inference(forward_demodulation,[],[f278,f243])).
fof(f278,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f168,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,sK16(X3,X2,X1,X0)))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK16])],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ? [X5] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X5))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(rectify,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4))))) & (! [X4] : ~p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))) | p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))))),
  inference(flattening,[],[f83])).
fof(f83,plain,(
  ! [X0,X1] : ! [X2,X3] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) <=> ? [X4] : p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,X4)))))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X7,X6] : ! [X10,X8] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X7,fun(X6,bool)),X10),s(X7,X8)))) <=> ? [X15] : p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X8))),s(X6,X15)))))),
  file('thm.relation.IN_RDOM_DELETE','thm.relation.RDOM_DEF')).
fof(f24753,plain,(
  ~p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))),s(sK11,sK12)))) | $spl1),
  inference(forward_demodulation,[],[f348,f1251])).
fof(f1251,plain,(
  ( ! [X6,X4,X7,X5] : (s(bool,'const.relation.RDOM_2'(s(fun(X4,fun(X5,bool)),X6),s(X4,X7))) = s(bool,'const.bool.IN_2'(s(X4,X7),s(fun(X4,bool),'const.relation.RDOM_1'(s(fun(X4,fun(X5,bool)),X6)))))) )),
  inference(superposition,[],[f275,f243])).
fof(f275,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))) )),
  inference(cnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0,X1,X2,X3] : s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0,X1] : ! [X2,X3] : s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))) = s(bool,happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X2))),s(X1,X3)))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6,X7] : ! [X21,X22] : s(bool,'const.relation.RDOM_2'(s(fun(X7,fun(X6,bool)),X21),s(X7,X22))) = s(bool,happ(s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),X21))),s(X7,X22)))),
  file('thm.relation.IN_RDOM_DELETE','ARITY_1_2_const.relation.RDOM')).
fof(f386,plain,(
  s(bool,'T_0') != s(bool,'F_0')),
  inference(unit_resulting_resolution,[],[f353,f216])).
fof(f216,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f51])).
fof(f51,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X14] : (s(bool,'T_0') = s(bool,X14) <=> p(s(bool,X14)))),
  file('thm.relation.IN_RDOM_DELETE','thm.bool.EQ_CLAUSES|split|1')).
fof(f26768,plain,(
  p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13),s(sK11,sK12),s(sK10,sK16(sK12,sK14,sK10,sK11))))) | ($spl2 | $spl5)),
  inference(unit_resulting_resolution,[],[f351,f24755,f357])).
fof(f357,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | s(X1,X3) = s(X1,X4)) )),
  inference(forward_demodulation,[],[f282,f243])).
fof(f282,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | s(X1,X3) = s(X1,X4) | ~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & s(X1,X3) != s(X1,X4))) & (~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | s(X1,X3) = s(X1,X4) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5))))))),
  inference(flattening,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5] : ((~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & s(X1,X3) != s(X1,X4))) & ((~p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | s(X1,X3) = s(X1,X4)) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5))))))),
  inference(nnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1,X2,X3,X4,X5] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & s(X1,X3) != s(X1,X4)))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) <=> (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) & ~s(X1,X3) = s(X1,X4)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X6,X7] : ! [X10,X8,X16,X17] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X7,fun(X6,bool)),X10),s(X7,X8),s(X7,X16),s(X6,X17)))) <=> (p(s(bool,happ(s(fun(X6,bool),happ(s(fun(X7,fun(X6,bool)),X10),s(X7,X16))),s(X6,X17)))) & ~s(X7,X8) = s(X7,X16)))),
  file('thm.relation.IN_RDOM_DELETE','thm.relation.RDOM_DELETE_DEF')).
fof(f24755,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK16(sK12,sK14,sK10,sK11)),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f24754,f356])).
fof(f356,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) | p(s(bool,'const.bool.IN_2'(s(X1,sK16(X3,X2,X1,X0)),s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))))) )),
  inference(forward_demodulation,[],[f277,f243])).
fof(f277,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X1,sK16(X3,X2,X1,X0))))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))))) )),
  inference(cnf_transformation,[],[f168])).
fof(f24754,plain,(
  p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)))) | $spl2),
  inference(forward_demodulation,[],[f349,f1251])).
fof(f27150,plain,(
  $false | ($spl1 | $spl2 | $spl5)),
  inference(subsumption_resolution,[],[f26783,f25818])).
fof(f25818,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))))))) ) | ($spl1 | $spl5)),
  inference(unit_resulting_resolution,[],[f351,f25215,f357])).
fof(f26783,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK16(sK12,sK14,sK10,sK11)),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)))))) | $spl2),
  inference(unit_resulting_resolution,[],[f201,f2223,f24755,f978])).
fof(f978,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,X51)) | ~p(s(bool,'const.bool.IN_2'(s(X48,X50),s(fun(X48,bool),X49)))) | p(s(bool,X52)) | ~sP4(X51,happ(s(fun(X48,bool),X49),s(X48,X50)),X52)) )),
  inference(superposition,[],[f290,f243])).
fof(f290,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP4(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2] : ((~sP4(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP4(X0,X1,X2)))),
  inference(rectify,[],[f175])).
fof(f175,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(flattening,[],[f174])).
fof(f174,plain,(
  ! [X2,X0,X1] : ((~sP4(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP4(X2,X0,X1)))),
  inference(nnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X2,X0,X1] : (sP4(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP4])])).
fof(f2223,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0',happ(s(fun(X0,bool),X1),s(X0,X2)),'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f959,f285])).
fof(f285,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP4(X0,X2,X1)) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2] : ((~sP5(X0,X1,X2) | (sP2(X2,X1,X0) & sP3(X1,X2,X0) & sP4(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP2(X2,X1,X0) | ~sP3(X1,X2,X0) | ~sP4(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP5(X0,X1,X2)))),
  inference(rectify,[],[f172])).
fof(f172,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP5(X2,X1,X0)))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X2,X1,X0] : ((~sP5(X2,X1,X0) | (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP2(X0,X1,X2) | ~sP3(X1,X0,X2) | ~sP4(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP5(X2,X1,X0)))),
  inference(nnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X2,X1,X0] : (sP5(X2,X1,X0) <=> (sP2(X0,X1,X2) & sP3(X1,X0,X2) & sP4(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP5])])).
fof(f959,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)),happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f201,f243,f303])).
fof(f303,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f182,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP5(X2,X1,X0)) & (~sP5(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP5(X2,X1,X0))),
  inference(definition_folding,[],[f90,f126,f125,f124,f123])).
fof(f123,plain,(
  ! [X0,X1,X2] : (sP2(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f124,plain,(
  ! [X1,X0,X2] : (sP3(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f90,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,X20)) <=> s(bool,X18) = s(bool,X19)) <=> ((p(s(bool,X20)) | p(s(bool,X19)) | p(s(bool,X18))) & (p(s(bool,X20)) | ~p(s(bool,X18)) | ~p(s(bool,X19))) & (p(s(bool,X19)) | ~p(s(bool,X18)) | ~p(s(bool,X20))) & (p(s(bool,X18)) | ~p(s(bool,X19)) | ~p(s(bool,X20)))))),
  file('thm.relation.IN_RDOM_DELETE','thm.sat.sat_2|49|_prove_ltb')).
fof(f24752,plain,(
  $false | ($spl1 | $spl96)),
  inference(subsumption_resolution,[],[f348,f22653])).
fof(f24751,plain,(
  $false | ($spl0 | $spl4)),
  inference(subsumption_resolution,[],[f24750,f353])).
fof(f24750,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl4)),
  inference(forward_demodulation,[],[f24749,f22568])).
fof(f22568,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f2392,f353,f386,f4530])).
fof(f4530,plain,(
  ( ! [X26,X27,X25] : (p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(subsumption_resolution,[],[f3532,f201])).
fof(f3532,plain,(
  ( ! [X26,X27,X25] : (~p(s(bool,'T_0')) | p(s(bool,X26)) | sP1(X25,X26) | s(bool,X25) = s(bool,X27) | s(bool,'T_0') = s(bool,X27)) )),
  inference(superposition,[],[f267,f509])).
fof(f509,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | s(bool,'T_0') = s(bool,X1) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f208,f208])).
fof(f267,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f163,plain,(
  ! [X0,X1] : ((~sP1(X0,X1) | ((p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | ~p(s(bool,X0))))) & ((~p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1)))),
  inference(rectify,[],[f162])).
fof(f162,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0)))),
  inference(flattening,[],[f161])).
fof(f161,plain,(
  ! [X1,X0] : ((~sP1(X1,X0) | ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1))))) & (((~p(s(bool,X1)) & ~p(s(bool,X0))) | (p(s(bool,X0)) & p(s(bool,X1)))) | sP1(X1,X0)))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X1,X0] : (sP1(X1,X0) <=> ((p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,X1)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f2392,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP1('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f265])).
fof(f265,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | p(s(bool,X0)) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f163])).
fof(f1023,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4))))) )),
  inference(equality_resolution,[],[f281])).
fof(f281,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X1,X3) != s(X1,X4) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f24749,plain,(
  p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12),s(sK11,sK12),s(sK10,sK16(sK12,'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)),sK10,sK11))))) | ($spl0 | $spl4)),
  inference(backward_demodulation,[],[f352,f23213])).
fof(f23213,plain,(
  p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13),s(sK11,sK12),s(sK10,sK16(sK12,'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)),sK10,sK11))))) | $spl0),
  inference(forward_demodulation,[],[f22990,f1518])).
fof(f22990,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK16(sK12,'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)),sK10,sK11)),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))),s(sK11,sK12)))))) | $spl0),
  inference(unit_resulting_resolution,[],[f22760,f356])).
fof(f22760,plain,(
  p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))),s(sK11,sK12)))) | $spl0),
  inference(forward_demodulation,[],[f347,f1251])).
fof(f24622,plain,(
  $false | ($spl0 | $spl3)),
  inference(subsumption_resolution,[],[f24621,f353])).
fof(f24621,plain,(
  p(s(bool,'F_0')) | ($spl0 | $spl3)),
  inference(backward_demodulation,[],[f24619,f23213])).
fof(f24619,plain,(
  ( ! [X0,X1] : (s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,X0),s(sK11,sK12),s(sK10,X1)))) ) | $spl3),
  inference(forward_demodulation,[],[f24572,f22929])).
fof(f22929,plain,(
  s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))) | $spl3),
  inference(unit_resulting_resolution,[],[f386,f22759,f606])).
fof(f22759,plain,(
  ~p(s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12)))) | $spl3),
  inference(forward_demodulation,[],[f350,f1251])).
fof(f24572,plain,(
  ( ! [X0,X1] : (s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))) = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,X0),s(sK11,sK12),s(sK10,X1)))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f22762,f23471,f606])).
fof(f23471,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,X0),s(sK11,sK12),s(sK10,X1))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f22761,f358])).
fof(f358,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) | p(s(bool,'const.bool.IN_2'(s(X0,X5),s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))))))) )),
  inference(forward_demodulation,[],[f280,f243])).
fof(f280,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (p(s(bool,happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),X2),s(X1,X4))),s(X0,X5)))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f22761,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),happ(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))))))) ) | $spl3),
  inference(unit_resulting_resolution,[],[f22759,f355])).
fof(f22762,plain,(
  s(bool,'T_0') != s(bool,'const.relation.RDOM_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK12))) | $spl3),
  inference(unit_resulting_resolution,[],[f22759,f216])).
fof(f22655,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39)) ) | $spl98),
  inference(cnf_transformation,[],[f22655_D])).
fof(f22655_D,plain,(
  ( ! [X39,X38] : (sP1(X38,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))) | s(bool,'T_0') = s(bool,X39) | s(bool,X38) = s(bool,X39)) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f22653,plain,(
  ( ! [X30] : (p(s(bool,X30))) ) | $spl96),
  inference(cnf_transformation,[],[f22653_D])).
fof(f22653_D,plain,(
  ( ! [X30] : (p(s(bool,X30))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f22654,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | sP1(X38,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39)) ) | $spl1),
  inference(subsumption_resolution,[],[f22612,f216])).
fof(f22612,plain,(
  ( ! [X39,X37,X38] : (p(s(bool,X37)) | sP1(X38,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))) | s(bool,X38) = s(bool,X39) | s(bool,'T_0') = s(bool,X39) | s(bool,'T_0') = s(bool,X37)) ) | $spl1),
  inference(superposition,[],[f4530,f2084])).
fof(f2084,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(subsumption_resolution,[],[f2079,f353])).
fof(f2079,plain,(
  ( ! [X5] : (p(s(bool,'F_0')) | s(bool,X5) = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))) | s(bool,'T_0') = s(bool,X5)) ) | $spl1),
  inference(superposition,[],[f606,f548])).
fof(f548,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f348,f221])).
fof(f221,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X14] : (s(bool,'F_0') = s(bool,X14) <=> ~p(s(bool,X14)))),
  file('thm.relation.IN_RDOM_DELETE','thm.bool.EQ_CLAUSES|split|2')).
fof(f22652,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) | $spl94),
  inference(cnf_transformation,[],[f22652_D])).
fof(f22652_D,plain,(
  ( ! [X31,X32] : (sP1(X31,'F_0') | s(bool,'T_0') = s(bool,X32) | s(bool,X31) = s(bool,X32)) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f22651,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32)) )),
  inference(subsumption_resolution,[],[f22609,f216])).
fof(f22609,plain,(
  ( ! [X30,X31,X32] : (p(s(bool,X30)) | sP1(X31,'F_0') | s(bool,X31) = s(bool,X32) | s(bool,'T_0') = s(bool,X32) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f4530,f208])).
fof(f8313,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'T_0'),s(bool,'F_0'),s(X65,X67))))) ) | $spl44),
  inference(cnf_transformation,[],[f8313_D])).
fof(f8313_D,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'T_0'),s(bool,'F_0'),s(X65,X67))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f21012,plain,(
  ( ! [X277,X279,X275,X283,X276,X278,X274] : (s(bool,X283) != s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277),s(X275,X278),s(X274,X279))) | s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277),s(X275,X278),s(X274,X279))) | s(bool,'T_0') = s(bool,X283)) ) | $spl88),
  inference(cnf_transformation,[],[f21012_D])).
fof(f21012_D,plain,(
  ( ! [X277,X279,X275,X283,X276,X278,X274] : (s(bool,X283) != s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277),s(X275,X278),s(X274,X279))) | s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277),s(X275,X278),s(X274,X279))) | s(bool,'T_0') = s(bool,X283)) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f21907,plain,(
  ( ! [X294,X300,X302,X296,X298,X295,X301,X303,X297,X299] : (s(bool,X303) != s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X295,fun(X294,bool)),X296),s(X295,X297),s(X295,X298),s(X294,X299))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X300,bool)),X301),s(bool,'T_0'),s(bool,'F_0'),s(X300,X302)))) | s(bool,'T_0') = s(bool,X303) | s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X295,fun(X294,bool)),X296),s(X295,X297),s(X295,X298),s(X294,X299)))) )),
  inference(forward_demodulation,[],[f12534,f279])).
fof(f12534,plain,(
  ( ! [X294,X300,X302,X296,X298,X295,X301,X303,X297,X299] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X300,bool)),X301),s(bool,'T_0'),s(bool,'F_0'),s(X300,X302)))) | s(bool,X303) != s(bool,happ(s(fun(X294,bool),happ(s(fun(X295,fun(X294,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X295,fun(X294,bool)),X296),s(X295,X297))),s(X295,X298))),s(X294,X299))) | s(bool,'T_0') = s(bool,X303) | s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X295,fun(X294,bool)),X296),s(X295,X297),s(X295,X298),s(X294,X299)))) )),
  inference(superposition,[],[f1020,f1519])).
fof(f1519,plain,(
  ( ! [X14,X12,X17,X15,X13,X16] : (s(bool,'T_0') = s(bool,happ(s(fun(X12,bool),happ(s(fun(X13,fun(X12,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X13,fun(X12,bool)),X14),s(X13,X15))),s(X13,X16))),s(X12,X17))) | s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X13,fun(X12,bool)),X14),s(X13,X15),s(X13,X16),s(X12,X17)))) )),
  inference(superposition,[],[f279,f208])).
fof(f1020,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X7,bool)),X8),s(bool,X6),s(bool,'F_0'),s(X7,X9)))) | s(bool,X5) != s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f281,f208])).
fof(f4491,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) | $spl8),
  inference(cnf_transformation,[],[f4491_D])).
fof(f4491_D,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0)) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f21899,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,X16),s(bool,'F_0'),s(bool,X16))))) ) | $spl92),
  inference(cnf_transformation,[],[f21899_D])).
fof(f21899_D,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,X16),s(bool,'F_0'),s(bool,X16))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f21900,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X20),s(bool,X19),s(bool,'F_0'),s(bool,X19)))) | s(bool,'T_0') = s(bool,X21)) ) | $spl44),
  inference(subsumption_resolution,[],[f21639,f1020])).
fof(f21639,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X20),s(bool,X19),s(bool,'F_0'),s(bool,X19)))) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl44),
  inference(superposition,[],[f8313,f509])).
fof(f21898,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,X16),s(bool,'F_0'),s(bool,X16)))) | s(bool,'T_0') = s(bool,X18)) ) | $spl44),
  inference(subsumption_resolution,[],[f21638,f1020])).
fof(f21638,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,X16),s(bool,'F_0'),s(bool,X16)))) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl44),
  inference(superposition,[],[f8313,f509])).
fof(f21888,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,X0),s(bool,'F_0'),s(X1,X3))))) ) | $spl90),
  inference(cnf_transformation,[],[f21888_D])).
fof(f21888_D,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,X0),s(bool,'F_0'),s(X1,X3))))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f21889,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X6,bool)),X7),s(bool,X5),s(bool,'F_0'),s(X6,X8)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl44),
  inference(subsumption_resolution,[],[f21621,f1020])).
fof(f21621,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X6,bool)),X7),s(bool,X5),s(bool,'F_0'),s(X6,X8)))) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl44),
  inference(superposition,[],[f8313,f509])).
fof(f21887,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,X0),s(bool,'F_0'),s(X1,X3)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl44),
  inference(subsumption_resolution,[],[f21620,f1020])).
fof(f21620,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,X0),s(bool,'F_0'),s(X1,X3)))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl44),
  inference(superposition,[],[f8313,f509])).
fof(f8155,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'F_0'),s(bool,'T_0'),s(X65,X67))))) ) | $spl40),
  inference(cnf_transformation,[],[f8155_D])).
fof(f8155_D,plain,(
  ( ! [X66,X67,X65] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'F_0'),s(bool,'T_0'),s(X65,X67))))) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f8772,plain,(
  ( ! [X218,X216,X219,X217,X223] : (s(bool,X223) != s(bool,'const.relation.RDOM_2'(s(fun(X216,fun(X217,bool)),X218),s(X216,X219))) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X216,fun(X217,bool)),X218),s(X216,X219))) | s(bool,'T_0') = s(bool,X223)) ) | $spl58),
  inference(cnf_transformation,[],[f8772_D])).
fof(f8772_D,plain,(
  ( ! [X218,X216,X219,X217,X223] : (s(bool,X223) != s(bool,'const.relation.RDOM_2'(s(fun(X216,fun(X217,bool)),X218),s(X216,X219))) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X216,fun(X217,bool)),X218),s(X216,X219))) | s(bool,'T_0') = s(bool,X223)) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f21014,plain,(
  ( ! [X206,X204,X202,X200,X207,X205,X203,X201] : (s(bool,X207) != s(bool,'const.relation.RDOM_2'(s(fun(X200,fun(X201,bool)),X202),s(X200,X203))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X204,bool)),X205),s(bool,'F_0'),s(bool,'T_0'),s(X204,X206)))) | s(bool,'T_0') = s(bool,X207) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X200,fun(X201,bool)),X202),s(X200,X203)))) )),
  inference(forward_demodulation,[],[f8616,f275])).
fof(f8616,plain,(
  ( ! [X206,X204,X202,X200,X207,X205,X203,X201] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X204,bool)),X205),s(bool,'F_0'),s(bool,'T_0'),s(X204,X206)))) | s(bool,X207) != s(bool,happ(s(fun(X200,bool),'const.relation.RDOM_1'(s(fun(X200,fun(X201,bool)),X202))),s(X200,X203))) | s(bool,'T_0') = s(bool,X207) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X200,fun(X201,bool)),X202),s(X200,X203)))) )),
  inference(superposition,[],[f1016,f1252])).
fof(f1252,plain,(
  ( ! [X10,X8,X11,X9] : (s(bool,'T_0') = s(bool,happ(s(fun(X8,bool),'const.relation.RDOM_1'(s(fun(X8,fun(X9,bool)),X10))),s(X8,X11))) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X8,fun(X9,bool)),X10),s(X8,X11)))) )),
  inference(superposition,[],[f275,f208])).
fof(f1016,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X7,bool)),X8),s(bool,'F_0'),s(bool,X6),s(X7,X9)))) | s(bool,X5) != s(bool,X6) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f281,f208])).
fof(f21011,plain,(
  ( ! [X280,X282,X277,X279,X275,X281,X283,X276,X278,X274] : (s(bool,X283) != s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277),s(X275,X278),s(X274,X279))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X280,bool)),X281),s(bool,'F_0'),s(bool,'T_0'),s(X280,X282)))) | s(bool,'T_0') = s(bool,X283) | s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277),s(X275,X278),s(X274,X279)))) )),
  inference(forward_demodulation,[],[f12532,f279])).
fof(f12532,plain,(
  ( ! [X280,X282,X277,X279,X275,X281,X283,X276,X278,X274] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X280,bool)),X281),s(bool,'F_0'),s(bool,'T_0'),s(X280,X282)))) | s(bool,X283) != s(bool,happ(s(fun(X274,bool),happ(s(fun(X275,fun(X274,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277))),s(X275,X278))),s(X274,X279))) | s(bool,'T_0') = s(bool,X283) | s(bool,'F_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X275,fun(X274,bool)),X276),s(X275,X277),s(X275,X278),s(X274,X279)))) )),
  inference(superposition,[],[f1016,f1519])).
fof(f21003,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,'F_0'),s(bool,X16),s(bool,X16))))) ) | $spl86),
  inference(cnf_transformation,[],[f21003_D])).
fof(f21003_D,plain,(
  ( ! [X17,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,'F_0'),s(bool,X16),s(bool,X16))))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f21004,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X20),s(bool,'F_0'),s(bool,X19),s(bool,X19)))) | s(bool,'T_0') = s(bool,X21)) ) | $spl40),
  inference(subsumption_resolution,[],[f20748,f1016])).
fof(f20748,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X20),s(bool,'F_0'),s(bool,X19),s(bool,X19)))) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) ) | $spl40),
  inference(superposition,[],[f8155,f509])).
fof(f21002,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,'F_0'),s(bool,X16),s(bool,X16)))) | s(bool,'T_0') = s(bool,X18)) ) | $spl40),
  inference(subsumption_resolution,[],[f20747,f1016])).
fof(f20747,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(bool,bool)),X17),s(bool,'F_0'),s(bool,X16),s(bool,X16)))) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) ) | $spl40),
  inference(superposition,[],[f8155,f509])).
fof(f21000,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,'F_0'),s(bool,X0),s(X1,X3))))) ) | $spl84),
  inference(cnf_transformation,[],[f21000_D])).
fof(f21000_D,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,'F_0'),s(bool,X0),s(X1,X3))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f21001,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X6,bool)),X7),s(bool,'F_0'),s(bool,X5),s(X6,X8)))) | s(bool,'T_0') = s(bool,X9)) ) | $spl40),
  inference(subsumption_resolution,[],[f20736,f1016])).
fof(f20736,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X6,bool)),X7),s(bool,'F_0'),s(bool,X5),s(X6,X8)))) | s(bool,X5) = s(bool,X9) | s(bool,'T_0') = s(bool,X9)) ) | $spl40),
  inference(superposition,[],[f8155,f509])).
fof(f20999,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,'F_0'),s(bool,X0),s(X1,X3)))) | s(bool,'T_0') = s(bool,X4)) ) | $spl40),
  inference(subsumption_resolution,[],[f20735,f1016])).
fof(f20735,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,'F_0'),s(bool,X0),s(X1,X3)))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) ) | $spl40),
  inference(superposition,[],[f8155,f509])).
fof(f13107,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6))) | p(s(bool,X7))) ) | $spl82),
  inference(cnf_transformation,[],[f13107_D])).
fof(f13107_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6))) | p(s(bool,X7))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f13109,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)))) )),
  inference(condensation,[],[f13108])).
fof(f13108,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP7(X60,'T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59))) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f13089,f606])).
fof(f13089,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP7(X60,'T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59))) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1537,f509])).
fof(f1537,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119,X117] : (~p(s(bool,X123)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X118,fun(X117,bool)),X119),s(X118,X120),s(X118,X121),s(X117,X122)))) | p(s(bool,X124)) | ~sP7(X124,X123,happ(s(fun(X117,bool),happ(s(fun(X118,fun(X117,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X118,fun(X117,bool)),X119),s(X118,X120))),s(X118,X121))),s(X117,X122)))) )),
  inference(superposition,[],[f322,f279])).
fof(f322,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP7(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2))))))),
  inference(flattening,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) | ~p(s(bool,X0))))) | sP7(X2,X0,X1)) & (~sP7(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) & p(s(bool,X0)))) & ((~p(s(bool,X1)) | ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> sP7(X2,X0,X1))),
  inference(definition_folding,[],[f92,f129,f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : (sP6(X0,X1,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP6])])).
fof(f129,plain,(
  ! [X2,X0,X1] : (sP7(X2,X0,X1) <=> (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP7])])).
fof(f92,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f91])).
fof(f91,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) & p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X19)) & p(s(bool,X18)))) <=> ((p(s(bool,X20)) | ~p(s(bool,X19)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | ~p(s(bool,X20))) & (p(s(bool,X18)) | ~p(s(bool,X20)))))),
  file('thm.relation.IN_RDOM_DELETE','thm.sat.sat_23_prove_ltb')).
fof(f13106,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP7(X7,'T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)))) )),
  inference(condensation,[],[f13105])).
fof(f13105,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP7(X51,'T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50))) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f13088,f606])).
fof(f13088,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP7(X51,'T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50))) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1537,f509])).
fof(f13049,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) | $spl80),
  inference(cnf_transformation,[],[f13049_D])).
fof(f13049_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f13051,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f13050])).
fof(f13050,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP6('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f13031,f606])).
fof(f13031,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP6('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1536,f509])).
fof(f1536,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (~p(s(bool,X115)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X110,fun(X109,bool)),X111),s(X110,X112),s(X110,X113),s(X109,X114)))) | p(s(bool,X116)) | ~sP6(X115,happ(s(fun(X109,bool),happ(s(fun(X110,fun(X109,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X110,fun(X109,bool)),X111),s(X110,X112))),s(X110,X113))),s(X109,X114)),X116)) )),
  inference(superposition,[],[f313,f279])).
fof(f313,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP6(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(flattening,[],[f186])).
fof(f186,plain,(
  ! [X0,X1,X2] : ((~sP6(X0,X1,X2) | (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0)))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP6(X0,X1,X2)))),
  inference(nnf_transformation,[],[f128])).
fof(f13048,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP6('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f13047])).
fof(f13047,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP6('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f13030,f606])).
fof(f13030,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP6('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1536,f509])).
fof(f12989,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) | $spl78),
  inference(cnf_transformation,[],[f12989_D])).
fof(f12989_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f12991,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f12990])).
fof(f12990,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP3('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f12971,f606])).
fof(f12971,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP3('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1535,f509])).
fof(f1535,plain,(
  ( ! [X103,X101,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X102,fun(X101,bool)),X103),s(X102,X104),s(X102,X105),s(X101,X106)))) | p(s(bool,X108)) | ~sP3(X107,happ(s(fun(X101,bool),happ(s(fun(X102,fun(X101,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X102,fun(X101,bool)),X103),s(X102,X104))),s(X102,X105))),s(X101,X106)),X108)) )),
  inference(superposition,[],[f294,f279])).
fof(f294,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | p(s(bool,X2)) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f179])).
fof(f179,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f178])).
fof(f178,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(flattening,[],[f177])).
fof(f177,plain,(
  ! [X1,X0,X2] : ((~sP3(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP3(X1,X0,X2)))),
  inference(nnf_transformation,[],[f124])).
fof(f12988,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP3('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f12987])).
fof(f12987,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP3('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f12970,f606])).
fof(f12970,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP3('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1535,f509])).
fof(f12942,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) | $spl76),
  inference(cnf_transformation,[],[f12942_D])).
fof(f12942_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7) | p(s(bool,X7))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f12944,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f12943])).
fof(f12943,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP4('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f12924,f606])).
fof(f12924,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | p(s(bool,X60)) | ~sP4('T_0',happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1534,f509])).
fof(f1534,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (~p(s(bool,X99)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X94,fun(X93,bool)),X95),s(X94,X96),s(X94,X97),s(X93,X98)))) | p(s(bool,X100)) | ~sP4(X99,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X94,fun(X93,bool)),X95),s(X94,X96))),s(X94,X97))),s(X93,X98)),X100)) )),
  inference(superposition,[],[f290,f279])).
fof(f12941,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | p(s(bool,X7)) | ~sP4('T_0',happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7)) )),
  inference(condensation,[],[f12940])).
fof(f12940,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP4('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f12923,f606])).
fof(f12923,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | p(s(bool,X51)) | ~sP4('T_0',happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1534,f509])).
fof(f12404,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) ) | $spl74),
  inference(cnf_transformation,[],[f12404_D])).
fof(f12404_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f12406,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) )),
  inference(condensation,[],[f12405])).
fof(f12405,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),'T_0',X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f12387,f606])).
fof(f12387,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),'T_0',X60) | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1540,f509])).
fof(f1540,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (~p(s(bool,X147)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X142,fun(X141,bool)),X143),s(X142,X144),s(X142,X145),s(X141,X146)))) | sP9(happ(s(fun(X141,bool),happ(s(fun(X142,fun(X141,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X142,fun(X141,bool)),X143),s(X142,X144))),s(X142,X145))),s(X141,X146)),X147,X148)) )),
  inference(superposition,[],[f336,f279])).
fof(f336,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X0)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f197,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(flattening,[],[f196])).
fof(f196,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & (~p(s(bool,X2)) | (~p(s(bool,X1)) & ~p(s(bool,X0))))) | sP9(X2,X0,X1)) & (~sP9(X2,X0,X1) | ((~p(s(bool,X2)) | (p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0))) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> sP9(X2,X0,X1))),
  inference(definition_folding,[],[f94,f132,f131])).
fof(f131,plain,(
  ! [X2,X0,X1] : (sP8(X2,X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP8])])).
fof(f132,plain,(
  ! [X2,X0,X1] : (sP9(X2,X0,X1) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1)))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP9])])).
fof(f94,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f93])).
fof(f93,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> (p(s(bool,X1)) | p(s(bool,X0)))) <=> ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X18,X19,X20] : ((p(s(bool,X20)) <=> (p(s(bool,X19)) | p(s(bool,X18)))) <=> ((p(s(bool,X20)) | ~p(s(bool,X19))) & (p(s(bool,X20)) | ~p(s(bool,X18))) & (p(s(bool,X19)) | p(s(bool,X18)) | ~p(s(bool,X20)))))),
  file('thm.relation.IN_RDOM_DELETE','thm.sat.sat_25_prove_ltb')).
fof(f12403,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),'T_0',X7)) )),
  inference(condensation,[],[f12402])).
fof(f12402,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),'T_0',X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f12386,f606])).
fof(f12386,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),'T_0',X51) | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1540,f509])).
fof(f12359,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) ) | $spl72),
  inference(cnf_transformation,[],[f12359_D])).
fof(f12359_D,plain,(
  ( ! [X6,X4,X2,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f12361,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) )),
  inference(condensation,[],[f12360])).
fof(f12360,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60,'T_0') | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(subsumption_resolution,[],[f12342,f606])).
fof(f12342,plain,(
  ( ! [X61,X59,X57,X54,X60,X58,X56,X55,X53] : (~p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57),s(X54,X58),s(X55,X59)))) | sP9(happ(s(fun(X55,bool),happ(s(fun(X54,fun(X55,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))),s(X54,X58))),s(X55,X59)),X60,'T_0') | s(bool,X53) = s(bool,X61) | s(bool,'T_0') = s(bool,X61)) )),
  inference(superposition,[],[f1539,f509])).
fof(f1539,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (~p(s(bool,X139)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X134,fun(X133,bool)),X135),s(X134,X136),s(X134,X137),s(X133,X138)))) | sP9(happ(s(fun(X133,bool),happ(s(fun(X134,fun(X133,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X134,fun(X133,bool)),X135),s(X134,X136))),s(X134,X137))),s(X133,X138)),X140,X139)) )),
  inference(superposition,[],[f335,f279])).
fof(f335,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | ~p(s(bool,X1)) | sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f12358,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X5),s(X2,X6)))) | sP9(happ(s(fun(X2,bool),happ(s(fun(X1,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4))),s(X1,X5))),s(X2,X6)),X7,'T_0')) )),
  inference(condensation,[],[f12357])).
fof(f12357,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51,'T_0') | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(subsumption_resolution,[],[f12341,f606])).
fof(f12341,plain,(
  ( ! [X47,X45,X52,X50,X48,X46,X44,X51,X49] : (~p(s(bool,X44)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48),s(X45,X49),s(X46,X50)))) | sP9(happ(s(fun(X46,bool),happ(s(fun(X45,fun(X46,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X45,fun(X46,bool)),X47),s(X45,X48))),s(X45,X49))),s(X46,X50)),X51,'T_0') | s(bool,X44) = s(bool,X52) | s(bool,'T_0') = s(bool,X52)) )),
  inference(superposition,[],[f1539,f509])).
fof(f9059,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4))) | p(s(bool,X5))) ) | $spl70),
  inference(cnf_transformation,[],[f9059_D])).
fof(f9059_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4))) | p(s(bool,X5))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f9061,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)))) )),
  inference(condensation,[],[f9060])).
fof(f9060,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP7(X46,'T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f9044,f606])).
fof(f9044,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP7(X46,'T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45))) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1270,f509])).
fof(f1270,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,X87)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X83,fun(X84,bool)),X85),s(X83,X86)))) | p(s(bool,X88)) | ~sP7(X88,X87,happ(s(fun(X83,bool),'const.relation.RDOM_1'(s(fun(X83,fun(X84,bool)),X85))),s(X83,X86)))) )),
  inference(superposition,[],[f322,f275])).
fof(f9058,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP7(X5,'T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)))) )),
  inference(condensation,[],[f9057])).
fof(f9057,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f9043,f606])).
fof(f9043,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38))) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1270,f509])).
fof(f9006,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl68),
  inference(cnf_transformation,[],[f9006_D])).
fof(f9006_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f9008,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f9007])).
fof(f9007,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP6('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f8991,f606])).
fof(f8991,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP6('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1269,f509])).
fof(f1269,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (~p(s(bool,X81)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X77,fun(X78,bool)),X79),s(X77,X80)))) | p(s(bool,X82)) | ~sP6(X81,happ(s(fun(X77,bool),'const.relation.RDOM_1'(s(fun(X77,fun(X78,bool)),X79))),s(X77,X80)),X82)) )),
  inference(superposition,[],[f313,f275])).
fof(f9005,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP6('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f9004])).
fof(f9004,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f8990,f606])).
fof(f8990,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1269,f509])).
fof(f8954,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl66),
  inference(cnf_transformation,[],[f8954_D])).
fof(f8954_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f8956,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f8955])).
fof(f8955,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP3('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f8939,f606])).
fof(f8939,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP3('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1268,f509])).
fof(f1268,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,X75)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X71,fun(X72,bool)),X73),s(X71,X74)))) | p(s(bool,X76)) | ~sP3(X75,happ(s(fun(X71,bool),'const.relation.RDOM_1'(s(fun(X71,fun(X72,bool)),X73))),s(X71,X74)),X76)) )),
  inference(superposition,[],[f294,f275])).
fof(f8953,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP3('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f8952])).
fof(f8952,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f8938,f606])).
fof(f8938,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1268,f509])).
fof(f8902,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) | $spl64),
  inference(cnf_transformation,[],[f8902_D])).
fof(f8902_D,plain,(
  ( ! [X4,X2,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5) | p(s(bool,X5))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f8904,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f8903])).
fof(f8903,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(subsumption_resolution,[],[f8887,f606])).
fof(f8887,plain,(
  ( ! [X47,X45,X43,X41,X46,X44,X42] : (~p(s(bool,X41)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X42,fun(X43,bool)),X44),s(X42,X45)))) | p(s(bool,X46)) | ~sP4('T_0',happ(s(fun(X42,bool),'const.relation.RDOM_1'(s(fun(X42,fun(X43,bool)),X44))),s(X42,X45)),X46) | s(bool,X41) = s(bool,X47) | s(bool,'T_0') = s(bool,X47)) )),
  inference(superposition,[],[f1267,f509])).
fof(f1267,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (~p(s(bool,X69)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X65,fun(X66,bool)),X67),s(X65,X68)))) | p(s(bool,X70)) | ~sP4(X69,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70)) )),
  inference(superposition,[],[f290,f275])).
fof(f8901,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4)))) | p(s(bool,X5)) | ~sP4('T_0',happ(s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X2,bool)),X3))),s(X1,X4)),X5)) )),
  inference(condensation,[],[f8900])).
fof(f8900,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f8886,f606])).
fof(f8886,plain,(
  ( ! [X39,X37,X35,X38,X36,X34,X40] : (~p(s(bool,X34)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X35,fun(X36,bool)),X37),s(X35,X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X35,bool),'const.relation.RDOM_1'(s(fun(X35,fun(X36,bool)),X37))),s(X35,X38)),X39) | s(bool,X34) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f1267,f509])).
fof(f8823,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool))))) ) | $spl60),
  inference(cnf_transformation,[],[f8823_D])).
fof(f8823_D,plain,(
  ( ! [X2,X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,X2) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f8825,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | $spl63),
  inference(cnf_transformation,[],[f8825_D])).
fof(f8825_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f8813,plain,(
  ( ! [X4,X5,X3] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X3,X4,bool,bool)))) | s(bool,X5) = s(bool,sK15(X3,X4,bool,bool)) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1317,f509])).
fof(f1317,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(X0,sK15(X1,X2,bool,X0)))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'T_0') = s(bool,g(s(X0,sK15(X1,X2,bool,X0))))) )),
  inference(superposition,[],[f276,f208])).
fof(f276,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK15(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK15(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK15])],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.relation.IN_RDOM_DELETE',aHL_EXT)).
fof(f8812,plain,(
  ( ! [X2,X0,X1] : (s(bool,'F_0') != s(bool,f(s(bool,'T_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,g(s(bool,sK15(X0,X1,bool,bool)))) | s(bool,X2) = s(bool,sK15(X0,X1,bool,bool)) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1317,f509])).
fof(f8771,plain,(
  ( ! [X218,X216,X222,X220,X219,X217,X223,X221] : (s(bool,X223) != s(bool,'const.relation.RDOM_2'(s(fun(X216,fun(X217,bool)),X218),s(X216,X219))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X220,bool)),X221),s(bool,'T_0'),s(bool,'F_0'),s(X220,X222)))) | s(bool,'T_0') = s(bool,X223) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X216,fun(X217,bool)),X218),s(X216,X219)))) )),
  inference(forward_demodulation,[],[f8618,f275])).
fof(f8618,plain,(
  ( ! [X218,X216,X222,X220,X219,X217,X223,X221] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X220,bool)),X221),s(bool,'T_0'),s(bool,'F_0'),s(X220,X222)))) | s(bool,X223) != s(bool,happ(s(fun(X216,bool),'const.relation.RDOM_1'(s(fun(X216,fun(X217,bool)),X218))),s(X216,X219))) | s(bool,'T_0') = s(bool,X223) | s(bool,'F_0') = s(bool,'const.relation.RDOM_2'(s(fun(X216,fun(X217,bool)),X218),s(X216,X219)))) )),
  inference(superposition,[],[f1020,f1252])).
fof(f8340,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | $spl52),
  inference(cnf_transformation,[],[f8340_D])).
fof(f8340_D,plain,(
  s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f8336,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool))) ) | $spl48),
  inference(cnf_transformation,[],[f8336_D])).
fof(f8336_D,plain,(
  ( ! [X0,X1] : (s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f8343,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | $spl57),
  inference(cnf_transformation,[],[f8343_D])).
fof(f8343_D,plain,(
  s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl57),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl57])])).
fof(f8334,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X12) = s(fun(bool,bool),X13) | s(bool,'T_0') = s(bool,sK15(X12,X13,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1316,f208])).
fof(f1316,plain,(
  ( ! [X2,X0,X1] : (s(X2,f(s(bool,'F_0'))) != s(X2,g(s(bool,'F_0'))) | s(fun(bool,X2),X0) = s(fun(bool,X2),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,X2,bool))) )),
  inference(superposition,[],[f276,f208])).
fof(f8341,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) | $spl54),
  inference(cnf_transformation,[],[f8341_D])).
fof(f8341_D,plain,(
  ( ! [X6] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f8333,plain,(
  ( ! [X10,X11,X9] : (s(bool,X9) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X10) = s(fun(bool,bool),X11) | s(bool,'T_0') = s(bool,sK15(X10,X11,bool,bool)) | s(bool,'T_0') = s(bool,X9) | s(bool,'T_0') = s(bool,g(s(bool,'F_0')))) )),
  inference(superposition,[],[f1316,f509])).
fof(f8332,plain,(
  ( ! [X6,X8,X7] : (s(bool,X6) != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X7) = s(fun(bool,bool),X8) | s(bool,'T_0') = s(bool,sK15(X7,X8,bool,bool)) | s(bool,'T_0') = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X6)) )),
  inference(superposition,[],[f1316,f509])).
fof(f8335,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) | $spl46),
  inference(cnf_transformation,[],[f8335_D])).
fof(f8335_D,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f8338,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | $spl51),
  inference(cnf_transformation,[],[f8338_D])).
fof(f8338_D,plain,(
  s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f8331,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X3) = s(fun(bool,bool),X4) | s(bool,'T_0') = s(bool,sK15(X3,X4,bool,bool)) | s(bool,X5) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f1316,f509])).
fof(f8330,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') != s(bool,f(s(bool,'F_0'))) | s(fun(bool,bool),X0) = s(fun(bool,bool),X1) | s(bool,'T_0') = s(bool,sK15(X0,X1,bool,bool)) | s(bool,X2) = s(bool,g(s(bool,'F_0'))) | s(bool,'T_0') = s(bool,X2)) )),
  inference(superposition,[],[f1316,f509])).
fof(f8156,plain,(
  ( ! [X68,X64,X62,X63] : (s(bool,X68) != s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63))) | s(bool,'T_0') = s(bool,X68)) ) | $spl42),
  inference(cnf_transformation,[],[f8156_D])).
fof(f8156_D,plain,(
  ( ! [X68,X64,X62,X63] : (s(bool,X68) != s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63))) | s(bool,'T_0') = s(bool,X68)) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f8312,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(bool,X68) != s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'T_0'),s(bool,'F_0'),s(X65,X67)))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63)))) )),
  inference(forward_demodulation,[],[f8246,f243])).
fof(f8246,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'T_0'),s(bool,'F_0'),s(X65,X67)))) | s(bool,X68) != s(bool,happ(s(fun(X62,bool),X63),s(X62,X64))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63)))) )),
  inference(superposition,[],[f1020,f964])).
fof(f964,plain,(
  ( ! [X4,X5,X3] : (s(bool,'T_0') = s(bool,happ(s(fun(X3,bool),X4),s(X3,X5))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X3,X5),s(fun(X3,bool),X4)))) )),
  inference(superposition,[],[f243,f208])).
fof(f8154,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (s(bool,X68) != s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'F_0'),s(bool,'T_0'),s(X65,X67)))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63)))) )),
  inference(forward_demodulation,[],[f8090,f243])).
fof(f8090,plain,(
  ( ! [X68,X66,X64,X62,X67,X65,X63] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X65,bool)),X66),s(bool,'F_0'),s(bool,'T_0'),s(X65,X67)))) | s(bool,X68) != s(bool,happ(s(fun(X62,bool),X63),s(X62,X64))) | s(bool,'T_0') = s(bool,X68) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(X62,X64),s(fun(X62,bool),X63)))) )),
  inference(superposition,[],[f1016,f964])).
fof(f7469,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2))) | p(s(bool,X4))) ) | $spl38),
  inference(cnf_transformation,[],[f7469_D])).
fof(f7469_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2))) | p(s(bool,X4))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7471,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2)))) )),
  inference(condensation,[],[f7470])).
fof(f7470,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X36,bool),X38),s(X36,X37))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7457,f606])).
fof(f7457,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP7(X39,'T_0',happ(s(fun(X36,bool),X38),s(X36,X37))) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f981,f509])).
fof(f981,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | ~p(s(bool,'const.bool.IN_2'(s(X63,X65),s(fun(X63,bool),X64)))) | p(s(bool,X67)) | ~sP7(X67,X66,happ(s(fun(X63,bool),X64),s(X63,X65)))) )),
  inference(superposition,[],[f322,f243])).
fof(f7468,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP7(X4,'T_0',happ(s(fun(X1,bool),X3),s(X1,X2)))) )),
  inference(condensation,[],[f7467])).
fof(f7467,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP7(X33,'T_0',happ(s(fun(X30,bool),X32),s(X30,X31))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7456,f606])).
fof(f7456,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP7(X33,'T_0',happ(s(fun(X30,bool),X32),s(X30,X31))) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f981,f509])).
fof(f7436,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) | $spl36),
  inference(cnf_transformation,[],[f7436_D])).
fof(f7436_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f7438,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f7437])).
fof(f7437,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7424,f606])).
fof(f7424,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP6('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f980,f509])).
fof(f980,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | ~p(s(bool,'const.bool.IN_2'(s(X58,X60),s(fun(X58,bool),X59)))) | p(s(bool,X62)) | ~sP6(X61,happ(s(fun(X58,bool),X59),s(X58,X60)),X62)) )),
  inference(superposition,[],[f313,f243])).
fof(f7435,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP6('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f7434])).
fof(f7434,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP6('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7423,f606])).
fof(f7423,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP6('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f980,f509])).
fof(f7405,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) | $spl34),
  inference(cnf_transformation,[],[f7405_D])).
fof(f7405_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f7407,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f7406])).
fof(f7406,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7393,f606])).
fof(f7393,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP3('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f979,f509])).
fof(f979,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | ~p(s(bool,'const.bool.IN_2'(s(X53,X55),s(fun(X53,bool),X54)))) | p(s(bool,X57)) | ~sP3(X56,happ(s(fun(X53,bool),X54),s(X53,X55)),X57)) )),
  inference(superposition,[],[f294,f243])).
fof(f7404,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP3('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f7403])).
fof(f7403,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP3('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7392,f606])).
fof(f7392,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP3('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f979,f509])).
fof(f7344,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) | $spl32),
  inference(cnf_transformation,[],[f7344_D])).
fof(f7344_D,plain,(
  ( ! [X4,X2,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4) | p(s(bool,X4))) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f7346,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f7345])).
fof(f7345,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(subsumption_resolution,[],[f7332,f606])).
fof(f7332,plain,(
  ( ! [X39,X37,X35,X38,X36,X40] : (~p(s(bool,X35)) | ~p(s(bool,'const.bool.IN_2'(s(X36,X37),s(fun(X36,bool),X38)))) | p(s(bool,X39)) | ~sP4('T_0',happ(s(fun(X36,bool),X38),s(X36,X37)),X39) | s(bool,X35) = s(bool,X40) | s(bool,'T_0') = s(bool,X40)) )),
  inference(superposition,[],[f978,f509])).
fof(f7343,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),X3)))) | p(s(bool,X4)) | ~sP4('T_0',happ(s(fun(X1,bool),X3),s(X1,X2)),X4)) )),
  inference(condensation,[],[f7342])).
fof(f7342,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP4('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f7331,f606])).
fof(f7331,plain,(
  ( ! [X30,X33,X31,X29,X34,X32] : (~p(s(bool,X29)) | ~p(s(bool,'const.bool.IN_2'(s(X30,X31),s(fun(X30,bool),X32)))) | p(s(bool,X33)) | ~sP4('T_0',happ(s(fun(X30,bool),X32),s(X30,X31)),X33) | s(bool,X29) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f978,f509])).
fof(f4836,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4835,f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821,f4823,f4825,f4827,f4829,f4831])).
fof(f4831,plain,(
  ( ! [X158,X163,X161,X159,X157,X164,X162,X160] : (~sP8(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X158,fun(X157,bool)),X159),s(X158,X160))),s(X158,X161))),s(X157,X162)),X163,X164) | sP9(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X158,fun(X157,bool)),X159),s(X158,X160))),s(X158,X161))),s(X157,X162)),X163,X164)) ) | $spl8),
  inference(subsumption_resolution,[],[f4698,f201])).
fof(f4698,plain,(
  ( ! [X158,X163,X161,X159,X157,X164,X162,X160] : (~p(s(bool,'T_0')) | ~sP8(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X158,fun(X157,bool)),X159),s(X158,X160))),s(X158,X161))),s(X157,X162)),X163,X164) | sP9(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X158,fun(X157,bool)),X159),s(X158,X160))),s(X158,X161))),s(X157,X162)),X163,X164)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1542])).
fof(f4829,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (p(s(bool,X155)) | p(s(bool,X156)) | ~sP9(happ(s(fun(X149,bool),happ(s(fun(X150,fun(X149,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X150,fun(X149,bool)),X151),s(X150,X152))),s(X150,X153))),s(X149,X154)),X156,X155)) ) | $spl8),
  inference(subsumption_resolution,[],[f4697,f201])).
fof(f4697,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (~p(s(bool,'T_0')) | p(s(bool,X155)) | p(s(bool,X156)) | ~sP9(happ(s(fun(X149,bool),happ(s(fun(X150,fun(X149,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X150,fun(X149,bool)),X151),s(X150,X152))),s(X150,X153))),s(X149,X154)),X156,X155)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1541])).
fof(f4827,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (~p(s(bool,X147)) | sP9(happ(s(fun(X141,bool),happ(s(fun(X142,fun(X141,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X142,fun(X141,bool)),X143),s(X142,X144))),s(X142,X145))),s(X141,X146)),X147,X148)) ) | $spl8),
  inference(subsumption_resolution,[],[f4696,f201])).
fof(f4696,plain,(
  ( ! [X146,X144,X142,X147,X145,X143,X141,X148] : (~p(s(bool,'T_0')) | ~p(s(bool,X147)) | sP9(happ(s(fun(X141,bool),happ(s(fun(X142,fun(X141,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X142,fun(X141,bool)),X143),s(X142,X144))),s(X142,X145))),s(X141,X146)),X147,X148)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1540])).
fof(f4825,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (~p(s(bool,X139)) | sP9(happ(s(fun(X133,bool),happ(s(fun(X134,fun(X133,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X134,fun(X133,bool)),X135),s(X134,X136))),s(X134,X137))),s(X133,X138)),X140,X139)) ) | $spl8),
  inference(subsumption_resolution,[],[f4695,f201])).
fof(f4695,plain,(
  ( ! [X134,X140,X138,X136,X135,X133,X139,X137] : (~p(s(bool,'T_0')) | ~p(s(bool,X139)) | sP9(happ(s(fun(X133,bool),happ(s(fun(X134,fun(X133,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X134,fun(X133,bool)),X135),s(X134,X136))),s(X134,X137))),s(X133,X138)),X140,X139)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1539])).
fof(f4823,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (p(s(bool,X131)) | p(s(bool,X132)) | ~sP8(happ(s(fun(X125,bool),happ(s(fun(X126,fun(X125,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X126,fun(X125,bool)),X127),s(X126,X128))),s(X126,X129))),s(X125,X130)),X131,X132)) ) | $spl8),
  inference(subsumption_resolution,[],[f4694,f201])).
fof(f4694,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (~p(s(bool,'T_0')) | p(s(bool,X131)) | p(s(bool,X132)) | ~sP8(happ(s(fun(X125,bool),happ(s(fun(X126,fun(X125,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X126,fun(X125,bool)),X127),s(X126,X128))),s(X126,X129))),s(X125,X130)),X131,X132)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1538])).
fof(f4821,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119,X117] : (~p(s(bool,X123)) | p(s(bool,X124)) | ~sP7(X124,X123,happ(s(fun(X117,bool),happ(s(fun(X118,fun(X117,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X118,fun(X117,bool)),X119),s(X118,X120))),s(X118,X121))),s(X117,X122)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4693,f201])).
fof(f4693,plain,(
  ( ! [X123,X121,X118,X124,X122,X120,X119,X117] : (~p(s(bool,'T_0')) | ~p(s(bool,X123)) | p(s(bool,X124)) | ~sP7(X124,X123,happ(s(fun(X117,bool),happ(s(fun(X118,fun(X117,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X118,fun(X117,bool)),X119),s(X118,X120))),s(X118,X121))),s(X117,X122)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1537])).
fof(f4819,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (~p(s(bool,X115)) | p(s(bool,X116)) | ~sP6(X115,happ(s(fun(X109,bool),happ(s(fun(X110,fun(X109,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X110,fun(X109,bool)),X111),s(X110,X112))),s(X110,X113))),s(X109,X114)),X116)) ) | $spl8),
  inference(subsumption_resolution,[],[f4692,f201])).
fof(f4692,plain,(
  ( ! [X111,X109,X116,X114,X112,X110,X115,X113] : (~p(s(bool,'T_0')) | ~p(s(bool,X115)) | p(s(bool,X116)) | ~sP6(X115,happ(s(fun(X109,bool),happ(s(fun(X110,fun(X109,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X110,fun(X109,bool)),X111),s(X110,X112))),s(X110,X113))),s(X109,X114)),X116)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1536])).
fof(f4817,plain,(
  ( ! [X103,X101,X107,X105,X102,X108,X106,X104] : (~p(s(bool,X107)) | p(s(bool,X108)) | ~sP3(X107,happ(s(fun(X101,bool),happ(s(fun(X102,fun(X101,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X102,fun(X101,bool)),X103),s(X102,X104))),s(X102,X105))),s(X101,X106)),X108)) ) | $spl8),
  inference(subsumption_resolution,[],[f4691,f201])).
fof(f4691,plain,(
  ( ! [X103,X101,X107,X105,X102,X108,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X107)) | p(s(bool,X108)) | ~sP3(X107,happ(s(fun(X101,bool),happ(s(fun(X102,fun(X101,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X102,fun(X101,bool)),X103),s(X102,X104))),s(X102,X105))),s(X101,X106)),X108)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1535])).
fof(f4815,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (~p(s(bool,X99)) | p(s(bool,X100)) | ~sP4(X99,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X94,fun(X93,bool)),X95),s(X94,X96))),s(X94,X97))),s(X93,X98)),X100)) ) | $spl8),
  inference(subsumption_resolution,[],[f4690,f201])).
fof(f4690,plain,(
  ( ! [X94,X99,X97,X95,X93,X100,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X99)) | p(s(bool,X100)) | ~sP4(X99,happ(s(fun(X93,bool),happ(s(fun(X94,fun(X93,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X94,fun(X93,bool)),X95),s(X94,X96))),s(X94,X97))),s(X93,X98)),X100)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1534])).
fof(f4813,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (sP5(X91,X92,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90))) | ~sP4(X91,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X92) | ~sP3(X92,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X91) | ~sP2(happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X92,X91)) ) | $spl8),
  inference(subsumption_resolution,[],[f4689,f201])).
fof(f4689,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (~p(s(bool,'T_0')) | sP5(X91,X92,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90))) | ~sP4(X91,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X92) | ~sP3(X92,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X91) | ~sP2(happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X92,X91)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1533])).
fof(f4811,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (p(s(bool,X84)) | sP1(X84,happ(s(fun(X78,bool),happ(s(fun(X79,fun(X78,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X79,fun(X78,bool)),X80),s(X79,X81))),s(X79,X82))),s(X78,X83)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4688,f201])).
fof(f4688,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (~p(s(bool,'T_0')) | p(s(bool,X84)) | sP1(X84,happ(s(fun(X78,bool),happ(s(fun(X79,fun(X78,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X79,fun(X78,bool)),X80),s(X79,X81))),s(X79,X82))),s(X78,X83)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1532])).
fof(f4809,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (p(s(bool,X77)) | sP1(happ(s(fun(X71,bool),happ(s(fun(X72,fun(X71,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X72,fun(X71,bool)),X73),s(X72,X74))),s(X72,X75))),s(X71,X76)),X77)) ) | $spl8),
  inference(subsumption_resolution,[],[f4687,f201])).
fof(f4687,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (~p(s(bool,'T_0')) | p(s(bool,X77)) | sP1(happ(s(fun(X71,bool),happ(s(fun(X72,fun(X71,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X72,fun(X71,bool)),X73),s(X72,X74))),s(X72,X75))),s(X71,X76)),X77)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1531])).
fof(f4807,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (sP0(happ(s(fun(X64,bool),happ(s(fun(X65,fun(X64,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X65,fun(X64,bool)),X66),s(X65,X67))),s(X65,X68))),s(X64,X69)),X70)) ) | $spl8),
  inference(subsumption_resolution,[],[f4686,f201])).
fof(f4686,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X64,bool),happ(s(fun(X65,fun(X64,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X65,fun(X64,bool)),X66),s(X65,X67))),s(X65,X68))),s(X64,X69)),X70)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1530])).
fof(f4805,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (sP0(X63,happ(s(fun(X57,bool),happ(s(fun(X58,fun(X57,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X58,fun(X57,bool)),X59),s(X58,X60))),s(X58,X61))),s(X57,X62)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4685,f201])).
fof(f4685,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (~p(s(bool,'T_0')) | sP0(X63,happ(s(fun(X57,bool),happ(s(fun(X58,fun(X57,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X58,fun(X57,bool)),X59),s(X58,X60))),s(X58,X61))),s(X57,X62)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1529])).
fof(f4803,plain,(
  ( ! [X4,X3,X1] : (s(X1,X3) != s(X1,X4)) ) | $spl8),
  inference(subsumption_resolution,[],[f4684,f201])).
fof(f4684,plain,(
  ( ! [X4,X3,X1] : (~p(s(bool,'T_0')) | s(X1,X3) != s(X1,X4)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f281])).
fof(f4797,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~sP8(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X117,X118) | sP9(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X117,X118)) ) | $spl8),
  inference(subsumption_resolution,[],[f4674,f201])).
fof(f4674,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (~p(s(bool,'T_0')) | ~sP8(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X117,X118) | sP9(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X117,X118)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1275])).
fof(f4795,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (p(s(bool,X111)) | p(s(bool,X112)) | ~sP9(happ(s(fun(X107,bool),'const.relation.RDOM_1'(s(fun(X107,fun(X108,bool)),X109))),s(X107,X110)),X112,X111)) ) | $spl8),
  inference(subsumption_resolution,[],[f4673,f201])).
fof(f4673,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (~p(s(bool,'T_0')) | p(s(bool,X111)) | p(s(bool,X112)) | ~sP9(happ(s(fun(X107,bool),'const.relation.RDOM_1'(s(fun(X107,fun(X108,bool)),X109))),s(X107,X110)),X112,X111)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1274])).
fof(f4793,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (~p(s(bool,X105)) | sP9(happ(s(fun(X101,bool),'const.relation.RDOM_1'(s(fun(X101,fun(X102,bool)),X103))),s(X101,X104)),X105,X106)) ) | $spl8),
  inference(subsumption_resolution,[],[f4672,f201])).
fof(f4672,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (~p(s(bool,'T_0')) | ~p(s(bool,X105)) | sP9(happ(s(fun(X101,bool),'const.relation.RDOM_1'(s(fun(X101,fun(X102,bool)),X103))),s(X101,X104)),X105,X106)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1273])).
fof(f4791,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (~p(s(bool,X99)) | sP9(happ(s(fun(X95,bool),'const.relation.RDOM_1'(s(fun(X95,fun(X96,bool)),X97))),s(X95,X98)),X100,X99)) ) | $spl8),
  inference(subsumption_resolution,[],[f4671,f201])).
fof(f4671,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (~p(s(bool,'T_0')) | ~p(s(bool,X99)) | sP9(happ(s(fun(X95,bool),'const.relation.RDOM_1'(s(fun(X95,fun(X96,bool)),X97))),s(X95,X98)),X100,X99)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1272])).
fof(f4789,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (p(s(bool,X93)) | p(s(bool,X94)) | ~sP8(happ(s(fun(X89,bool),'const.relation.RDOM_1'(s(fun(X89,fun(X90,bool)),X91))),s(X89,X92)),X93,X94)) ) | $spl8),
  inference(subsumption_resolution,[],[f4670,f201])).
fof(f4670,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~p(s(bool,'T_0')) | p(s(bool,X93)) | p(s(bool,X94)) | ~sP8(happ(s(fun(X89,bool),'const.relation.RDOM_1'(s(fun(X89,fun(X90,bool)),X91))),s(X89,X92)),X93,X94)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1271])).
fof(f4787,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,X87)) | p(s(bool,X88)) | ~sP7(X88,X87,happ(s(fun(X83,bool),'const.relation.RDOM_1'(s(fun(X83,fun(X84,bool)),X85))),s(X83,X86)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4669,f201])).
fof(f4669,plain,(
  ( ! [X88,X87,X85,X83,X86,X84] : (~p(s(bool,'T_0')) | ~p(s(bool,X87)) | p(s(bool,X88)) | ~sP7(X88,X87,happ(s(fun(X83,bool),'const.relation.RDOM_1'(s(fun(X83,fun(X84,bool)),X85))),s(X83,X86)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1270])).
fof(f4785,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (~p(s(bool,X81)) | p(s(bool,X82)) | ~sP6(X81,happ(s(fun(X77,bool),'const.relation.RDOM_1'(s(fun(X77,fun(X78,bool)),X79))),s(X77,X80)),X82)) ) | $spl8),
  inference(subsumption_resolution,[],[f4668,f201])).
fof(f4668,plain,(
  ( ! [X80,X78,X81,X79,X77,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | p(s(bool,X82)) | ~sP6(X81,happ(s(fun(X77,bool),'const.relation.RDOM_1'(s(fun(X77,fun(X78,bool)),X79))),s(X77,X80)),X82)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1269])).
fof(f4783,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,X75)) | p(s(bool,X76)) | ~sP3(X75,happ(s(fun(X71,bool),'const.relation.RDOM_1'(s(fun(X71,fun(X72,bool)),X73))),s(X71,X74)),X76)) ) | $spl8),
  inference(subsumption_resolution,[],[f4667,f201])).
fof(f4667,plain,(
  ( ! [X76,X74,X72,X71,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X75)) | p(s(bool,X76)) | ~sP3(X75,happ(s(fun(X71,bool),'const.relation.RDOM_1'(s(fun(X71,fun(X72,bool)),X73))),s(X71,X74)),X76)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1268])).
fof(f4781,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (~p(s(bool,X69)) | p(s(bool,X70)) | ~sP4(X69,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70)) ) | $spl8),
  inference(subsumption_resolution,[],[f4666,f201])).
fof(f4666,plain,(
  ( ! [X70,X68,X66,X69,X67,X65] : (~p(s(bool,'T_0')) | ~p(s(bool,X69)) | p(s(bool,X70)) | ~sP4(X69,happ(s(fun(X65,bool),'const.relation.RDOM_1'(s(fun(X65,fun(X66,bool)),X67))),s(X65,X68)),X70)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1267])).
fof(f4779,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (sP5(X63,X64,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62))) | ~sP4(X63,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64) | ~sP3(X64,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X63) | ~sP2(happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64,X63)) ) | $spl8),
  inference(subsumption_resolution,[],[f4665,f201])).
fof(f4665,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (~p(s(bool,'T_0')) | sP5(X63,X64,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62))) | ~sP4(X63,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64) | ~sP3(X64,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X63) | ~sP2(happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64,X63)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1266])).
fof(f4777,plain,(
  ( ! [X57,X54,X58,X56,X55] : (p(s(bool,X58)) | sP1(X58,happ(s(fun(X54,bool),'const.relation.RDOM_1'(s(fun(X54,fun(X55,bool)),X56))),s(X54,X57)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4664,f201])).
fof(f4664,plain,(
  ( ! [X57,X54,X58,X56,X55] : (~p(s(bool,'T_0')) | p(s(bool,X58)) | sP1(X58,happ(s(fun(X54,bool),'const.relation.RDOM_1'(s(fun(X54,fun(X55,bool)),X56))),s(X54,X57)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1265])).
fof(f4775,plain,(
  ( ! [X52,X50,X53,X51,X49] : (p(s(bool,X53)) | sP1(happ(s(fun(X49,bool),'const.relation.RDOM_1'(s(fun(X49,fun(X50,bool)),X51))),s(X49,X52)),X53)) ) | $spl8),
  inference(subsumption_resolution,[],[f4663,f201])).
fof(f4663,plain,(
  ( ! [X52,X50,X53,X51,X49] : (~p(s(bool,'T_0')) | p(s(bool,X53)) | sP1(happ(s(fun(X49,bool),'const.relation.RDOM_1'(s(fun(X49,fun(X50,bool)),X51))),s(X49,X52)),X53)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1264])).
fof(f4773,plain,(
  ( ! [X47,X45,X48,X46,X44] : (sP0(happ(s(fun(X44,bool),'const.relation.RDOM_1'(s(fun(X44,fun(X45,bool)),X46))),s(X44,X47)),X48)) ) | $spl8),
  inference(subsumption_resolution,[],[f4662,f201])).
fof(f4662,plain,(
  ( ! [X47,X45,X48,X46,X44] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X44,bool),'const.relation.RDOM_1'(s(fun(X44,fun(X45,bool)),X46))),s(X44,X47)),X48)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1263])).
fof(f4771,plain,(
  ( ! [X39,X43,X41,X42,X40] : (sP0(X43,happ(s(fun(X39,bool),'const.relation.RDOM_1'(s(fun(X39,fun(X40,bool)),X41))),s(X39,X42)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4661,f201])).
fof(f4661,plain,(
  ( ! [X39,X43,X41,X42,X40] : (~p(s(bool,'T_0')) | sP0(X43,happ(s(fun(X39,bool),'const.relation.RDOM_1'(s(fun(X39,fun(X40,bool)),X41))),s(X39,X42)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f1262])).
fof(f4760,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~sP8(happ(s(fun(X88,bool),X89),s(X88,X90)),X91,X92) | sP9(happ(s(fun(X88,bool),X89),s(X88,X90)),X91,X92)) ) | $spl8),
  inference(subsumption_resolution,[],[f4623,f201])).
fof(f4623,plain,(
  ( ! [X92,X90,X88,X91,X89] : (~p(s(bool,'T_0')) | ~sP8(happ(s(fun(X88,bool),X89),s(X88,X90)),X91,X92) | sP9(happ(s(fun(X88,bool),X89),s(X88,X90)),X91,X92)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f986])).
fof(f4758,plain,(
  ( ! [X87,X85,X83,X86,X84] : (p(s(bool,X86)) | p(s(bool,X87)) | ~sP9(happ(s(fun(X83,bool),X84),s(X83,X85)),X87,X86)) ) | $spl8),
  inference(subsumption_resolution,[],[f4622,f201])).
fof(f4622,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~p(s(bool,'T_0')) | p(s(bool,X86)) | p(s(bool,X87)) | ~sP9(happ(s(fun(X83,bool),X84),s(X83,X85)),X87,X86)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f985])).
fof(f4756,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,X81)) | sP9(happ(s(fun(X78,bool),X79),s(X78,X80)),X81,X82)) ) | $spl8),
  inference(subsumption_resolution,[],[f4621,f201])).
fof(f4621,plain,(
  ( ! [X80,X78,X81,X79,X82] : (~p(s(bool,'T_0')) | ~p(s(bool,X81)) | sP9(happ(s(fun(X78,bool),X79),s(X78,X80)),X81,X82)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f984])).
fof(f4754,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,X76)) | sP9(happ(s(fun(X73,bool),X74),s(X73,X75)),X77,X76)) ) | $spl8),
  inference(subsumption_resolution,[],[f4620,f201])).
fof(f4620,plain,(
  ( ! [X76,X74,X77,X75,X73] : (~p(s(bool,'T_0')) | ~p(s(bool,X76)) | sP9(happ(s(fun(X73,bool),X74),s(X73,X75)),X77,X76)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f983])).
fof(f4752,plain,(
  ( ! [X70,X68,X72,X71,X69] : (p(s(bool,X71)) | p(s(bool,X72)) | ~sP8(happ(s(fun(X68,bool),X69),s(X68,X70)),X71,X72)) ) | $spl8),
  inference(subsumption_resolution,[],[f4619,f201])).
fof(f4619,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~p(s(bool,'T_0')) | p(s(bool,X71)) | p(s(bool,X72)) | ~sP8(happ(s(fun(X68,bool),X69),s(X68,X70)),X71,X72)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f982])).
fof(f4750,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X67,X66,happ(s(fun(X63,bool),X64),s(X63,X65)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4618,f201])).
fof(f4618,plain,(
  ( ! [X66,X64,X67,X65,X63] : (~p(s(bool,'T_0')) | ~p(s(bool,X66)) | p(s(bool,X67)) | ~sP7(X67,X66,happ(s(fun(X63,bool),X64),s(X63,X65)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f981])).
fof(f4748,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,X61)) | p(s(bool,X62)) | ~sP6(X61,happ(s(fun(X58,bool),X59),s(X58,X60)),X62)) ) | $spl8),
  inference(subsumption_resolution,[],[f4617,f201])).
fof(f4617,plain,(
  ( ! [X61,X59,X62,X60,X58] : (~p(s(bool,'T_0')) | ~p(s(bool,X61)) | p(s(bool,X62)) | ~sP6(X61,happ(s(fun(X58,bool),X59),s(X58,X60)),X62)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f980])).
fof(f4746,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,X56)) | p(s(bool,X57)) | ~sP3(X56,happ(s(fun(X53,bool),X54),s(X53,X55)),X57)) ) | $spl8),
  inference(subsumption_resolution,[],[f4616,f201])).
fof(f4616,plain,(
  ( ! [X57,X54,X56,X55,X53] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | p(s(bool,X57)) | ~sP3(X56,happ(s(fun(X53,bool),X54),s(X53,X55)),X57)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f979])).
fof(f4744,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,X51)) | p(s(bool,X52)) | ~sP4(X51,happ(s(fun(X48,bool),X49),s(X48,X50)),X52)) ) | $spl8),
  inference(subsumption_resolution,[],[f4615,f201])).
fof(f4615,plain,(
  ( ! [X52,X50,X48,X51,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X51)) | p(s(bool,X52)) | ~sP4(X51,happ(s(fun(X48,bool),X49),s(X48,X50)),X52)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f978])).
fof(f4742,plain,(
  ( ! [X47,X45,X43,X46,X44] : (sP5(X46,X47,happ(s(fun(X43,bool),X44),s(X43,X45))) | ~sP4(X46,happ(s(fun(X43,bool),X44),s(X43,X45)),X47) | ~sP3(X47,happ(s(fun(X43,bool),X44),s(X43,X45)),X46) | ~sP2(happ(s(fun(X43,bool),X44),s(X43,X45)),X47,X46)) ) | $spl8),
  inference(subsumption_resolution,[],[f4614,f201])).
fof(f4614,plain,(
  ( ! [X47,X45,X43,X46,X44] : (~p(s(bool,'T_0')) | sP5(X46,X47,happ(s(fun(X43,bool),X44),s(X43,X45))) | ~sP4(X46,happ(s(fun(X43,bool),X44),s(X43,X45)),X47) | ~sP3(X47,happ(s(fun(X43,bool),X44),s(X43,X45)),X46) | ~sP2(happ(s(fun(X43,bool),X44),s(X43,X45)),X47,X46)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f977])).
fof(f4740,plain,(
  ( ! [X39,X41,X42,X40] : (p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4613,f201])).
fof(f4613,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),X40),s(X39,X41)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f976])).
fof(f4738,plain,(
  ( ! [X37,X35,X38,X36] : (p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),X36),s(X35,X37)),X38)) ) | $spl8),
  inference(subsumption_resolution,[],[f4612,f201])).
fof(f4612,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),X36),s(X35,X37)),X38)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f975])).
fof(f4736,plain,(
  ( ! [X33,X31,X34,X32] : (sP0(happ(s(fun(X31,bool),X32),s(X31,X33)),X34)) ) | $spl8),
  inference(subsumption_resolution,[],[f4611,f201])).
fof(f4611,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP0(happ(s(fun(X31,bool),X32),s(X31,X33)),X34)) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f974])).
fof(f4734,plain,(
  ( ! [X30,X28,X29,X27] : (sP0(X30,happ(s(fun(X27,bool),X28),s(X27,X29)))) ) | $spl8),
  inference(subsumption_resolution,[],[f4610,f201])).
fof(f4610,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'T_0')) | sP0(X30,happ(s(fun(X27,bool),X28),s(X27,X29)))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f973])).
fof(f4728,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl8),
  inference(subsumption_resolution,[],[f4577,f4491])).
fof(f4577,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | ~p(s(bool,X0))) ) | $spl8),
  inference(backward_demodulation,[],[f4491,f220])).
fof(f3686,plain,(
  ( ! [X652,X654,X653,X651] : (s(X653,f(s(bool,'T_0'))) != s(X653,g(s(bool,'T_0'))) | s(fun(bool,X653),X651) = s(fun(bool,X653),X652) | s(bool,X654) = s(bool,sK15(X651,X652,X653,bool)) | s(bool,'T_0') = s(bool,X654)) )),
  inference(superposition,[],[f276,f509])).
fof(f3685,plain,(
  ( ! [X648,X650,X649,X647] : (s(bool,'T_0') != s(bool,f(s(X647,sK15(X648,X649,bool,X647)))) | s(fun(X647,bool),X648) = s(fun(X647,bool),X649) | s(bool,X650) = s(bool,g(s(X647,sK15(X648,X649,bool,X647)))) | s(bool,'T_0') = s(bool,X650)) )),
  inference(superposition,[],[f276,f509])).
fof(f3627,plain,(
  ( ! [X441,X438,X440,X442,X437,X439] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X438,fun(bool,bool)),X439),s(X438,X440),s(X438,X441),s(bool,'T_0')))) | p(s(bool,'const.bool.IN_2'(s(bool,X437),s(fun(bool,bool),happ(s(fun(X438,fun(bool,bool)),X439),s(X438,X441)))))) | s(bool,X437) = s(bool,X442) | s(bool,'T_0') = s(bool,X442)) )),
  inference(superposition,[],[f358,f509])).
fof(f3626,plain,(
  ( ! [X436,X432,X434,X431,X433,X435] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X432,bool)),X433),s(bool,X434),s(bool,'T_0'),s(X432,X435)))) | p(s(bool,'const.bool.IN_2'(s(X432,X435),s(fun(X432,bool),happ(s(fun(bool,fun(X432,bool)),X433),s(bool,X431)))))) | s(bool,X431) = s(bool,X436) | s(bool,'T_0') = s(bool,X436)) )),
  inference(superposition,[],[f358,f509])).
fof(f3624,plain,(
  ( ! [X420,X422,X424,X421,X423,X419] : (~p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X420,fun(bool,bool)),X421),s(X420,X422)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X420,fun(bool,bool)),X421),s(X420,X423),s(X420,X422),s(bool,X419)))) | s(X420,X422) = s(X420,X423) | s(bool,X419) = s(bool,X424) | s(bool,'T_0') = s(bool,X424)) )),
  inference(superposition,[],[f357,f509])).
fof(f3623,plain,(
  ( ! [X414,X416,X418,X413,X415,X417] : (~p(s(bool,'const.bool.IN_2'(s(X414,X415),s(fun(X414,bool),happ(s(fun(bool,fun(X414,bool)),X416),s(bool,'T_0')))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X414,bool)),X416),s(bool,X417),s(bool,X413),s(X414,X415)))) | s(bool,X413) = s(bool,X417) | s(bool,X413) = s(bool,X418) | s(bool,'T_0') = s(bool,X418)) )),
  inference(superposition,[],[f357,f509])).
fof(f3622,plain,(
  ( ! [X412,X410,X409,X411] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X410,bool)),X411),s(bool,'T_0')))) | p(s(bool,'const.bool.IN_2'(s(X410,sK16(X409,X411,X410,bool)),s(fun(X410,bool),happ(s(fun(bool,fun(X410,bool)),X411),s(bool,X409)))))) | s(bool,X409) = s(bool,X412) | s(bool,'T_0') = s(bool,X412)) )),
  inference(superposition,[],[f356,f509])).
fof(f3620,plain,(
  ( ! [X399,X401,X403,X400,X402] : (~p(s(bool,'const.bool.IN_2'(s(X400,X401),s(fun(X400,bool),happ(s(fun(bool,fun(X400,bool)),X402),s(bool,'T_0')))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X400,bool)),X402),s(bool,X399)))) | s(bool,X399) = s(bool,X403) | s(bool,'T_0') = s(bool,X403)) )),
  inference(superposition,[],[f355,f509])).
fof(f3619,plain,(
  ( ! [X397,X393,X395,X396,X398,X394] : (s(bool,'T_0') != s(bool,X394) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X395,bool)),X396),s(bool,X394),s(bool,X393),s(X395,X397)))) | s(bool,X393) = s(bool,X398) | s(bool,'T_0') = s(bool,X398)) )),
  inference(superposition,[],[f281,f509])).
fof(f3618,plain,(
  ( ! [X389,X391,X387,X388,X390,X392] : (s(bool,'T_0') != s(bool,X388) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X389,bool)),X390),s(bool,X387),s(bool,X388),s(X389,X391)))) | s(bool,X387) = s(bool,X392) | s(bool,'T_0') = s(bool,X392)) )),
  inference(superposition,[],[f281,f509])).
fof(f4547,plain,(
  ( ! [X344,X346,X341,X343,X345,X347,X342] : (p(s(bool,'const.bool.IN_2'(s(X342,X346),s(fun(X342,bool),happ(s(fun(X341,fun(X342,bool)),X343),s(X341,X345)))))) | s(bool,X347) = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X341,fun(X342,bool)),X343),s(X341,X344),s(X341,X345),s(X342,X346))) | s(bool,'T_0') = s(bool,X347)) )),
  inference(subsumption_resolution,[],[f3609,f201])).
fof(f3609,plain,(
  ( ! [X344,X346,X341,X343,X345,X347,X342] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X342,X346),s(fun(X342,bool),happ(s(fun(X341,fun(X342,bool)),X343),s(X341,X345)))))) | s(bool,X347) = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X341,fun(X342,bool)),X343),s(X341,X344),s(X341,X345),s(X342,X346))) | s(bool,'T_0') = s(bool,X347)) )),
  inference(superposition,[],[f358,f509])).
fof(f4546,plain,(
  ( ! [X337,X339,X340,X336,X338] : (p(s(bool,'const.bool.IN_2'(s(X337,sK16(X339,X338,X337,X336)),s(fun(X337,bool),happ(s(fun(X336,fun(X337,bool)),X338),s(X336,X339)))))) | s(bool,X340) = s(bool,'const.relation.RDOM_2'(s(fun(X336,fun(X337,bool)),X338),s(X336,X339))) | s(bool,'T_0') = s(bool,X340)) )),
  inference(subsumption_resolution,[],[f3608,f201])).
fof(f3608,plain,(
  ( ! [X337,X339,X340,X336,X338] : (~p(s(bool,'T_0')) | p(s(bool,'const.bool.IN_2'(s(X337,sK16(X339,X338,X337,X336)),s(fun(X337,bool),happ(s(fun(X336,fun(X337,bool)),X338),s(X336,X339)))))) | s(bool,X340) = s(bool,'const.relation.RDOM_2'(s(fun(X336,fun(X337,bool)),X338),s(X336,X339))) | s(bool,'T_0') = s(bool,X340)) )),
  inference(superposition,[],[f356,f509])).
fof(f4545,plain,(
  ( ! [X315,X316,X318,X314,X317,X319] : (p(s(bool,'const.relation.RDOM_2'(s(fun(X316,fun(X314,bool)),X317),s(X316,X318)))) | s(bool,X319) = s(bool,'const.bool.IN_2'(s(X314,X315),s(fun(X314,bool),happ(s(fun(X316,fun(X314,bool)),X317),s(X316,X318))))) | s(bool,'T_0') = s(bool,X319)) )),
  inference(subsumption_resolution,[],[f3604,f201])).
fof(f3604,plain,(
  ( ! [X315,X316,X318,X314,X317,X319] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RDOM_2'(s(fun(X316,fun(X314,bool)),X317),s(X316,X318)))) | s(bool,X319) = s(bool,'const.bool.IN_2'(s(X314,X315),s(fun(X314,bool),happ(s(fun(X316,fun(X314,bool)),X317),s(X316,X318))))) | s(bool,'T_0') = s(bool,X319)) )),
  inference(superposition,[],[f355,f509])).
fof(f4544,plain,(
  ( ! [X313,X308,X310,X312,X309,X311,X307] : (p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X309,fun(X307,bool)),X310),s(X309,X312),s(X309,X311),s(X307,X308)))) | s(X309,X311) = s(X309,X312) | s(bool,X313) = s(bool,'const.bool.IN_2'(s(X307,X308),s(fun(X307,bool),happ(s(fun(X309,fun(X307,bool)),X310),s(X309,X311))))) | s(bool,'T_0') = s(bool,X313)) )),
  inference(subsumption_resolution,[],[f3603,f201])).
fof(f3603,plain,(
  ( ! [X313,X308,X310,X312,X309,X311,X307] : (~p(s(bool,'T_0')) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X309,fun(X307,bool)),X310),s(X309,X312),s(X309,X311),s(X307,X308)))) | s(X309,X311) = s(X309,X312) | s(bool,X313) = s(bool,'const.bool.IN_2'(s(X307,X308),s(fun(X307,bool),happ(s(fun(X309,fun(X307,bool)),X310),s(X309,X311))))) | s(bool,'T_0') = s(bool,X313)) )),
  inference(superposition,[],[f357,f509])).
fof(f4541,plain,(
  ( ! [X70,X68,X69,X67] : (~sP8(X67,X68,X69) | sP9(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(subsumption_resolution,[],[f3543,f201])).
fof(f3543,plain,(
  ( ! [X70,X68,X69,X67] : (~p(s(bool,'T_0')) | ~sP8(X67,X68,X69) | sP9(X67,X68,X69) | s(bool,X67) = s(bool,X70) | s(bool,'T_0') = s(bool,X70)) )),
  inference(superposition,[],[f340,f509])).
fof(f4540,plain,(
  ( ! [X66,X64,X65,X63] : (p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(subsumption_resolution,[],[f3542,f201])).
fof(f3542,plain,(
  ( ! [X66,X64,X65,X63] : (~p(s(bool,'T_0')) | p(s(bool,X64)) | p(s(bool,X65)) | ~sP9(X63,X65,X64) | s(bool,X63) = s(bool,X66) | s(bool,'T_0') = s(bool,X66)) )),
  inference(superposition,[],[f337,f509])).
fof(f4539,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(subsumption_resolution,[],[f3541,f201])).
fof(f3541,plain,(
  ( ! [X61,X59,X62,X60] : (~p(s(bool,'T_0')) | ~p(s(bool,X60)) | sP9(X59,X60,X61) | s(bool,X59) = s(bool,X62) | s(bool,'T_0') = s(bool,X62)) )),
  inference(superposition,[],[f336,f509])).
fof(f4538,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(subsumption_resolution,[],[f3540,f201])).
fof(f3540,plain,(
  ( ! [X57,X58,X56,X55] : (~p(s(bool,'T_0')) | ~p(s(bool,X56)) | sP9(X55,X57,X56) | s(bool,X55) = s(bool,X58) | s(bool,'T_0') = s(bool,X58)) )),
  inference(superposition,[],[f335,f509])).
fof(f4537,plain,(
  ( ! [X54,X52,X53,X51] : (p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(subsumption_resolution,[],[f3539,f201])).
fof(f3539,plain,(
  ( ! [X54,X52,X53,X51] : (~p(s(bool,'T_0')) | p(s(bool,X52)) | p(s(bool,X53)) | ~sP8(X51,X52,X53) | s(bool,X51) = s(bool,X54) | s(bool,'T_0') = s(bool,X54)) )),
  inference(superposition,[],[f330,f509])).
fof(f4536,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(subsumption_resolution,[],[f3538,f201])).
fof(f3538,plain,(
  ( ! [X47,X50,X48,X49] : (~p(s(bool,'T_0')) | ~p(s(bool,X48)) | p(s(bool,X49)) | ~sP7(X49,X48,X47) | s(bool,X47) = s(bool,X50) | s(bool,'T_0') = s(bool,X50)) )),
  inference(superposition,[],[f322,f509])).
fof(f4535,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(subsumption_resolution,[],[f3537,f201])).
fof(f3537,plain,(
  ( ! [X45,X43,X46,X44] : (~p(s(bool,'T_0')) | ~p(s(bool,X44)) | p(s(bool,X45)) | ~sP6(X44,X43,X45) | s(bool,X43) = s(bool,X46) | s(bool,'T_0') = s(bool,X46)) )),
  inference(superposition,[],[f313,f509])).
fof(f4534,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(subsumption_resolution,[],[f3536,f201])).
fof(f3536,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'T_0')) | ~p(s(bool,X40)) | p(s(bool,X41)) | ~sP3(X40,X39,X41) | s(bool,X39) = s(bool,X42) | s(bool,'T_0') = s(bool,X42)) )),
  inference(superposition,[],[f294,f509])).
fof(f4533,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(subsumption_resolution,[],[f3535,f201])).
fof(f3535,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'T_0')) | ~p(s(bool,X36)) | p(s(bool,X37)) | ~sP4(X36,X35,X37) | s(bool,X35) = s(bool,X38) | s(bool,'T_0') = s(bool,X38)) )),
  inference(superposition,[],[f290,f509])).
fof(f4532,plain,(
  ( ! [X33,X31,X34,X32] : (sP5(X32,X33,X31) | ~sP4(X32,X31,X33) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(subsumption_resolution,[],[f3534,f201])).
fof(f3534,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'T_0')) | sP5(X32,X33,X31) | ~sP4(X32,X31,X33) | ~sP3(X33,X31,X32) | ~sP2(X31,X33,X32) | s(bool,X31) = s(bool,X34) | s(bool,'T_0') = s(bool,X34)) )),
  inference(superposition,[],[f287,f509])).
fof(f4531,plain,(
  ( ! [X30,X28,X29] : (p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(subsumption_resolution,[],[f3533,f201])).
fof(f3533,plain,(
  ( ! [X30,X28,X29] : (~p(s(bool,'T_0')) | p(s(bool,X29)) | sP1(X29,X28) | s(bool,X28) = s(bool,X30) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f270,f509])).
fof(f4529,plain,(
  ( ! [X24,X23,X22] : (sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(subsumption_resolution,[],[f3531,f201])).
fof(f3531,plain,(
  ( ! [X24,X23,X22] : (~p(s(bool,'T_0')) | sP0(X22,X23) | s(bool,X22) = s(bool,X24) | s(bool,'T_0') = s(bool,X24)) )),
  inference(superposition,[],[f252,f509])).
fof(f4528,plain,(
  ( ! [X21,X19,X20] : (sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(subsumption_resolution,[],[f3530,f201])).
fof(f3530,plain,(
  ( ! [X21,X19,X20] : (~p(s(bool,'T_0')) | sP0(X20,X19) | s(bool,X19) = s(bool,X21) | s(bool,'T_0') = s(bool,X21)) )),
  inference(superposition,[],[f251,f509])).
fof(f4527,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(subsumption_resolution,[],[f3529,f201])).
fof(f3529,plain,(
  ( ! [X17,X18,X16] : (~p(s(bool,'T_0')) | ~p(s(bool,X17)) | s(bool,X16) = s(bool,X17) | s(bool,X16) = s(bool,X18) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f247,f509])).
fof(f4526,plain,(
  ( ! [X6,X7] : (s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(subsumption_resolution,[],[f3525,f201])).
fof(f3525,plain,(
  ( ! [X6,X7] : (~p(s(bool,'T_0')) | s(bool,'F_0') != s(bool,X6) | s(bool,X6) = s(bool,X7) | s(bool,'T_0') = s(bool,X7)) )),
  inference(superposition,[],[f220,f509])).
fof(f3464,plain,(
  ( ! [X428,X430,X432,X429,X431,X427] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X428,fun(bool,bool)),X429),s(X428,X430),s(X428,X431),s(bool,X427)))) | p(s(bool,'const.bool.IN_2'(s(bool,'T_0'),s(fun(bool,bool),happ(s(fun(X428,fun(bool,bool)),X429),s(X428,X431)))))) | s(bool,X427) = s(bool,X432) | s(bool,'T_0') = s(bool,X432)) )),
  inference(superposition,[],[f358,f509])).
fof(f3463,plain,(
  ( ! [X422,X424,X426,X421,X423,X425] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X422,bool)),X423),s(bool,X424),s(bool,X421),s(X422,X425)))) | p(s(bool,'const.bool.IN_2'(s(X422,X425),s(fun(X422,bool),happ(s(fun(bool,fun(X422,bool)),X423),s(bool,'T_0')))))) | s(bool,X421) = s(bool,X426) | s(bool,'T_0') = s(bool,X426)) )),
  inference(superposition,[],[f358,f509])).
fof(f3461,plain,(
  ( ! [X412,X414,X410,X413,X409,X411] : (~p(s(bool,'const.bool.IN_2'(s(bool,X409),s(fun(bool,bool),happ(s(fun(X410,fun(bool,bool)),X411),s(X410,X412)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X410,fun(bool,bool)),X411),s(X410,X413),s(X410,X412),s(bool,'T_0')))) | s(X410,X412) = s(X410,X413) | s(bool,X409) = s(bool,X414) | s(bool,'T_0') = s(bool,X414)) )),
  inference(superposition,[],[f357,f509])).
fof(f3460,plain,(
  ( ! [X408,X405,X407,X403,X404,X406] : (~p(s(bool,'const.bool.IN_2'(s(X404,X405),s(fun(X404,bool),happ(s(fun(bool,fun(X404,bool)),X406),s(bool,X403)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X404,bool)),X406),s(bool,X407),s(bool,'T_0'),s(X404,X405)))) | s(bool,'T_0') = s(bool,X407) | s(bool,X403) = s(bool,X408) | s(bool,'T_0') = s(bool,X408)) )),
  inference(superposition,[],[f357,f509])).
fof(f3459,plain,(
  ( ! [X399,X401,X400,X402] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X400,bool)),X401),s(bool,X399)))) | p(s(bool,'const.bool.IN_2'(s(X400,sK16('T_0',X401,X400,bool)),s(fun(X400,bool),happ(s(fun(bool,fun(X400,bool)),X401),s(bool,'T_0')))))) | s(bool,X399) = s(bool,X402) | s(bool,'T_0') = s(bool,X402)) )),
  inference(superposition,[],[f356,f509])).
fof(f3457,plain,(
  ( ! [X389,X391,X393,X390,X392] : (~p(s(bool,'const.bool.IN_2'(s(X390,X391),s(fun(X390,bool),happ(s(fun(bool,fun(X390,bool)),X392),s(bool,X389)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X390,bool)),X392),s(bool,'T_0')))) | s(bool,X389) = s(bool,X393) | s(bool,'T_0') = s(bool,X393)) )),
  inference(superposition,[],[f355,f509])).
fof(f3456,plain,(
  ( ! [X383,X385,X387,X388,X384,X386] : (s(bool,X383) != s(bool,X384) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X385,bool)),X386),s(bool,X384),s(bool,'T_0'),s(X385,X387)))) | s(bool,X383) = s(bool,X388) | s(bool,'T_0') = s(bool,X388)) )),
  inference(superposition,[],[f281,f509])).
fof(f3455,plain,(
  ( ! [X377,X379,X380,X382,X378,X381] : (s(bool,X377) != s(bool,X378) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X379,bool)),X380),s(bool,'T_0'),s(bool,X378),s(X379,X381)))) | s(bool,X377) = s(bool,X382) | s(bool,'T_0') = s(bool,X382)) )),
  inference(superposition,[],[f281,f509])).
fof(f3376,plain,(
  ( ! [X4,X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | s(bool,X2) = s(bool,X4) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f216,f509])).
fof(f3208,plain,(
  ( ! [X656,X658,X657,X655] : (s(X657,f(s(bool,X658))) != s(X657,g(s(bool,X658))) | s(fun(bool,X657),X655) = s(fun(bool,X657),X656) | s(bool,'T_0') = s(bool,X658) | s(bool,'T_0') = s(bool,sK15(X655,X656,X657,bool))) )),
  inference(superposition,[],[f276,f509])).
fof(f3207,plain,(
  ( ! [X652,X654,X653,X651] : (s(bool,X654) != s(bool,f(s(X651,sK15(X652,X653,bool,X651)))) | s(fun(X651,bool),X652) = s(fun(X651,bool),X653) | s(bool,'T_0') = s(bool,X654) | s(bool,'T_0') = s(bool,g(s(X651,sK15(X652,X653,bool,X651))))) )),
  inference(superposition,[],[f276,f509])).
fof(f3149,plain,(
  ( ! [X445,X441,X443,X444,X440,X442] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X442,fun(bool,bool)),X443),s(X442,X444),s(X442,X445),s(bool,X441)))) | p(s(bool,'const.bool.IN_2'(s(bool,X440),s(fun(bool,bool),happ(s(fun(X442,fun(bool,bool)),X443),s(X442,X445)))))) | s(bool,'T_0') = s(bool,X441) | s(bool,'T_0') = s(bool,X440)) )),
  inference(superposition,[],[f358,f509])).
fof(f3148,plain,(
  ( ! [X436,X438,X434,X437,X439,X435] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X436,bool)),X437),s(bool,X438),s(bool,X435),s(X436,X439)))) | p(s(bool,'const.bool.IN_2'(s(X436,X439),s(fun(X436,bool),happ(s(fun(bool,fun(X436,bool)),X437),s(bool,X434)))))) | s(bool,'T_0') = s(bool,X435) | s(bool,'T_0') = s(bool,X434)) )),
  inference(superposition,[],[f358,f509])).
fof(f3146,plain,(
  ( ! [X422,X424,X426,X423,X425,X427] : (~p(s(bool,'const.bool.IN_2'(s(bool,X423),s(fun(bool,bool),happ(s(fun(X424,fun(bool,bool)),X425),s(X424,X426)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X424,fun(bool,bool)),X425),s(X424,X427),s(X424,X426),s(bool,X422)))) | s(X424,X426) = s(X424,X427) | s(bool,'T_0') = s(bool,X423) | s(bool,'T_0') = s(bool,X422)) )),
  inference(superposition,[],[f357,f509])).
fof(f3145,plain,(
  ( ! [X420,X416,X418,X421,X417,X419] : (~p(s(bool,'const.bool.IN_2'(s(X418,X419),s(fun(X418,bool),happ(s(fun(bool,fun(X418,bool)),X420),s(bool,X417)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X418,bool)),X420),s(bool,X421),s(bool,X416),s(X418,X419)))) | s(bool,X416) = s(bool,X421) | s(bool,'T_0') = s(bool,X417) | s(bool,'T_0') = s(bool,X416)) )),
  inference(superposition,[],[f357,f509])).
fof(f3144,plain,(
  ( ! [X412,X414,X413,X415] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X414,bool)),X415),s(bool,X413)))) | p(s(bool,'const.bool.IN_2'(s(X414,sK16(X412,X415,X414,bool)),s(fun(X414,bool),happ(s(fun(bool,fun(X414,bool)),X415),s(bool,X412)))))) | s(bool,'T_0') = s(bool,X413) | s(bool,'T_0') = s(bool,X412)) )),
  inference(superposition,[],[f356,f509])).
fof(f3142,plain,(
  ( ! [X405,X403,X404,X406,X402] : (~p(s(bool,'const.bool.IN_2'(s(X404,X405),s(fun(X404,bool),happ(s(fun(bool,fun(X404,bool)),X406),s(bool,X403)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X404,bool)),X406),s(bool,X402)))) | s(bool,'T_0') = s(bool,X403) | s(bool,'T_0') = s(bool,X402)) )),
  inference(superposition,[],[f355,f509])).
fof(f3141,plain,(
  ( ! [X397,X399,X401,X396,X398,X400] : (s(bool,X397) != s(bool,X398) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X399,bool)),X400),s(bool,X398),s(bool,X396),s(X399,X401)))) | s(bool,'T_0') = s(bool,X397) | s(bool,'T_0') = s(bool,X396)) )),
  inference(superposition,[],[f281,f509])).
fof(f3140,plain,(
  ( ! [X391,X393,X395,X390,X392,X394] : (s(bool,X391) != s(bool,X392) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X393,bool)),X394),s(bool,X390),s(bool,X392),s(X393,X395)))) | s(bool,'T_0') = s(bool,X391) | s(bool,'T_0') = s(bool,X390)) )),
  inference(superposition,[],[f281,f509])).
fof(f3024,plain,(
  ( ! [X30,X28,X26,X24,X29,X27,X25] : (s(bool,'T_0') = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X25,fun(X24,bool)),X26),s(X25,X27),s(X25,X28),s(X24,X29))) | s(bool,X30) = s(bool,happ(s(fun(X24,bool),happ(s(fun(X25,fun(X24,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X25,fun(X24,bool)),X26),s(X25,X27))),s(X25,X28))),s(X24,X29))) | s(bool,'T_0') = s(bool,X30)) )),
  inference(superposition,[],[f509,f279])).
fof(f3023,plain,(
  ( ! [X23,X21,X19,X22,X20] : (s(bool,'T_0') = s(bool,'const.relation.RDOM_2'(s(fun(X19,fun(X20,bool)),X21),s(X19,X22))) | s(bool,X23) = s(bool,happ(s(fun(X19,bool),'const.relation.RDOM_1'(s(fun(X19,fun(X20,bool)),X21))),s(X19,X22))) | s(bool,'T_0') = s(bool,X23)) )),
  inference(superposition,[],[f509,f275])).
fof(f3022,plain,(
  ( ! [X17,X15,X18,X16] : (s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(X15,X17),s(fun(X15,bool),X16))) | s(bool,X18) = s(bool,happ(s(fun(X15,bool),X16),s(X15,X17))) | s(bool,'T_0') = s(bool,X18)) )),
  inference(superposition,[],[f509,f243])).
fof(f3011,plain,(
  ( ! [X2,X0,X3,X1] : (s(bool,X0) = s(bool,X1) | s(bool,X1) = s(bool,X2) | s(bool,'T_0') = s(bool,X2) | s(bool,X0) = s(bool,X3) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f509,f509])).
fof(f2993,plain,(
  ( ! [X37,X35,X33,X31,X36,X34,X32] : (s(bool,X37) = s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X32,fun(X31,bool)),X33),s(X32,X34),s(X32,X35),s(X31,X36))) | s(bool,'T_0') = s(bool,X37) | s(bool,'T_0') = s(bool,happ(s(fun(X31,bool),happ(s(fun(X32,fun(X31,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X32,fun(X31,bool)),X33),s(X32,X34))),s(X32,X35))),s(X31,X36)))) )),
  inference(superposition,[],[f509,f279])).
fof(f2992,plain,(
  ( ! [X30,X28,X26,X29,X27] : (s(bool,X30) = s(bool,'const.relation.RDOM_2'(s(fun(X26,fun(X27,bool)),X28),s(X26,X29))) | s(bool,'T_0') = s(bool,X30) | s(bool,'T_0') = s(bool,happ(s(fun(X26,bool),'const.relation.RDOM_1'(s(fun(X26,fun(X27,bool)),X28))),s(X26,X29)))) )),
  inference(superposition,[],[f509,f275])).
fof(f2991,plain,(
  ( ! [X24,X23,X25,X22] : (s(bool,X25) = s(bool,'const.bool.IN_2'(s(X22,X24),s(fun(X22,bool),X23))) | s(bool,'T_0') = s(bool,X25) | s(bool,'T_0') = s(bool,happ(s(fun(X22,bool),X23),s(X22,X24)))) )),
  inference(superposition,[],[f509,f243])).
fof(f2384,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'T_0') != s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f1023,f216])).
fof(f2395,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP1('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f265])).
fof(f2396,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~sP1('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f1023,f1023,f265])).
fof(f2398,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP1('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f201,f1023,f267])).
fof(f2400,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP1('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f1023,f270])).
fof(f2402,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP5('T_0','T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f201,f201,f1023,f286])).
fof(f2429,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f367,f362,f365,f1023,f288])).
fof(f2469,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP4('T_0','T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f201,f201,f1023,f290])).
fof(f2476,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP4(X0,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X4),s(X2,X5)),X6)) )),
  inference(unit_resulting_resolution,[],[f1023,f292])).
fof(f2477,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP4('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5,X6)) )),
  inference(unit_resulting_resolution,[],[f1023,f293])).
fof(f2478,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP3('T_0','T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f201,f201,f1023,f294])).
fof(f2485,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP3(X0,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X4),s(X2,X5)),X6)) )),
  inference(unit_resulting_resolution,[],[f1023,f296])).
fof(f2486,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP3('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5,X6)) )),
  inference(unit_resulting_resolution,[],[f1023,f297])).
fof(f2487,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP2('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f353,f1023,f298])).
fof(f2495,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP2('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f353,f1023,f298])).
fof(f2496,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~sP2('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1023,f353,f1023,f298])).
fof(f2503,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP2('F_0','F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f353,f353,f1023,f298])).
fof(f2504,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~sP2('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0','const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f1023,f353,f1023,f298])).
fof(f2505,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~sP2('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f1023,f298])).
fof(f2506,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (~sP2('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)),'const.relation.RDOM_DELETE_4'(s(fun(X10,fun(X11,bool)),X12),s(X10,X13),s(X10,X13),s(X11,X14)))) )),
  inference(unit_resulting_resolution,[],[f1023,f1023,f1023,f298])).
fof(f2511,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f386,f1023,f302])).
fof(f2514,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'T_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f386,f1023,f302])).
fof(f2522,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),happ(s(fun(X5,bool),X6),s(X5,X7)),'const.bool.IN_2'(s(X5,X7),s(fun(X5,bool),X6)))) )),
  inference(unit_resulting_resolution,[],[f243,f1023,f305])).
fof(f2523,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),happ(s(fun(X5,bool),'const.relation.RDOM_1'(s(fun(X5,fun(X6,bool)),X7))),s(X5,X8)),'const.relation.RDOM_2'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8)))) )),
  inference(unit_resulting_resolution,[],[f275,f1023,f305])).
fof(f2524,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (~sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),happ(s(fun(X5,bool),happ(s(fun(X6,fun(X5,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X6,fun(X5,bool)),X7),s(X6,X8))),s(X6,X9))),s(X5,X10)),'const.relation.RDOM_DELETE_4'(s(fun(X6,fun(X5,bool)),X7),s(X6,X8),s(X6,X9),s(X5,X10)))) )),
  inference(unit_resulting_resolution,[],[f279,f1023,f305])).
fof(f2528,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (~sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),X7)),happ(s(fun(X5,bool),X7),s(X5,X6)))) )),
  inference(unit_resulting_resolution,[],[f243,f1023,f305])).
fof(f2529,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_2'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8)),happ(s(fun(X5,bool),'const.relation.RDOM_1'(s(fun(X5,fun(X6,bool)),X7))),s(X5,X8)))) )),
  inference(unit_resulting_resolution,[],[f275,f1023,f305])).
fof(f2530,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (~sP5('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X9),s(X6,X10)),happ(s(fun(X6,bool),happ(s(fun(X5,fun(X6,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8))),s(X5,X9))),s(X6,X10)))) )),
  inference(unit_resulting_resolution,[],[f279,f1023,f305])).
fof(f2533,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('T_0',X0,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X4),s(X2,X5)))) )),
  inference(unit_resulting_resolution,[],[f201,f1023,f307])).
fof(f2539,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP7('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5)) )),
  inference(unit_resulting_resolution,[],[f201,f1023,f308])).
fof(f2545,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP6('T_0','T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f201,f201,f1023,f313])).
fof(f2552,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6(X0,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X4),s(X2,X5)),X6)) )),
  inference(unit_resulting_resolution,[],[f1023,f315])).
fof(f2553,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP6('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5,X6)) )),
  inference(unit_resulting_resolution,[],[f1023,f316])).
fof(f2554,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('F_0',X0,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X2,bool)),X3),s(X1,X4),s(X1,X4),s(X2,X5)))) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f317])).
fof(f2561,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (sP7('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5,'const.relation.RDOM_DELETE_4'(s(fun(X6,fun(X7,bool)),X8),s(X6,X9),s(X6,X9),s(X7,X10)))) )),
  inference(unit_resulting_resolution,[],[f1023,f1023,f317])).
fof(f2563,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5)) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f318])).
fof(f2570,plain,(
  ( ! [X6,X4,X2,X0,X10,X8,X7,X5,X3,X1,X9] : (sP7('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)),X10)) )),
  inference(unit_resulting_resolution,[],[f1023,f1023,f318])).
fof(f2572,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP7('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'T_0','T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f201,f1023,f322])).
fof(f2575,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP9('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f1023,f323])).
fof(f2577,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP9('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'T_0',X5)) )),
  inference(unit_resulting_resolution,[],[f201,f1023,f324])).
fof(f2619,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f201,f353,f1023,f330])).
fof(f2627,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP8('T_0','F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f353,f201,f1023,f330])).
fof(f2628,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~sP8('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f1023,f201,f1023,f330])).
fof(f2635,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (sP8('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5,X6)) )),
  inference(unit_resulting_resolution,[],[f1023,f333])).
fof(f2636,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f353,f1023,f334])).
fof(f2643,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('F_0','F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f353,f353,f1023,f334])).
fof(f2645,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (sP9('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f1023,f334])).
fof(f2650,plain,(
  ( ! [X4,X2,X0,X3,X1] : (sP9('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0','F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f353,f1023,f334])).
fof(f2651,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (sP9('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f1023,f353,f1023,f334])).
fof(f2652,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (sP9('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0','const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f1023,f334])).
fof(f2653,plain,(
  ( ! [X6,X4,X2,X0,X14,X12,X10,X8,X7,X5,X3,X1,X13,X11,X9] : (sP9('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)),'const.relation.RDOM_DELETE_4'(s(fun(X10,fun(X11,bool)),X12),s(X10,X13),s(X10,X13),s(X11,X14)))) )),
  inference(unit_resulting_resolution,[],[f1023,f1023,f1023,f334])).
fof(f2656,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP9('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f201,f1023,f337])).
fof(f2661,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP9('T_0','F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f201,f353,f1023,f337])).
fof(f2662,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~sP9('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f201,f1023,f1023,f337])).
fof(f2666,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),X5,'F_0')) )),
  inference(unit_resulting_resolution,[],[f371,f1023,f341])).
fof(f2669,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP7('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0',X5)) )),
  inference(unit_resulting_resolution,[],[f372,f1023,f341])).
fof(f2673,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f354])).
fof(f2677,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~sP0('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)))) )),
  inference(unit_resulting_resolution,[],[f353,f1023,f354])).
fof(f2678,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1,X9] : (~sP0('const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X3),s(X1,X4)),'const.relation.RDOM_DELETE_4'(s(fun(X5,fun(X6,bool)),X7),s(X5,X8),s(X5,X8),s(X6,X9)))) )),
  inference(unit_resulting_resolution,[],[f1023,f1023,f354])).
fof(f2149,plain,(
  ( ! [X2,X0,X1] : (sP4('T_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)),happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f957,f285])).
fof(f1595,plain,(
  ( ! [X239,X237,X235,X233,X238,X236,X234,X232] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X238,bool)),X239),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235),s(X233,X236),s(X232,X237)))))) | p(s(bool,'const.bool.IN_2'(s(X238,sK16(happ(s(fun(X232,bool),happ(s(fun(X233,fun(X232,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235))),s(X233,X236))),s(X232,X237)),X239,X238,bool)),s(fun(X238,bool),happ(s(fun(bool,fun(X238,bool)),X239),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235),s(X233,X236),s(X232,X237))))))))) )),
  inference(forward_demodulation,[],[f1594,f1518])).
fof(f1594,plain,(
  ( ! [X239,X237,X235,X233,X238,X236,X234,X232] : (p(s(bool,'const.bool.IN_2'(s(X238,sK16(happ(s(fun(X232,bool),happ(s(fun(X233,fun(X232,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235))),s(X233,X236))),s(X232,X237)),X239,X238,bool)),s(fun(X238,bool),happ(s(fun(bool,fun(X238,bool)),X239),s(bool,'const.bool.IN_2'(s(X232,X237),s(fun(X232,bool),happ(s(fun(X233,fun(X232,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235))),s(X233,X236)))))))))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X238,bool)),X239),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235),s(X233,X236),s(X232,X237))))))) )),
  inference(forward_demodulation,[],[f1551,f243])).
fof(f1551,plain,(
  ( ! [X239,X237,X235,X233,X238,X236,X234,X232] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X238,bool)),X239),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235),s(X233,X236),s(X232,X237)))))) | p(s(bool,'const.bool.IN_2'(s(X238,sK16(happ(s(fun(X232,bool),happ(s(fun(X233,fun(X232,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235))),s(X233,X236))),s(X232,X237)),X239,X238,bool)),s(fun(X238,bool),happ(s(fun(bool,fun(X238,bool)),X239),s(bool,happ(s(fun(X232,bool),happ(s(fun(X233,fun(X232,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X233,fun(X232,bool)),X234),s(X233,X235))),s(X233,X236))),s(X232,X237))))))))) )),
  inference(superposition,[],[f356,f279])).
fof(f1542,plain,(
  ( ! [X158,X163,X161,X159,X157,X164,X162,X160] : (sP9(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X158,fun(X157,bool)),X159),s(X158,X160))),s(X158,X161))),s(X157,X162)),X163,X164) | ~sP8(happ(s(fun(X157,bool),happ(s(fun(X158,fun(X157,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X158,fun(X157,bool)),X159),s(X158,X160))),s(X158,X161))),s(X157,X162)),X163,X164) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X158,fun(X157,bool)),X159),s(X158,X160),s(X158,X161),s(X157,X162))))) )),
  inference(superposition,[],[f340,f279])).
fof(f1541,plain,(
  ( ! [X156,X154,X152,X151,X149,X155,X153,X150] : (~sP9(happ(s(fun(X149,bool),happ(s(fun(X150,fun(X149,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X150,fun(X149,bool)),X151),s(X150,X152))),s(X150,X153))),s(X149,X154)),X156,X155) | p(s(bool,X155)) | p(s(bool,X156)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X150,fun(X149,bool)),X151),s(X150,X152),s(X150,X153),s(X149,X154))))) )),
  inference(superposition,[],[f337,f279])).
fof(f1538,plain,(
  ( ! [X127,X125,X132,X130,X128,X126,X131,X129] : (~sP8(happ(s(fun(X125,bool),happ(s(fun(X126,fun(X125,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X126,fun(X125,bool)),X127),s(X126,X128))),s(X126,X129))),s(X125,X130)),X131,X132) | p(s(bool,X131)) | p(s(bool,X132)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X126,fun(X125,bool)),X127),s(X126,X128),s(X126,X129),s(X125,X130))))) )),
  inference(superposition,[],[f330,f279])).
fof(f1533,plain,(
  ( ! [X92,X90,X88,X87,X85,X91,X89,X86] : (sP5(X91,X92,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88),s(X86,X89),s(X85,X90)))) | ~sP4(X91,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X92) | ~sP3(X92,happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X91) | ~sP2(happ(s(fun(X85,bool),happ(s(fun(X86,fun(X85,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X86,fun(X85,bool)),X87),s(X86,X88))),s(X86,X89))),s(X85,X90)),X92,X91)) )),
  inference(superposition,[],[f287,f279])).
fof(f1532,plain,(
  ( ! [X80,X78,X83,X81,X79,X84,X82] : (sP1(X84,happ(s(fun(X78,bool),happ(s(fun(X79,fun(X78,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X79,fun(X78,bool)),X80),s(X79,X81))),s(X79,X82))),s(X78,X83))) | p(s(bool,X84)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X79,fun(X78,bool)),X80),s(X79,X81),s(X79,X82),s(X78,X83))))) )),
  inference(superposition,[],[f270,f279])).
fof(f1531,plain,(
  ( ! [X76,X74,X72,X71,X77,X75,X73] : (sP1(happ(s(fun(X71,bool),happ(s(fun(X72,fun(X71,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X72,fun(X71,bool)),X73),s(X72,X74))),s(X72,X75))),s(X71,X76)),X77) | p(s(bool,X77)) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X72,fun(X71,bool)),X73),s(X72,X74),s(X72,X75),s(X71,X76))))) )),
  inference(superposition,[],[f267,f279])).
fof(f1530,plain,(
  ( ! [X70,X68,X66,X64,X69,X67,X65] : (sP0(happ(s(fun(X64,bool),happ(s(fun(X65,fun(X64,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X65,fun(X64,bool)),X66),s(X65,X67))),s(X65,X68))),s(X64,X69)),X70) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X65,fun(X64,bool)),X66),s(X65,X67),s(X65,X68),s(X64,X69))))) )),
  inference(superposition,[],[f252,f279])).
fof(f1529,plain,(
  ( ! [X61,X59,X57,X62,X60,X58,X63] : (sP0(X63,happ(s(fun(X57,bool),happ(s(fun(X58,fun(X57,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X58,fun(X57,bool)),X59),s(X58,X60))),s(X58,X61))),s(X57,X62))) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X58,fun(X57,bool)),X59),s(X58,X60),s(X58,X61),s(X57,X62))))) )),
  inference(superposition,[],[f251,f279])).
fof(f1494,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X1,X4))),s(X0,X5)),'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) )),
  inference(unit_resulting_resolution,[],[f353,f279,f305])).
fof(f1495,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),happ(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3))),s(X1,X4))),s(X0,X5)),'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(X0,bool)),X2),s(X1,X3),s(X1,X4),s(X0,X5)))) )),
  inference(unit_resulting_resolution,[],[f201,f279,f303])).
fof(f1571,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5),s(X2,X6),s(X3,X7))),s(bool,'const.bool.IN_2'(s(X3,X7),s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5))),s(X2,X6))))),s(X0,X8))))) )),
  inference(forward_demodulation,[],[f1496,f243])).
fof(f1496,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5),s(X2,X6),s(X3,X7))),s(bool,happ(s(fun(X3,bool),happ(s(fun(X2,fun(X3,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5))),s(X2,X6))),s(X3,X7))),s(X0,X8))))) )),
  inference(unit_resulting_resolution,[],[f279,f281])).
fof(f1497,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~sP5('F_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5)),happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(unit_resulting_resolution,[],[f353,f279,f305])).
fof(f1498,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (sP5('T_0','const.relation.RDOM_DELETE_4'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3),s(X0,X4),s(X1,X5)),happ(s(fun(X1,bool),happ(s(fun(X0,fun(X1,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3))),s(X0,X4))),s(X1,X5)))) )),
  inference(unit_resulting_resolution,[],[f201,f279,f303])).
fof(f1570,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,'const.bool.IN_2'(s(X2,X7),s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X3,fun(X2,bool)),X4),s(X3,X5))),s(X3,X6))))),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X3,fun(X2,bool)),X4),s(X3,X5),s(X3,X6),s(X2,X7))),s(X0,X8))))) )),
  inference(forward_demodulation,[],[f1499,f243])).
fof(f1499,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,happ(s(fun(X2,bool),happ(s(fun(X3,fun(X2,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X3,fun(X2,bool)),X4),s(X3,X5))),s(X3,X6))),s(X2,X7))),s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X3,fun(X2,bool)),X4),s(X3,X5),s(X3,X6),s(X2,X7))),s(X0,X8))))) )),
  inference(unit_resulting_resolution,[],[f279,f281])).
fof(f1446,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),happ(s(fun(bool,fun(X6,bool)),X8),s(bool,X5)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X6,bool)),X8),s(bool,X9),s(bool,'F_0'),s(X6,X7)))) | s(bool,'F_0') = s(bool,X9) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f357,f208])).
fof(f1445,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X3),s(bool,'F_0')))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X3),s(bool,X4),s(bool,X0),s(X1,X2)))) | s(bool,X0) = s(bool,X4) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f357,f208])).
fof(f1441,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.bool.IN_2'(s(bool,X5),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),X7),s(X6,X8)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X6,fun(bool,bool)),X7),s(X6,X9),s(X6,X8),s(bool,'F_0')))) | s(X6,X8) = s(X6,X9) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f357,f208])).
fof(f1440,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X1,fun(bool,bool)),X2),s(X1,X3)))))) | p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(bool,bool)),X2),s(X1,X4),s(X1,X3),s(bool,X0)))) | s(X1,X3) = s(X1,X4) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f357,f208])).
fof(f1403,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X6,fun(bool,bool)),X7),s(X6,X8),s(X6,X9),s(bool,X5)))) | p(s(bool,'const.bool.IN_2'(s(bool,'F_0'),s(fun(bool,bool),happ(s(fun(X6,fun(bool,bool)),X7),s(X6,X9)))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f358,f208])).
fof(f1402,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X1,fun(bool,bool)),X2),s(X1,X3),s(X1,X4),s(bool,'F_0')))) | p(s(bool,'const.bool.IN_2'(s(bool,X0),s(fun(bool,bool),happ(s(fun(X1,fun(bool,bool)),X2),s(X1,X4)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f358,f208])).
fof(f1398,plain,(
  ( ! [X6,X8,X7,X5,X9] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X6,bool)),X7),s(bool,X8),s(bool,X5),s(X6,X9)))) | p(s(bool,'const.bool.IN_2'(s(X6,X9),s(fun(X6,bool),happ(s(fun(bool,fun(X6,bool)),X7),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X5)) )),
  inference(superposition,[],[f358,f208])).
fof(f1397,plain,(
  ( ! [X4,X2,X0,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X1,bool)),X2),s(bool,X3),s(bool,'F_0'),s(X1,X4)))) | p(s(bool,'const.bool.IN_2'(s(X1,X4),s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X2),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f358,f208])).
fof(f1372,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X17,bool)),X18),s(bool,'const.relation.RDOM_2'(s(fun(X13,fun(X14,bool)),X15),s(X13,X16)))))) | p(s(bool,'const.bool.IN_2'(s(X17,sK16(happ(s(fun(X13,bool),'const.relation.RDOM_1'(s(fun(X13,fun(X14,bool)),X15))),s(X13,X16)),X18,X17,bool)),s(fun(X17,bool),happ(s(fun(bool,fun(X17,bool)),X18),s(bool,'const.relation.RDOM_2'(s(fun(X13,fun(X14,bool)),X15),s(X13,X16))))))))) )),
  inference(forward_demodulation,[],[f1371,f1251])).
fof(f1371,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (p(s(bool,'const.bool.IN_2'(s(X17,sK16(happ(s(fun(X13,bool),'const.relation.RDOM_1'(s(fun(X13,fun(X14,bool)),X15))),s(X13,X16)),X18,X17,bool)),s(fun(X17,bool),happ(s(fun(bool,fun(X17,bool)),X18),s(bool,'const.bool.IN_2'(s(X13,X16),s(fun(X13,bool),'const.relation.RDOM_1'(s(fun(X13,fun(X14,bool)),X15)))))))))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X17,bool)),X18),s(bool,'const.relation.RDOM_2'(s(fun(X13,fun(X14,bool)),X15),s(X13,X16))))))) )),
  inference(forward_demodulation,[],[f1367,f243])).
fof(f1367,plain,(
  ( ! [X14,X17,X15,X13,X18,X16] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X17,bool)),X18),s(bool,'const.relation.RDOM_2'(s(fun(X13,fun(X14,bool)),X15),s(X13,X16)))))) | p(s(bool,'const.bool.IN_2'(s(X17,sK16(happ(s(fun(X13,bool),'const.relation.RDOM_1'(s(fun(X13,fun(X14,bool)),X15))),s(X13,X16)),X18,X17,bool)),s(fun(X17,bool),happ(s(fun(bool,fun(X17,bool)),X18),s(bool,happ(s(fun(X13,bool),'const.relation.RDOM_1'(s(fun(X13,fun(X14,bool)),X15))),s(X13,X16))))))))) )),
  inference(superposition,[],[f356,f275])).
fof(f1370,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X11,bool)),X12),s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),X9)))))) | p(s(bool,'const.bool.IN_2'(s(X11,sK16(happ(s(fun(X8,bool),X9),s(X8,X10)),X12,X11,bool)),s(fun(X11,bool),happ(s(fun(bool,fun(X11,bool)),X12),s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),X9))))))))) )),
  inference(forward_demodulation,[],[f1366,f243])).
fof(f1366,plain,(
  ( ! [X12,X10,X8,X11,X9] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X11,bool)),X12),s(bool,'const.bool.IN_2'(s(X8,X10),s(fun(X8,bool),X9)))))) | p(s(bool,'const.bool.IN_2'(s(X11,sK16(happ(s(fun(X8,bool),X9),s(X8,X10)),X12,X11,bool)),s(fun(X11,bool),happ(s(fun(bool,fun(X11,bool)),X12),s(bool,happ(s(fun(X8,bool),X9),s(X8,X10))))))))) )),
  inference(superposition,[],[f356,f243])).
fof(f1364,plain,(
  ( ! [X4,X5,X3] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X4,bool)),X5),s(bool,X3)))) | p(s(bool,'const.bool.IN_2'(s(X4,sK16('F_0',X5,X4,bool)),s(fun(X4,bool),happ(s(fun(bool,fun(X4,bool)),X5),s(bool,'F_0')))))) | s(bool,'T_0') = s(bool,X3)) )),
  inference(superposition,[],[f356,f208])).
fof(f1363,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X1,bool)),X2),s(bool,'F_0')))) | p(s(bool,'const.bool.IN_2'(s(X1,sK16(X0,X2,X1,bool)),s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X2),s(bool,X0)))))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f356,f208])).
fof(f1337,plain,(
  ( ! [X6,X4,X7,X5] : (~p(s(bool,'const.bool.IN_2'(s(X5,X6),s(fun(X5,bool),happ(s(fun(bool,fun(X5,bool)),X7),s(bool,X4)))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X5,bool)),X7),s(bool,'F_0')))) | s(bool,'T_0') = s(bool,X4)) )),
  inference(superposition,[],[f355,f208])).
fof(f1336,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),happ(s(fun(bool,fun(X1,bool)),X3),s(bool,'F_0')))))) | p(s(bool,'const.relation.RDOM_2'(s(fun(bool,fun(X1,bool)),X3),s(bool,X0)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f355,f208])).
fof(f1275,plain,(
  ( ! [X118,X116,X114,X117,X115,X113] : (sP9(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X117,X118) | ~sP8(happ(s(fun(X113,bool),'const.relation.RDOM_1'(s(fun(X113,fun(X114,bool)),X115))),s(X113,X116)),X117,X118) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X113,fun(X114,bool)),X115),s(X113,X116))))) )),
  inference(superposition,[],[f340,f275])).
fof(f1274,plain,(
  ( ! [X111,X109,X107,X112,X110,X108] : (~sP9(happ(s(fun(X107,bool),'const.relation.RDOM_1'(s(fun(X107,fun(X108,bool)),X109))),s(X107,X110)),X112,X111) | p(s(bool,X111)) | p(s(bool,X112)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X107,fun(X108,bool)),X109),s(X107,X110))))) )),
  inference(superposition,[],[f337,f275])).
fof(f1273,plain,(
  ( ! [X103,X101,X105,X102,X106,X104] : (sP9(happ(s(fun(X101,bool),'const.relation.RDOM_1'(s(fun(X101,fun(X102,bool)),X103))),s(X101,X104)),X105,X106) | ~p(s(bool,X105)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X101,fun(X102,bool)),X103),s(X101,X104))))) )),
  inference(superposition,[],[f336,f275])).
fof(f1272,plain,(
  ( ! [X99,X97,X95,X100,X98,X96] : (sP9(happ(s(fun(X95,bool),'const.relation.RDOM_1'(s(fun(X95,fun(X96,bool)),X97))),s(X95,X98)),X100,X99) | ~p(s(bool,X99)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X95,fun(X96,bool)),X97),s(X95,X98))))) )),
  inference(superposition,[],[f335,f275])).
fof(f1271,plain,(
  ( ! [X94,X92,X90,X93,X91,X89] : (~sP8(happ(s(fun(X89,bool),'const.relation.RDOM_1'(s(fun(X89,fun(X90,bool)),X91))),s(X89,X92)),X93,X94) | p(s(bool,X93)) | p(s(bool,X94)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X89,fun(X90,bool)),X91),s(X89,X92))))) )),
  inference(superposition,[],[f330,f275])).
fof(f1266,plain,(
  ( ! [X61,X59,X64,X62,X60,X63] : (sP5(X63,X64,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X59,fun(X60,bool)),X61),s(X59,X62)))) | ~sP4(X63,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64) | ~sP3(X64,happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X63) | ~sP2(happ(s(fun(X59,bool),'const.relation.RDOM_1'(s(fun(X59,fun(X60,bool)),X61))),s(X59,X62)),X64,X63)) )),
  inference(superposition,[],[f287,f275])).
fof(f1265,plain,(
  ( ! [X57,X54,X58,X56,X55] : (sP1(X58,happ(s(fun(X54,bool),'const.relation.RDOM_1'(s(fun(X54,fun(X55,bool)),X56))),s(X54,X57))) | p(s(bool,X58)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X54,fun(X55,bool)),X56),s(X54,X57))))) )),
  inference(superposition,[],[f270,f275])).
fof(f1264,plain,(
  ( ! [X52,X50,X53,X51,X49] : (sP1(happ(s(fun(X49,bool),'const.relation.RDOM_1'(s(fun(X49,fun(X50,bool)),X51))),s(X49,X52)),X53) | p(s(bool,X53)) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X49,fun(X50,bool)),X51),s(X49,X52))))) )),
  inference(superposition,[],[f267,f275])).
fof(f1263,plain,(
  ( ! [X47,X45,X48,X46,X44] : (sP0(happ(s(fun(X44,bool),'const.relation.RDOM_1'(s(fun(X44,fun(X45,bool)),X46))),s(X44,X47)),X48) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X44,fun(X45,bool)),X46),s(X44,X47))))) )),
  inference(superposition,[],[f252,f275])).
fof(f1262,plain,(
  ( ! [X39,X43,X41,X42,X40] : (sP0(X43,happ(s(fun(X39,bool),'const.relation.RDOM_1'(s(fun(X39,fun(X40,bool)),X41))),s(X39,X42))) | ~p(s(bool,'const.relation.RDOM_2'(s(fun(X39,fun(X40,bool)),X41),s(X39,X42))))) )),
  inference(superposition,[],[f251,f275])).
fof(f1240,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('F_0',happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)),'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f353,f275,f305])).
fof(f1241,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0',happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)),'const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f201,f275,f303])).
fof(f1283,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,'const.relation.RDOM_2'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5))),s(bool,'const.bool.IN_2'(s(X2,X5),s(fun(X2,bool),'const.relation.RDOM_1'(s(fun(X2,fun(X3,bool)),X4))))),s(X0,X6))))) )),
  inference(forward_demodulation,[],[f1242,f243])).
fof(f1242,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,'const.relation.RDOM_2'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5))),s(bool,happ(s(fun(X2,bool),'const.relation.RDOM_1'(s(fun(X2,fun(X3,bool)),X4))),s(X2,X5))),s(X0,X6))))) )),
  inference(unit_resulting_resolution,[],[f275,f281])).
fof(f1243,plain,(
  ( ! [X2,X0,X3,X1] : (~sP5('F_0','const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)),happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f353,f275,f305])).
fof(f1244,plain,(
  ( ! [X2,X0,X3,X1] : (sP5('T_0','const.relation.RDOM_2'(s(fun(X0,fun(X1,bool)),X2),s(X0,X3)),happ(s(fun(X0,bool),'const.relation.RDOM_1'(s(fun(X0,fun(X1,bool)),X2))),s(X0,X3)))) )),
  inference(unit_resulting_resolution,[],[f201,f275,f303])).
fof(f1282,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,'const.bool.IN_2'(s(X2,X5),s(fun(X2,bool),'const.relation.RDOM_1'(s(fun(X2,fun(X3,bool)),X4))))),s(bool,'const.relation.RDOM_2'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5))),s(X0,X6))))) )),
  inference(forward_demodulation,[],[f1245,f243])).
fof(f1245,plain,(
  ( ! [X6,X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,happ(s(fun(X2,bool),'const.relation.RDOM_1'(s(fun(X2,fun(X3,bool)),X4))),s(X2,X5))),s(bool,'const.relation.RDOM_2'(s(fun(X2,fun(X3,bool)),X4),s(X2,X5))),s(X0,X6))))) )),
  inference(unit_resulting_resolution,[],[f275,f281])).
fof(f1019,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X2,bool)),X3),s(bool,X1),s(bool,X0),s(X2,X4)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f281,f208])).
fof(f1015,plain,(
  ( ! [X4,X2,X0,X3,X1] : (s(bool,'F_0') != s(bool,X1) | ~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X2,bool)),X3),s(bool,X0),s(bool,X1),s(X2,X4)))) | s(bool,'T_0') = s(bool,X0)) )),
  inference(superposition,[],[f281,f208])).
fof(f1024,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),X3))),s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),X3))),s(X0,X5))))) )),
  inference(forward_demodulation,[],[f1014,f243])).
fof(f1014,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(bool,fun(X0,bool)),X1),s(bool,happ(s(fun(X2,bool),X3),s(X2,X4))),s(bool,'const.bool.IN_2'(s(X2,X4),s(fun(X2,bool),X3))),s(X0,X5))))) )),
  inference(unit_resulting_resolution,[],[f243,f281])).
fof(f986,plain,(
  ( ! [X92,X90,X88,X91,X89] : (sP9(happ(s(fun(X88,bool),X89),s(X88,X90)),X91,X92) | ~sP8(happ(s(fun(X88,bool),X89),s(X88,X90)),X91,X92) | ~p(s(bool,'const.bool.IN_2'(s(X88,X90),s(fun(X88,bool),X89))))) )),
  inference(superposition,[],[f340,f243])).
fof(f985,plain,(
  ( ! [X87,X85,X83,X86,X84] : (~sP9(happ(s(fun(X83,bool),X84),s(X83,X85)),X87,X86) | p(s(bool,X86)) | p(s(bool,X87)) | ~p(s(bool,'const.bool.IN_2'(s(X83,X85),s(fun(X83,bool),X84))))) )),
  inference(superposition,[],[f337,f243])).
fof(f984,plain,(
  ( ! [X80,X78,X81,X79,X82] : (sP9(happ(s(fun(X78,bool),X79),s(X78,X80)),X81,X82) | ~p(s(bool,X81)) | ~p(s(bool,'const.bool.IN_2'(s(X78,X80),s(fun(X78,bool),X79))))) )),
  inference(superposition,[],[f336,f243])).
fof(f983,plain,(
  ( ! [X76,X74,X77,X75,X73] : (sP9(happ(s(fun(X73,bool),X74),s(X73,X75)),X77,X76) | ~p(s(bool,X76)) | ~p(s(bool,'const.bool.IN_2'(s(X73,X75),s(fun(X73,bool),X74))))) )),
  inference(superposition,[],[f335,f243])).
fof(f982,plain,(
  ( ! [X70,X68,X72,X71,X69] : (~sP8(happ(s(fun(X68,bool),X69),s(X68,X70)),X71,X72) | p(s(bool,X71)) | p(s(bool,X72)) | ~p(s(bool,'const.bool.IN_2'(s(X68,X70),s(fun(X68,bool),X69))))) )),
  inference(superposition,[],[f330,f243])).
fof(f977,plain,(
  ( ! [X47,X45,X43,X46,X44] : (sP5(X46,X47,happ(s(fun(X43,bool),X44),s(X43,X45))) | ~p(s(bool,'const.bool.IN_2'(s(X43,X45),s(fun(X43,bool),X44)))) | ~sP4(X46,happ(s(fun(X43,bool),X44),s(X43,X45)),X47) | ~sP3(X47,happ(s(fun(X43,bool),X44),s(X43,X45)),X46) | ~sP2(happ(s(fun(X43,bool),X44),s(X43,X45)),X47,X46)) )),
  inference(superposition,[],[f287,f243])).
fof(f976,plain,(
  ( ! [X39,X41,X42,X40] : (~p(s(bool,'const.bool.IN_2'(s(X39,X41),s(fun(X39,bool),X40)))) | p(s(bool,X42)) | sP1(X42,happ(s(fun(X39,bool),X40),s(X39,X41)))) )),
  inference(superposition,[],[f270,f243])).
fof(f975,plain,(
  ( ! [X37,X35,X38,X36] : (~p(s(bool,'const.bool.IN_2'(s(X35,X37),s(fun(X35,bool),X36)))) | p(s(bool,X38)) | sP1(happ(s(fun(X35,bool),X36),s(X35,X37)),X38)) )),
  inference(superposition,[],[f267,f243])).
fof(f974,plain,(
  ( ! [X33,X31,X34,X32] : (~p(s(bool,'const.bool.IN_2'(s(X31,X33),s(fun(X31,bool),X32)))) | sP0(happ(s(fun(X31,bool),X32),s(X31,X33)),X34)) )),
  inference(superposition,[],[f252,f243])).
fof(f973,plain,(
  ( ! [X30,X28,X29,X27] : (~p(s(bool,'const.bool.IN_2'(s(X27,X29),s(fun(X27,bool),X28)))) | sP0(X30,happ(s(fun(X27,bool),X28),s(X27,X29)))) )),
  inference(superposition,[],[f251,f243])).
fof(f956,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0',happ(s(fun(X0,bool),X1),s(X0,X2)),'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f353,f243,f305])).
fof(f957,plain,(
  ( ! [X2,X0,X1] : (sP5('T_0',happ(s(fun(X0,bool),X1),s(X0,X2)),'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) )),
  inference(unit_resulting_resolution,[],[f201,f243,f303])).
fof(f958,plain,(
  ( ! [X2,X0,X1] : (~sP5('F_0','const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)),happ(s(fun(X0,bool),X2),s(X0,X1)))) )),
  inference(unit_resulting_resolution,[],[f353,f243,f305])).
fof(f912,plain,(
  sP5('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f367,f362,f365,f353,f288])).
fof(f868,plain,(
  sP5('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f369,f364,f361,f201,f287])).
fof(f824,plain,(
  ~sP9('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f814,f325])).
fof(f814,plain,(
  ~sP8('T_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f201,f330])).
fof(f801,plain,(
  sP9('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f375,f353,f353,f329])).
fof(f741,plain,(
  ~sP7('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f732,f306])).
fof(f732,plain,(
  ~sP6('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f201,f313])).
fof(f722,plain,(
  sP7('T_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f370,f201,f201,f309])).
fof(f718,plain,(
  ~sP5('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f709,f283])).
fof(f709,plain,(
  ~sP2('F_0','F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f353,f298])).
fof(f703,plain,(
  ~sP5('F_0','T_0','T_0')),
  inference(unit_resulting_resolution,[],[f694,f284])).
fof(f694,plain,(
  ~sP3('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f201,f294])).
fof(f674,plain,(
  ~sP4('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f201,f290])).
fof(f645,plain,(
  ~sP5('T_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f201,f386,f304])).
fof(f647,plain,(
  ~sP5('T_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f201,f386,f304])).
fof(f626,plain,(
  sP5('F_0','F_0','T_0')),
  inference(unit_resulting_resolution,[],[f353,f386,f302])).
fof(f627,plain,(
  sP5('F_0','T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f386,f302])).
fof(f467,plain,(
  ( ! [X0] : (sP9('T_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f201,f336])).
fof(f456,plain,(
  ( ! [X0] : (sP9('T_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f201,f335])).
fof(f451,plain,(
  ( ! [X0] : (~sP9('F_0','T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f353,f324])).
fof(f449,plain,(
  ( ! [X0] : (~sP9('F_0',X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f353,f323])).
fof(f443,plain,(
  ( ! [X0] : (sP7('F_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f353,f353,f318])).
fof(f436,plain,(
  ( ! [X0] : (sP7('F_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f353,f353,f317])).
fof(f417,plain,(
  ( ! [X0] : (~sP7('T_0','F_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f353,f308])).
fof(f414,plain,(
  ( ! [X0] : (~sP7('T_0',X0,'F_0')) )),
  inference(unit_resulting_resolution,[],[f201,f353,f307])).
fof(f411,plain,(
  ~sP0('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f354])).
fof(f404,plain,(
  sP1('F_0','T_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f270])).
fof(f399,plain,(
  sP1('T_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f201,f267])).
fof(f396,plain,(
  ~sP1('T_0','T_0')),
  inference(unit_resulting_resolution,[],[f201,f201,f266])).
fof(f395,plain,(
  ~sP1('F_0','F_0')),
  inference(unit_resulting_resolution,[],[f353,f353,f265])).
fof(f375,plain,(
  ( ! [X0,X1] : (sP8('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f333])).
fof(f374,plain,(
  ( ! [X0,X1] : (sP8(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f201,f332])).
fof(f373,plain,(
  ( ! [X0,X1] : (sP8(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f331])).
fof(f372,plain,(
  ( ! [X0,X1] : (sP6('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f316])).
fof(f371,plain,(
  ( ! [X0,X1] : (sP6(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f353,f315])).
fof(f370,plain,(
  ( ! [X0,X1] : (sP6(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f314])).
fof(f369,plain,(
  ( ! [X0,X1] : (sP2('T_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f201,f301])).
fof(f368,plain,(
  ( ! [X0,X1] : (sP2(X0,'T_0',X1)) )),
  inference(unit_resulting_resolution,[],[f201,f300])).
fof(f367,plain,(
  ( ! [X0,X1] : (sP2(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f299])).
fof(f366,plain,(
  ( ! [X0,X1] : (sP3('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f297])).
fof(f365,plain,(
  ( ! [X0,X1] : (sP3(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f353,f296])).
fof(f364,plain,(
  ( ! [X0,X1] : (sP3(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f295])).
fof(f363,plain,(
  ( ! [X0,X1] : (sP4('F_0',X0,X1)) )),
  inference(unit_resulting_resolution,[],[f353,f293])).
fof(f362,plain,(
  ( ! [X0,X1] : (sP4(X0,'F_0',X1)) )),
  inference(unit_resulting_resolution,[],[f353,f292])).
fof(f361,plain,(
  ( ! [X0,X1] : (sP4(X0,X1,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f291])).
fof(f360,plain,(
  ( ! [X0] : (sP0('T_0',X0)) )),
  inference(unit_resulting_resolution,[],[f201,f252])).
fof(f359,plain,(
  ( ! [X0] : (sP0(X0,'T_0')) )),
  inference(unit_resulting_resolution,[],[f201,f251])).
fof(f334,plain,(
  ( ! [X2,X0,X1] : (sP9(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f197])).
fof(f337,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)) | ~sP9(X2,X0,X1)) )),
  inference(cnf_transformation,[],[f197])).
fof(f330,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f195])).
fof(f195,plain,(
  ! [X0,X1,X2] : ((~sP8(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & p(s(bool,X0))) | sP8(X0,X1,X2)))),
  inference(rectify,[],[f194])).
fof(f194,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(flattening,[],[f193])).
fof(f193,plain,(
  ! [X2,X0,X1] : ((~sP8(X2,X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & p(s(bool,X2))) | sP8(X2,X0,X1)))),
  inference(nnf_transformation,[],[f131])).
fof(f331,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f195])).
fof(f332,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f195])).
fof(f333,plain,(
  ( ! [X2,X0,X1] : (sP8(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f195])).
fof(f323,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f192,plain,(
  ! [X0,X1,X2] : ((~sP9(X0,X1,X2) | ((p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1))) & sP8(X0,X1,X2))) & ((~p(s(bool,X0)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X1))) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)))),
  inference(rectify,[],[f191])).
fof(f191,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & ((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1) | sP9(X2,X0,X1)))),
  inference(flattening,[],[f190])).
fof(f190,plain,(
  ! [X2,X0,X1] : ((~sP9(X2,X0,X1) | ((p(s(bool,X2)) | ~p(s(bool,X1))) & (p(s(bool,X2)) | ~p(s(bool,X0))) & sP8(X2,X0,X1))) & (((~p(s(bool,X2)) & p(s(bool,X1))) | (~p(s(bool,X2)) & p(s(bool,X0))) | ~sP8(X2,X0,X1)) | sP9(X2,X0,X1)))),
  inference(nnf_transformation,[],[f132])).
fof(f324,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f192])).
fof(f325,plain,(
  ( ! [X2,X0,X1] : (~sP9(X0,X1,X2) | sP8(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f192])).
fof(f340,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~sP8(X0,X1,X2) | sP9(X0,X1,X2)) )),
  inference(duplicate_literal_removal,[],[f326])).
fof(f326,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f327,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | ~p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f328,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f192])).
fof(f329,plain,(
  ( ! [X2,X0,X1] : (sP9(X0,X1,X2) | ~sP8(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f192])).
fof(f317,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X1)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f318,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | p(s(bool,X0)) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f319,plain,(
  ( ! [X2,X0,X1] : (sP7(X2,X0,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f189])).
fof(f314,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f187])).
fof(f315,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f187])).
fof(f316,plain,(
  ( ! [X2,X0,X1] : (sP6(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f187])).
fof(f306,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f185,plain,(
  ! [X0,X1,X2] : ((~sP7(X0,X1,X2) | (sP6(X1,X2,X0) & (p(s(bool,X2)) | ~p(s(bool,X0))) & (p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP6(X1,X2,X0) | (~p(s(bool,X2)) & p(s(bool,X0))) | (~p(s(bool,X1)) & p(s(bool,X0))) | sP7(X0,X1,X2)))),
  inference(rectify,[],[f184])).
fof(f184,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & (~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2))) | sP7(X2,X0,X1)))),
  inference(flattening,[],[f183])).
fof(f183,plain,(
  ! [X2,X0,X1] : ((~sP7(X2,X0,X1) | (sP6(X0,X1,X2) & (p(s(bool,X1)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X2))))) & ((~sP6(X0,X1,X2) | (~p(s(bool,X1)) & p(s(bool,X2))) | (~p(s(bool,X0)) & p(s(bool,X2)))) | sP7(X2,X0,X1)))),
  inference(nnf_transformation,[],[f129])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (~sP7(X0,X1,X2) | p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f185])).
fof(f309,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | ~p(s(bool,X1)) | ~p(s(bool,X2)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f341,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(duplicate_literal_removal,[],[f312])).
fof(f312,plain,(
  ( ! [X2,X0,X1] : (sP7(X0,X1,X2) | p(s(bool,X0)) | p(s(bool,X0)) | ~sP6(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f185])).
fof(f302,plain,(
  ( ! [X2,X0,X1] : (sP5(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f304,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | ~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f182])).
fof(f305,plain,(
  ( ! [X2,X0,X1] : (~sP5(X2,X1,X0) | s(bool,X0) != s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f182])).
fof(f298,plain,(
  ( ! [X2,X0,X1] : (~sP2(X0,X1,X2) | p(s(bool,X1)) | p(s(bool,X2)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f181,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(flattening,[],[f180])).
fof(f180,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(nnf_transformation,[],[f123])).
fof(f299,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f181])).
fof(f300,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f181])).
fof(f301,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f181])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f179])).
fof(f296,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f179])).
fof(f297,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f179])).
fof(f291,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f176])).
fof(f292,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f176])).
fof(f293,plain,(
  ( ! [X2,X0,X1] : (sP4(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f176])).
fof(f283,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f284,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | sP3(X1,X2,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f286,plain,(
  ( ! [X2,X0,X1] : (~sP5(X0,X1,X2) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f173])).
fof(f287,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X2)) | sP5(X0,X1,X2) | ~sP4(X0,X2,X1) | ~sP3(X1,X2,X0) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f288,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X1)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f289,plain,(
  ( ! [X2,X0,X1] : (~sP3(X1,X2,X0) | p(s(bool,X0)) | ~sP4(X0,X2,X1) | sP5(X0,X1,X2) | ~sP2(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f173])).
fof(f266,plain,(
  ( ! [X0,X1] : (~sP1(X0,X1) | ~p(s(bool,X1)) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f163])).
fof(f270,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | p(s(bool,X0)) | sP1(X0,X1)) )),
  inference(cnf_transformation,[],[f163])).
fof(f354,plain,(
  ( ! [X0,X1] : (~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(subsumption_resolution,[],[f250,f353])).
fof(f250,plain,(
  ( ! [X0,X1] : (p(s(bool,'F_0')) | p(s(bool,X0)) | p(s(bool,X1)) | ~sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(flattening,[],[f155])).
fof(f155,plain,(
  ! [X0,X1] : ((~sP0(X0,X1) | (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0')))) & ((~p(s(bool,X1)) & ~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | sP0(X0,X1)))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0,X1] : (sP0(X0,X1) <=> (p(s(bool,X1)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f251,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f252,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP0(X0,X1)) )),
  inference(cnf_transformation,[],[f156])).
fof(f247,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f109])).
fof(f220,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f141])).
fof(f217,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f139])).
fof(f4835,plain,(
  ( ! [X177,X175,X173,X178,X176,X174,X172] : (sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X172,bool),happ(s(fun(X173,fun(X172,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X173,fun(X172,bool)),X174),s(X173,X175))),s(X173,X176))),s(X172,X177)),X178)) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4700,f201])).
fof(f4700,plain,(
  ( ! [X177,X175,X173,X178,X176,X174,X172] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X172,bool),happ(s(fun(X173,fun(X172,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X173,fun(X172,bool)),X174),s(X173,X175))),s(X173,X176))),s(X172,X177)),X178)) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4491,f1544])).
fof(f1544,plain,(
  ( ! [X177,X175,X173,X178,X176,X174,X172] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X173,fun(X172,bool)),X174),s(X173,X175),s(X173,X176),s(X172,X177)))) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X172,bool),happ(s(fun(X173,fun(X172,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X173,fun(X172,bool)),X174),s(X173,X175))),s(X173,X176))),s(X172,X177)),X178)) ) | $spl2),
  inference(superposition,[],[f469,f279])).
fof(f469,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X0,X1)) ) | $spl2),
  inference(subsumption_resolution,[],[f468,f201])).
fof(f468,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X0,X1)) ) | $spl2),
  inference(superposition,[],[f336,f389])).
fof(f389,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14))))) | $spl2),
  inference(unit_resulting_resolution,[],[f349,f217])).
fof(f4834,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4833,f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821,f4823,f4825,f4827,f4829,f4831])).
fof(f4833,plain,(
  ( ! [X167,X165,X171,X169,X166,X170,X168] : (sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X171,happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X166,fun(X165,bool)),X167),s(X166,X168))),s(X166,X169))),s(X165,X170)))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4699,f201])).
fof(f4699,plain,(
  ( ! [X167,X165,X171,X169,X166,X170,X168] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X171,happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X166,fun(X165,bool)),X167),s(X166,X168))),s(X166,X169))),s(X165,X170)))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4491,f1543])).
fof(f1543,plain,(
  ( ! [X167,X165,X171,X169,X166,X170,X168] : (~p(s(bool,'const.relation.RDOM_DELETE_4'(s(fun(X166,fun(X165,bool)),X167),s(X166,X168),s(X166,X169),s(X165,X170)))) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X171,happ(s(fun(X165,bool),happ(s(fun(X166,fun(X165,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X166,fun(X165,bool)),X167),s(X166,X168))),s(X166,X169))),s(X165,X170)))) ) | $spl2),
  inference(superposition,[],[f458,f279])).
fof(f458,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X1,X0)) ) | $spl2),
  inference(subsumption_resolution,[],[f457,f201])).
fof(f457,plain,(
  ( ! [X0,X1] : (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X1,X0)) ) | $spl2),
  inference(superposition,[],[f335,f389])).
fof(f4832,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821,f4823,f4825,f4827,f4829,f4831])).
fof(f4830,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821,f4823,f4825,f4827,f4829])).
fof(f4828,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821,f4823,f4825,f4827])).
fof(f4826,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821,f4823,f4825])).
fof(f4824,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821,f4823])).
fof(f4822,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819,f4821])).
fof(f4820,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817,f4819])).
fof(f4818,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815,f4817])).
fof(f4816,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813,f4815])).
fof(f4814,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811,f4813])).
fof(f4812,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809,f4811])).
fof(f4810,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807,f4809])).
fof(f4808,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805,f4807])).
fof(f4806,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803,f4805])).
fof(f4804,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797,f4803])).
fof(f4802,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4801,f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797])).
fof(f4801,plain,(
  ( ! [X127,X125,X128,X126,X124] : (sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X124,bool),'const.relation.RDOM_1'(s(fun(X124,fun(X125,bool)),X126))),s(X124,X127)),X128)) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4676,f201])).
fof(f4676,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X124,bool),'const.relation.RDOM_1'(s(fun(X124,fun(X125,bool)),X126))),s(X124,X127)),X128)) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4491,f1277])).
fof(f1277,plain,(
  ( ! [X127,X125,X128,X126,X124] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X124,fun(X125,bool)),X126),s(X124,X127)))) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X124,bool),'const.relation.RDOM_1'(s(fun(X124,fun(X125,bool)),X126))),s(X124,X127)),X128)) ) | $spl2),
  inference(superposition,[],[f469,f275])).
fof(f4800,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4799,f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797])).
fof(f4799,plain,(
  ( ! [X123,X121,X122,X120,X119] : (sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X123,happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4675,f201])).
fof(f4675,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X123,happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4491,f1276])).
fof(f1276,plain,(
  ( ! [X123,X121,X122,X120,X119] : (~p(s(bool,'const.relation.RDOM_2'(s(fun(X119,fun(X120,bool)),X121),s(X119,X122)))) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X123,happ(s(fun(X119,bool),'const.relation.RDOM_1'(s(fun(X119,fun(X120,bool)),X121))),s(X119,X122)))) ) | $spl2),
  inference(superposition,[],[f458,f275])).
fof(f4798,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795,f4797])).
fof(f4796,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793,f4795])).
fof(f4794,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791,f4793])).
fof(f4792,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789,f4791])).
fof(f4790,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787,f4789])).
fof(f4788,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785,f4787])).
fof(f4786,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783,f4785])).
fof(f4784,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781,f4783])).
fof(f4782,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779,f4781])).
fof(f4780,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777,f4779])).
fof(f4778,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775,f4777])).
fof(f4776,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773,f4775])).
fof(f4774,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771,f4773])).
fof(f4772,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760,f4771])).
fof(f4765,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4764,f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760])).
fof(f4764,plain,(
  ( ! [X99,X97,X100,X98] : (sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X97,bool),X98),s(X97,X99)),X100)) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4625,f201])).
fof(f4625,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X97,bool),X98),s(X97,X99)),X100)) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4491,f988])).
fof(f988,plain,(
  ( ! [X99,X97,X100,X98] : (~p(s(bool,'const.bool.IN_2'(s(X97,X99),s(fun(X97,bool),X98)))) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),happ(s(fun(X97,bool),X98),s(X97,X99)),X100)) ) | $spl2),
  inference(superposition,[],[f469,f243])).
fof(f4763,plain,(
  $false | ($spl2 | $spl8)),
  inference(global_subsumption,[],[f4762,f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760])).
fof(f4762,plain,(
  ( ! [X94,X95,X93,X96] : (sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X96,happ(s(fun(X93,bool),X94),s(X93,X95)))) ) | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f4624,f201])).
fof(f4624,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'T_0')) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X96,happ(s(fun(X93,bool),X94),s(X93,X95)))) ) | ($spl2 | $spl8)),
  inference(backward_demodulation,[],[f4491,f987])).
fof(f987,plain,(
  ( ! [X94,X95,X93,X96] : (~p(s(bool,'const.bool.IN_2'(s(X93,X95),s(fun(X93,bool),X94)))) | sP9('const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))),X96,happ(s(fun(X93,bool),X94),s(X93,X95)))) ) | $spl2),
  inference(superposition,[],[f458,f243])).
fof(f4761,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758,f4760])).
fof(f4759,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756,f4758])).
fof(f4757,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754,f4756])).
fof(f4755,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752,f4754])).
fof(f4753,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750,f4752])).
fof(f4751,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748,f4750])).
fof(f4749,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746,f4748])).
fof(f4747,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744,f4746])).
fof(f4745,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742,f4744])).
fof(f4743,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740,f4742])).
fof(f4741,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738,f4740])).
fof(f4739,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736,f4738])).
fof(f4737,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734,f4736])).
fof(f4735,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728,f4734])).
fof(f4730,plain,(
  $false | $spl8),
  inference(subsumption_resolution,[],[f4579,f201])).
fof(f4579,plain,(
  ~p(s(bool,'T_0')) | $spl8),
  inference(backward_demodulation,[],[f4491,f353])).
fof(f4729,plain,(
  $false | $spl8),
  inference(global_subsumption,[],[f200,f199,f198,f201,f353,f208,f217,f216,f221,f220,f243,f247,f244,f252,f251,f354,f270,f267,f266,f265,f275,f276,f355,f356,f279,f357,f281,f358,f289,f288,f287,f286,f285,f284,f283,f293,f292,f291,f290,f297,f296,f295,f294,f301,f300,f299,f298,f305,f304,f303,f302,f341,f309,f308,f307,f306,f316,f315,f314,f313,f322,f319,f318,f317,f329,f328,f327,f340,f325,f324,f323,f333,f332,f331,f330,f337,f336,f335,f334,f359,f360,f361,f362,f363,f364,f365,f366,f367,f368,f369,f370,f371,f372,f373,f374,f375,f386,f395,f396,f399,f404,f411,f414,f417,f436,f443,f449,f451,f456,f467,f509,f606,f627,f626,f647,f645,f674,f694,f703,f709,f718,f722,f732,f741,f801,f814,f824,f868,f912,f959,f958,f957,f956,f964,f973,f974,f975,f976,f977,f978,f979,f980,f981,f982,f983,f984,f985,f986,f1024,f1015,f1016,f1019,f1020,f1023,f1282,f1244,f1243,f1283,f1241,f1240,f1251,f1252,f1262,f1263,f1264,f1265,f1266,f1267,f1268,f1269,f1270,f1271,f1272,f1273,f1274,f1275,f1316,f1317,f1336,f1337,f1363,f1364,f1370,f1372,f1397,f1398,f1402,f1403,f1440,f1441,f1445,f1446,f1570,f1498,f1497,f1571,f1495,f1494,f1518,f1519,f1529,f1530,f1531,f1532,f1533,f1534,f1535,f1536,f1537,f1538,f1539,f1540,f1541,f1542,f1595,f2149,f2223,f2678,f2677,f2673,f2669,f2666,f2662,f2661,f2656,f2653,f2652,f2651,f2650,f2645,f2643,f2636,f2635,f2628,f2627,f2619,f2577,f2575,f2572,f2570,f2563,f2561,f2554,f2553,f2552,f2545,f2539,f2533,f2530,f2529,f2528,f2524,f2523,f2522,f2514,f2511,f2506,f2505,f2504,f2503,f2496,f2495,f2487,f2486,f2485,f2478,f2477,f2476,f2469,f2429,f2402,f2400,f2398,f2396,f2395,f2392,f2384,f2991,f2992,f2993,f3011,f3022,f3023,f3024,f3140,f3141,f3142,f3144,f3145,f3146,f3148,f3149,f3207,f3208,f3376,f4489,f4493,f4497,f4500,f4503,f4506,f4509,f4512,f4515,f4518,f4521,f4524,f3455,f3456,f3457,f3459,f3460,f3461,f3463,f3464,f4526,f4527,f4528,f4529,f4530,f4531,f4532,f4533,f4534,f4535,f4536,f4537,f4538,f4539,f4540,f4541,f4544,f4545,f4546,f4547,f3618,f3619,f3620,f3622,f3623,f3624,f3626,f3627,f3685,f3686,f4553,f4555,f4558,f4560,f4562,f4564,f4566,f4568,f4570,f4572,f4574,f4576,f4491,f4728])).
fof(f4525,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) | $spl30),
  inference(cnf_transformation,[],[f4525_D])).
fof(f4525_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~sP8('T_0',X1,X2)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f4576,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4575])).
fof(f4575,plain,(
  ( ! [X341,X339,X340,X342] : (~sP8('T_0',X340,X341) | sP9('T_0',X340,X341) | s(bool,X339) = s(bool,X342) | s(bool,'T_0') = s(bool,X342)) )),
  inference(subsumption_resolution,[],[f3758,f606])).
fof(f3758,plain,(
  ( ! [X341,X339,X340,X342] : (~p(s(bool,X339)) | ~sP8('T_0',X340,X341) | sP9('T_0',X340,X341) | s(bool,X339) = s(bool,X342) | s(bool,'T_0') = s(bool,X342)) )),
  inference(superposition,[],[f340,f509])).
fof(f4522,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) | $spl28),
  inference(cnf_transformation,[],[f4522_D])).
fof(f4522_D,plain,(
  ( ! [X2,X1] : (~sP9('T_0',X2,X1) | p(s(bool,X2)) | p(s(bool,X1))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f4574,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4573])).
fof(f4573,plain,(
  ( ! [X335,X337,X336,X338] : (p(s(bool,X336)) | p(s(bool,X337)) | ~sP9('T_0',X337,X336) | s(bool,X335) = s(bool,X338) | s(bool,'T_0') = s(bool,X338)) )),
  inference(subsumption_resolution,[],[f3757,f606])).
fof(f3757,plain,(
  ( ! [X335,X337,X336,X338] : (~p(s(bool,X335)) | p(s(bool,X336)) | p(s(bool,X337)) | ~sP9('T_0',X337,X336) | s(bool,X335) = s(bool,X338) | s(bool,'T_0') = s(bool,X338)) )),
  inference(superposition,[],[f337,f509])).
fof(f4519,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) | $spl26),
  inference(cnf_transformation,[],[f4519_D])).
fof(f4519_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X1,X2) | ~p(s(bool,X1))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f4572,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4571])).
fof(f4571,plain,(
  ( ! [X333,X331,X332,X334] : (~p(s(bool,X332)) | sP9('T_0',X332,X333) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X334)) )),
  inference(subsumption_resolution,[],[f3756,f606])).
fof(f3756,plain,(
  ( ! [X333,X331,X332,X334] : (~p(s(bool,X331)) | ~p(s(bool,X332)) | sP9('T_0',X332,X333) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X334)) )),
  inference(superposition,[],[f336,f509])).
fof(f4516,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f4516_D])).
fof(f4516_D,plain,(
  ( ! [X2,X1] : (sP9('T_0',X2,X1) | ~p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f4570,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4569])).
fof(f4569,plain,(
  ( ! [X327,X329,X328,X330] : (~p(s(bool,X328)) | sP9('T_0',X329,X328) | s(bool,X327) = s(bool,X330) | s(bool,'T_0') = s(bool,X330)) )),
  inference(subsumption_resolution,[],[f3755,f606])).
fof(f3755,plain,(
  ( ! [X327,X329,X328,X330] : (~p(s(bool,X327)) | ~p(s(bool,X328)) | sP9('T_0',X329,X328) | s(bool,X327) = s(bool,X330) | s(bool,'T_0') = s(bool,X330)) )),
  inference(superposition,[],[f335,f509])).
fof(f4513,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) | $spl22),
  inference(cnf_transformation,[],[f4513_D])).
fof(f4513_D,plain,(
  ( ! [X2,X1] : (~sP8('T_0',X1,X2) | p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f4568,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4567])).
fof(f4567,plain,(
  ( ! [X325,X323,X324,X326] : (p(s(bool,X324)) | p(s(bool,X325)) | ~sP8('T_0',X324,X325) | s(bool,X323) = s(bool,X326) | s(bool,'T_0') = s(bool,X326)) )),
  inference(subsumption_resolution,[],[f3754,f606])).
fof(f3754,plain,(
  ( ! [X325,X323,X324,X326] : (~p(s(bool,X323)) | p(s(bool,X324)) | p(s(bool,X325)) | ~sP8('T_0',X324,X325) | s(bool,X323) = s(bool,X326) | s(bool,'T_0') = s(bool,X326)) )),
  inference(superposition,[],[f330,f509])).
fof(f4510,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) | $spl20),
  inference(cnf_transformation,[],[f4510_D])).
fof(f4510_D,plain,(
  ( ! [X2,X1] : (~sP7(X2,X1,'T_0') | p(s(bool,X2)) | ~p(s(bool,X1))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f4566,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4565])).
fof(f4565,plain,(
  ( ! [X321,X320,X322,X319] : (~p(s(bool,X320)) | p(s(bool,X321)) | ~sP7(X321,X320,'T_0') | s(bool,X319) = s(bool,X322) | s(bool,'T_0') = s(bool,X322)) )),
  inference(subsumption_resolution,[],[f3753,f606])).
fof(f3753,plain,(
  ( ! [X321,X320,X322,X319] : (~p(s(bool,X319)) | ~p(s(bool,X320)) | p(s(bool,X321)) | ~sP7(X321,X320,'T_0') | s(bool,X319) = s(bool,X322) | s(bool,'T_0') = s(bool,X322)) )),
  inference(superposition,[],[f322,f509])).
fof(f4507,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl18),
  inference(cnf_transformation,[],[f4507_D])).
fof(f4507_D,plain,(
  ( ! [X2,X1] : (~sP6(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f4564,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4563])).
fof(f4563,plain,(
  ( ! [X315,X316,X318,X317] : (~p(s(bool,X316)) | p(s(bool,X317)) | ~sP6(X316,'T_0',X317) | s(bool,X315) = s(bool,X318) | s(bool,'T_0') = s(bool,X318)) )),
  inference(subsumption_resolution,[],[f3752,f606])).
fof(f3752,plain,(
  ( ! [X315,X316,X318,X317] : (~p(s(bool,X315)) | ~p(s(bool,X316)) | p(s(bool,X317)) | ~sP6(X316,'T_0',X317) | s(bool,X315) = s(bool,X318) | s(bool,'T_0') = s(bool,X318)) )),
  inference(superposition,[],[f313,f509])).
fof(f4504,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl16),
  inference(cnf_transformation,[],[f4504_D])).
fof(f4504_D,plain,(
  ( ! [X2,X1] : (~sP3(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f4562,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4561])).
fof(f4561,plain,(
  ( ! [X313,X312,X314,X311] : (~p(s(bool,X312)) | p(s(bool,X313)) | ~sP3(X312,'T_0',X313) | s(bool,X311) = s(bool,X314) | s(bool,'T_0') = s(bool,X314)) )),
  inference(subsumption_resolution,[],[f3751,f606])).
fof(f3751,plain,(
  ( ! [X313,X312,X314,X311] : (~p(s(bool,X311)) | ~p(s(bool,X312)) | p(s(bool,X313)) | ~sP3(X312,'T_0',X313) | s(bool,X311) = s(bool,X314) | s(bool,'T_0') = s(bool,X314)) )),
  inference(superposition,[],[f294,f509])).
fof(f4501,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) | $spl14),
  inference(cnf_transformation,[],[f4501_D])).
fof(f4501_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~p(s(bool,X1)) | p(s(bool,X2))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f4560,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4559])).
fof(f4559,plain,(
  ( ! [X308,X310,X309,X307] : (~p(s(bool,X308)) | p(s(bool,X309)) | ~sP4(X308,'T_0',X309) | s(bool,X307) = s(bool,X310) | s(bool,'T_0') = s(bool,X310)) )),
  inference(subsumption_resolution,[],[f3750,f606])).
fof(f3750,plain,(
  ( ! [X308,X310,X309,X307] : (~p(s(bool,X307)) | ~p(s(bool,X308)) | p(s(bool,X309)) | ~sP4(X308,'T_0',X309) | s(bool,X307) = s(bool,X310) | s(bool,'T_0') = s(bool,X310)) )),
  inference(superposition,[],[f290,f509])).
fof(f4498,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) | $spl12),
  inference(cnf_transformation,[],[f4498_D])).
fof(f4498_D,plain,(
  ( ! [X2,X1] : (~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | sP5(X1,X2,'T_0')) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f4558,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f4557,f369])).
fof(f4557,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f4556])).
fof(f4556,plain,(
  ( ! [X304,X306,X303,X305] : (sP5(X304,X305,'T_0') | ~sP4(X304,'T_0',X305) | ~sP3(X305,'T_0',X304) | ~sP2('T_0',X305,X304) | s(bool,X303) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) )),
  inference(subsumption_resolution,[],[f3749,f606])).
fof(f3749,plain,(
  ( ! [X304,X306,X303,X305] : (~p(s(bool,X303)) | sP5(X304,X305,'T_0') | ~sP4(X304,'T_0',X305) | ~sP3(X305,'T_0',X304) | ~sP2('T_0',X305,X304) | s(bool,X303) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) )),
  inference(superposition,[],[f287,f509])).
fof(f4494,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) | $spl10),
  inference(cnf_transformation,[],[f4494_D])).
fof(f4494_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1(X1,'T_0')) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f4555,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4554])).
fof(f4554,plain,(
  ( ! [X300,X302,X301] : (p(s(bool,X301)) | sP1(X301,'T_0') | s(bool,X300) = s(bool,X302) | s(bool,'T_0') = s(bool,X302)) )),
  inference(subsumption_resolution,[],[f3748,f606])).
fof(f3748,plain,(
  ( ! [X300,X302,X301] : (~p(s(bool,X300)) | p(s(bool,X301)) | sP1(X301,'T_0') | s(bool,X300) = s(bool,X302) | s(bool,'T_0') = s(bool,X302)) )),
  inference(superposition,[],[f270,f509])).
fof(f4490,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) | $spl6),
  inference(cnf_transformation,[],[f4490_D])).
fof(f4490_D,plain,(
  ( ! [X1] : (p(s(bool,X1)) | sP1('T_0',X1)) ) <=> ~$spl6),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl6])])).
fof(f4553,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4552])).
fof(f4552,plain,(
  ( ! [X298,X297,X299] : (p(s(bool,X298)) | sP1('T_0',X298) | s(bool,X297) = s(bool,X299) | s(bool,'T_0') = s(bool,X299)) )),
  inference(subsumption_resolution,[],[f3747,f606])).
fof(f3747,plain,(
  ( ! [X298,X297,X299] : (~p(s(bool,X297)) | p(s(bool,X298)) | sP1('T_0',X298) | s(bool,X297) = s(bool,X299) | s(bool,'T_0') = s(bool,X299)) )),
  inference(superposition,[],[f267,f509])).
fof(f4524,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~sP8('T_0',X1,X2) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4523])).
fof(f4523,plain,(
  ( ! [X341,X339,X340,X342] : (~sP8('T_0',X340,X341) | sP9('T_0',X340,X341) | s(bool,X339) = s(bool,X342) | s(bool,'T_0') = s(bool,X342)) )),
  inference(subsumption_resolution,[],[f3446,f606])).
fof(f3446,plain,(
  ( ! [X341,X339,X340,X342] : (~p(s(bool,X339)) | ~sP8('T_0',X340,X341) | sP9('T_0',X340,X341) | s(bool,X339) = s(bool,X342) | s(bool,'T_0') = s(bool,X342)) )),
  inference(superposition,[],[f340,f509])).
fof(f4521,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4520])).
fof(f4520,plain,(
  ( ! [X335,X337,X336,X338] : (p(s(bool,X336)) | p(s(bool,X337)) | ~sP9('T_0',X337,X336) | s(bool,X335) = s(bool,X338) | s(bool,'T_0') = s(bool,X338)) )),
  inference(subsumption_resolution,[],[f3445,f606])).
fof(f3445,plain,(
  ( ! [X335,X337,X336,X338] : (~p(s(bool,X335)) | p(s(bool,X336)) | p(s(bool,X337)) | ~sP9('T_0',X337,X336) | s(bool,X335) = s(bool,X338) | s(bool,'T_0') = s(bool,X338)) )),
  inference(superposition,[],[f337,f509])).
fof(f4518,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X1,X2)) )),
  inference(condensation,[],[f4517])).
fof(f4517,plain,(
  ( ! [X333,X331,X332,X334] : (~p(s(bool,X332)) | sP9('T_0',X332,X333) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X334)) )),
  inference(subsumption_resolution,[],[f3444,f606])).
fof(f3444,plain,(
  ( ! [X333,X331,X332,X334] : (~p(s(bool,X331)) | ~p(s(bool,X332)) | sP9('T_0',X332,X333) | s(bool,X331) = s(bool,X334) | s(bool,'T_0') = s(bool,X334)) )),
  inference(superposition,[],[f336,f509])).
fof(f4515,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | sP9('T_0',X2,X1)) )),
  inference(condensation,[],[f4514])).
fof(f4514,plain,(
  ( ! [X327,X329,X328,X330] : (~p(s(bool,X328)) | sP9('T_0',X329,X328) | s(bool,X327) = s(bool,X330) | s(bool,'T_0') = s(bool,X330)) )),
  inference(subsumption_resolution,[],[f3443,f606])).
fof(f3443,plain,(
  ( ! [X327,X329,X328,X330] : (~p(s(bool,X327)) | ~p(s(bool,X328)) | sP9('T_0',X329,X328) | s(bool,X327) = s(bool,X330) | s(bool,'T_0') = s(bool,X330)) )),
  inference(superposition,[],[f335,f509])).
fof(f4512,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | p(s(bool,X2)) | ~sP8('T_0',X1,X2)) )),
  inference(condensation,[],[f4511])).
fof(f4511,plain,(
  ( ! [X325,X323,X324,X326] : (p(s(bool,X324)) | p(s(bool,X325)) | ~sP8('T_0',X324,X325) | s(bool,X323) = s(bool,X326) | s(bool,'T_0') = s(bool,X326)) )),
  inference(subsumption_resolution,[],[f3442,f606])).
fof(f3442,plain,(
  ( ! [X325,X323,X324,X326] : (~p(s(bool,X323)) | p(s(bool,X324)) | p(s(bool,X325)) | ~sP8('T_0',X324,X325) | s(bool,X323) = s(bool,X326) | s(bool,'T_0') = s(bool,X326)) )),
  inference(superposition,[],[f330,f509])).
fof(f4509,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP7(X2,X1,'T_0')) )),
  inference(condensation,[],[f4508])).
fof(f4508,plain,(
  ( ! [X321,X320,X322,X319] : (~p(s(bool,X320)) | p(s(bool,X321)) | ~sP7(X321,X320,'T_0') | s(bool,X319) = s(bool,X322) | s(bool,'T_0') = s(bool,X322)) )),
  inference(subsumption_resolution,[],[f3441,f606])).
fof(f3441,plain,(
  ( ! [X321,X320,X322,X319] : (~p(s(bool,X319)) | ~p(s(bool,X320)) | p(s(bool,X321)) | ~sP7(X321,X320,'T_0') | s(bool,X319) = s(bool,X322) | s(bool,'T_0') = s(bool,X322)) )),
  inference(superposition,[],[f322,f509])).
fof(f4506,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP6(X1,'T_0',X2)) )),
  inference(condensation,[],[f4505])).
fof(f4505,plain,(
  ( ! [X315,X316,X318,X317] : (~p(s(bool,X316)) | p(s(bool,X317)) | ~sP6(X316,'T_0',X317) | s(bool,X315) = s(bool,X318) | s(bool,'T_0') = s(bool,X318)) )),
  inference(subsumption_resolution,[],[f3440,f606])).
fof(f3440,plain,(
  ( ! [X315,X316,X318,X317] : (~p(s(bool,X315)) | ~p(s(bool,X316)) | p(s(bool,X317)) | ~sP6(X316,'T_0',X317) | s(bool,X315) = s(bool,X318) | s(bool,'T_0') = s(bool,X318)) )),
  inference(superposition,[],[f313,f509])).
fof(f4503,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP3(X1,'T_0',X2)) )),
  inference(condensation,[],[f4502])).
fof(f4502,plain,(
  ( ! [X313,X312,X314,X311] : (~p(s(bool,X312)) | p(s(bool,X313)) | ~sP3(X312,'T_0',X313) | s(bool,X311) = s(bool,X314) | s(bool,'T_0') = s(bool,X314)) )),
  inference(subsumption_resolution,[],[f3439,f606])).
fof(f3439,plain,(
  ( ! [X313,X312,X314,X311] : (~p(s(bool,X311)) | ~p(s(bool,X312)) | p(s(bool,X313)) | ~sP3(X312,'T_0',X313) | s(bool,X311) = s(bool,X314) | s(bool,'T_0') = s(bool,X314)) )),
  inference(superposition,[],[f294,f509])).
fof(f4500,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP4(X1,'T_0',X2)) )),
  inference(condensation,[],[f4499])).
fof(f4499,plain,(
  ( ! [X308,X310,X309,X307] : (~p(s(bool,X308)) | p(s(bool,X309)) | ~sP4(X308,'T_0',X309) | s(bool,X307) = s(bool,X310) | s(bool,'T_0') = s(bool,X310)) )),
  inference(subsumption_resolution,[],[f3438,f606])).
fof(f3438,plain,(
  ( ! [X308,X310,X309,X307] : (~p(s(bool,X307)) | ~p(s(bool,X308)) | p(s(bool,X309)) | ~sP4(X308,'T_0',X309) | s(bool,X307) = s(bool,X310) | s(bool,'T_0') = s(bool,X310)) )),
  inference(superposition,[],[f290,f509])).
fof(f4497,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1)) )),
  inference(subsumption_resolution,[],[f4496,f369])).
fof(f4496,plain,(
  ( ! [X2,X0,X1] : (s(bool,'T_0') = s(bool,X0) | sP5(X1,X2,'T_0') | ~sP4(X1,'T_0',X2) | ~sP3(X2,'T_0',X1) | ~sP2('T_0',X2,X1)) )),
  inference(condensation,[],[f4495])).
fof(f4495,plain,(
  ( ! [X304,X306,X303,X305] : (sP5(X304,X305,'T_0') | ~sP4(X304,'T_0',X305) | ~sP3(X305,'T_0',X304) | ~sP2('T_0',X305,X304) | s(bool,X303) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) )),
  inference(subsumption_resolution,[],[f3437,f606])).
fof(f3437,plain,(
  ( ! [X304,X306,X303,X305] : (~p(s(bool,X303)) | sP5(X304,X305,'T_0') | ~sP4(X304,'T_0',X305) | ~sP3(X305,'T_0',X304) | ~sP2('T_0',X305,X304) | s(bool,X303) = s(bool,X306) | s(bool,'T_0') = s(bool,X306)) )),
  inference(superposition,[],[f287,f509])).
fof(f4493,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1(X1,'T_0')) )),
  inference(condensation,[],[f4492])).
fof(f4492,plain,(
  ( ! [X300,X302,X301] : (p(s(bool,X301)) | sP1(X301,'T_0') | s(bool,X300) = s(bool,X302) | s(bool,'T_0') = s(bool,X302)) )),
  inference(subsumption_resolution,[],[f3436,f606])).
fof(f3436,plain,(
  ( ! [X300,X302,X301] : (~p(s(bool,X300)) | p(s(bool,X301)) | sP1(X301,'T_0') | s(bool,X300) = s(bool,X302) | s(bool,'T_0') = s(bool,X302)) )),
  inference(superposition,[],[f270,f509])).
fof(f4489,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') = s(bool,X0) | p(s(bool,X1)) | sP1('T_0',X1)) )),
  inference(condensation,[],[f4488])).
fof(f4488,plain,(
  ( ! [X298,X297,X299] : (p(s(bool,X298)) | sP1('T_0',X298) | s(bool,X297) = s(bool,X299) | s(bool,'T_0') = s(bool,X299)) )),
  inference(subsumption_resolution,[],[f3435,f606])).
fof(f3435,plain,(
  ( ! [X298,X297,X299] : (~p(s(bool,X297)) | p(s(bool,X298)) | sP1('T_0',X298) | s(bool,X297) = s(bool,X299) | s(bool,'T_0') = s(bool,X299)) )),
  inference(superposition,[],[f267,f509])).
fof(f347,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)))))))) | $spl0),
  inference(cnf_transformation,[],[f347_D])).
fof(f347_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)))))))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f349,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | $spl2),
  inference(cnf_transformation,[],[f349_D])).
fof(f349_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f198,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))))),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  (p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)))))))) | (p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) & s(sK11,sK12) != s(sK11,sK13))) & (~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)))))))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | s(sK11,sK12) = s(sK11,sK13))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13,sK14])],[f135])).
fof(f135,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & s(X1,X2) != s(X1,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | s(X1,X2) = s(X1,X3)))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ? [X0,X1,X2,X3,X4] : ((p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) | (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & s(X1,X2) != s(X1,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) | (~p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) | s(X1,X2) = s(X1,X3))))),
  inference(nnf_transformation,[],[f95])).
fof(f95,plain,(
  ? [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) <~> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & s(X1,X2) != s(X1,X3)))),
  inference(ennf_transformation,[],[f40])).
fof(f40,plain,(
  ~! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & s(X1,X2) != s(X1,X3)))),
  inference(flattening,[],[f39])).
fof(f39,plain,(
  ~! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X1,fun(X0,bool)),X4),s(X1,X3)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X1,X2),s(fun(X1,bool),'const.relation.RDOM_1'(s(fun(X1,fun(X0,bool)),X4)))))) & ~s(X1,X2) = s(X1,X3)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X7,fun(X6,bool)),X10),s(X7,X9)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),X10)))))) & ~s(X7,X8) = s(X7,X9)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9,X10] : (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),'const.relation.RDOM_DELETE_2'(s(fun(X7,fun(X6,bool)),X10),s(X7,X9)))))))) <=> (p(s(bool,'const.bool.IN_2'(s(X7,X8),s(fun(X7,bool),'const.relation.RDOM_1'(s(fun(X7,fun(X6,bool)),X10)))))) & ~s(X7,X8) = s(X7,X9)))),
  file('thm.relation.IN_RDOM_DELETE','thm.relation.IN_RDOM_DELETE')).
fof(f351,plain,(
  s(sK11,sK12) != s(sK11,sK13) | $spl5),
  inference(cnf_transformation,[],[f351_D])).
fof(f351_D,plain,(
  s(sK11,sK12) != s(sK11,sK13) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f199,plain,(
  s(sK11,sK12) != s(sK11,sK13) | p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))))),
  inference(cnf_transformation,[],[f136])).
fof(f348,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)))))))) | $spl1),
  inference(cnf_transformation,[],[f348_D])).
fof(f348_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13)))))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f350,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | $spl3),
  inference(cnf_transformation,[],[f350_D])).
fof(f350_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f352,plain,(
  s(sK11,sK12) = s(sK11,sK13) | $spl4),
  inference(cnf_transformation,[],[f352_D])).
fof(f352_D,plain,(
  s(sK11,sK12) = s(sK11,sK13) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f200,plain,(
  s(sK11,sK12) = s(sK11,sK13) | ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),sK14)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,sK12),s(fun(sK11,bool),'const.relation.RDOM_1'(s(fun(sK11,fun(sK10,bool)),'const.relation.RDOM_DELETE_2'(s(fun(sK11,fun(sK10,bool)),sK14),s(sK11,sK13))))))))),
  inference(cnf_transformation,[],[f136])).
% SZS output end Proof for thm.relation
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 929
% Passive clauses: 3476
% Generated clauses: 27213
% Final active clauses: 430
% Final passive clauses: 1256
% Input formulas: 37
% Initial clauses: 142
% 
% Duplicate literals: 355
% Trivial inequalities: 47
% Fw subsumption resolutions: 619
% Fw demodulations: 4940
% Bw demodulations: 238
% Condensations: 1021
% Global subsumptions: 53
% 
% Simple tautologies: 77
% Equational tautologies: 137
% Forward subsumptions: 16354
% Bw demodulations to eq. taut.: 3
% 
% Unit resulting resolution: 8470
% Factoring: 1
% Forward superposition: 5818
% Backward superposition: 4878
% Self superposition: 48
% Equality factoring: 28
% Equality resolution: 6
% 
% Split clauses: 78
% Split components: 162
% SAT solver clauses: 3743
% SAT solver unit clauses: 208
% SAT solver binary clauses: 1408
% 
% Sat splits: 78
% Sat splitting refutations: 56
% 
% Memory used [KB]: 19701
% Time elapsed: 1.221 s
% ------------------------------
% ------------------------------
% Success in time 2.525 s
