Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.integer
% (9589)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 348
% Passive clauses: 5339
% Generated clauses: 39104
% Final active clauses: 252
% Final passive clauses: 4537
% Input formulas: 133
% Initial clauses: 215
% 
% Duplicate literals: 971
% Trivial inequalities: 735
% Fw subsumption resolutions: 4075
% Fw demodulations: 8215
% Bw demodulations: 90
% 
% Simple tautologies: 651
% Equational tautologies: 49
% Forward subsumptions: 17903
% 
% Binary resolution: 1248
% Factoring: 2
% Forward superposition: 11573
% Backward superposition: 11493
% Self superposition: 141
% Equality factoring: 77
% Equality resolution: 61
% 
% Split clauses: 494
% Split components: 989
% SAT solver clauses: 1479
% SAT solver unit clauses: 33
% SAT solver binary clauses: 275
% 
% Sat splits: 494
% Sat splitting refutations: 27
% 
% TWLsolver clauses: 1475
% TWLsolver calls for satisfiability: 226
% 
% Memory used [KB]: 19189
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 143 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.integer
% (9663)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 221
% Passive clauses: 1917
% Generated clauses: 22092
% Final active clauses: 192
% Final passive clauses: 1647
% Input formulas: 133
% Initial clauses: 215
% 
% Duplicate literals: 269
% Trivial inequalities: 198
% Fw subsumption resolutions: 1813
% Fw demodulations: 5077
% Bw demodulations: 76
% Condensations: 473
% Global subsumptions: 12
% 
% Simple tautologies: 214
% Equational tautologies: 19
% Forward subsumptions: 9183
% 
% Binary resolution: 907
% Factoring: 2
% Forward superposition: 7093
% Backward superposition: 3727
% Self superposition: 92
% Equality factoring: 18
% Equality resolution: 37
% 
% Split clauses: 92
% Split components: 187
% SAT solver clauses: 3035
% SAT solver unit clauses: 264
% SAT solver binary clauses: 98
% SAT solver learnt clauses: 24
% 
% Sat splits: 92
% Sat splitting refutations: 2
% 
% TWLsolver clauses: 2987
% TWLsolver calls for satisfiability: 1231
% 
% Memory used [KB]: 20596
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.integer
% (9720)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 906
% Passive clauses: 12626
% Generated clauses: 174831
% Final active clauses: 201
% Final passive clauses: 1860
% Input formulas: 133
% Initial clauses: 215
% Discarded non-redundant clauses: 5114
% 
% Duplicate literals: 2861
% Trivial inequalities: 1334
% Fw subsumption resolutions: 12602
% Fw demodulations: 46138
% Bw demodulations: 5781
% Condensations: 3244
% Global subsumptions: 290
% 
% Simple tautologies: 734
% Equational tautologies: 446
% Forward subsumptions: 93161
% Fw demodulations to eq. taut.: 267
% Bw demodulations to eq. taut.: 217
% 
% Unit resulting resolution: 7084
% Factoring: 6
% Forward superposition: 49098
% Backward superposition: 36539
% Self superposition: 602
% Equality factoring: 117
% Equality resolution: 97
% 
% Split clauses: 96
% Split components: 201
% SAT solver clauses: 13683
% SAT solver unit clauses: 1539
% SAT solver binary clauses: 3434
% 
% Sat splits: 96
% Sat splitting refutations: 573
% 
% Memory used [KB]: 130232
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 32 next slice time: 32
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.integer
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.integer
% SZS output start Proof for thm.integer
fof(f42286,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f261,f308,f262,f271,f272,f287,f288,f322,f310,f297,f298,f301,f302,f303,f304,f387,f388,f419,f389,f390,f420,f541,f306,f551,f552,f567,f577,f578,f1149,f1151,f1148,f1147,f1146,f1145,f1152,f1153,f1202,f1203,f1332,f1363,f1486,f1428,f1427,f1490,f1488,f1637,f1818,f1874,f2166,f2168,f2508,f2579,f1673,f1664,f2686,f2687,f2690,f2691,f2692,f2693,f2740,f2741,f1242,f1240,f2768,f2769,f2779,f2811,f2942,f2943,f2959,f3444,f3669,f3670,f3990,f3991,f4981,f4983,f5065,f5066,f5272,f312,f5273,f5429,f5432,f5431,f7051,f7052,f7065,f7066,f7373,f7374,f7415,f7413,f7602,f8052,f8053,f8102,f8100,f8302,f8483,f8488,f8893,f8896,f9260,f9262,f9261,f9783,f9785,f9784,f10088,f10132,f10089,f10134,f10135,f10991,f8912,f8906,f8905,f8904,f10992,f8715,f10996,f10997,f10999,f8672,f8671,f11004,f11000,f8504,f8498,f10993,f8497,f10994,f8496,f11007,f8486,f8389,f11009,f10998,f2027,f2018,f11029,f11144,f11145,f11171,f11169,f11448,f11469,f11458,f11459,f11472,f11462,f11463,f11475,f11464,f11465,f11474,f11945,f4829,f4820,f1820,f11975,f12241,f12244,f12245,f12233,f12232,f12247,f12248,f11755,f1175,f12275,f12274,f12273,f12272,f12267,f12293,f12388,f12389,f12392,f12393,f12394,f12395,f12421,f12422,f12448,f12446,f12659,f12704,f12669,f12670,f12706,f12695,f12696,f12708,f12699,f12700,f12710,f12701,f12702,f12709,f12778,f12847,f12848,f316,f12873,f323,f12874,f325,f326,f12877,f12878,f329,f328,f12879,f12880,f330,f327,f14136,f14651,f14652,f14712,f14710,f14689,f591,f14699,f14700,f14825,f11467,f14835,f14836,f14918,f15104,f11470,f15105,f11473,f16369,f16370,f16407,f16408,f594,f16409,f16410,f11477,f593,f16707,f16714,f6616,f16717,f16744,f16745,f16746,f16749,f16592,f16761,f16762,f16589,f16588,f16587,f16766,f16751,f16750,f11476,f16712,f16710,f16769,f16770,f16768,f16796,f16590,f16591,f16798,f16568,f16800,f16347,f14809,f16804,f14184,f12362,f11078,f9775,f16806,f9774,f16807,f5043,f16802,f3251,f2660,f2428,f16810,f16813,f313,f9705,f16814,f9704,f16815,f9182,f16818,f9181,f315,f16819,f18570,f18619,f18618,f18617,f18616,f18615,f18614,f18613,f18612,f18611,f18610,f18609,f18608,f18607,f18606,f18605,f18603,f18602,f18638,f18640,f18533,f18532,f18637,f18636,f12242,f18645,f22688,f22695,f22382,f22381,f22700,f16104,f16103,f15788,f15787,f9292,f7103,f9291,f7102,f9150,f22780,f9149,f22781,f8726,f22696,f8725,f22699,f7842,f7841,f7723,f7722,f7403,f22778,f7402,f22779,f21497,f22783,f21605,f19021,f18973,f2958,f4719,f2778,f22799,f22801,f22802,f22796,f16853,f16436,f15624,f15437,f15120,f22816,f14766,f13770,f13769,f22842,f22843,f22844,f9740,f22845,f9544,f22812,f9217,f22846,f22847,f7101,f9136,f7095,f9135,f7092,f8871,f22848,f8461,f22849,f8393,f22806,f7109,f7108,f22834,f317,f22851,f22855,f6589,f22856,f6588,f22857,f5458,f22819,f5410,f22814,f22861,f21621,f21620,f21513,f21512,f20749,f20747,f20740,f20739,f22800,f22880,f22882,f15352,f16306,f16300,f16302,f16285,f16287,f22884,f22858,f22831,f22754,f22752,f22886,f22889,f22888,f22744,f22742,f22723,f22890,f22703,f22892,f22894,f22896,f22897,f22898,f22900,f22901,f22903,f17604,f22887,f17601,f22908,f17600,f22909,f17599,f22910,f17598,f22911,f22912,f22914,f22913,f22915,f22698,f9278,f22916,f9277,f7094,f22917,f19073,f19075,f16078,f16079,f12712,f18518,f22920,f16077,f16076,f12711,f22921,f22895,f2676,f3267,f22854,f22850,f22852,f11094,f12378,f12863,f14200,f22954,f320,f23198,f23275,f23276,f23277,f23278,f22798,f22864,f22865,f22808,f22810,f22797,f587,f23541,f23540,f23539,f23538,f23537,f23536,f23535,f23534,f23533,f23948,f23949,f1150,f592,f26042,f26043,f22918,f26093,f26094,f22881,f26536,f28216,f28217,f28248,f16809,f28689,f28696,f28926,f28930,f30964,f30965,f31020,f31018,f31260,f31262,f31938,f31946,f31964,f31911,f31910,f31969,f32286,f32287,f32334,f32332,f34093,f34094,f34135,f34133,f34359,f34372,f34373,f34434,f34779,f34780,f34827,f34825,f35058,f35071,f35072,f35133,f37113,f37117,f37128,f37133,f37135,f37131,f37525,f37527,f37115,f37538,f37541,f37543,f37539,f37130,f38850,f38851,f38854,f38855,f39436,f39437,f39438,f38853,f39439,f41344,f41410,f41415,f41421,f41423,f41424,f41426,f41427,f22702,f11170,f41428,f41429,f41098,f41430,f41097,f31968,f41431,f41505,f41506,f41509,f41512,f41820,f41821,f41822,f41824,f41827,f41828,f41829,f25575,f25576,f1178,f41519,f41517,f25673,f25674,f12246,f41411,f41917,f39807,f41919,f39806,f41921,f39805,f41923,f39819,f39799,f39636,f41927,f39635,f41929,f39634,f41931,f41932,f41934,f39646,f41935,f39091,f41943,f39090,f41945,f39089,f41947,f39101,f41949,f38974,f8485,f38973,f41951,f38972,f41953,f41954,f37114,f38986,f38966,f41955,f38518,f41957,f38517,f41959,f38516,f41961,f38530,f38510,f2597,f41963,f26337,f41965,f26336,f41967,f26335,f41969,f41970,f22891,f26157,f41971,f26156,f26155,f41972,f41973,f41974,f26152,f26151,f41975,f26150,f26149,f26148,f18644,f26145,f41977,f26144,f41978,f26143,f41979,f41980,f41981,f41982,f41983,f589,f26137,f26136,f590,f25790,f41984,f25789,f25788,f41985,f25785,f25784,f41986,f25783,f25782,f25781,f1241,f25778,f41987,f25777,f41988,f25776,f41989,f25775,f41990,f25769,f25768,f588,f41991,f41992,f25685,f22879,f25585,f25584,f41993,f41994,f41995,f25581,f25580,f41996,f25579,f25578,f25577,f1177,f25574,f41997,f25573,f41998,f25572,f41999,f25571,f42000,f42001,f42002,f18639,f417,f25567,f25566,f41976,f42003,f22803,f42005,f22804,f25401,f16760,f42010,f42011,f16759,f25398,f16763,f25397,f16764,f25396,f16765,f42012,f16767,f1638,f39506,f42013,f39505,f42014,f38755,f42015,f38756,f42016,f42017,f42021,f42022,f42025,f42026,f28404,f28403,f42024,f318,f41477,f41478,f1176,f39800,f418,f41925,f39409,f42027,f39405,f42030,f39402,f42031,f38823,f42032,f38819,f42033,f38816,f42029,f42034,f37526,f8897,f37524,f8894,f37116,f8489,f37112,f41950,f8484,f41112,f42051,f41111,f42052,f42050,f42053,f42054,f42055,f22833,f39492,f31971,f39491,f28927,f31970,f42056,f41417,f28928,f42248,f1639,f42260,f1641,f12447,f42285])).
fof(f42285,plain,(
  $false | ($spl33 | $spl48)),
  inference(subsumption_resolution,[],[f42284,f3477])).
fof(f3477,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) ) | $spl33),
  inference(superposition,[],[f2982,f2982])).
fof(f2982,plain,(
  ( ! [X56] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X56),s('type.integer.int',X56)))) ) | $spl33),
  inference(resolution,[],[f2817,f237])).
fof(f237,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))) )),
  inference(cnf_transformation,[],[f157])).
fof(f157,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X0] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X0))))),
  inference(rectify,[],[f37])).
fof(f37,axiom,(
  ! [X6] : ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X6))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_LT_REFL')).
fof(f2817,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | $spl33),
  inference(resolution,[],[f417,f201])).
fof(f201,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f166])).
fof(f166,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f73])).
fof(f73,axiom,(
  ! [X31,X32] : ((p(s(bool,X31)) => p(s(bool,X32))) => ((p(s(bool,X32)) => p(s(bool,X31))) => s(bool,X31) = s(bool,X32)))),
  file('thm.integer.INT_LT_LADD','thm.bool.IMP_ANTISYM_AX')).
fof(f42284,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl48),
  inference(forward_demodulation,[],[f199,f1176])).
fof(f199,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))),
  inference(cnf_transformation,[],[f178])).
fof(f178,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK1,sK2,sK3])],[f165])).
fof(f165,plain,(
  ? [X0,X1,X2] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(ennf_transformation,[],[f135])).
fof(f135,plain,(
  ~! [X0,X1,X2] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7,X8] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X8))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7,X8] : s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X8))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_LT_LADD')).
fof(f12447,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | $spl169),
  inference(cnf_transformation,[],[f12447_D])).
fof(f12447_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) <=> ~$spl169),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl169])])).
fof(f1641,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | $spl62),
  inference(cnf_transformation,[],[f1641_D])).
fof(f1641_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f42260,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl33 | $spl48)),
  inference(forward_demodulation,[],[f42143,f1176])).
fof(f42143,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | ($spl33 | $spl48)),
  inference(superposition,[],[f9379,f1176])).
fof(f9379,plain,(
  ( ! [X7] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X7)))) | s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X7))))) ) | $spl33),
  inference(resolution,[],[f3595,f238])).
fof(f238,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0'))))) )),
  inference(cnf_transformation,[],[f191])).
fof(f191,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X0)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK12])],[f173])).
fof(f173,plain,(
  ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) | ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(ennf_transformation,[],[f159])).
fof(f159,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(flattening,[],[f158])).
fof(f158,plain,(
  ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X1] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X1))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X46] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X46),s('type.integer.int','const.integer.int_|48|_0')))) => ? [X47] : s('type.integer.int',X46) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',X47))))),
  file('thm.integer.INT_LT_LADD','thm.integer.NUM_POSINT_EX')).
fof(f3595,plain,(
  ( ! [X6,X7] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6))))) ) | $spl33),
  inference(subsumption_resolution,[],[f3552,f417])).
fof(f3552,plain,(
  ( ! [X6,X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7))))) ) | $spl33),
  inference(superposition,[],[f234,f2982])).
fof(f234,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(flattening,[],[f171])).
fof(f171,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(ennf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0,X1,X2] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6,X7,X8] : ((p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) & p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X8))))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_LT_TRANS')).
fof(f1639,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl60),
  inference(cnf_transformation,[],[f1639_D])).
fof(f1639_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f42248,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl48),
  inference(forward_demodulation,[],[f42119,f1176])).
fof(f42119,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | $spl48),
  inference(superposition,[],[f238,f1176])).
fof(f28928,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl316),
  inference(cnf_transformation,[],[f28928_D])).
fof(f28928_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl316),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl316])])).
fof(f41417,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | $spl360),
  inference(cnf_transformation,[],[f41417_D])).
fof(f41417_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) <=> ~$spl360),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl360])])).
fof(f42056,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl318),
  inference(forward_demodulation,[],[f36719,f28930])).
fof(f36719,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | $spl318),
  inference(superposition,[],[f238,f28930])).
fof(f31970,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X43)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl326),
  inference(cnf_transformation,[],[f31970_D])).
fof(f31970_D,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X43)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl326),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl326])])).
fof(f28927,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl317),
  inference(cnf_transformation,[],[f28927_D])).
fof(f28927_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl317),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl317])])).
fof(f39491,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl69),
  inference(superposition,[],[f11146,f225])).
fof(f225,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f189])).
fof(f189,plain,(
  ! [X0,X1] : ((s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) <=> s('type.integer.int',X0) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f125])).
fof(f125,axiom,(
  ! [X6,X7] : (s('type.integer.int',X7) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) <=> s('type.integer.int',X6) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_ADD_LID_UNIQ')).
fof(f11146,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X0))))) ) | $spl69),
  inference(resolution,[],[f2596,f234])).
fof(f2596,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | $spl69),
  inference(cnf_transformation,[],[f2596_D])).
fof(f2596_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f31971,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X41)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl328),
  inference(cnf_transformation,[],[f31971_D])).
fof(f31971_D,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X41)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) <=> ~$spl328),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl328])])).
fof(f39492,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl69),
  inference(superposition,[],[f11146,f227])).
fof(f227,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) )),
  inference(cnf_transformation,[],[f190])).
fof(f190,plain,(
  ! [X0,X1] : ((s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) & (s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))))),
  inference(nnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) <=> s('type.integer.int',X1) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  inference(rectify,[],[f58])).
fof(f58,axiom,(
  ! [X6,X7] : (s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) <=> s('type.integer.int',X7) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_ADD_RID_UNIQ')).
fof(f22833,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl261),
  inference(cnf_transformation,[],[f22833_D])).
fof(f22833_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl261),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl261])])).
fof(f42055,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl69),
  inference(forward_demodulation,[],[f39504,f231])).
fof(f231,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f150])).
fof(f150,plain,(
  ! [X0] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int',X0)))),
  inference(rectify,[],[f118])).
fof(f118,axiom,(
  ! [X6] : s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))),s('type.integer.int',X6)))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_ADD_LINV')).
fof(f39504,plain,(
  ( ! [X40] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))))))) ) | $spl69),
  inference(superposition,[],[f11146,f231])).
fof(f42054,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl69 | $spl318)),
  inference(forward_demodulation,[],[f39516,f28930])).
fof(f39516,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | ($spl69 | $spl318)),
  inference(superposition,[],[f11146,f28930])).
fof(f42053,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl69 | $spl212)),
  inference(forward_demodulation,[],[f39520,f16809])).
fof(f39520,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl69 | $spl212)),
  inference(superposition,[],[f11146,f16809])).
fof(f42050,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl498),
  inference(cnf_transformation,[],[f42050_D])).
fof(f42050_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl498),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl498])])).
fof(f42052,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X41)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl502),
  inference(cnf_transformation,[],[f42052_D])).
fof(f42052_D,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X41)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) <=> ~$spl502),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl502])])).
fof(f41111,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X41)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl33),
  inference(superposition,[],[f9379,f227])).
fof(f42051,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X43)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl500),
  inference(cnf_transformation,[],[f42051_D])).
fof(f42051_D,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X43)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl500),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl500])])).
fof(f41112,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X43)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl33),
  inference(superposition,[],[f9379,f225])).
fof(f8484,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl100),
  inference(cnf_transformation,[],[f8484_D])).
fof(f8484_D,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f41950,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl407),
  inference(cnf_transformation,[],[f41950_D])).
fof(f41950_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl407),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl407])])).
fof(f37112,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(forward_demodulation,[],[f37032,f233])).
fof(f233,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  inference(rectify,[],[f55])).
fof(f55,axiom,(
  ! [X6] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_ADD_RID')).
fof(f37032,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f255,f225])).
fof(f255,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f199,f220])).
fof(f220,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X1) = s('type.integer.int',X2) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f187])).
fof(f187,plain,(
  ! [X0,X1,X2] : ((s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) | s('type.integer.int',X1) = s('type.integer.int',X2)) & (s('type.integer.int',X1) != s('type.integer.int',X2) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))),
  inference(nnf_transformation,[],[f143])).
fof(f143,plain,(
  ! [X0,X1,X2] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) <=> s('type.integer.int',X1) = s('type.integer.int',X2))),
  inference(rectify,[],[f54])).
fof(f54,axiom,(
  ! [X6,X7,X8] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8))) <=> s('type.integer.int',X7) = s('type.integer.int',X8))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_EQ_LADD')).
fof(f8489,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl104),
  inference(cnf_transformation,[],[f8489_D])).
fof(f8489_D,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f37116,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(forward_demodulation,[],[f37033,f233])).
fof(f37033,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f255,f227])).
fof(f8894,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl106),
  inference(cnf_transformation,[],[f8894_D])).
fof(f8894_D,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f37524,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(forward_demodulation,[],[f37444,f233])).
fof(f37444,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f257,f225])).
fof(f257,plain,(
  ( ! [X4,X5] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X4))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X5))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5)))) )),
  inference(superposition,[],[f199,f222])).
fof(f222,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) )),
  inference(cnf_transformation,[],[f188])).
fof(f188,plain,(
  ! [X0,X1,X2] : ((s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X0) = s('type.integer.int',X1)) & (s('type.integer.int',X0) != s('type.integer.int',X1) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2)))))),
  inference(nnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0,X1,X2] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) <=> s('type.integer.int',X0) = s('type.integer.int',X1))),
  inference(rectify,[],[f57])).
fof(f57,axiom,(
  ! [X6,X7,X8] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X8))) <=> s('type.integer.int',X6) = s('type.integer.int',X7))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_EQ_RADD')).
fof(f8897,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl108),
  inference(cnf_transformation,[],[f8897_D])).
fof(f8897_D,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f37526,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(forward_demodulation,[],[f37445,f233])).
fof(f37445,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f257,f227])).
fof(f42034,plain,(
  ( ! [X70] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X70),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X70),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) ) | $spl496),
  inference(cnf_transformation,[],[f42034_D])).
fof(f42034_D,plain,(
  ( ! [X70] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X70),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X70),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) ) <=> ~$spl496),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl496])])).
fof(f42029,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl487),
  inference(cnf_transformation,[],[f42029_D])).
fof(f42029_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl487),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl487])])).
fof(f38816,plain,(
  ( ! [X70] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X70),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X70),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) ) | $spl69),
  inference(superposition,[],[f281,f14594])).
fof(f14594,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) | $spl69),
  inference(forward_demodulation,[],[f14593,f13410])).
fof(f13410,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | $spl69),
  inference(forward_demodulation,[],[f13409,f11143])).
fof(f11143,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | $spl69),
  inference(resolution,[],[f2596,f238])).
fof(f13409,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | $spl69),
  inference(subsumption_resolution,[],[f13311,f2596])).
fof(f13311,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | $spl69),
  inference(superposition,[],[f238,f11143])).
fof(f14593,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) | $spl69),
  inference(subsumption_resolution,[],[f14495,f2596])).
fof(f14495,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) | $spl69),
  inference(superposition,[],[f238,f13410])).
fof(f281,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X0))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) )),
  inference(superposition,[],[f199,f220])).
fof(f42033,plain,(
  ( ! [X73] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X73),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X73),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) ) | $spl494),
  inference(cnf_transformation,[],[f42033_D])).
fof(f42033_D,plain,(
  ( ! [X73] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X73),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X73),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) ) <=> ~$spl494),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl494])])).
fof(f38819,plain,(
  ( ! [X73] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X73),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X73),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) ) | $spl69),
  inference(superposition,[],[f281,f13410])).
fof(f42032,plain,(
  ( ! [X77] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X77),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X77),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) ) | $spl492),
  inference(cnf_transformation,[],[f42032_D])).
fof(f42032_D,plain,(
  ( ! [X77] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X77),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X77),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) ) <=> ~$spl492),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl492])])).
fof(f38823,plain,(
  ( ! [X77] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X77),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X77),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) ) | $spl69),
  inference(superposition,[],[f281,f11143])).
fof(f42031,plain,(
  ( ! [X70] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X70))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))),s('type.integer.int',X70)))) ) | $spl490),
  inference(cnf_transformation,[],[f42031_D])).
fof(f42031_D,plain,(
  ( ! [X70] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X70))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))),s('type.integer.int',X70)))) ) <=> ~$spl490),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl490])])).
fof(f39402,plain,(
  ( ! [X70] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X70))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))),s('type.integer.int',X70)))) ) | $spl69),
  inference(superposition,[],[f283,f14594])).
fof(f283,plain,(
  ( ! [X4,X5] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X4))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X4))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X5))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',X5)))) )),
  inference(superposition,[],[f199,f222])).
fof(f42030,plain,(
  ( ! [X73] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X73))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))),s('type.integer.int',X73)))) ) | $spl488),
  inference(cnf_transformation,[],[f42030_D])).
fof(f42030_D,plain,(
  ( ! [X73] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X73))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))),s('type.integer.int',X73)))) ) <=> ~$spl488),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl488])])).
fof(f39405,plain,(
  ( ! [X73] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X73))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))),s('type.integer.int',X73)))) ) | $spl69),
  inference(superposition,[],[f283,f13410])).
fof(f42027,plain,(
  ( ! [X77] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X77))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))),s('type.integer.int',X77)))) ) | $spl484),
  inference(cnf_transformation,[],[f42027_D])).
fof(f42027_D,plain,(
  ( ! [X77] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X77))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))),s('type.integer.int',X77)))) ) <=> ~$spl484),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl484])])).
fof(f39409,plain,(
  ( ! [X77] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X77))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))),s('type.integer.int',X77)))) ) | $spl69),
  inference(superposition,[],[f283,f11143])).
fof(f41925,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) ) | $spl380),
  inference(cnf_transformation,[],[f41925_D])).
fof(f41925_D,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) ) <=> ~$spl380),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl380])])).
fof(f418,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | $spl32),
  inference(cnf_transformation,[],[f418_D])).
fof(f418_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f39800,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) ) | ($spl2 | $spl33 | $spl86)),
  inference(superposition,[],[f23954,f16514])).
fof(f16514,plain,(
  ( ! [X47] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X47),s('type.integer.int',X47)))) ) | ($spl2 | $spl33)),
  inference(resolution,[],[f15342,f237])).
fof(f15342,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | ($spl2 | $spl33)),
  inference(forward_demodulation,[],[f15341,f307])).
fof(f307,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl2),
  inference(cnf_transformation,[],[f307_D])).
fof(f307_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f15341,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X6))) ) | ($spl2 | $spl33)),
  inference(subsumption_resolution,[],[f15302,f417])).
fof(f15302,plain,(
  ( ! [X6] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X6))) ) | $spl2),
  inference(superposition,[],[f201,f307])).
fof(f23954,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) ) | ($spl33 | $spl86)),
  inference(resolution,[],[f23848,f201])).
fof(f23848,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | ($spl33 | $spl86)),
  inference(resolution,[],[f7100,f3595])).
fof(f7100,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl86),
  inference(cnf_transformation,[],[f7100_D])).
fof(f7100_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f1176,plain,(
  s('type.integer.int',sK2) = s('type.integer.int',sK3) | $spl48),
  inference(cnf_transformation,[],[f1176_D])).
fof(f1176_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int',sK3) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f41478,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | $spl59),
  inference(resolution,[],[f41346,f235])).
fof(f235,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | s('type.integer.int',X0) = s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int',X1) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X0)))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X7] : (s('type.integer.int',X6) = s('type.integer.int',X7) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X6),s('type.integer.int',X7)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X6)))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_LT_TOTAL')).
fof(f41346,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | $spl59),
  inference(resolution,[],[f1489,f224])).
fof(f224,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2))))) )),
  inference(cnf_transformation,[],[f170])).
fof(f170,plain,(
  ! [X0,X1,X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))))),
  inference(ennf_transformation,[],[f145])).
fof(f145,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X2)))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2)))))))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  ! [X6,X7,X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X8)))) => p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X8)))))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_LT_LADD_IMP')).
fof(f1489,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl59),
  inference(cnf_transformation,[],[f1489_D])).
fof(f1489_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f41477,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | $spl59),
  inference(resolution,[],[f41346,f235])).
fof(f318,plain,(
  s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl15),
  inference(cnf_transformation,[],[f318_D])).
fof(f318_D,plain,(
  s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f42024,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | $spl481),
  inference(cnf_transformation,[],[f42024_D])).
fof(f42024_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) <=> ~$spl481),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl481])])).
fof(f28403,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl234),
  inference(superposition,[],[f27734,f228])).
fof(f228,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f190])).
fof(f27734,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))))) | $spl234),
  inference(superposition,[],[f22778,f231])).
fof(f28404,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl234),
  inference(superposition,[],[f27734,f226])).
fof(f226,plain,(
  ( ! [X0,X1] : (s('type.integer.int',X1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',X0) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(cnf_transformation,[],[f189])).
fof(f42026,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl482),
  inference(cnf_transformation,[],[f42026_D])).
fof(f42026_D,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) <=> ~$spl482),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl482])])).
fof(f42025,plain,(
  ( ! [X12,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl15),
  inference(subsumption_resolution,[],[f38712,f225])).
fof(f38712,plain,(
  ( ! [X12,X13] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl15),
  inference(superposition,[],[f37421,f225])).
fof(f37421,plain,(
  ( ! [X17] : (s('type.integer.int',X17) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | s('type.integer.int',X17) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl15),
  inference(superposition,[],[f30781,f231])).
fof(f30781,plain,(
  ( ! [X2,X3] : (s('type.integer.int',X2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X3),s('type.integer.int',sK3))) | s('type.integer.int',X2) != s('type.integer.int',X3)) ) | $spl15),
  inference(superposition,[],[f30681,f223])).
fof(f223,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X2))) | s('type.integer.int',X0) != s('type.integer.int',X1)) )),
  inference(cnf_transformation,[],[f188])).
fof(f30681,plain,(
  ( ! [X0] : (s('type.integer.int',X0) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK3)))) ) | $spl15),
  inference(equality_resolution,[],[f22974])).
fof(f22974,plain,(
  ( ! [X0,X1] : (s('type.integer.int',sK3) != s('type.integer.int',X0) | s('type.integer.int',X1) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) ) | $spl15),
  inference(superposition,[],[f318,f227])).
fof(f42022,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl478),
  inference(cnf_transformation,[],[f42022_D])).
fof(f42022_D,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) <=> ~$spl478),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl478])])).
fof(f42021,plain,(
  ( ! [X14,X15] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl15),
  inference(subsumption_resolution,[],[f38713,f227])).
fof(f38713,plain,(
  ( ! [X14,X15] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))) | s('type.integer.int',X14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl15),
  inference(superposition,[],[f37421,f227])).
fof(f42017,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))),s('type.integer.int',X13)))) ) | $spl476),
  inference(cnf_transformation,[],[f42017_D])).
fof(f42017_D,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))),s('type.integer.int',X13)))) ) <=> ~$spl476),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl476])])).
fof(f42016,plain,(
  ( ! [X14] : (s('type.integer.int',X14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) | $spl474),
  inference(cnf_transformation,[],[f42016_D])).
fof(f42016_D,plain,(
  ( ! [X14] : (s('type.integer.int',X14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) ) <=> ~$spl474),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl474])])).
fof(f38756,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))),s('type.integer.int',X13)))) ) | $spl15),
  inference(duplicate_literal_removal,[],[f38750])).
fof(f38750,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))),s('type.integer.int',X13)))) ) | $spl15),
  inference(superposition,[],[f37421,f225])).
fof(f42015,plain,(
  ( ! [X15] : (s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3)))))) ) | $spl472),
  inference(cnf_transformation,[],[f42015_D])).
fof(f42015_D,plain,(
  ( ! [X15] : (s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3)))))) ) <=> ~$spl472),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl472])])).
fof(f38755,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3)))))) ) | $spl15),
  inference(duplicate_literal_removal,[],[f38751])).
fof(f38751,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X14) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3)))))) ) | $spl15),
  inference(superposition,[],[f37421,f227])).
fof(f42014,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl470),
  inference(cnf_transformation,[],[f42014_D])).
fof(f42014_D,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) <=> ~$spl470),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl470])])).
fof(f39505,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl69),
  inference(superposition,[],[f11146,f227])).
fof(f42013,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl468),
  inference(cnf_transformation,[],[f42013_D])).
fof(f42013_D,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl468),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl468])])).
fof(f39506,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl69),
  inference(superposition,[],[f11146,f225])).
fof(f1638,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl61),
  inference(cnf_transformation,[],[f1638_D])).
fof(f1638_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f16767,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X0))))) ) | $spl196),
  inference(cnf_transformation,[],[f16767_D])).
fof(f16767_D,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X0))))) ) <=> ~$spl196),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl196])])).
fof(f42012,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl150),
  inference(forward_demodulation,[],[f25395,f233])).
fof(f25395,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl150),
  inference(superposition,[],[f224,f11471])).
fof(f11471,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl150),
  inference(cnf_transformation,[],[f11471_D])).
fof(f11471_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl150),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl150])])).
fof(f16765,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))))) ) | $spl194),
  inference(cnf_transformation,[],[f16765_D])).
fof(f16765_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))))) ) <=> ~$spl194),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl194])])).
fof(f25396,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))))) ) | $spl150),
  inference(superposition,[],[f234,f11471])).
fof(f16764,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl192),
  inference(cnf_transformation,[],[f16764_D])).
fof(f16764_D,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) <=> ~$spl192),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl192])])).
fof(f25397,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl150),
  inference(superposition,[],[f234,f11471])).
fof(f16763,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3))))) ) | $spl190),
  inference(cnf_transformation,[],[f16763_D])).
fof(f16763_D,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3))))) ) <=> ~$spl190),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl190])])).
fof(f25398,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3))))) ) | $spl150),
  inference(superposition,[],[f234,f11471])).
fof(f16759,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl188),
  inference(cnf_transformation,[],[f16759_D])).
fof(f16759_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl188),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl188])])).
fof(f42011,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl8 | $spl150)),
  inference(subsumption_resolution,[],[f25399,f802])).
fof(f802,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl8),
  inference(superposition,[],[f313,f233])).
fof(f25399,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl150),
  inference(superposition,[],[f235,f11471])).
fof(f42010,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl8 | $spl150)),
  inference(subsumption_resolution,[],[f25400,f802])).
fof(f25400,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl150),
  inference(superposition,[],[f235,f11471])).
fof(f16760,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl189),
  inference(cnf_transformation,[],[f16760_D])).
fof(f16760_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl189),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl189])])).
fof(f25401,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl33 | $spl150)),
  inference(superposition,[],[f3595,f11471])).
fof(f22804,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | $spl249),
  inference(cnf_transformation,[],[f22804_D])).
fof(f22804_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) <=> ~$spl249),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl249])])).
fof(f42005,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl30 | $spl150)),
  inference(forward_demodulation,[],[f25410,f11471])).
fof(f25410,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl30 | $spl150)),
  inference(superposition,[],[f330,f11471])).
fof(f22803,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | $spl248),
  inference(cnf_transformation,[],[f22803_D])).
fof(f22803_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) <=> ~$spl248),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl248])])).
fof(f42003,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl33 | $spl150)),
  inference(forward_demodulation,[],[f25416,f11471])).
fof(f25416,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl2 | $spl33 | $spl150)),
  inference(superposition,[],[f15342,f11471])).
fof(f41976,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X5))) ) | $spl432),
  inference(cnf_transformation,[],[f41976_D])).
fof(f41976_D,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X5))) ) <=> ~$spl432),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl432])])).
fof(f25566,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | p(s(bool,X1))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f25354,f201])).
fof(f25354,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl22 | $spl150)),
  inference(backward_demodulation,[],[f11471,f324])).
fof(f324,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl22),
  inference(cnf_transformation,[],[f324_D])).
fof(f324_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f25567,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X2)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f25354,f201])).
fof(f417,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | $spl33),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f18639,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) | $spl222),
  inference(cnf_transformation,[],[f18639_D])).
fof(f18639_D,plain,(
  ( ! [X2] : (~p(s(bool,X2))) ) <=> ~$spl222),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl222])])).
fof(f42002,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f25568,f25550])).
fof(f25550,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) | ($spl22 | $spl30 | $spl150)),
  inference(backward_demodulation,[],[f25354,f330])).
fof(f25568,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X3)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f25354,f204])).
fof(f204,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f167])).
fof(f42001,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X4))) ) | ($spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f25569,f25550])).
fof(f25569,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X4))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f25354,f204])).
fof(f42000,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3))))))) ) | $spl466),
  inference(cnf_transformation,[],[f42000_D])).
fof(f42000_D,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3))))))) ) <=> ~$spl466),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl466])])).
fof(f25571,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK3))))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f224,f25354])).
fof(f41999,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X2))))) ) | $spl464),
  inference(cnf_transformation,[],[f41999_D])).
fof(f41999_D,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X2))))) ) <=> ~$spl464),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl464])])).
fof(f25572,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X2))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f234,f25354])).
fof(f41998,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',sK2))))) ) | $spl462),
  inference(cnf_transformation,[],[f41998_D])).
fof(f41998_D,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',sK2))))) ) <=> ~$spl462),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl462])])).
fof(f25573,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',sK2))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f234,f25354])).
fof(f41997,plain,(
  ( ! [X4] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X4)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X4))))) ) | $spl460),
  inference(cnf_transformation,[],[f41997_D])).
fof(f41997_D,plain,(
  ( ! [X4] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X4)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X4))))) ) <=> ~$spl460),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl460])])).
fof(f25574,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X4)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X4))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f234,f25354])).
fof(f1177,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | $spl51),
  inference(cnf_transformation,[],[f1177_D])).
fof(f1177_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) <=> ~$spl51),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl51])])).
fof(f25577,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | ($spl22 | $spl33 | $spl150)),
  inference(superposition,[],[f3595,f25354])).
fof(f25578,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X5)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f201,f25354])).
fof(f25579,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | p(s(bool,X6))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f201,f25354])).
fof(f41996,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X8))) ) | $spl458),
  inference(cnf_transformation,[],[f41996_D])).
fof(f41996_D,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X8))) ) <=> ~$spl458),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl458])])).
fof(f25580,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X7))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f201,f25354])).
fof(f25581,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X8))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f201,f25354])).
fof(f41995,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X9))) ) | ($spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f25582,f25550])).
fof(f25582,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X9))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f204,f25354])).
fof(f41994,plain,(
  ( ! [X10] : (~p(s(bool,X10)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f25583,f25550])).
fof(f25583,plain,(
  ( ! [X10] : (s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X10)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f204,f25354])).
fof(f41993,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X12))) ) | $spl456),
  inference(cnf_transformation,[],[f41993_D])).
fof(f41993_D,plain,(
  ( ! [X12] : (s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X12))) ) <=> ~$spl456),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl456])])).
fof(f25584,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X11))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f204,f25354])).
fof(f25585,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X12) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X12))) ) | ($spl22 | $spl150)),
  inference(superposition,[],[f204,f25354])).
fof(f22879,plain,(
  ( ! [X33] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X33),s('type.integer.int',X33)))) ) | $spl282),
  inference(cnf_transformation,[],[f22879_D])).
fof(f22879_D,plain,(
  ( ! [X33] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X33),s('type.integer.int',X33)))) ) <=> ~$spl282),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl282])])).
fof(f25685,plain,(
  ( ! [X16] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X16),s('type.integer.int',X16)))) ) | ($spl2 | $spl22 | $spl30 | $spl33 | $spl150)),
  inference(forward_demodulation,[],[f25659,f25354])).
fof(f25659,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X16),s('type.integer.int',X16))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl2 | $spl22 | $spl30 | $spl33 | $spl150)),
  inference(superposition,[],[f25550,f16514])).
fof(f41992,plain,(
  ( ! [X4,X5] : (~p(s(bool,X5)) | ~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f25675,f204])).
fof(f25675,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl22 | $spl30 | $spl150)),
  inference(superposition,[],[f25550,f204])).
fof(f41991,plain,(
  ( ! [X6,X7] : (~p(s(bool,X7)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X6))) ) | ($spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f25676,f204])).
fof(f25676,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X6))) ) | ($spl22 | $spl30 | $spl150)),
  inference(superposition,[],[f25550,f204])).
fof(f588,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl38),
  inference(cnf_transformation,[],[f588_D])).
fof(f588_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f25768,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f25546,f201])).
fof(f25546,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl2 | $spl22 | $spl150)),
  inference(backward_demodulation,[],[f25354,f307])).
fof(f25769,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f25546,f201])).
fof(f41990,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))))) ) | $spl454),
  inference(cnf_transformation,[],[f41990_D])).
fof(f41990_D,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))))) ) <=> ~$spl454),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl454])])).
fof(f25775,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK1))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f224,f25546])).
fof(f41989,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X1))))) ) | $spl452),
  inference(cnf_transformation,[],[f41989_D])).
fof(f41989_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X1))))) ) <=> ~$spl452),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl452])])).
fof(f25776,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X1))))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f234,f25546])).
fof(f41988,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK1))))) ) | $spl450),
  inference(cnf_transformation,[],[f41988_D])).
fof(f41988_D,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK1))))) ) <=> ~$spl450),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl450])])).
fof(f25777,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK1))))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f234,f25546])).
fof(f41987,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X3))))) ) | $spl448),
  inference(cnf_transformation,[],[f41987_D])).
fof(f41987_D,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X3))))) ) <=> ~$spl448),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl448])])).
fof(f25778,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X3))))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f234,f25546])).
fof(f1241,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) | $spl55),
  inference(cnf_transformation,[],[f1241_D])).
fof(f1241_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f25781,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) | ($spl2 | $spl22 | $spl33 | $spl150)),
  inference(superposition,[],[f3595,f25546])).
fof(f25782,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X4)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25546])).
fof(f25783,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X5))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25546])).
fof(f41986,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X7))) ) | $spl446),
  inference(cnf_transformation,[],[f41986_D])).
fof(f41986_D,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X7))) ) <=> ~$spl446),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl446])])).
fof(f25784,plain,(
  ( ! [X6] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X6))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25546])).
fof(f25785,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X7))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25546])).
fof(f41985,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X11))) ) | $spl444),
  inference(cnf_transformation,[],[f41985_D])).
fof(f41985_D,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X11))) ) <=> ~$spl444),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl444])])).
fof(f25788,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X10))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f204,f25546])).
fof(f25789,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X11))) ) | ($spl2 | $spl22 | $spl150)),
  inference(superposition,[],[f204,f25546])).
fof(f41984,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl442),
  inference(cnf_transformation,[],[f41984_D])).
fof(f41984_D,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) <=> ~$spl442),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl442])])).
fof(f25790,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl2 | $spl22 | $spl33 | $spl150)),
  inference(superposition,[],[f3588,f25546])).
fof(f3588,plain,(
  ( ! [X2,X3] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',X3)))) ) | $spl33),
  inference(subsumption_resolution,[],[f3478,f417])).
fof(f3478,plain,(
  ( ! [X2,X3] : (s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',X3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | p(s(bool,X2))) ) | $spl33),
  inference(superposition,[],[f2982,f201])).
fof(f590,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl40),
  inference(cnf_transformation,[],[f590_D])).
fof(f590_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f26136,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f25548,f201])).
fof(f25548,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl10 | $spl22 | $spl150)),
  inference(backward_demodulation,[],[f25354,f314])).
fof(f314,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl10),
  inference(cnf_transformation,[],[f314_D])).
fof(f314_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f26137,plain,(
  ( ! [X1] : (s(bool,X1) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f25548,f201])).
fof(f589,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl41),
  inference(cnf_transformation,[],[f589_D])).
fof(f589_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f41983,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f26138,f25550])).
fof(f26138,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f25548,f204])).
fof(f41982,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl10 | $spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f26139,f25550])).
fof(f26139,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f25548,f204])).
fof(f41981,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))))) ) | $spl440),
  inference(cnf_transformation,[],[f41981_D])).
fof(f41981_D,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))))) ) <=> ~$spl440),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl440])])).
fof(f41980,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(forward_demodulation,[],[f26142,f233])).
fof(f26142,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f224,f25548])).
fof(f41979,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) ) | $spl438),
  inference(cnf_transformation,[],[f41979_D])).
fof(f41979_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) ) <=> ~$spl438),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl438])])).
fof(f26143,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X1))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f234,f25548])).
fof(f41978,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl436),
  inference(cnf_transformation,[],[f41978_D])).
fof(f41978_D,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) <=> ~$spl436),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl436])])).
fof(f26144,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f234,f25548])).
fof(f41977,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X3))))) ) | $spl434),
  inference(cnf_transformation,[],[f41977_D])).
fof(f41977_D,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X3))))) ) <=> ~$spl434),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl434])])).
fof(f26145,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X3))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f234,f25548])).
fof(f18644,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl225),
  inference(cnf_transformation,[],[f18644_D])).
fof(f18644_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl225),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl225])])).
fof(f26148,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl10 | $spl22 | $spl33 | $spl150)),
  inference(superposition,[],[f3595,f25548])).
fof(f26149,plain,(
  ( ! [X4] : (s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X4)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25548])).
fof(f26150,plain,(
  ( ! [X5] : (s(bool,X5) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X5))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25548])).
fof(f41975,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X7))) ) | $spl430),
  inference(cnf_transformation,[],[f41975_D])).
fof(f41975_D,plain,(
  ( ! [X7] : (s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X7))) ) <=> ~$spl430),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl430])])).
fof(f26151,plain,(
  ( ! [X6] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X6))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25548])).
fof(f26152,plain,(
  ( ! [X7] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X7) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | p(s(bool,X7))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f201,f25548])).
fof(f41974,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X8))) ) | ($spl10 | $spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f26153,f25550])).
fof(f26153,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X8))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f204,f25548])).
fof(f41973,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl22 | $spl30 | $spl150)),
  inference(subsumption_resolution,[],[f26154,f25550])).
fof(f26154,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X9)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f204,f25548])).
fof(f41972,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X11))) ) | $spl428),
  inference(cnf_transformation,[],[f41972_D])).
fof(f41972_D,plain,(
  ( ! [X11] : (s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X11))) ) <=> ~$spl428),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl428])])).
fof(f26155,plain,(
  ( ! [X10] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X10) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X10))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f204,f25548])).
fof(f26156,plain,(
  ( ! [X11] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X11) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ~p(s(bool,X11))) ) | ($spl10 | $spl22 | $spl150)),
  inference(superposition,[],[f204,f25548])).
fof(f41971,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl426),
  inference(cnf_transformation,[],[f41971_D])).
fof(f41971_D,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) <=> ~$spl426),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl426])])).
fof(f26157,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl10 | $spl22 | $spl33 | $spl150)),
  inference(superposition,[],[f3588,f25548])).
fof(f22891,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | $spl290),
  inference(cnf_transformation,[],[f22891_D])).
fof(f22891_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl290),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl290])])).
fof(f41970,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl69 | $spl150)),
  inference(forward_demodulation,[],[f26334,f11471])).
fof(f26334,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl69 | $spl150)),
  inference(superposition,[],[f11148,f11471])).
fof(f11148,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl69),
  inference(resolution,[],[f2596,f201])).
fof(f41969,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl424),
  inference(cnf_transformation,[],[f41969_D])).
fof(f41969_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl424),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl424])])).
fof(f26335,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl10 | $spl22 | $spl69 | $spl150)),
  inference(superposition,[],[f11148,f25548])).
fof(f41967,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl422),
  inference(cnf_transformation,[],[f41967_D])).
fof(f41967_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl422),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl422])])).
fof(f26336,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl2 | $spl22 | $spl69 | $spl150)),
  inference(superposition,[],[f11148,f25546])).
fof(f41965,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | $spl420),
  inference(cnf_transformation,[],[f41965_D])).
fof(f41965_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl420),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl420])])).
fof(f26337,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl22 | $spl69 | $spl150)),
  inference(superposition,[],[f11148,f25354])).
fof(f41963,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl418),
  inference(cnf_transformation,[],[f41963_D])).
fof(f41963_D,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl418),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl418])])).
fof(f2597,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | $spl68),
  inference(cnf_transformation,[],[f2597_D])).
fof(f2597_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f38510,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X11))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | ($spl33 | $spl69 | $spl317)),
  inference(superposition,[],[f31983,f11147])).
fof(f11147,plain,(
  ( ! [X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) ) | ($spl33 | $spl69)),
  inference(resolution,[],[f2596,f3588])).
fof(f31983,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl317),
  inference(resolution,[],[f28927,f201])).
fof(f38530,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | ($spl2 | $spl22 | $spl33 | $spl150 | $spl317)),
  inference(forward_demodulation,[],[f38511,f25354])).
fof(f38511,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) ) | ($spl2 | $spl33 | $spl317)),
  inference(superposition,[],[f31983,f16514])).
fof(f41961,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl416),
  inference(cnf_transformation,[],[f41961_D])).
fof(f41961_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl416),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl416])])).
fof(f38516,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl10 | $spl22 | $spl150 | $spl317)),
  inference(superposition,[],[f31983,f25548])).
fof(f41959,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl414),
  inference(cnf_transformation,[],[f41959_D])).
fof(f41959_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl414),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl414])])).
fof(f38517,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl2 | $spl22 | $spl150 | $spl317)),
  inference(superposition,[],[f31983,f25546])).
fof(f41957,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | $spl412),
  inference(cnf_transformation,[],[f41957_D])).
fof(f41957_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl412),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl412])])).
fof(f38518,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl22 | $spl150 | $spl317)),
  inference(superposition,[],[f31983,f25354])).
fof(f41955,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl410),
  inference(cnf_transformation,[],[f41955_D])).
fof(f41955_D,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) <=> ~$spl410),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl410])])).
fof(f38966,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X11))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | ($spl33 | $spl69 | $spl83)),
  inference(superposition,[],[f23818,f11147])).
fof(f23818,plain,(
  ( ! [X2] : (p(s(bool,X2)) | s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl83),
  inference(resolution,[],[f7094,f201])).
fof(f38986,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | ($spl2 | $spl22 | $spl33 | $spl83 | $spl150)),
  inference(forward_demodulation,[],[f38967,f25354])).
fof(f38967,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | ($spl2 | $spl33 | $spl83)),
  inference(superposition,[],[f23818,f16514])).
fof(f37114,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | $spl342),
  inference(cnf_transformation,[],[f37114_D])).
fof(f37114_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) <=> ~$spl342),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl342])])).
fof(f41954,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl83 | $spl150)),
  inference(forward_demodulation,[],[f38970,f11471])).
fof(f38970,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl83 | $spl150)),
  inference(superposition,[],[f23818,f11471])).
fof(f41953,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl408),
  inference(cnf_transformation,[],[f41953_D])).
fof(f41953_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl408),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl408])])).
fof(f38972,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl10 | $spl22 | $spl83 | $spl150)),
  inference(superposition,[],[f23818,f25548])).
fof(f41951,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl406),
  inference(cnf_transformation,[],[f41951_D])).
fof(f41951_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl406),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl406])])).
fof(f38973,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl2 | $spl22 | $spl83 | $spl150)),
  inference(superposition,[],[f23818,f25546])).
fof(f8485,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | $spl102),
  inference(cnf_transformation,[],[f8485_D])).
fof(f8485_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f38974,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | ($spl22 | $spl83 | $spl150)),
  inference(superposition,[],[f23818,f25354])).
fof(f41949,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl404),
  inference(cnf_transformation,[],[f41949_D])).
fof(f41949_D,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl404),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl404])])).
fof(f39101,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl2 | $spl22 | $spl33 | $spl86 | $spl150)),
  inference(forward_demodulation,[],[f39084,f25354])).
fof(f39084,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl2 | $spl33 | $spl86)),
  inference(superposition,[],[f23854,f16514])).
fof(f23854,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl86),
  inference(resolution,[],[f7100,f204])).
fof(f41947,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl402),
  inference(cnf_transformation,[],[f41947_D])).
fof(f41947_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl402),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl402])])).
fof(f39089,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl10 | $spl22 | $spl86 | $spl150)),
  inference(superposition,[],[f23854,f25548])).
fof(f41945,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl400),
  inference(cnf_transformation,[],[f41945_D])).
fof(f41945_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl400),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl400])])).
fof(f39090,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl2 | $spl22 | $spl86 | $spl150)),
  inference(superposition,[],[f23854,f25546])).
fof(f41943,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl398),
  inference(cnf_transformation,[],[f41943_D])).
fof(f41943_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl398),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl398])])).
fof(f39091,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl22 | $spl86 | $spl150)),
  inference(superposition,[],[f23854,f25354])).
fof(f41935,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl390),
  inference(cnf_transformation,[],[f41935_D])).
fof(f41935_D,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) <=> ~$spl390),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl390])])).
fof(f39646,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl2 | $spl15 | $spl22 | $spl33 | $spl83 | $spl150)),
  inference(forward_demodulation,[],[f39629,f25354])).
fof(f39629,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl2 | $spl15 | $spl33 | $spl83)),
  inference(superposition,[],[f23889,f16514])).
fof(f23889,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | ($spl15 | $spl83)),
  inference(resolution,[],[f7137,f204])).
fof(f7137,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl15 | $spl83)),
  inference(subsumption_resolution,[],[f7108,f318])).
fof(f41934,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl388),
  inference(cnf_transformation,[],[f41934_D])).
fof(f41934_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl388),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl388])])).
fof(f41932,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl15 | $spl83 | $spl150)),
  inference(forward_demodulation,[],[f39632,f11471])).
fof(f39632,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl15 | $spl83 | $spl150)),
  inference(superposition,[],[f23889,f11471])).
fof(f41931,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl386),
  inference(cnf_transformation,[],[f41931_D])).
fof(f41931_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl386),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl386])])).
fof(f39634,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl10 | $spl15 | $spl22 | $spl83 | $spl150)),
  inference(superposition,[],[f23889,f25548])).
fof(f41929,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl384),
  inference(cnf_transformation,[],[f41929_D])).
fof(f41929_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl384),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl384])])).
fof(f39635,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl2 | $spl15 | $spl22 | $spl83 | $spl150)),
  inference(superposition,[],[f23889,f25546])).
fof(f41927,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl382),
  inference(cnf_transformation,[],[f41927_D])).
fof(f41927_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl382),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl382])])).
fof(f39636,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl15 | $spl22 | $spl83 | $spl150)),
  inference(superposition,[],[f23889,f25354])).
fof(f39799,plain,(
  ( ! [X11] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X11),s('type.integer.int',X11))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) ) | ($spl33 | $spl69 | $spl86)),
  inference(superposition,[],[f23954,f11147])).
fof(f39819,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',X12))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) ) | ($spl2 | $spl22 | $spl33 | $spl86 | $spl150)),
  inference(forward_demodulation,[],[f39800,f25354])).
fof(f41923,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl378),
  inference(cnf_transformation,[],[f41923_D])).
fof(f41923_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl378),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl378])])).
fof(f39805,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl10 | $spl22 | $spl33 | $spl86 | $spl150)),
  inference(superposition,[],[f23954,f25548])).
fof(f41921,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl376),
  inference(cnf_transformation,[],[f41921_D])).
fof(f41921_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl376),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl376])])).
fof(f39806,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl2 | $spl22 | $spl33 | $spl86 | $spl150)),
  inference(superposition,[],[f23954,f25546])).
fof(f41919,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) | $spl374),
  inference(cnf_transformation,[],[f41919_D])).
fof(f41919_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) <=> ~$spl374),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl374])])).
fof(f39807,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))) | ($spl22 | $spl33 | $spl86 | $spl150)),
  inference(superposition,[],[f23954,f25354])).
fof(f41917,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',X41)))) ) | $spl372),
  inference(cnf_transformation,[],[f41917_D])).
fof(f41917_D,plain,(
  ( ! [X41] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',X41)))) ) <=> ~$spl372),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl372])])).
fof(f41411,plain,(
  ( ! [X41] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',X41)))) ) | ($spl2 | $spl22 | $spl33 | $spl142 | $spl150)),
  inference(forward_demodulation,[],[f19385,f25354])).
fof(f19385,plain,(
  ( ! [X41] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',X41)))) ) | ($spl2 | $spl33 | $spl142)),
  inference(superposition,[],[f13020,f16514])).
fof(f13020,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) ) | $spl142),
  inference(resolution,[],[f11171,f204])).
fof(f12246,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X6)) | p(s(bool,X7))) ) | $spl164),
  inference(cnf_transformation,[],[f12246_D])).
fof(f12246_D,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X6)) | p(s(bool,X7))) ) <=> ~$spl164),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl164])])).
fof(f25674,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | ~p(s(bool,X3)) | p(s(bool,X2)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl22 | $spl30 | $spl150)),
  inference(superposition,[],[f25550,f201])).
fof(f25673,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | ~p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X0))) ) | ($spl22 | $spl30 | $spl150)),
  inference(superposition,[],[f25550,f201])).
fof(f41517,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | ($spl22 | $spl59 | $spl150)),
  inference(forward_demodulation,[],[f41478,f25354])).
fof(f41519,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | ($spl22 | $spl59 | $spl150)),
  inference(forward_demodulation,[],[f41477,f25354])).
fof(f1178,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | $spl50),
  inference(cnf_transformation,[],[f1178_D])).
fof(f1178_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f25576,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | ($spl22 | $spl150)),
  inference(superposition,[],[f235,f25354])).
fof(f25575,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | ($spl22 | $spl150)),
  inference(superposition,[],[f235,f25354])).
fof(f41829,plain,(
  $false | ($spl22 | $spl59 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f41639,f25503])).
fof(f25503,plain,(
  ( ! [X49] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X49),s('type.integer.int',X49)))) ) | ($spl61 | $spl150)),
  inference(resolution,[],[f16642,f237])).
fof(f16642,plain,(
  ( ! [X6] : (p(s(bool,X6)) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) ) | ($spl61 | $spl150)),
  inference(forward_demodulation,[],[f16641,f11471])).
fof(f16641,plain,(
  ( ! [X6] : (s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X6))) ) | ($spl61 | $spl150)),
  inference(subsumption_resolution,[],[f16595,f1638])).
fof(f16595,plain,(
  ( ! [X6] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X6) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | p(s(bool,X6))) ) | $spl150),
  inference(superposition,[],[f201,f11471])).
fof(f41639,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | ($spl22 | $spl59 | $spl61 | $spl150)),
  inference(backward_demodulation,[],[f41565,f25545])).
fof(f25545,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl22 | $spl150)),
  inference(backward_demodulation,[],[f25354,f199])).
fof(f41565,plain,(
  s('type.integer.int',sK2) = s('type.integer.int',sK3) | ($spl22 | $spl59 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f41564,f41346])).
fof(f41564,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | ($spl22 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f25576,f1638])).
fof(f41828,plain,(
  $false | ($spl22 | $spl59 | $spl61 | $spl137 | $spl150)),
  inference(subsumption_resolution,[],[f41609,f233])).
fof(f41609,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | ($spl22 | $spl59 | $spl61 | $spl137 | $spl150)),
  inference(backward_demodulation,[],[f41565,f11009])).
fof(f41827,plain,(
  $false | ($spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f41605,f231])).
fof(f41605,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))),s('type.integer.int',sK2))) | ($spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(backward_demodulation,[],[f41565,f8549])).
fof(f8549,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3))),s('type.integer.int',sK2))) | $spl33),
  inference(superposition,[],[f8507,f231])).
fof(f8507,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',sK3)))) ) | $spl33),
  inference(subsumption_resolution,[],[f8477,f2982])).
fof(f8477,plain,(
  ( ! [X4] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',sK3)))) ) | $spl33),
  inference(superposition,[],[f255,f2982])).
fof(f41824,plain,(
  $false | ($spl10 | $spl20 | $spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f41823,f1638])).
fof(f41823,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl10 | $spl20 | $spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(forward_demodulation,[],[f41594,f11471])).
fof(f41594,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl10 | $spl20 | $spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(backward_demodulation,[],[f41565,f4668])).
fof(f4668,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl10 | $spl20 | $spl33)),
  inference(subsumption_resolution,[],[f4667,f1015])).
fof(f1015,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl20),
  inference(superposition,[],[f323,f233])).
fof(f4667,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | ($spl10 | $spl33)),
  inference(subsumption_resolution,[],[f4621,f417])).
fof(f4621,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl10),
  inference(superposition,[],[f235,f314])).
fof(f41822,plain,(
  $false | ($spl22 | $spl26 | $spl59 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f41589,f237])).
fof(f41589,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl22 | $spl26 | $spl59 | $spl61 | $spl150)),
  inference(backward_demodulation,[],[f41565,f328])).
fof(f41821,plain,(
  $false | ($spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(trivial_inequality_removal,[],[f41604])).
fof(f41604,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X4),s('type.integer.int',sK2)))) ) | ($spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(backward_demodulation,[],[f41565,f8507])).
fof(f41820,plain,(
  $false | ($spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(trivial_inequality_removal,[],[f41606])).
fof(f41606,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X4)))) ) | ($spl22 | $spl33 | $spl59 | $spl61 | $spl150)),
  inference(backward_demodulation,[],[f41565,f8915])).
fof(f8915,plain,(
  ( ! [X4] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X4)))) ) | $spl33),
  inference(subsumption_resolution,[],[f8887,f2982])).
fof(f8887,plain,(
  ( ! [X4] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X4)))) ) | $spl33),
  inference(superposition,[],[f257,f2982])).
fof(f41512,plain,(
  $false | ($spl22 | $spl49 | $spl59 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f41511,f1638])).
fof(f41511,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl49 | $spl59 | $spl150)),
  inference(forward_demodulation,[],[f41510,f25354])).
fof(f41510,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ($spl49 | $spl59)),
  inference(subsumption_resolution,[],[f41478,f1175])).
fof(f41509,plain,(
  $false | ($spl22 | $spl49 | $spl59 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f41508,f1638])).
fof(f41508,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl49 | $spl59 | $spl150)),
  inference(forward_demodulation,[],[f41507,f25354])).
fof(f41507,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ($spl49 | $spl59)),
  inference(subsumption_resolution,[],[f41477,f1175])).
fof(f41506,plain,(
  $false | ($spl50 | $spl59)),
  inference(subsumption_resolution,[],[f41476,f237])).
fof(f41476,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))) | ($spl50 | $spl59)),
  inference(resolution,[],[f41346,f11138])).
fof(f11138,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X8))))) ) | $spl50),
  inference(subsumption_resolution,[],[f11124,f2840])).
fof(f2840,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',sK2))))) ) | $spl50),
  inference(resolution,[],[f1178,f234])).
fof(f11124,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X8)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int',sK3))))) ) | $spl50),
  inference(superposition,[],[f1178,f235])).
fof(f41505,plain,(
  $false | ($spl50 | $spl59)),
  inference(subsumption_resolution,[],[f41474,f237])).
fof(f41474,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK2)))) | ($spl50 | $spl59)),
  inference(resolution,[],[f41346,f11138])).
fof(f41431,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl370),
  inference(cnf_transformation,[],[f41431_D])).
fof(f41431_D,plain,(
  ( ! [X12,X13] : (s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) <=> ~$spl370),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl370])])).
fof(f31968,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl325),
  inference(cnf_transformation,[],[f31968_D])).
fof(f31968_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl325),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl325])])).
fof(f41097,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl33),
  inference(superposition,[],[f9379,f225])).
fof(f41430,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl368),
  inference(cnf_transformation,[],[f41430_D])).
fof(f41430_D,plain,(
  ( ! [X14,X15] : (s('type.integer.int',X14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) <=> ~$spl368),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl368])])).
fof(f41098,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X14) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl33),
  inference(superposition,[],[f9379,f227])).
fof(f41429,plain,(
  ( ! [X40] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))))))) ) | $spl366),
  inference(cnf_transformation,[],[f41429_D])).
fof(f41429_D,plain,(
  ( ! [X40] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))))))) ) <=> ~$spl366),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl366])])).
fof(f41428,plain,(
  ( ! [X40] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40)))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl33),
  inference(forward_demodulation,[],[f41110,f231])).
fof(f41110,plain,(
  ( ! [X40] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))))))) ) | $spl33),
  inference(superposition,[],[f9379,f231])).
fof(f11170,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | $spl143),
  inference(cnf_transformation,[],[f11170_D])).
fof(f11170_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) <=> ~$spl143),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl143])])).
fof(f22702,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | $spl232),
  inference(cnf_transformation,[],[f22702_D])).
fof(f22702_D,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) <=> ~$spl232),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl232])])).
fof(f41427,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | ($spl33 | $spl69)),
  inference(forward_demodulation,[],[f41117,f14594])).
fof(f41117,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | ($spl33 | $spl69)),
  inference(superposition,[],[f9379,f14594])).
fof(f41426,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))))))))) | $spl364),
  inference(cnf_transformation,[],[f41426_D])).
fof(f41426_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))))))))) <=> ~$spl364),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl364])])).
fof(f41424,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl33 | $spl61)),
  inference(forward_demodulation,[],[f41118,f24164])).
fof(f24164,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2)))))))))) | $spl61),
  inference(forward_demodulation,[],[f24163,f23785])).
fof(f23785,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))) | $spl61),
  inference(forward_demodulation,[],[f23784,f23279])).
fof(f23279,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK2)))) | $spl61),
  inference(resolution,[],[f1638,f238])).
fof(f23784,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK2)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))) | $spl61),
  inference(subsumption_resolution,[],[f23664,f1638])).
fof(f23664,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK2)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))) | $spl61),
  inference(superposition,[],[f238,f23279])).
fof(f24163,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2)))))))))) | $spl61),
  inference(subsumption_resolution,[],[f24051,f1638])).
fof(f24051,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2)))))))))) | $spl61),
  inference(superposition,[],[f238,f23785])).
fof(f41118,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2)))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK2))))))))))))) | ($spl33 | $spl61)),
  inference(superposition,[],[f9379,f24164])).
fof(f41423,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) | $spl362),
  inference(cnf_transformation,[],[f41423_D])).
fof(f41423_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) <=> ~$spl362),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl362])])).
fof(f41421,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl33 | $spl317)),
  inference(forward_demodulation,[],[f41119,f32760])).
fof(f32760,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl317),
  inference(forward_demodulation,[],[f32759,f31977])).
fof(f31977,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl317),
  inference(resolution,[],[f28927,f238])).
fof(f32759,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl317),
  inference(subsumption_resolution,[],[f32599,f28927])).
fof(f32599,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) | $spl317),
  inference(superposition,[],[f238,f31977])).
fof(f41119,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))))))) | ($spl33 | $spl317)),
  inference(superposition,[],[f9379,f32760])).
fof(f41415,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl33 | $spl318)),
  inference(forward_demodulation,[],[f41123,f28930])).
fof(f41123,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))))) | ($spl33 | $spl318)),
  inference(superposition,[],[f9379,f28930])).
fof(f41410,plain,(
  $false | ($spl50 | $spl59 | $spl142 | $spl168 | $spl178)),
  inference(subsumption_resolution,[],[f41409,f23333])).
fof(f23333,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl142 | $spl168 | $spl178)),
  inference(equality_resolution,[],[f21384])).
fof(f21384,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,X0))) ) | ($spl142 | $spl168 | $spl178)),
  inference(backward_demodulation,[],[f18145,f12711])).
fof(f18145,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ($spl142 | $spl168)),
  inference(resolution,[],[f12995,f11171])).
fof(f12995,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) ) | $spl168),
  inference(resolution,[],[f12448,f204])).
fof(f41409,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl50 | $spl59 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f41346,f25937])).
fof(f25937,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2))) | ($spl50 | $spl142 | $spl168 | $spl178)),
  inference(resolution,[],[f23351,f1178])).
fof(f23351,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) ) | ($spl142 | $spl168 | $spl178)),
  inference(resolution,[],[f23333,f204])).
fof(f41344,plain,(
  $false | ($spl26 | $spl33 | $spl58)),
  inference(subsumption_resolution,[],[f41281,f328])).
fof(f41281,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ($spl33 | $spl58)),
  inference(resolution,[],[f1490,f3595])).
fof(f39439,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl358),
  inference(cnf_transformation,[],[f39439_D])).
fof(f39439_D,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) <=> ~$spl358),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl358])])).
fof(f38853,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | $spl353),
  inference(cnf_transformation,[],[f38853_D])).
fof(f38853_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) <=> ~$spl353),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl353])])).
fof(f39438,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(forward_demodulation,[],[f39383,f233])).
fof(f39383,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f283,f227])).
fof(f39437,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl356),
  inference(cnf_transformation,[],[f39437_D])).
fof(f39437_D,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) <=> ~$spl356),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl356])])).
fof(f39436,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(forward_demodulation,[],[f39382,f233])).
fof(f39382,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f283,f225])).
fof(f38855,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl354),
  inference(cnf_transformation,[],[f38855_D])).
fof(f38855_D,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) <=> ~$spl354),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl354])])).
fof(f38854,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(forward_demodulation,[],[f38797,f233])).
fof(f38797,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f281,f227])).
fof(f38851,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl350),
  inference(cnf_transformation,[],[f38851_D])).
fof(f38851_D,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) <=> ~$spl350),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl350])])).
fof(f38850,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(forward_demodulation,[],[f38796,f233])).
fof(f38796,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK3))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f281,f225])).
fof(f37130,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl345),
  inference(cnf_transformation,[],[f37130_D])).
fof(f37130_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl345),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl345])])).
fof(f37539,plain,(
  ( ! [X67] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X67))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X67)))) ) | $spl348),
  inference(cnf_transformation,[],[f37539_D])).
fof(f37539_D,plain,(
  ( ! [X67] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X67))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X67)))) ) <=> ~$spl348),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl348])])).
fof(f37543,plain,(
  ( ! [X69] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X69))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X69))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37542,f4492])).
fof(f4492,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))))))))) | $spl33),
  inference(forward_demodulation,[],[f4491,f4313])).
fof(f4313,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))) | $spl33),
  inference(forward_demodulation,[],[f4312,f4143])).
fof(f4143,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))) | $spl33),
  inference(subsumption_resolution,[],[f4062,f237])).
fof(f4062,plain,(
  ( ! [X20] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X20),s('type.integer.int',X20)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))) ) | $spl33),
  inference(superposition,[],[f238,f3477])).
fof(f4312,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))) | $spl33),
  inference(subsumption_resolution,[],[f4217,f237])).
fof(f4217,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))) | $spl33),
  inference(superposition,[],[f238,f4143])).
fof(f4491,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))))))))) | $spl33),
  inference(subsumption_resolution,[],[f4404,f237])).
fof(f4404,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))))))))) | $spl33),
  inference(superposition,[],[f238,f4313])).
fof(f37542,plain,(
  ( ! [X69] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X69))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))))))))),s('type.integer.int',X69)))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37463,f25875])).
fof(f25875,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ($spl142 | $spl168 | $spl178)),
  inference(resolution,[],[f23178,f23333])).
fof(f23178,plain,(
  ( ! [X4] : (~p(s(bool,X4)) | s(bool,X4) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) ) | $spl178),
  inference(resolution,[],[f23164,f204])).
fof(f23164,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | $spl178),
  inference(equality_resolution,[],[f12711])).
fof(f37463,plain,(
  ( ! [X69] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X69))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))))))))),s('type.integer.int',X69)))) ) | $spl33),
  inference(superposition,[],[f257,f4492])).
fof(f37541,plain,(
  ( ! [X68] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X68))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X68))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37540,f4313])).
fof(f37540,plain,(
  ( ! [X68] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X68))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))),s('type.integer.int',X68)))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37462,f25875])).
fof(f37462,plain,(
  ( ! [X68] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X68))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))),s('type.integer.int',X68)))) ) | $spl33),
  inference(superposition,[],[f257,f4313])).
fof(f37538,plain,(
  ( ! [X67] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X67))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X67))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37537,f4143])).
fof(f37537,plain,(
  ( ! [X67] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X67))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))),s('type.integer.int',X67)))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37461,f25875])).
fof(f37461,plain,(
  ( ! [X67] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X67))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))),s('type.integer.int',X67)))) ) | $spl33),
  inference(superposition,[],[f257,f4143])).
fof(f37115,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | $spl343),
  inference(cnf_transformation,[],[f37115_D])).
fof(f37115_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) <=> ~$spl343),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl343])])).
fof(f37527,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | ($spl2 | $spl22 | $spl150)),
  inference(forward_demodulation,[],[f37526,f25546])).
fof(f37525,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | ($spl2 | $spl22 | $spl150)),
  inference(forward_demodulation,[],[f37524,f25546])).
fof(f37131,plain,(
  ( ! [X67] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int',sK2)))) ) | $spl346),
  inference(cnf_transformation,[],[f37131_D])).
fof(f37131_D,plain,(
  ( ! [X67] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int',sK2)))) ) <=> ~$spl346),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl346])])).
fof(f37135,plain,(
  ( ! [X69] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X69),s('type.integer.int','const.integer.int_|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X69),s('type.integer.int',sK2))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37134,f4492])).
fof(f37134,plain,(
  ( ! [X69] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X69),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X69),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))))))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37051,f25875])).
fof(f37051,plain,(
  ( ! [X69] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X69),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X69),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))))))))) ) | $spl33),
  inference(superposition,[],[f255,f4492])).
fof(f37133,plain,(
  ( ! [X68] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X68),s('type.integer.int','const.integer.int_|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X68),s('type.integer.int',sK2))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37132,f4313])).
fof(f37132,plain,(
  ( ! [X68] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X68),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X68),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))))))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37050,f25875])).
fof(f37050,plain,(
  ( ! [X68] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X68),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X68),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0')))))))))) ) | $spl33),
  inference(superposition,[],[f255,f4313])).
fof(f37128,plain,(
  ( ! [X67] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_|48|_0'))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int',sK2))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37127,f4143])).
fof(f37127,plain,(
  ( ! [X67] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))) ) | ($spl33 | $spl142 | $spl168 | $spl178)),
  inference(forward_demodulation,[],[f37049,f25875])).
fof(f37049,plain,(
  ( ! [X67] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X67),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_|48|_0'))))))) ) | $spl33),
  inference(superposition,[],[f255,f4143])).
fof(f37117,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | ($spl2 | $spl22 | $spl150)),
  inference(forward_demodulation,[],[f37116,f25546])).
fof(f37113,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | ($spl2 | $spl22 | $spl150)),
  inference(forward_demodulation,[],[f37112,f25546])).
fof(f35133,plain,(
  $false | ($spl15 | $spl33 | $spl50 | $spl83 | $spl142 | $spl168 | $spl178 | $spl338)),
  inference(subsumption_resolution,[],[f35132,f23333])).
fof(f35132,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl15 | $spl33 | $spl50 | $spl83 | $spl142 | $spl168 | $spl178 | $spl338)),
  inference(forward_demodulation,[],[f35131,f25937])).
fof(f35131,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | ($spl15 | $spl33 | $spl83 | $spl178 | $spl338)),
  inference(subsumption_resolution,[],[f34991,f24331])).
fof(f24331,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK3))))))) ) | ($spl15 | $spl33 | $spl83)),
  inference(resolution,[],[f23913,f3595])).
fof(f23913,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',sK3))),s('type.integer.int',X2))))) ) | ($spl15 | $spl83)),
  inference(forward_demodulation,[],[f23886,f233])).
fof(f23886,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl15 | $spl83)),
  inference(resolution,[],[f7137,f224])).
fof(f34991,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))))) | ($spl178 | $spl338)),
  inference(superposition,[],[f23173,f34825])).
fof(f23173,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X0)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',X0))))) ) | $spl178),
  inference(resolution,[],[f23164,f234])).
fof(f35072,plain,(
  $false | ($spl15 | $spl61 | $spl83 | $spl338)),
  inference(subsumption_resolution,[],[f34863,f1638])).
fof(f34863,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl15 | $spl83 | $spl338)),
  inference(superposition,[],[f23913,f34825])).
fof(f35071,plain,(
  $false | ($spl15 | $spl33 | $spl83 | $spl142 | $spl168 | $spl178 | $spl338)),
  inference(subsumption_resolution,[],[f34862,f23333])).
fof(f34862,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl15 | $spl33 | $spl83 | $spl338)),
  inference(superposition,[],[f24331,f34825])).
fof(f35058,plain,(
  $false | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178 | $spl338)),
  inference(subsumption_resolution,[],[f35057,f1638])).
fof(f35057,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl15 | $spl22 | $spl83 | $spl150 | $spl178 | $spl338)),
  inference(forward_demodulation,[],[f34831,f25354])).
fof(f34831,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ($spl15 | $spl83 | $spl178 | $spl338)),
  inference(backward_demodulation,[],[f34825,f33936])).
fof(f33936,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))),s('type.integer.int',sK3)))) | ($spl15 | $spl83 | $spl178)),
  inference(resolution,[],[f23174,f23913])).
fof(f23174,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',sK3))))) ) | $spl178),
  inference(resolution,[],[f23164,f234])).
fof(f34825,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | $spl338),
  inference(cnf_transformation,[],[f34825_D])).
fof(f34825_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) <=> ~$spl338),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl338])])).
fof(f34827,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) | $spl340),
  inference(cnf_transformation,[],[f34827_D])).
fof(f34827_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) <=> ~$spl340),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl340])])).
fof(f34780,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f34604,f235])).
fof(f34604,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f33936,f25616])).
fof(f25616,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X2))))) ) | ($spl22 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f25572,f1638])).
fof(f34779,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f34604,f235])).
fof(f34434,plain,(
  $false | ($spl15 | $spl33 | $spl50 | $spl83 | $spl142 | $spl168 | $spl178 | $spl334)),
  inference(subsumption_resolution,[],[f34433,f23333])).
fof(f34433,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl15 | $spl33 | $spl50 | $spl83 | $spl142 | $spl168 | $spl178 | $spl334)),
  inference(forward_demodulation,[],[f34432,f25937])).
fof(f34432,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | ($spl15 | $spl33 | $spl83 | $spl178 | $spl334)),
  inference(subsumption_resolution,[],[f34292,f32193])).
fof(f32193,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_|48|_0'))))))) ) | ($spl15 | $spl33 | $spl83 | $spl178)),
  inference(resolution,[],[f32059,f3595])).
fof(f32059,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X2))))) ) | ($spl15 | $spl83 | $spl178)),
  inference(forward_demodulation,[],[f32032,f233])).
fof(f32032,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl15 | $spl83 | $spl178)),
  inference(resolution,[],[f31945,f224])).
fof(f31945,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl15 | $spl83 | $spl178)),
  inference(forward_demodulation,[],[f31944,f231])).
fof(f31944,plain,(
  ( ! [X40] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))))))) ) | ($spl15 | $spl83 | $spl178)),
  inference(subsumption_resolution,[],[f31909,f7137])).
fof(f31909,plain,(
  ( ! [X40] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))))))) ) | $spl178),
  inference(superposition,[],[f23173,f231])).
fof(f34292,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl178 | $spl334)),
  inference(superposition,[],[f23173,f34133])).
fof(f34373,plain,(
  $false | ($spl15 | $spl61 | $spl83 | $spl178 | $spl334)),
  inference(subsumption_resolution,[],[f34164,f1638])).
fof(f34164,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl15 | $spl83 | $spl178 | $spl334)),
  inference(superposition,[],[f32059,f34133])).
fof(f34372,plain,(
  $false | ($spl15 | $spl33 | $spl83 | $spl142 | $spl168 | $spl178 | $spl334)),
  inference(subsumption_resolution,[],[f34163,f23333])).
fof(f34163,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl15 | $spl33 | $spl83 | $spl178 | $spl334)),
  inference(superposition,[],[f32193,f34133])).
fof(f34359,plain,(
  $false | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178 | $spl334)),
  inference(subsumption_resolution,[],[f34358,f1638])).
fof(f34358,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl15 | $spl22 | $spl83 | $spl150 | $spl178 | $spl334)),
  inference(forward_demodulation,[],[f34139,f25354])).
fof(f34139,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ($spl15 | $spl83 | $spl178 | $spl334)),
  inference(backward_demodulation,[],[f34133,f33937])).
fof(f33937,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK3)))) | ($spl15 | $spl83 | $spl178)),
  inference(resolution,[],[f23174,f32059])).
fof(f34133,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | $spl334),
  inference(cnf_transformation,[],[f34133_D])).
fof(f34133_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl334),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl334])])).
fof(f34135,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) | $spl336),
  inference(cnf_transformation,[],[f34135_D])).
fof(f34135_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) <=> ~$spl336),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl336])])).
fof(f34094,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f34001,f235])).
fof(f34001,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f33937,f25616])).
fof(f34093,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f34001,f235])).
fof(f32332,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | $spl330),
  inference(cnf_transformation,[],[f32332_D])).
fof(f32332_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl330),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl330])])).
fof(f32334,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) | $spl332),
  inference(cnf_transformation,[],[f32334_D])).
fof(f32334_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) <=> ~$spl332),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl332])])).
fof(f32287,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f32197,f235])).
fof(f32197,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f32059,f25616])).
fof(f32286,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl15 | $spl22 | $spl61 | $spl83 | $spl150 | $spl178)),
  inference(resolution,[],[f32197,f235])).
fof(f31969,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl324),
  inference(cnf_transformation,[],[f31969_D])).
fof(f31969_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl324),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl324])])).
fof(f31910,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X41)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl178),
  inference(superposition,[],[f23173,f227])).
fof(f31911,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X43)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl178),
  inference(superposition,[],[f23173,f225])).
fof(f31964,plain,(
  $false | ($spl15 | $spl33 | $spl83 | $spl178 | $spl212 | $spl316)),
  inference(subsumption_resolution,[],[f31963,f28956])).
fof(f28956,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl33 | $spl316)),
  inference(resolution,[],[f28928,f3595])).
fof(f31963,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl15 | $spl83 | $spl178 | $spl212)),
  inference(forward_demodulation,[],[f31962,f16809])).
fof(f31962,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl15 | $spl83 | $spl178 | $spl212)),
  inference(subsumption_resolution,[],[f31923,f7137])).
fof(f31923,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | ($spl178 | $spl212)),
  inference(superposition,[],[f23173,f16809])).
fof(f31946,plain,(
  $false | ($spl15 | $spl33 | $spl83 | $spl178 | $spl316)),
  inference(subsumption_resolution,[],[f31945,f28956])).
fof(f31938,plain,(
  $false | ($spl15 | $spl33 | $spl83 | $spl178 | $spl316)),
  inference(subsumption_resolution,[],[f31874,f28956])).
fof(f31874,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl15 | $spl83 | $spl178)),
  inference(resolution,[],[f23173,f7137])).
fof(f31262,plain,(
  $false | ($spl22 | $spl61 | $spl150 | $spl316 | $spl320)),
  inference(subsumption_resolution,[],[f31261,f1638])).
fof(f31261,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl22 | $spl150 | $spl316 | $spl320)),
  inference(forward_demodulation,[],[f31062,f25354])).
fof(f31062,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ($spl316 | $spl320)),
  inference(superposition,[],[f29131,f31018])).
fof(f29131,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))),s('type.integer.int',X2))))) ) | $spl316),
  inference(forward_demodulation,[],[f29096,f233])).
fof(f29096,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl316),
  inference(resolution,[],[f29077,f224])).
fof(f29077,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl316),
  inference(superposition,[],[f28986,f231])).
fof(f28986,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_|48|_0'))))))) ) | $spl316),
  inference(forward_demodulation,[],[f28959,f233])).
fof(f28959,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_|48|_0'))))))) ) | $spl316),
  inference(resolution,[],[f28928,f224])).
fof(f31260,plain,(
  $false | ($spl33 | $spl50 | $spl142 | $spl168 | $spl178 | $spl316 | $spl320)),
  inference(subsumption_resolution,[],[f31259,f23333])).
fof(f31259,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl33 | $spl50 | $spl142 | $spl168 | $spl178 | $spl316 | $spl320)),
  inference(forward_demodulation,[],[f31061,f25937])).
fof(f31061,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | ($spl33 | $spl316 | $spl320)),
  inference(superposition,[],[f29418,f31018])).
fof(f29418,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))))))) ) | ($spl33 | $spl316)),
  inference(resolution,[],[f29131,f3595])).
fof(f31018,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))) | $spl320),
  inference(cnf_transformation,[],[f31018_D])).
fof(f31018_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))) <=> ~$spl320),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl320])])).
fof(f31020,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))),s('type.integer.int',sK2)))) | $spl322),
  inference(cnf_transformation,[],[f31020_D])).
fof(f31020_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))),s('type.integer.int',sK2)))) <=> ~$spl322),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl322])])).
fof(f30965,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))) | ($spl22 | $spl61 | $spl150 | $spl316)),
  inference(resolution,[],[f30801,f235])).
fof(f30801,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0')))))))) | ($spl22 | $spl61 | $spl150 | $spl316)),
  inference(resolution,[],[f25616,f29131])).
fof(f30964,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_|48|_0'))))) | ($spl22 | $spl61 | $spl150 | $spl316)),
  inference(resolution,[],[f30801,f235])).
fof(f28930,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl318),
  inference(cnf_transformation,[],[f28930_D])).
fof(f28930_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl318),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl318])])).
fof(f28926,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl212),
  inference(forward_demodulation,[],[f28780,f16809])).
fof(f28780,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl212),
  inference(superposition,[],[f238,f16809])).
fof(f28696,plain,(
  $false | $spl314),
  inference(trivial_inequality_removal,[],[f28695])).
fof(f28695,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) ) | $spl314),
  inference(forward_demodulation,[],[f28671,f232])).
fof(f232,plain,(
  ( ! [X0] : (s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : s('type.integer.int',X0) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))),
  inference(rectify,[],[f113])).
fof(f113,axiom,(
  ! [X6] : s('type.integer.int',X6) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X6)))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_ADD_LID')).
fof(f28671,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))))) ) | $spl314),
  inference(superposition,[],[f28583,f230])).
fof(f230,plain,(
  ( ! [X0,X1] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))) )),
  inference(cnf_transformation,[],[f149])).
fof(f149,plain,(
  ! [X0,X1] : s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X1)))))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6,X7] : s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X6),s('type.integer.int',X7))))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X6))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X7)))))),
  file('thm.integer.INT_LT_LADD','thm.integer.INT_NEG_ADD')).
fof(f28583,plain,(
  ( ! [X8] : (s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X8)))) ) | $spl314),
  inference(subsumption_resolution,[],[f28562,f237])).
fof(f28562,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X8)))) ) | $spl314),
  inference(superposition,[],[f28248,f225])).
fof(f28689,plain,(
  $false | $spl314),
  inference(trivial_inequality_removal,[],[f28672])).
fof(f28672,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl314),
  inference(superposition,[],[f28583,f231])).
fof(f16809,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl212),
  inference(cnf_transformation,[],[f16809_D])).
fof(f16809_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl212),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl212])])).
fof(f28248,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl314),
  inference(cnf_transformation,[],[f28248_D])).
fof(f28248_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl314),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl314])])).
fof(f28217,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl133),
  inference(resolution,[],[f11003,f235])).
fof(f11003,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl133),
  inference(cnf_transformation,[],[f11003_D])).
fof(f11003_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl133),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl133])])).
fof(f28216,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl133),
  inference(resolution,[],[f11003,f235])).
fof(f26536,plain,(
  $false | ($spl22 | $spl150 | $spl249)),
  inference(subsumption_resolution,[],[f22804,f25354])).
fof(f22881,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) ) | $spl284),
  inference(cnf_transformation,[],[f22881_D])).
fof(f22881_D,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) ) <=> ~$spl284),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl284])])).
fof(f26094,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl173 | $spl178)),
  inference(subsumption_resolution,[],[f26088,f21384])).
fof(f26088,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl168 | $spl173)),
  inference(superposition,[],[f21383,f204])).
fof(f21383,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl142 | $spl168 | $spl173)),
  inference(backward_demodulation,[],[f18145,f12706])).
fof(f26093,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl173 | $spl178)),
  inference(subsumption_resolution,[],[f26087,f21384])).
fof(f26087,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl173)),
  inference(superposition,[],[f21383,f204])).
fof(f22918,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) ) | $spl310),
  inference(cnf_transformation,[],[f22918_D])).
fof(f22918_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) ) <=> ~$spl310),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl310])])).
fof(f26043,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl173 | $spl178)),
  inference(subsumption_resolution,[],[f26037,f12711])).
fof(f26037,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | $spl173),
  inference(superposition,[],[f12706,f204])).
fof(f26042,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl173 | $spl178)),
  inference(subsumption_resolution,[],[f26036,f12711])).
fof(f26036,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl173),
  inference(superposition,[],[f12706,f204])).
fof(f592,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | $spl42),
  inference(cnf_transformation,[],[f592_D])).
fof(f592_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f1150,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | $spl46),
  inference(cnf_transformation,[],[f1150_D])).
fof(f1150_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f23949,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl17 | $spl30)),
  inference(subsumption_resolution,[],[f23945,f330])).
fof(f23945,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | $spl17),
  inference(superposition,[],[f320,f201])).
fof(f23948,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl17 | $spl30)),
  inference(subsumption_resolution,[],[f23944,f330])).
fof(f23944,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl17),
  inference(superposition,[],[f320,f201])).
fof(f23533,plain,(
  $false | ($spl26 | $spl304)),
  inference(resolution,[],[f23506,f328])).
fof(f23506,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl304),
  inference(equality_resolution,[],[f22913])).
fof(f23534,plain,(
  $false | ($spl154 | $spl304)),
  inference(resolution,[],[f23506,f11475])).
fof(f23535,plain,(
  $false | ($spl176 | $spl304)),
  inference(resolution,[],[f23506,f12710])).
fof(f23536,plain,(
  $false | ($spl10 | $spl20 | $spl33 | $spl304)),
  inference(resolution,[],[f23506,f4668])).
fof(f23537,plain,(
  $false | ($spl54 | $spl304)),
  inference(resolution,[],[f23506,f1242])).
fof(f23538,plain,(
  $false | ($spl178 | $spl304)),
  inference(resolution,[],[f23506,f23164])).
fof(f23539,plain,(
  $false | ($spl142 | $spl168 | $spl178 | $spl304)),
  inference(resolution,[],[f23506,f23333])).
fof(f23540,plain,(
  $false | ($spl78 | $spl304)),
  inference(resolution,[],[f23506,f5431])).
fof(f23541,plain,(
  $false | ($spl50 | $spl304)),
  inference(resolution,[],[f23506,f1178])).
fof(f587,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl39),
  inference(cnf_transformation,[],[f587_D])).
fof(f587_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f22797,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl168 | $spl171)),
  inference(subsumption_resolution,[],[f21515,f12995])).
fof(f21515,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl168 | $spl171)),
  inference(superposition,[],[f21382,f204])).
fof(f21382,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | ($spl142 | $spl168 | $spl171)),
  inference(backward_demodulation,[],[f18145,f12704])).
fof(f22810,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl171)),
  inference(subsumption_resolution,[],[f22809,f12995])).
fof(f22809,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl171)),
  inference(forward_demodulation,[],[f16078,f18145])).
fof(f22808,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl168 | $spl171)),
  inference(subsumption_resolution,[],[f22807,f12995])).
fof(f22807,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl168 | $spl171)),
  inference(forward_demodulation,[],[f16079,f18145])).
fof(f22865,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl173)),
  inference(subsumption_resolution,[],[f21622,f12995])).
fof(f21622,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl173)),
  inference(superposition,[],[f21383,f204])).
fof(f22864,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl168 | $spl173)),
  inference(subsumption_resolution,[],[f21623,f12995])).
fof(f21623,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl168 | $spl173)),
  inference(superposition,[],[f21383,f204])).
fof(f22798,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl171)),
  inference(subsumption_resolution,[],[f21514,f12995])).
fof(f21514,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl171)),
  inference(superposition,[],[f21382,f204])).
fof(f23278,plain,(
  $false | ($spl8 | $spl33 | $spl44 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f16636,f9387])).
fof(f9387,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl33 | $spl44)),
  inference(resolution,[],[f3595,f594])).
fof(f16636,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl8 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f16635,f802])).
fof(f16635,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | ($spl61 | $spl150)),
  inference(subsumption_resolution,[],[f16590,f1638])).
fof(f23277,plain,(
  $false | ($spl8 | $spl33 | $spl44 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f16638,f9387])).
fof(f16638,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl8 | $spl61 | $spl150)),
  inference(subsumption_resolution,[],[f16637,f802])).
fof(f16637,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | ($spl61 | $spl150)),
  inference(subsumption_resolution,[],[f16591,f1638])).
fof(f23276,plain,(
  $false | ($spl33 | $spl44 | $spl188)),
  inference(subsumption_resolution,[],[f16759,f9387])).
fof(f23275,plain,(
  $false | ($spl50 | $spl60 | $spl69)),
  inference(subsumption_resolution,[],[f11142,f1639])).
fof(f11142,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl50 | $spl69)),
  inference(resolution,[],[f2596,f2839])).
fof(f2839,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X0))))) ) | $spl50),
  inference(resolution,[],[f1178,f234])).
fof(f23198,plain,(
  $false | ($spl60 | $spl158)),
  inference(subsumption_resolution,[],[f1639,f16870])).
fof(f16870,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl158),
  inference(equality_resolution,[],[f11477])).
fof(f320,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | $spl17),
  inference(cnf_transformation,[],[f320_D])).
fof(f320_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f22954,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f228])).
fof(f14200,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl26),
  inference(superposition,[],[f328,f228])).
fof(f12863,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f228])).
fof(f12378,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f228])).
fof(f11094,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f228])).
fof(f22852,plain,(
  ( ! [X41,X42] : (s('type.integer.int',sK3) != s('type.integer.int',X41) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl270),
  inference(cnf_transformation,[],[f22852_D])).
fof(f22852_D,plain,(
  ( ! [X41,X42] : (s('type.integer.int',sK3) != s('type.integer.int',X41) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) <=> ~$spl270),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl270])])).
fof(f22850,plain,(
  ( ! [X41,X42] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int',X41) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) )),
  inference(forward_demodulation,[],[f6602,f233])).
fof(f6602,plain,(
  ( ! [X41,X42] : (s('type.integer.int',sK3) != s('type.integer.int',X41) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) )),
  inference(superposition,[],[f291,f227])).
fof(f291,plain,(
  ( ! [X2] : (s('type.integer.int',sK3) != s('type.integer.int',X2) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X2)))))) )),
  inference(superposition,[],[f199,f221])).
fof(f221,plain,(
  ( ! [X2,X0,X1] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X2))) | s('type.integer.int',X1) != s('type.integer.int',X2)) )),
  inference(cnf_transformation,[],[f187])).
fof(f22854,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f22853,f233])).
fof(f22853,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(forward_demodulation,[],[f6600,f231])).
fof(f6600,plain,(
  ( ! [X39] : (s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X39))),s('type.integer.int',X39)))))))) )),
  inference(superposition,[],[f291,f231])).
fof(f3267,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl26),
  inference(superposition,[],[f328,f228])).
fof(f2676,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f228])).
fof(f22895,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | $spl293),
  inference(cnf_transformation,[],[f22895_D])).
fof(f22895_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl293),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl293])])).
fof(f22921,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl146 | $spl178)),
  inference(forward_demodulation,[],[f18692,f11468])).
fof(f11468,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | $spl146),
  inference(cnf_transformation,[],[f11468_D])).
fof(f11468_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) <=> ~$spl146),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl146])])).
fof(f18692,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | ($spl146 | $spl178)),
  inference(superposition,[],[f12711,f11468])).
fof(f12711,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,X0))) ) | $spl178),
  inference(cnf_transformation,[],[f12711_D])).
fof(f12711_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,X0))) ) <=> ~$spl178),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl178])])).
fof(f16076,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | $spl171),
  inference(superposition,[],[f12704,f201])).
fof(f16077,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl171),
  inference(superposition,[],[f12704,f201])).
fof(f22920,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | $spl313),
  inference(cnf_transformation,[],[f22920_D])).
fof(f22920_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) <=> ~$spl313),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl313])])).
fof(f18518,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl173),
  inference(superposition,[],[f12706,f228])).
fof(f12712,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X2))) ) | $spl180),
  inference(cnf_transformation,[],[f12712_D])).
fof(f12712_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X2))) ) <=> ~$spl180),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl180])])).
fof(f16079,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | $spl171),
  inference(superposition,[],[f12704,f204])).
fof(f16078,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl171),
  inference(superposition,[],[f12704,f204])).
fof(f19075,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl2 | $spl171 | $spl178)),
  inference(subsumption_resolution,[],[f19072,f12711])).
fof(f19072,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X3))) ) | ($spl2 | $spl171)),
  inference(superposition,[],[f16075,f204])).
fof(f16075,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl2 | $spl171)),
  inference(superposition,[],[f12704,f307])).
fof(f19073,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl2 | $spl171 | $spl178)),
  inference(subsumption_resolution,[],[f19071,f12711])).
fof(f19071,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl2 | $spl171)),
  inference(superposition,[],[f16075,f204])).
fof(f22917,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK2)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl308),
  inference(cnf_transformation,[],[f22917_D])).
fof(f22917_D,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK2)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) <=> ~$spl308),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl308])])).
fof(f7094,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) | $spl83),
  inference(cnf_transformation,[],[f7094_D])).
fof(f7094_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) <=> ~$spl83),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl83])])).
fof(f9277,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X12),s('type.integer.int',sK2)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl50),
  inference(superposition,[],[f2840,f225])).
fof(f22916,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X14),s('type.integer.int',sK2)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl306),
  inference(cnf_transformation,[],[f22916_D])).
fof(f22916_D,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X14),s('type.integer.int',sK2)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) <=> ~$spl306),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl306])])).
fof(f9278,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X14),s('type.integer.int',sK2)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl50),
  inference(superposition,[],[f2840,f227])).
fof(f22698,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | $spl228),
  inference(cnf_transformation,[],[f22698_D])).
fof(f22698_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) <=> ~$spl228),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl228])])).
fof(f22915,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) | $spl50),
  inference(forward_demodulation,[],[f9290,f231])).
fof(f9290,plain,(
  ( ! [X40] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))),s('type.integer.int',sK2))))) ) | $spl50),
  inference(superposition,[],[f2840,f231])).
fof(f22913,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X6))) ) | $spl304),
  inference(cnf_transformation,[],[f22913_D])).
fof(f22913_D,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X6))) ) <=> ~$spl304),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl304])])).
fof(f22914,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl142 | $spl158 | $spl178)),
  inference(subsumption_resolution,[],[f16868,f13616])).
fof(f13616,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,X4))) ) | ($spl142 | $spl178)),
  inference(subsumption_resolution,[],[f13608,f11171])).
fof(f13608,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))))) ) | $spl178),
  inference(superposition,[],[f12711,f204])).
fof(f16868,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | ~p(s(bool,X5)) | ~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl158),
  inference(superposition,[],[f11477,f204])).
fof(f22912,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X6))) ) | ($spl142 | $spl158 | $spl178)),
  inference(subsumption_resolution,[],[f16869,f13616])).
fof(f16869,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | ~p(s(bool,X7)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X6))) ) | $spl158),
  inference(superposition,[],[f11477,f204])).
fof(f22911,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK1))))))) ) | $spl302),
  inference(cnf_transformation,[],[f22911_D])).
fof(f22911_D,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK1))))))) ) <=> ~$spl302),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl302])])).
fof(f17598,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK1))))))) ) | $spl146),
  inference(superposition,[],[f224,f11468])).
fof(f22910,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))))) ) | $spl300),
  inference(cnf_transformation,[],[f22910_D])).
fof(f22910_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))))) ) <=> ~$spl300),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl300])])).
fof(f17599,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))))) ) | $spl146),
  inference(superposition,[],[f234,f11468])).
fof(f22909,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl298),
  inference(cnf_transformation,[],[f22909_D])).
fof(f22909_D,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) <=> ~$spl298),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl298])])).
fof(f17600,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int',sK1)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl146),
  inference(superposition,[],[f234,f11468])).
fof(f22908,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X3))))) ) | $spl296),
  inference(cnf_transformation,[],[f22908_D])).
fof(f22908_D,plain,(
  ( ! [X3] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X3))))) ) <=> ~$spl296),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl296])])).
fof(f17601,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',X3))))) ) | $spl146),
  inference(superposition,[],[f234,f11468])).
fof(f22887,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl287),
  inference(cnf_transformation,[],[f22887_D])).
fof(f22887_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl287),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl287])])).
fof(f17604,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl33 | $spl146)),
  inference(superposition,[],[f3595,f11468])).
fof(f22903,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl30 | $spl146)),
  inference(forward_demodulation,[],[f17613,f11468])).
fof(f17613,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | ($spl30 | $spl146)),
  inference(superposition,[],[f330,f11468])).
fof(f22901,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl2 | $spl33 | $spl146)),
  inference(forward_demodulation,[],[f17619,f11468])).
fof(f17619,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | ($spl2 | $spl33 | $spl146)),
  inference(superposition,[],[f15342,f11468])).
fof(f22900,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2))) | $spl294),
  inference(cnf_transformation,[],[f22900_D])).
fof(f22900_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2))) <=> ~$spl294),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl294])])).
fof(f22898,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl50 | $spl146)),
  inference(forward_demodulation,[],[f17919,f11468])).
fof(f17919,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | ($spl50 | $spl146)),
  inference(superposition,[],[f2842,f11468])).
fof(f2842,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) ) | $spl50),
  inference(resolution,[],[f1178,f204])).
fof(f22897,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl69 | $spl146)),
  inference(forward_demodulation,[],[f18100,f11468])).
fof(f18100,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | ($spl69 | $spl146)),
  inference(superposition,[],[f11148,f11468])).
fof(f22896,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | $spl292),
  inference(cnf_transformation,[],[f22896_D])).
fof(f22896_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl292),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl292])])).
fof(f22894,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl142 | $spl146)),
  inference(forward_demodulation,[],[f18260,f11468])).
fof(f18260,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | ($spl142 | $spl146)),
  inference(superposition,[],[f13020,f11468])).
fof(f22892,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | $spl291),
  inference(cnf_transformation,[],[f22892_D])).
fof(f22892_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) <=> ~$spl291),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl291])])).
fof(f22703,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | ($spl33 | $spl69 | $spl142 | $spl146 | $spl178)),
  inference(forward_demodulation,[],[f20696,f11468])).
fof(f20696,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | ($spl33 | $spl69 | $spl142 | $spl146 | $spl178)),
  inference(superposition,[],[f20685,f11468])).
fof(f20685,plain,(
  ( ! [X48] : (s(bool,X48) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X48))) ) | ($spl33 | $spl69 | $spl142 | $spl178)),
  inference(subsumption_resolution,[],[f20654,f237])).
fof(f20654,plain,(
  ( ! [X47,X48] : (s(bool,X48) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X47),s('type.integer.int',X47)))) | ~p(s(bool,X48))) ) | ($spl33 | $spl69 | $spl142 | $spl178)),
  inference(superposition,[],[f13616,f11147])).
fof(f22890,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) ) | $spl288),
  inference(cnf_transformation,[],[f22890_D])).
fof(f22890_D,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) ) <=> ~$spl288),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl288])])).
fof(f22723,plain,(
  ( ! [X21] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) ) | ($spl142 | $spl146 | $spl158 | $spl178)),
  inference(forward_demodulation,[],[f22722,f11468])).
fof(f22722,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl142 | $spl146 | $spl158 | $spl178)),
  inference(subsumption_resolution,[],[f18661,f11477])).
fof(f18661,plain,(
  ( ! [X21] : (s(bool,X21) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X21)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl142 | $spl146 | $spl178)),
  inference(superposition,[],[f13616,f11468])).
fof(f22742,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X9))) ) | ($spl146 | $spl158)),
  inference(forward_demodulation,[],[f22741,f11468])).
fof(f22741,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl146 | $spl158)),
  inference(subsumption_resolution,[],[f17610,f11477])).
fof(f17610,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X9)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | $spl146),
  inference(superposition,[],[f204,f11468])).
fof(f22744,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X8))) ) | ($spl146 | $spl158)),
  inference(forward_demodulation,[],[f22743,f11468])).
fof(f22743,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | ~p(s(bool,X8))) ) | ($spl146 | $spl158)),
  inference(subsumption_resolution,[],[f17609,f11477])).
fof(f17609,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | ~p(s(bool,X8))) ) | $spl146),
  inference(superposition,[],[f204,f11468])).
fof(f22888,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl286),
  inference(cnf_transformation,[],[f22888_D])).
fof(f22888_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl286),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl286])])).
fof(f22889,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl134 | $spl146)),
  inference(subsumption_resolution,[],[f17603,f11007])).
fof(f17603,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl146),
  inference(superposition,[],[f235,f11468])).
fof(f22886,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl134 | $spl146)),
  inference(subsumption_resolution,[],[f17602,f11007])).
fof(f17602,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl146),
  inference(superposition,[],[f235,f11468])).
fof(f22752,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X3))) ) | ($spl146 | $spl158)),
  inference(forward_demodulation,[],[f22751,f11468])).
fof(f22751,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | ~p(s(bool,X3))) ) | ($spl146 | $spl158)),
  inference(subsumption_resolution,[],[f17596,f11477])).
fof(f17596,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | ~p(s(bool,X3))) ) | $spl146),
  inference(superposition,[],[f11468,f204])).
fof(f22754,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X2))) ) | ($spl146 | $spl158)),
  inference(forward_demodulation,[],[f22753,f11468])).
fof(f22753,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl146 | $spl158)),
  inference(subsumption_resolution,[],[f17595,f11477])).
fof(f17595,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | $spl146),
  inference(superposition,[],[f11468,f204])).
fof(f22831,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl26 | $spl146)),
  inference(forward_demodulation,[],[f14200,f11468])).
fof(f22858,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl26 | $spl146)),
  inference(forward_demodulation,[],[f3267,f11468])).
fof(f22884,plain,(
  $false | ($spl42 | $spl61 | $spl146)),
  inference(subsumption_resolution,[],[f22883,f1638])).
fof(f22883,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl42 | $spl146)),
  inference(forward_demodulation,[],[f592,f11468])).
fof(f16287,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X3))) ) | ($spl10 | $spl30)),
  inference(forward_demodulation,[],[f16286,f314])).
fof(f16286,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f16238,f330])).
fof(f16238,plain,(
  ( ! [X3] : (s(bool,X3) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | $spl10),
  inference(superposition,[],[f314,f204])).
fof(f16285,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X2))) ) | ($spl10 | $spl30)),
  inference(forward_demodulation,[],[f16284,f314])).
fof(f16284,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f16237,f330])).
fof(f16237,plain,(
  ( ! [X2] : (s(bool,X2) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl10),
  inference(superposition,[],[f314,f204])).
fof(f16302,plain,(
  ( ! [X9] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X9))) ) | ($spl10 | $spl30)),
  inference(forward_demodulation,[],[f16301,f314])).
fof(f16301,plain,(
  ( ! [X9] : (~p(s(bool,X9)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f16251,f330])).
fof(f16251,plain,(
  ( ! [X9] : (s(bool,X9) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X9)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl10),
  inference(superposition,[],[f204,f314])).
fof(f16300,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X8))) ) | ($spl10 | $spl30)),
  inference(forward_demodulation,[],[f16299,f314])).
fof(f16299,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X8))) ) | ($spl10 | $spl30)),
  inference(subsumption_resolution,[],[f16250,f330])).
fof(f16250,plain,(
  ( ! [X8] : (s(bool,X8) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X8))) ) | $spl10),
  inference(superposition,[],[f204,f314])).
fof(f16306,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | ($spl10 | $spl30 | $spl142 | $spl178)),
  inference(forward_demodulation,[],[f16305,f314])).
fof(f16305,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl30 | $spl142 | $spl178)),
  inference(subsumption_resolution,[],[f16257,f330])).
fof(f16257,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X13)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl10 | $spl142 | $spl178)),
  inference(superposition,[],[f13616,f314])).
fof(f15352,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | ($spl2 | $spl30 | $spl142 | $spl178)),
  inference(forward_demodulation,[],[f15351,f307])).
fof(f15351,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl2 | $spl30 | $spl142 | $spl178)),
  inference(subsumption_resolution,[],[f15311,f330])).
fof(f15311,plain,(
  ( ! [X13] : (s(bool,X13) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X13)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl2 | $spl142 | $spl178)),
  inference(superposition,[],[f13616,f307])).
fof(f22882,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl2 | $spl142 | $spl168 | $spl171 | $spl178)),
  inference(forward_demodulation,[],[f19073,f18145])).
fof(f22880,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl2 | $spl142 | $spl168 | $spl171 | $spl178)),
  inference(forward_demodulation,[],[f19075,f18145])).
fof(f22800,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,X0))) ) | $spl246),
  inference(cnf_transformation,[],[f22800_D])).
fof(f22800_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,X0))) ) <=> ~$spl246),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl246])])).
fof(f20739,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | p(s(bool,X1))) ) | ($spl30 | $spl168)),
  inference(superposition,[],[f12994,f201])).
fof(f12994,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl30 | $spl168)),
  inference(resolution,[],[f12448,f330])).
fof(f20740,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,X2)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl30 | $spl168)),
  inference(superposition,[],[f12994,f201])).
fof(f20747,plain,(
  ( ! [X3] : (~p(s(bool,X3)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl30 | $spl168)),
  inference(subsumption_resolution,[],[f20741,f12995])).
fof(f20741,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,X3)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | ($spl30 | $spl168)),
  inference(superposition,[],[f12994,f204])).
fof(f20749,plain,(
  ( ! [X4] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X4))) ) | ($spl30 | $spl168)),
  inference(subsumption_resolution,[],[f20742,f12995])).
fof(f20742,plain,(
  ( ! [X4] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X4))) ) | ($spl30 | $spl168)),
  inference(superposition,[],[f12994,f204])).
fof(f21512,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | ($spl142 | $spl168 | $spl171)),
  inference(superposition,[],[f21382,f201])).
fof(f21513,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl171)),
  inference(superposition,[],[f21382,f201])).
fof(f21620,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | ($spl142 | $spl168 | $spl173)),
  inference(superposition,[],[f21383,f201])).
fof(f21621,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl173)),
  inference(superposition,[],[f21383,f201])).
fof(f22861,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl146),
  inference(forward_demodulation,[],[f2676,f11468])).
fof(f22814,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | $spl255),
  inference(cnf_transformation,[],[f22814_D])).
fof(f22814_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) <=> ~$spl255),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl255])])).
fof(f5410,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl7),
  inference(superposition,[],[f312,f228])).
fof(f22819,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1)))) | $spl258),
  inference(cnf_transformation,[],[f22819_D])).
fof(f22819_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1)))) <=> ~$spl258),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl258])])).
fof(f5458,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl78),
  inference(superposition,[],[f5431,f228])).
fof(f22857,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X12))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl274),
  inference(cnf_transformation,[],[f22857_D])).
fof(f22857_D,plain,(
  ( ! [X12,X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X12))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) <=> ~$spl274),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl274])])).
fof(f6588,plain,(
  ( ! [X12,X13] : (s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X12))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) )),
  inference(superposition,[],[f291,f225])).
fof(f22856,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X14))))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl272),
  inference(cnf_transformation,[],[f22856_D])).
fof(f22856_D,plain,(
  ( ! [X14,X15] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X14))))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) <=> ~$spl272),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl272])])).
fof(f6589,plain,(
  ( ! [X14,X15] : (s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X14))))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) )),
  inference(superposition,[],[f291,f227])).
fof(f22855,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl146),
  inference(forward_demodulation,[],[f22854,f11468])).
fof(f22851,plain,(
  ( ! [X41,X42] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int',X41) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl146),
  inference(forward_demodulation,[],[f22850,f11468])).
fof(f317,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl14),
  inference(cnf_transformation,[],[f317_D])).
fof(f317_D,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f22834,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl260),
  inference(cnf_transformation,[],[f22834_D])).
fof(f22834_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl260),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl260])])).
fof(f7108,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl83),
  inference(resolution,[],[f7094,f235])).
fof(f7109,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl83),
  inference(resolution,[],[f7094,f235])).
fof(f22806,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))) | $spl251),
  inference(cnf_transformation,[],[f22806_D])).
fof(f22806_D,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))) <=> ~$spl251),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl251])])).
fof(f8393,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl97),
  inference(superposition,[],[f8099,f226])).
fof(f8099,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))) | $spl97),
  inference(cnf_transformation,[],[f8099_D])).
fof(f8099_D,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))) <=> ~$spl97),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl97])])).
fof(f22849,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) ) | $spl268),
  inference(cnf_transformation,[],[f22849_D])).
fof(f22849_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0)))) ) <=> ~$spl268),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl268])])).
fof(f8461,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f255,f228])).
fof(f22848,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) ) | $spl266),
  inference(cnf_transformation,[],[f22848_D])).
fof(f22848_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1)))) ) <=> ~$spl266),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl266])])).
fof(f8871,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X0))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f257,f228])).
fof(f7092,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl80),
  inference(cnf_transformation,[],[f7092_D])).
fof(f7092_D,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f9135,plain,(
  ( ! [X12,X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl50),
  inference(superposition,[],[f2839,f225])).
fof(f7095,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl84),
  inference(cnf_transformation,[],[f7095_D])).
fof(f7095_D,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f9136,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl50),
  inference(superposition,[],[f2839,f227])).
fof(f7101,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl87),
  inference(cnf_transformation,[],[f7101_D])).
fof(f7101_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f22847,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl50),
  inference(forward_demodulation,[],[f9148,f231])).
fof(f9148,plain,(
  ( ! [X40] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X40))),s('type.integer.int',X40))))))) ) | $spl50),
  inference(superposition,[],[f2839,f231])).
fof(f22846,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl264),
  inference(cnf_transformation,[],[f22846_D])).
fof(f22846_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) <=> ~$spl264),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl264])])).
fof(f9217,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f267,f228])).
fof(f267,plain,(
  ( ! [X4,X5] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X4),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X5),s('type.integer.int',X4))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X5),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) )),
  inference(superposition,[],[f199,f220])).
fof(f22812,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))) | $spl253),
  inference(cnf_transformation,[],[f22812_D])).
fof(f22812_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))) <=> ~$spl253),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl253])])).
fof(f9544,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl33 | $spl78)),
  inference(superposition,[],[f9397,f228])).
fof(f9397,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) | ($spl33 | $spl78)),
  inference(resolution,[],[f3595,f5431])).
fof(f22845,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1)))) ) | $spl262),
  inference(cnf_transformation,[],[f22845_D])).
fof(f22845_D,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1)))) ) <=> ~$spl262),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl262])])).
fof(f9740,plain,(
  ( ! [X0,X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',X1))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))) )),
  inference(superposition,[],[f269,f228])).
fof(f269,plain,(
  ( ! [X8,X9] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X8),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X9)))) )),
  inference(superposition,[],[f199,f222])).
fof(f22844,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl146),
  inference(forward_demodulation,[],[f11094,f11468])).
fof(f22843,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl146),
  inference(forward_demodulation,[],[f12378,f11468])).
fof(f22842,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl146),
  inference(forward_demodulation,[],[f12863,f11468])).
fof(f13769,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl83),
  inference(resolution,[],[f7094,f235])).
fof(f13770,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl83),
  inference(resolution,[],[f7094,f235])).
fof(f14766,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl78),
  inference(superposition,[],[f5431,f228])).
fof(f22816,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK1)))) | $spl256),
  inference(cnf_transformation,[],[f22816_D])).
fof(f22816_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK1)))) <=> ~$spl256),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl256])])).
fof(f15120,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl176),
  inference(superposition,[],[f12710,f228])).
fof(f15437,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK1))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl7),
  inference(superposition,[],[f312,f228])).
fof(f15624,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl33 | $spl78)),
  inference(superposition,[],[f9397,f228])).
fof(f16436,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl97),
  inference(superposition,[],[f8099,f226])).
fof(f16853,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ($spl2 | $spl158)),
  inference(superposition,[],[f11477,f307])).
fof(f22796,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | $spl245),
  inference(cnf_transformation,[],[f22796_D])).
fof(f22796_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) <=> ~$spl245),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl245])])).
fof(f22802,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl142 | $spl168 | $spl173)),
  inference(forward_demodulation,[],[f18518,f18145])).
fof(f22801,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl168 | $spl171)),
  inference(forward_demodulation,[],[f16077,f18145])).
fof(f22799,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | ($spl142 | $spl168 | $spl171)),
  inference(forward_demodulation,[],[f16076,f18145])).
fof(f2778,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl70),
  inference(cnf_transformation,[],[f2778_D])).
fof(f2778_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f4719,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl10 | $spl20 | $spl33)),
  inference(superposition,[],[f4668,f228])).
fof(f2958,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl73),
  inference(cnf_transformation,[],[f2958_D])).
fof(f2958_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) <=> ~$spl73),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl73])])).
fof(f18973,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl41),
  inference(superposition,[],[f589,f228])).
fof(f19021,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl10 | $spl20 | $spl33)),
  inference(superposition,[],[f4668,f228])).
fof(f21605,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl142 | $spl168 | $spl173)),
  inference(superposition,[],[f21383,f228])).
fof(f22783,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) | $spl243),
  inference(cnf_transformation,[],[f22783_D])).
fof(f22783_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) <=> ~$spl243),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl243])])).
fof(f21497,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | ($spl142 | $spl168 | $spl171)),
  inference(superposition,[],[f21382,f228])).
fof(f22779,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl236),
  inference(cnf_transformation,[],[f22779_D])).
fof(f22779_D,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) <=> ~$spl236),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl236])])).
fof(f7402,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | ($spl15 | $spl33 | $spl83)),
  inference(superposition,[],[f7249,f225])).
fof(f7249,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) | ($spl15 | $spl33 | $spl83)),
  inference(resolution,[],[f7207,f2815])).
fof(f2815,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X0),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X0))))) ) | $spl33),
  inference(resolution,[],[f417,f234])).
fof(f7207,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',sK3))),s('type.integer.int',X2))))) ) | ($spl15 | $spl83)),
  inference(forward_demodulation,[],[f7180,f233])).
fof(f7180,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl15 | $spl83)),
  inference(resolution,[],[f7137,f224])).
fof(f22778,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) ) | $spl234),
  inference(cnf_transformation,[],[f22778_D])).
fof(f22778_D,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) ) <=> ~$spl234),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl234])])).
fof(f7403,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) ) | ($spl15 | $spl33 | $spl83)),
  inference(superposition,[],[f7249,f227])).
fof(f7722,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | ($spl15 | $spl33 | $spl83)),
  inference(superposition,[],[f7249,f225])).
fof(f7723,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) ) | ($spl15 | $spl33 | $spl83)),
  inference(superposition,[],[f7249,f227])).
fof(f7841,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl94),
  inference(superposition,[],[f7415,f225])).
fof(f7842,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) ) | $spl94),
  inference(superposition,[],[f7415,f227])).
fof(f22699,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',X12)))) ) | $spl230),
  inference(cnf_transformation,[],[f22699_D])).
fof(f22699_D,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',X12)))) ) <=> ~$spl230),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl230])])).
fof(f8725,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',X12)))) ) | $spl98),
  inference(superposition,[],[f8102,f225])).
fof(f22696,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))))) ) | $spl226),
  inference(cnf_transformation,[],[f22696_D])).
fof(f22696_D,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))))) ) <=> ~$spl226),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl226])])).
fof(f8726,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))))) ) | $spl98),
  inference(superposition,[],[f8102,f227])).
fof(f22781,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X41)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl240),
  inference(cnf_transformation,[],[f22781_D])).
fof(f22781_D,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X41)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) <=> ~$spl240),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl240])])).
fof(f9149,plain,(
  ( ! [X41,X42] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X41)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl50),
  inference(superposition,[],[f2839,f227])).
fof(f22780,plain,(
  ( ! [X43,X44] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X43)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl238),
  inference(cnf_transformation,[],[f22780_D])).
fof(f22780_D,plain,(
  ( ! [X43,X44] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X43)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl238),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl238])])).
fof(f9150,plain,(
  ( ! [X43,X44] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',X43)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl50),
  inference(superposition,[],[f2839,f225])).
fof(f7102,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',sK3)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl88),
  inference(cnf_transformation,[],[f7102_D])).
fof(f7102_D,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',sK3)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f9291,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl50),
  inference(superposition,[],[f2840,f227])).
fof(f7103,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int',sK3)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl90),
  inference(cnf_transformation,[],[f7103_D])).
fof(f7103_D,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int',sK3)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f9292,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int',sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl50),
  inference(superposition,[],[f2840,f225])).
fof(f15787,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | ($spl15 | $spl33 | $spl83)),
  inference(superposition,[],[f7249,f225])).
fof(f15788,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) ) | ($spl15 | $spl33 | $spl83)),
  inference(superposition,[],[f7249,f227])).
fof(f16103,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl94),
  inference(superposition,[],[f7415,f225])).
fof(f16104,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3)))))) ) | $spl94),
  inference(superposition,[],[f7415,f227])).
fof(f22700,plain,(
  s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | $spl69),
  inference(forward_demodulation,[],[f21881,f14594])).
fof(f21881,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | $spl69),
  inference(superposition,[],[f238,f14594])).
fof(f22381,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',X12)))) ) | $spl98),
  inference(superposition,[],[f8102,f225])).
fof(f22382,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))))) ) | $spl98),
  inference(superposition,[],[f8102,f227])).
fof(f22695,plain,(
  $false | $spl132),
  inference(trivial_inequality_removal,[],[f22694])).
fof(f22694,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0)))) ) | $spl132),
  inference(forward_demodulation,[],[f22670,f232])).
fof(f22670,plain,(
  ( ! [X0] : (s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X0))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X0)))))) ) | $spl132),
  inference(superposition,[],[f22533,f230])).
fof(f22533,plain,(
  ( ! [X8] : (s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X8)))) ) | $spl132),
  inference(subsumption_resolution,[],[f22518,f237])).
fof(f22518,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',X8)))) ) | $spl132),
  inference(superposition,[],[f11004,f225])).
fof(f22688,plain,(
  $false | $spl132),
  inference(trivial_inequality_removal,[],[f22671])).
fof(f22671,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl132),
  inference(superposition,[],[f22533,f231])).
fof(f18645,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl224),
  inference(cnf_transformation,[],[f18645_D])).
fof(f18645_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl224),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl224])])).
fof(f12242,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl162),
  inference(cnf_transformation,[],[f12242_D])).
fof(f12242_D,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) <=> ~$spl162),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl162])])).
fof(f18636,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl10 | $spl20 | $spl33)),
  inference(duplicate_literal_removal,[],[f4743])).
fof(f4743,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | ($spl10 | $spl20 | $spl33)),
  inference(superposition,[],[f4668,f201])).
fof(f18637,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl10 | $spl20 | $spl33)),
  inference(duplicate_literal_removal,[],[f4744])).
fof(f4744,plain,(
  ( ! [X1] : (p(s(bool,X1)) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl10 | $spl20 | $spl33)),
  inference(superposition,[],[f4668,f201])).
fof(f18532,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | $spl173),
  inference(superposition,[],[f12706,f201])).
fof(f18533,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl173),
  inference(superposition,[],[f12706,f201])).
fof(f18640,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | ($spl142 | $spl173)),
  inference(subsumption_resolution,[],[f18535,f13020])).
fof(f18535,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | $spl173),
  inference(superposition,[],[f12706,f204])).
fof(f18638,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl142 | $spl173)),
  inference(subsumption_resolution,[],[f18534,f13020])).
fof(f18534,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl173),
  inference(superposition,[],[f12706,f204])).
fof(f18602,plain,(
  $false | ($spl1 | $spl40 | $spl87 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f7241])).
fof(f7241,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',sK2))))))) ) | ($spl1 | $spl87)),
  inference(forward_demodulation,[],[f7214,f233])).
fof(f7214,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int',sK2))))))) ) | ($spl1 | $spl87)),
  inference(resolution,[],[f7172,f224])).
fof(f7172,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | ($spl1 | $spl87)),
  inference(subsumption_resolution,[],[f7143,f306])).
fof(f7143,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl87),
  inference(resolution,[],[f7101,f235])).
fof(f18574,plain,(
  ( ! [X3] : (~p(s(bool,X3))) ) | ($spl40 | $spl142 | $spl173)),
  inference(subsumption_resolution,[],[f18573,f13020])).
fof(f18573,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X3))) ) | ($spl40 | $spl173)),
  inference(subsumption_resolution,[],[f18535,f590])).
fof(f18603,plain,(
  $false | ($spl15 | $spl40 | $spl83 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f7207])).
fof(f18605,plain,(
  $false | ($spl26 | $spl40 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f328])).
fof(f18606,plain,(
  $false | ($spl40 | $spl142 | $spl154 | $spl173)),
  inference(resolution,[],[f18574,f11475])).
fof(f18607,plain,(
  $false | ($spl40 | $spl142 | $spl173 | $spl176)),
  inference(resolution,[],[f18574,f12710])).
fof(f18608,plain,(
  $false | ($spl40 | $spl44 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f594])).
fof(f18609,plain,(
  $false | ($spl40 | $spl54 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f1242])).
fof(f18610,plain,(
  $false | ($spl40 | $spl94 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f7415])).
fof(f18611,plain,(
  $false | ($spl40 | $spl142 | $spl168 | $spl173)),
  inference(resolution,[],[f18574,f12448])).
fof(f18612,plain,(
  $false | ($spl40 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f11171])).
fof(f18613,plain,(
  $false | ($spl1 | $spl40 | $spl87 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f7172])).
fof(f18614,plain,(
  $false | ($spl15 | $spl40 | $spl83 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f7266])).
fof(f7266,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK3)))))) | ($spl15 | $spl83)),
  inference(superposition,[],[f7207,f231])).
fof(f18615,plain,(
  $false | ($spl40 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f590])).
fof(f18616,plain,(
  $false | ($spl1 | $spl40 | $spl87 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f7357])).
fof(f7357,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl1 | $spl87)),
  inference(superposition,[],[f7241,f231])).
fof(f18617,plain,(
  $false | ($spl40 | $spl78 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f5431])).
fof(f18618,plain,(
  $false | ($spl40 | $spl50 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f1178])).
fof(f18619,plain,(
  $false | ($spl15 | $spl40 | $spl83 | $spl142 | $spl173)),
  inference(resolution,[],[f18574,f7137])).
fof(f18570,plain,(
  $false | ($spl40 | $spl173 | $spl178)),
  inference(equality_resolution,[],[f18541])).
fof(f18541,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) ) | ($spl40 | $spl173 | $spl178)),
  inference(subsumption_resolution,[],[f18540,f590])).
fof(f18540,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl173 | $spl178)),
  inference(subsumption_resolution,[],[f18534,f12711])).
fof(f16819,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl220),
  inference(cnf_transformation,[],[f16819_D])).
fof(f16819_D,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) <=> ~$spl220),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl220])])).
fof(f315,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl11),
  inference(cnf_transformation,[],[f315_D])).
fof(f315_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f9181,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f267,f225])).
fof(f16818,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl218),
  inference(cnf_transformation,[],[f16818_D])).
fof(f16818_D,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) <=> ~$spl218),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl218])])).
fof(f9182,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f267,f227])).
fof(f16815,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl216),
  inference(cnf_transformation,[],[f16815_D])).
fof(f16815_D,plain,(
  ( ! [X17,X18,X16] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) <=> ~$spl216),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl216])])).
fof(f9704,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f269,f225])).
fof(f16814,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl214),
  inference(cnf_transformation,[],[f16814_D])).
fof(f16814_D,plain,(
  ( ! [X21,X19,X20] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) <=> ~$spl214),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl214])])).
fof(f9705,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f269,f227])).
fof(f313,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl8),
  inference(cnf_transformation,[],[f313_D])).
fof(f313_D,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f16813,plain,(
  ( ! [X58] : (s('type.integer.int',X58) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X58))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) )),
  inference(forward_demodulation,[],[f16812,f232])).
fof(f16812,plain,(
  ( ! [X58] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X58))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X58))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) )),
  inference(forward_demodulation,[],[f9717,f231])).
fof(f9717,plain,(
  ( ! [X57,X58] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X58))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X57))),s('type.integer.int',X58)))) )),
  inference(superposition,[],[f269,f231])).
fof(f16810,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl213),
  inference(cnf_transformation,[],[f16810_D])).
fof(f16810_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl213),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl213])])).
fof(f2428,plain,(
  ( ! [X12] : (s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl12),
  inference(superposition,[],[f950,f225])).
fof(f950,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | $spl12),
  inference(superposition,[],[f316,f231])).
fof(f2660,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f3251,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl26),
  inference(superposition,[],[f328,f225])).
fof(f16802,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl205),
  inference(cnf_transformation,[],[f16802_D])).
fof(f16802_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl205),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl205])])).
fof(f5043,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl23),
  inference(superposition,[],[f325,f225])).
fof(f16807,plain,(
  ( ! [X61,X59] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) ) | $spl210),
  inference(cnf_transformation,[],[f16807_D])).
fof(f16807_D,plain,(
  ( ! [X61,X59] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) ) <=> ~$spl210),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl210])])).
fof(f9774,plain,(
  ( ! [X61,X59,X60] : (s('type.integer.int',X60) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X60))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) )),
  inference(forward_demodulation,[],[f9718,f232])).
fof(f9718,plain,(
  ( ! [X61,X59,X60] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X60))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X60))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) )),
  inference(superposition,[],[f269,f227])).
fof(f16806,plain,(
  ( ! [X64,X62] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) ) | $spl208),
  inference(cnf_transformation,[],[f16806_D])).
fof(f16806_D,plain,(
  ( ! [X64,X62] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) ) <=> ~$spl208),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl208])])).
fof(f9775,plain,(
  ( ! [X64,X62,X63] : (s('type.integer.int',X63) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X63))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) )),
  inference(forward_demodulation,[],[f9719,f232])).
fof(f9719,plain,(
  ( ! [X64,X62,X63] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X63))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X63))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f269,f225])).
fof(f11078,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f12362,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f14184,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl26),
  inference(superposition,[],[f328,f225])).
fof(f16804,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl206),
  inference(cnf_transformation,[],[f16804_D])).
fof(f16804_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) <=> ~$spl206),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl206])])).
fof(f14809,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl154),
  inference(superposition,[],[f11475,f225])).
fof(f16347,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl23),
  inference(superposition,[],[f325,f225])).
fof(f16800,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl202),
  inference(cnf_transformation,[],[f16800_D])).
fof(f16800_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl202),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl202])])).
fof(f16568,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl150),
  inference(superposition,[],[f11471,f225])).
fof(f16798,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl200),
  inference(cnf_transformation,[],[f16798_D])).
fof(f16798_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) <=> ~$spl200),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl200])])).
fof(f16591,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl150),
  inference(superposition,[],[f235,f11471])).
fof(f16590,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl150),
  inference(superposition,[],[f235,f11471])).
fof(f16796,plain,(
  $false | ($spl33 | $spl60 | $spl168)),
  inference(subsumption_resolution,[],[f16771,f12448])).
fof(f16771,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | ($spl33 | $spl60)),
  inference(resolution,[],[f1639,f3595])).
fof(f16768,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3))) ) | $spl198),
  inference(cnf_transformation,[],[f16768_D])).
fof(f16768_D,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3))) ) <=> ~$spl198),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl198])])).
fof(f16770,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl16 | $spl33 | $spl142 | $spl147 | $spl178)),
  inference(subsumption_resolution,[],[f16694,f13616])).
fof(f16694,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X0))) ) | ($spl16 | $spl33 | $spl147)),
  inference(superposition,[],[f15873,f201])).
fof(f15873,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X0))) ) | ($spl16 | $spl33 | $spl147)),
  inference(subsumption_resolution,[],[f15872,f417])).
fof(f15872,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X0))) ) | ($spl16 | $spl147)),
  inference(forward_demodulation,[],[f15868,f319])).
fof(f319,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | $spl16),
  inference(cnf_transformation,[],[f319_D])).
fof(f319_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f15868,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl147),
  inference(superposition,[],[f11469,f201])).
fof(f16769,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl16 | $spl33 | $spl142 | $spl147 | $spl178)),
  inference(subsumption_resolution,[],[f16695,f13616])).
fof(f16695,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,X2)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl16 | $spl33 | $spl147)),
  inference(superposition,[],[f15873,f201])).
fof(f16710,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl16 | $spl33 | $spl142 | $spl147 | $spl178)),
  inference(subsumption_resolution,[],[f16694,f13616])).
fof(f16712,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | ($spl16 | $spl33 | $spl142 | $spl147 | $spl178)),
  inference(subsumption_resolution,[],[f16695,f13616])).
fof(f11476,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X0))) ) | $spl156),
  inference(cnf_transformation,[],[f11476_D])).
fof(f11476_D,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X0))) ) <=> ~$spl156),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl156])])).
fof(f16750,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl16 | $spl33 | $spl147)),
  inference(duplicate_literal_removal,[],[f16671])).
fof(f16671,plain,(
  ( ! [X0,X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,X1))) ) | ($spl16 | $spl33 | $spl147)),
  inference(superposition,[],[f15873,f201])).
fof(f16751,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X2)) | p(s(bool,X3))) ) | ($spl16 | $spl33 | $spl147)),
  inference(duplicate_literal_removal,[],[f16672])).
fof(f16672,plain,(
  ( ! [X2,X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X2)) | p(s(bool,X3)) | p(s(bool,X2))) ) | ($spl16 | $spl33 | $spl147)),
  inference(superposition,[],[f15873,f201])).
fof(f16766,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int',X0)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl150),
  inference(forward_demodulation,[],[f16586,f233])).
fof(f16586,plain,(
  ( ! [X0] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | $spl150),
  inference(superposition,[],[f224,f11471])).
fof(f16587,plain,(
  ( ! [X1] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X1))))) ) | $spl150),
  inference(superposition,[],[f234,f11471])).
fof(f16588,plain,(
  ( ! [X2] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))))) ) | $spl150),
  inference(superposition,[],[f234,f11471])).
fof(f16589,plain,(
  ( ! [X3] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X3))))) ) | $spl150),
  inference(superposition,[],[f234,f11471])).
fof(f16762,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl8 | $spl150)),
  inference(subsumption_resolution,[],[f16590,f802])).
fof(f16761,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl8 | $spl150)),
  inference(subsumption_resolution,[],[f16591,f802])).
fof(f16592,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl33 | $spl150)),
  inference(superposition,[],[f3595,f11471])).
fof(f16749,plain,(
  $false | ($spl1 | $spl33 | $spl42 | $spl87)),
  inference(subsumption_resolution,[],[f592,f9380])).
fof(f9380,plain,(
  ( ! [X8] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X8),s('type.integer.int',sK2))),s('type.integer.int',X8))))) ) | ($spl1 | $spl33 | $spl87)),
  inference(resolution,[],[f3595,f7241])).
fof(f16746,plain,(
  $false | ($spl10 | $spl46)),
  inference(trivial_inequality_removal,[],[f16726])).
fof(f16726,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl10 | $spl46)),
  inference(superposition,[],[f1150,f314])).
fof(f16745,plain,(
  $false | ($spl2 | $spl46)),
  inference(trivial_inequality_removal,[],[f16727])).
fof(f16727,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl2 | $spl46)),
  inference(superposition,[],[f1150,f307])).
fof(f16744,plain,(
  $false | $spl46),
  inference(equality_resolution,[],[f1150])).
fof(f16717,plain,(
  ( ! [X43,X44] : (s('type.integer.int',sK3) != s('type.integer.int',X43) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl186),
  inference(cnf_transformation,[],[f16717_D])).
fof(f16717_D,plain,(
  ( ! [X43,X44] : (s('type.integer.int',sK3) != s('type.integer.int',X43) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl186),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl186])])).
fof(f6616,plain,(
  ( ! [X43,X44] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int',X43) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) )),
  inference(forward_demodulation,[],[f6603,f233])).
fof(f6603,plain,(
  ( ! [X43,X44] : (s('type.integer.int',sK3) != s('type.integer.int',X43) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) )),
  inference(superposition,[],[f291,f225])).
fof(f16714,plain,(
  $false | ($spl16 | $spl33 | $spl61 | $spl147)),
  inference(subsumption_resolution,[],[f16698,f1638])).
fof(f16698,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl16 | $spl33 | $spl147)),
  inference(equality_resolution,[],[f15873])).
fof(f16707,plain,(
  $false | ($spl16 | $spl33 | $spl61 | $spl147 | $spl150)),
  inference(subsumption_resolution,[],[f16706,f1638])).
fof(f16706,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl16 | $spl33 | $spl147 | $spl150)),
  inference(forward_demodulation,[],[f16699,f11471])).
fof(f16699,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl16 | $spl33 | $spl147 | $spl150)),
  inference(trivial_inequality_removal,[],[f16679])).
fof(f16679,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ($spl16 | $spl33 | $spl147 | $spl150)),
  inference(superposition,[],[f15873,f11471])).
fof(f593,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl45),
  inference(cnf_transformation,[],[f593_D])).
fof(f593_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f11477,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) | $spl158),
  inference(cnf_transformation,[],[f11477_D])).
fof(f11477_D,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2))) ) <=> ~$spl158),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl158])])).
fof(f16410,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | ~p(s(bool,X3))) ) | $spl151),
  inference(superposition,[],[f11472,f204])).
fof(f16409,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl151),
  inference(superposition,[],[f11472,f204])).
fof(f594,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl44),
  inference(cnf_transformation,[],[f594_D])).
fof(f594_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f16408,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl151),
  inference(superposition,[],[f11472,f201])).
fof(f16407,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl151),
  inference(superposition,[],[f11472,f201])).
fof(f16370,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl23 | $spl30)),
  inference(subsumption_resolution,[],[f16362,f330])).
fof(f16362,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl23),
  inference(superposition,[],[f325,f201])).
fof(f16369,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl23 | $spl30)),
  inference(subsumption_resolution,[],[f16361,f330])).
fof(f16361,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl23),
  inference(superposition,[],[f325,f201])).
fof(f11473,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl152),
  inference(cnf_transformation,[],[f11473_D])).
fof(f11473_D,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) ) <=> ~$spl152),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl152])])).
fof(f15105,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl176),
  inference(superposition,[],[f12710,f227])).
fof(f11470,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X12)))) ) | $spl148),
  inference(cnf_transformation,[],[f11470_D])).
fof(f11470_D,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X12)))) ) <=> ~$spl148),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl148])])).
fof(f15104,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X12)))) ) | $spl176),
  inference(superposition,[],[f12710,f225])).
fof(f14918,plain,(
  $false | ($spl142 | $spl177)),
  inference(subsumption_resolution,[],[f14859,f11171])).
fof(f14859,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | $spl177),
  inference(resolution,[],[f12709,f224])).
fof(f14836,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl154),
  inference(superposition,[],[f11475,f227])).
fof(f14835,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X12)))) ) | $spl154),
  inference(superposition,[],[f11475,f225])).
fof(f11467,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl145),
  inference(cnf_transformation,[],[f11467_D])).
fof(f11467_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl145),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl145])])).
fof(f14825,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl154),
  inference(superposition,[],[f11475,f228])).
fof(f14700,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl155),
  inference(superposition,[],[f11474,f227])).
fof(f14699,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X12)))) ) | $spl155),
  inference(superposition,[],[f11474,f225])).
fof(f591,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | $spl43),
  inference(cnf_transformation,[],[f591_D])).
fof(f591_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) <=> ~$spl43),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl43])])).
fof(f14689,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl155),
  inference(superposition,[],[f11474,f228])).
fof(f14710,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl182),
  inference(cnf_transformation,[],[f14710_D])).
fof(f14710_D,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) <=> ~$spl182),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl182])])).
fof(f14712,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl184),
  inference(cnf_transformation,[],[f14712_D])).
fof(f14712_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl184),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl184])])).
fof(f14652,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl155),
  inference(resolution,[],[f11474,f235])).
fof(f14651,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) | $spl155),
  inference(resolution,[],[f11474,f235])).
fof(f14136,plain,(
  $false | ($spl69 | $spl138)),
  inference(subsumption_resolution,[],[f14135,f2596])).
fof(f14135,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl69 | $spl138)),
  inference(forward_demodulation,[],[f11029,f11143])).
fof(f327,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl27),
  inference(cnf_transformation,[],[f327_D])).
fof(f327_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f330,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | $spl30),
  inference(cnf_transformation,[],[f330_D])).
fof(f330_D,plain,(
  ( ! [X2] : (~p(s(bool,X2)) | s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f12880,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f199,f204])).
fof(f12879,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f204])).
fof(f328,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl26),
  inference(cnf_transformation,[],[f328_D])).
fof(f328_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f329,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | $spl28),
  inference(cnf_transformation,[],[f329_D])).
fof(f329_D,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f12878,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f201])).
fof(f12877,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f199,f201])).
fof(f326,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl24),
  inference(cnf_transformation,[],[f326_D])).
fof(f326_D,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f325,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl23),
  inference(cnf_transformation,[],[f325_D])).
fof(f325_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f12874,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) )),
  inference(superposition,[],[f199,f227])).
fof(f323,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl20),
  inference(cnf_transformation,[],[f323_D])).
fof(f323_D,plain,(
  ( ! [X12] : (s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f12873,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f316,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl12),
  inference(cnf_transformation,[],[f316_D])).
fof(f316_D,plain,(
  ( ! [X13] : (s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f12848,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) )),
  inference(superposition,[],[f199,f227])).
fof(f12847,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f12778,plain,(
  $false | ($spl50 | $spl69 | $spl166)),
  inference(subsumption_resolution,[],[f2596,f12460])).
fof(f12460,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl50 | $spl166)),
  inference(backward_demodulation,[],[f12446,f1178])).
fof(f12709,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl177),
  inference(cnf_transformation,[],[f12709_D])).
fof(f12709_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl177),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl177])])).
fof(f12702,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | $spl166),
  inference(superposition,[],[f12449,f204])).
fof(f12449,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl166),
  inference(backward_demodulation,[],[f12446,f199])).
fof(f12701,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl166),
  inference(superposition,[],[f12449,f204])).
fof(f12710,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl176),
  inference(cnf_transformation,[],[f12710_D])).
fof(f12710_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl176),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl176])])).
fof(f12700,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl166),
  inference(superposition,[],[f12449,f201])).
fof(f12699,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | $spl166),
  inference(superposition,[],[f12449,f201])).
fof(f12708,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl175),
  inference(cnf_transformation,[],[f12708_D])).
fof(f12708_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl175),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl175])])).
fof(f12696,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl166),
  inference(superposition,[],[f12449,f227])).
fof(f12695,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl166),
  inference(superposition,[],[f12449,f225])).
fof(f12706,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl173),
  inference(cnf_transformation,[],[f12706_D])).
fof(f12706_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl173),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl173])])).
fof(f12670,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl166),
  inference(superposition,[],[f12449,f227])).
fof(f12669,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X12)))) ) | $spl166),
  inference(superposition,[],[f12449,f225])).
fof(f12704,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl171),
  inference(cnf_transformation,[],[f12704_D])).
fof(f12704_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl171),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl171])])).
fof(f12659,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl166),
  inference(superposition,[],[f12449,f228])).
fof(f12446,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK2) | $spl166),
  inference(cnf_transformation,[],[f12446_D])).
fof(f12446_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK2) <=> ~$spl166),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl166])])).
fof(f12448,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | $spl168),
  inference(cnf_transformation,[],[f12448_D])).
fof(f12448_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) <=> ~$spl168),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl168])])).
fof(f12422,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK2) | $spl61),
  inference(resolution,[],[f1638,f235])).
fof(f12421,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK2)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK2) | $spl61),
  inference(resolution,[],[f1638,f235])).
fof(f12395,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f199,f204])).
fof(f12394,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f204])).
fof(f12393,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f201])).
fof(f12392,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f199,f201])).
fof(f12389,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) )),
  inference(superposition,[],[f199,f227])).
fof(f12388,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f12293,plain,(
  $false | ($spl33 | $spl162)),
  inference(subsumption_resolution,[],[f12268,f12242])).
fof(f12268,plain,(
  ( ! [X47,X46] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X46),s('type.integer.int',X47))))) ) | ($spl33 | $spl162)),
  inference(resolution,[],[f12242,f3595])).
fof(f12267,plain,(
  $false | $spl162),
  inference(resolution,[],[f12242,f237])).
fof(f12272,plain,(
  $false | ($spl1 | $spl33 | $spl87 | $spl162)),
  inference(resolution,[],[f12242,f9380])).
fof(f12273,plain,(
  $false | ($spl33 | $spl44 | $spl162)),
  inference(resolution,[],[f12242,f9387])).
fof(f12274,plain,(
  $false | ($spl1 | $spl33 | $spl87 | $spl162)),
  inference(resolution,[],[f12242,f9394])).
fof(f9394,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))) | ($spl1 | $spl33 | $spl87)),
  inference(resolution,[],[f3595,f7357])).
fof(f12275,plain,(
  $false | ($spl87 | $spl162)),
  inference(resolution,[],[f12242,f7101])).
fof(f1175,plain,(
  s('type.integer.int',sK2) != s('type.integer.int',sK3) | $spl49),
  inference(cnf_transformation,[],[f1175_D])).
fof(f1175_D,plain,(
  s('type.integer.int',sK2) != s('type.integer.int',sK3) <=> ~$spl49),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl49])])).
fof(f11755,plain,(
  s('type.integer.int',sK2) != s('type.integer.int',sK3) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | $spl63),
  inference(superposition,[],[f1640,f238])).
fof(f1640,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | $spl63),
  inference(cnf_transformation,[],[f1640_D])).
fof(f1640_D,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f12248,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X0))) ) | $spl156),
  inference(subsumption_resolution,[],[f12230,f201])).
fof(f12230,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | p(s(bool,X0))) ) | $spl156),
  inference(superposition,[],[f11476,f201])).
fof(f12247,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | p(s(bool,X2)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl156),
  inference(subsumption_resolution,[],[f12231,f201])).
fof(f12231,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,X2)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl156),
  inference(superposition,[],[f11476,f201])).
fof(f12232,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,X5) | p(s(bool,X5)) | ~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))))) ) | $spl156),
  inference(superposition,[],[f11476,f204])).
fof(f12233,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,X7) | p(s(bool,X7)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ~p(s(bool,X6))) ) | $spl156),
  inference(superposition,[],[f11476,f204])).
fof(f12245,plain,(
  $false | ($spl33 | $spl140 | $spl156)),
  inference(subsumption_resolution,[],[f12234,f11179])).
fof(f11179,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | ($spl33 | $spl140)),
  inference(backward_demodulation,[],[f11169,f417])).
fof(f12234,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl156),
  inference(equality_resolution,[],[f11476])).
fof(f12244,plain,(
  ( ! [X2,X3] : (p(s(bool,X3)) | p(s(bool,X2))) ) | ($spl33 | $spl140 | $spl156)),
  inference(subsumption_resolution,[],[f12243,f201])).
fof(f12243,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,X3) | p(s(bool,X3)) | p(s(bool,X2))) ) | ($spl33 | $spl140 | $spl156)),
  inference(subsumption_resolution,[],[f12231,f11179])).
fof(f12241,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl33 | $spl140 | $spl156)),
  inference(subsumption_resolution,[],[f12240,f201])).
fof(f12240,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) ) | ($spl33 | $spl140 | $spl156)),
  inference(subsumption_resolution,[],[f12230,f11179])).
fof(f11975,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl160),
  inference(cnf_transformation,[],[f11975_D])).
fof(f11975_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl160),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl160])])).
fof(f1820,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | $spl64),
  inference(cnf_transformation,[],[f1820_D])).
fof(f1820_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f4820,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl66),
  inference(superposition,[],[f2168,f238])).
fof(f4829,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl66),
  inference(superposition,[],[f238,f2168])).
fof(f11945,plain,(
  $false | ($spl49 | $spl64 | $spl69 | $spl140)),
  inference(subsumption_resolution,[],[f11944,f11184])).
fof(f11184,plain,(
  s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int',sK2) | ($spl49 | $spl140)),
  inference(backward_demodulation,[],[f11169,f1175])).
fof(f11944,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK2) | ($spl64 | $spl69 | $spl140)),
  inference(forward_demodulation,[],[f11943,f11765])).
fof(f11765,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | ($spl69 | $spl140)),
  inference(forward_demodulation,[],[f11143,f11169])).
fof(f11943,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | ($spl64 | $spl69 | $spl140)),
  inference(forward_demodulation,[],[f11942,f1820])).
fof(f11942,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | ($spl69 | $spl140)),
  inference(subsumption_resolution,[],[f11837,f237])).
fof(f11837,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | ($spl69 | $spl140)),
  inference(superposition,[],[f238,f11765])).
fof(f11474,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl155),
  inference(cnf_transformation,[],[f11474_D])).
fof(f11474_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) <=> ~$spl155),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl155])])).
fof(f11465,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) | ~p(s(bool,X3))) ) | $spl140),
  inference(superposition,[],[f11172,f204])).
fof(f11172,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))))) | $spl140),
  inference(backward_demodulation,[],[f11169,f199])).
fof(f11464,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))))))) ) | $spl140),
  inference(superposition,[],[f11172,f204])).
fof(f11475,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) | $spl154),
  inference(cnf_transformation,[],[f11475_D])).
fof(f11475_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) <=> ~$spl154),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl154])])).
fof(f11463,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))))))) ) | $spl140),
  inference(superposition,[],[f11172,f201])).
fof(f11462,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) | p(s(bool,X0))) ) | $spl140),
  inference(superposition,[],[f11172,f201])).
fof(f11472,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl151),
  inference(cnf_transformation,[],[f11472_D])).
fof(f11472_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl151),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl151])])).
fof(f11459,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0')))))) ) | $spl140),
  inference(superposition,[],[f11172,f227])).
fof(f11458,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_|48|_0'))),s('type.integer.int',X12)))) ) | $spl140),
  inference(superposition,[],[f11172,f225])).
fof(f11469,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | $spl147),
  inference(cnf_transformation,[],[f11469_D])).
fof(f11469_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) <=> ~$spl147),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl147])])).
fof(f11448,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0'))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int','const.integer.int_|48|_0') != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl140),
  inference(superposition,[],[f11172,f228])).
fof(f11169,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK3) | $spl140),
  inference(cnf_transformation,[],[f11169_D])).
fof(f11169_D,plain,(
  s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK3) <=> ~$spl140),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl140])])).
fof(f11171,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | $spl142),
  inference(cnf_transformation,[],[f11171_D])).
fof(f11171_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) <=> ~$spl142),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl142])])).
fof(f11145,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK3) | $spl69),
  inference(resolution,[],[f2596,f235])).
fof(f11144,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | s('type.integer.int','const.integer.int_|48|_0') = s('type.integer.int',sK3) | $spl69),
  inference(resolution,[],[f2596,f235])).
fof(f11029,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl138),
  inference(cnf_transformation,[],[f11029_D])).
fof(f11029_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))),s('type.integer.int','const.integer.int_|48|_0')))) <=> ~$spl138),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl138])])).
fof(f2018,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl64),
  inference(superposition,[],[f1820,f238])).
fof(f2027,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))),s('type.integer.int','const.integer.int_|48|_0')))) | $spl64),
  inference(superposition,[],[f238,f1820])).
fof(f10998,plain,(
  ( ! [X9] : (s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))) ) | $spl128),
  inference(cnf_transformation,[],[f10998_D])).
fof(f10998_D,plain,(
  ( ! [X9] : (s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))) ) <=> ~$spl128),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl128])])).
fof(f11009,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | $spl137),
  inference(cnf_transformation,[],[f11009_D])).
fof(f11009_D,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) <=> ~$spl137),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl137])])).
fof(f8389,plain,(
  ( ! [X9] : (s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))) ) | $spl97),
  inference(superposition,[],[f8099,f227])).
fof(f8486,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | $spl103),
  inference(cnf_transformation,[],[f8486_D])).
fof(f8486_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) <=> ~$spl103),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl103])])).
fof(f11007,plain,(
  ( ! [X16] : (s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',sK2)))) ) | $spl134),
  inference(cnf_transformation,[],[f11007_D])).
fof(f11007_D,plain,(
  ( ! [X16] : (s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',sK2)))) ) <=> ~$spl134),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl134])])).
fof(f8496,plain,(
  ( ! [X58] : (s('type.integer.int',X58) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int',sK2))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl2),
  inference(forward_demodulation,[],[f8495,f233])).
fof(f8495,plain,(
  ( ! [X58] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl2),
  inference(forward_demodulation,[],[f8494,f231])).
fof(f8494,plain,(
  ( ! [X57,X58] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X57)))))) ) | $spl2),
  inference(forward_demodulation,[],[f8493,f307])).
fof(f8493,plain,(
  ( ! [X57,X58] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X57)))))) )),
  inference(forward_demodulation,[],[f8423,f233])).
fof(f8423,plain,(
  ( ! [X57,X58] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X58),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X57)))))) )),
  inference(superposition,[],[f255,f231])).
fof(f10994,plain,(
  ( ! [X61,X59] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X59))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) ) | $spl124),
  inference(cnf_transformation,[],[f10994_D])).
fof(f10994_D,plain,(
  ( ! [X61,X59] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X59))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) ) <=> ~$spl124),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl124])])).
fof(f8497,plain,(
  ( ! [X61,X59,X60] : (s('type.integer.int',X60) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X60),s('type.integer.int',sK2))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X59))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) )),
  inference(forward_demodulation,[],[f8424,f233])).
fof(f8424,plain,(
  ( ! [X61,X59,X60] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X59))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X60),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X60),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) )),
  inference(superposition,[],[f255,f227])).
fof(f10993,plain,(
  ( ! [X64,X62] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X62))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) ) | $spl122),
  inference(cnf_transformation,[],[f10993_D])).
fof(f10993_D,plain,(
  ( ! [X64,X62] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X62))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) ) <=> ~$spl122),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl122])])).
fof(f8498,plain,(
  ( ! [X64,X62,X63] : (s('type.integer.int',X63) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X63),s('type.integer.int',sK2))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X62))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) )),
  inference(forward_demodulation,[],[f8425,f233])).
fof(f8425,plain,(
  ( ! [X64,X62,X63] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X62))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X63),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X63),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f255,f225])).
fof(f8504,plain,(
  ( ! [X16] : (s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',sK2))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl2),
  inference(forward_demodulation,[],[f8503,f233])).
fof(f8503,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) ) | $spl2),
  inference(forward_demodulation,[],[f8440,f307])).
fof(f8440,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) )),
  inference(superposition,[],[f255,f233])).
fof(f11000,plain,(
  ( ! [X8] : (s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))),s('type.integer.int',X8)))) ) | $spl130),
  inference(cnf_transformation,[],[f11000_D])).
fof(f11000_D,plain,(
  ( ! [X8] : (s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))),s('type.integer.int',X8)))) ) <=> ~$spl130),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl130])])).
fof(f11004,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | $spl132),
  inference(cnf_transformation,[],[f11004_D])).
fof(f11004_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) <=> ~$spl132),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl132])])).
fof(f8671,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))),s('type.integer.int',X8)))) ) | ($spl1 | $spl87)),
  inference(superposition,[],[f7959,f225])).
fof(f7959,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))))) | ($spl1 | $spl87)),
  inference(superposition,[],[f7813,f231])).
fof(f7813,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',X2))))) ) | ($spl1 | $spl87)),
  inference(forward_demodulation,[],[f7778,f233])).
fof(f7778,plain,(
  ( ! [X2] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))))) ) | ($spl1 | $spl87)),
  inference(resolution,[],[f7357,f224])).
fof(f8672,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))))) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))) ) | ($spl1 | $spl87)),
  inference(superposition,[],[f7959,f227])).
fof(f10999,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))),s('type.integer.int',X8)))) ) | $spl98),
  inference(forward_demodulation,[],[f8711,f233])).
fof(f8711,plain,(
  ( ! [X8] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK2)))) | s('type.integer.int',X8) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))),s('type.integer.int',X8)))) ) | $spl98),
  inference(superposition,[],[f8102,f225])).
fof(f10997,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))) ) | $spl98),
  inference(forward_demodulation,[],[f8712,f233])).
fof(f8712,plain,(
  ( ! [X9] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int',sK2)))) | s('type.integer.int',X9) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))) ) | $spl98),
  inference(superposition,[],[f8102,f227])).
fof(f10996,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))) | $spl127),
  inference(cnf_transformation,[],[f10996_D])).
fof(f10996_D,plain,(
  s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))) <=> ~$spl127),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl127])])).
fof(f8715,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) != s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))) | $spl98),
  inference(superposition,[],[f8102,f228])).
fof(f10992,plain,(
  ( ! [X16] : (s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X16)))) ) | $spl120),
  inference(cnf_transformation,[],[f10992_D])).
fof(f10992_D,plain,(
  ( ! [X16] : (s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X16)))) ) <=> ~$spl120),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl120])])).
fof(f8904,plain,(
  ( ! [X58] : (s('type.integer.int',X58) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X58))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl2),
  inference(forward_demodulation,[],[f8903,f232])).
fof(f8903,plain,(
  ( ! [X58] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X58))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X58))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl2),
  inference(forward_demodulation,[],[f8902,f231])).
fof(f8902,plain,(
  ( ! [X57,X58] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X58))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X57))),s('type.integer.int',X58)))) ) | $spl2),
  inference(forward_demodulation,[],[f8901,f307])).
fof(f8901,plain,(
  ( ! [X57,X58] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X58))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X57))),s('type.integer.int',X58)))) )),
  inference(forward_demodulation,[],[f8833,f233])).
fof(f8833,plain,(
  ( ! [X57,X58] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X58))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',X57))),s('type.integer.int',X57))),s('type.integer.int',X58)))) )),
  inference(superposition,[],[f257,f231])).
fof(f8905,plain,(
  ( ! [X61,X59,X60] : (s('type.integer.int',X60) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X60))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X59))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) )),
  inference(forward_demodulation,[],[f8834,f232])).
fof(f8834,plain,(
  ( ! [X61,X59,X60] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X59),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X59))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X60))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X60))) | s('type.integer.int',X61) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X61),s('type.integer.int',X59)))) )),
  inference(superposition,[],[f257,f227])).
fof(f8906,plain,(
  ( ! [X64,X62,X63] : (s('type.integer.int',X63) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X63))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X62))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) )),
  inference(forward_demodulation,[],[f8835,f232])).
fof(f8835,plain,(
  ( ! [X64,X62,X63] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X62),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',X62))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X63))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X63))) | s('type.integer.int',X64) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X62),s('type.integer.int',X64)))) )),
  inference(superposition,[],[f257,f225])).
fof(f8912,plain,(
  ( ! [X16] : (s('type.integer.int',X16) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X16))) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) ) | $spl2),
  inference(forward_demodulation,[],[f8911,f232])).
fof(f8911,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X16))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))) ) | $spl2),
  inference(forward_demodulation,[],[f8850,f307])).
fof(f8850,plain,(
  ( ! [X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X16))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',X16)))) )),
  inference(superposition,[],[f257,f233])).
fof(f10991,plain,(
  $false | ($spl33 | $spl68 | $spl117)),
  inference(subsumption_resolution,[],[f10134,f9504])).
fof(f9504,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl33 | $spl68)),
  inference(resolution,[],[f9398,f2817])).
fof(f9398,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3)))) | ($spl33 | $spl68)),
  inference(resolution,[],[f3595,f2597])).
fof(f10135,plain,(
  ( ! [X18] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl118),
  inference(cnf_transformation,[],[f10135_D])).
fof(f10135_D,plain,(
  ( ! [X18] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) <=> ~$spl118),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl118])])).
fof(f10134,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | $spl117),
  inference(cnf_transformation,[],[f10134_D])).
fof(f10134_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) <=> ~$spl117),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl117])])).
fof(f10089,plain,(
  ( ! [X18] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X18),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | ($spl33 | $spl68)),
  inference(superposition,[],[f267,f9503])).
fof(f9503,plain,(
  ( ! [X1] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) ) | ($spl33 | $spl68)),
  inference(resolution,[],[f9398,f3588])).
fof(f10132,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X17)))) ) | $spl114),
  inference(cnf_transformation,[],[f10132_D])).
fof(f10132_D,plain,(
  ( ! [X17] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X17)))) ) <=> ~$spl114),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl114])])).
fof(f10088,plain,(
  ( ! [X17] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_|48|_0'),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X17)))) ) | ($spl33 | $spl68)),
  inference(superposition,[],[f269,f9503])).
fof(f9784,plain,(
  ( ! [X7,X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X9))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl112),
  inference(cnf_transformation,[],[f9784_D])).
fof(f9784_D,plain,(
  ( ! [X7,X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X9))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) <=> ~$spl112),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl112])])).
fof(f9785,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X10),s('type.integer.int',X12))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f9760,f330])).
fof(f9760,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X10),s('type.integer.int',X12))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12))) | p(s(bool,X11)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f269,f201])).
fof(f9783,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X9))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f9759,f330])).
fof(f9759,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X7),s('type.integer.int',X9))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X9))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X8))) )),
  inference(superposition,[],[f269,f201])).
fof(f9261,plain,(
  ( ! [X7,X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',X7))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl110),
  inference(cnf_transformation,[],[f9261_D])).
fof(f9261_D,plain,(
  ( ! [X7,X9] : (s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',X7))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f9262,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X10))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f9237,f330])).
fof(f9237,plain,(
  ( ! [X12,X10,X11] : (s(bool,X11) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X10))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | p(s(bool,X11)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X10),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f267,f201])).
fof(f9260,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',X7))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl30),
  inference(subsumption_resolution,[],[f9236,f330])).
fof(f9236,plain,(
  ( ! [X8,X7,X9] : (s(bool,X8) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int',X7))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X9),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X8))) )),
  inference(superposition,[],[f267,f201])).
fof(f8896,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl2),
  inference(forward_demodulation,[],[f8895,f307])).
fof(f8895,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(forward_demodulation,[],[f8821,f233])).
fof(f8821,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X20))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X19),s('type.integer.int',X20))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f257,f227])).
fof(f8893,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl2),
  inference(forward_demodulation,[],[f8892,f307])).
fof(f8892,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(forward_demodulation,[],[f8820,f233])).
fof(f8820,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK2),s('type.integer.int',X17))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X17))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f257,f225])).
fof(f8488,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) ) | $spl2),
  inference(forward_demodulation,[],[f8487,f307])).
fof(f8487,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(forward_demodulation,[],[f8411,f233])).
fof(f8411,plain,(
  ( ! [X21,X19,X20] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X20),s('type.integer.int',X19))) | s('type.integer.int',X21) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X21),s('type.integer.int',X19)))) )),
  inference(superposition,[],[f255,f227])).
fof(f8483,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) ) | $spl2),
  inference(forward_demodulation,[],[f8482,f307])).
fof(f8482,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(forward_demodulation,[],[f8410,f233])).
fof(f8410,plain,(
  ( ! [X17,X18,X16] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',sK2))) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X17),s('type.integer.int',X16))) | s('type.integer.int',X18) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X16),s('type.integer.int',X18)))) )),
  inference(superposition,[],[f255,f225])).
fof(f8302,plain,(
  $false | ($spl1 | $spl33 | $spl87 | $spl96)),
  inference(subsumption_resolution,[],[f8140,f417])).
fof(f8140,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ($spl1 | $spl87 | $spl96)),
  inference(superposition,[],[f7813,f8100])).
fof(f8100,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))) | $spl96),
  inference(cnf_transformation,[],[f8100_D])).
fof(f8100_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f8102,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',sK2)))) | $spl98),
  inference(cnf_transformation,[],[f8102_D])).
fof(f8102_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',sK2)))) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f8053,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))) | ($spl1 | $spl33 | $spl87)),
  inference(resolution,[],[f7932,f235])).
fof(f7932,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2)))))))) | ($spl1 | $spl33 | $spl87)),
  inference(resolution,[],[f7813,f2815])).
fof(f8052,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_neg_1'(s('type.integer.int',sK2))))) | ($spl1 | $spl33 | $spl87)),
  inference(resolution,[],[f7932,f235])).
fof(f7602,plain,(
  $false | ($spl15 | $spl33 | $spl83 | $spl92)),
  inference(subsumption_resolution,[],[f7439,f417])).
fof(f7439,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ($spl15 | $spl83 | $spl92)),
  inference(superposition,[],[f7207,f7413])).
fof(f7413,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))) | $spl92),
  inference(cnf_transformation,[],[f7413_D])).
fof(f7413_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f7415,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) | $spl94),
  inference(cnf_transformation,[],[f7415_D])).
fof(f7415_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f7374,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))) | ($spl15 | $spl33 | $spl83)),
  inference(resolution,[],[f7249,f235])).
fof(f7373,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK3),s('type.integer.int',sK3))) | ($spl15 | $spl33 | $spl83)),
  inference(resolution,[],[f7249,f235])).
fof(f7066,plain,(
  ( ! [X43,X44] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X43),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl33),
  inference(superposition,[],[f2815,f225])).
fof(f7065,plain,(
  ( ! [X41,X42] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',X41),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X42) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X42),s('type.integer.int',X41)))) ) | $spl33),
  inference(superposition,[],[f2815,f227])).
fof(f7052,plain,(
  ( ! [X14,X15] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X14)))) | s('type.integer.int',X15) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X15),s('type.integer.int',X14)))) ) | $spl33),
  inference(superposition,[],[f2815,f227])).
fof(f7051,plain,(
  ( ! [X12,X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK3)))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',X12)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X12),s('type.integer.int',X13)))) ) | $spl33),
  inference(superposition,[],[f2815,f225])).
fof(f5431,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | $spl78),
  inference(cnf_transformation,[],[f5431_D])).
fof(f5431_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f5432,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl7 | $spl30)),
  inference(subsumption_resolution,[],[f5425,f330])).
fof(f5425,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl7),
  inference(superposition,[],[f312,f201])).
fof(f5429,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl7 | $spl30)),
  inference(subsumption_resolution,[],[f5424,f330])).
fof(f5424,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | $spl7),
  inference(superposition,[],[f312,f201])).
fof(f5273,plain,(
  ( ! [X43,X44] : (s('type.integer.int',sK1) != s('type.integer.int',X43) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) | $spl76),
  inference(cnf_transformation,[],[f5273_D])).
fof(f5273_D,plain,(
  ( ! [X43,X44] : (s('type.integer.int',sK1) != s('type.integer.int',X43) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f312,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl7),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f5272,plain,(
  ( ! [X43,X44] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',sK1) != s('type.integer.int',X43) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) )),
  inference(forward_demodulation,[],[f5260,f232])).
fof(f5260,plain,(
  ( ! [X43,X44] : (s('type.integer.int',sK1) != s('type.integer.int',X43) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X44) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X43),s('type.integer.int',X44)))) )),
  inference(superposition,[],[f263,f225])).
fof(f263,plain,(
  ( ! [X0] : (s('type.integer.int',sK1) != s('type.integer.int',X0) | s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X0),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) )),
  inference(superposition,[],[f199,f223])).
fof(f5066,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl23 | $spl30)),
  inference(subsumption_resolution,[],[f5058,f330])).
fof(f5058,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | $spl23),
  inference(superposition,[],[f325,f201])).
fof(f5065,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))))) ) | ($spl23 | $spl30)),
  inference(subsumption_resolution,[],[f5057,f330])).
fof(f5057,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | p(s(bool,X0))) ) | $spl23),
  inference(superposition,[],[f325,f201])).
fof(f4983,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | $spl74),
  inference(cnf_transformation,[],[f4983_D])).
fof(f4983_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f4981,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl66),
  inference(forward_demodulation,[],[f4889,f2168])).
fof(f4889,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))))))))) | $spl66),
  inference(superposition,[],[f238,f2168])).
fof(f3991,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl11 | $spl30)),
  inference(subsumption_resolution,[],[f3983,f330])).
fof(f3983,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl11),
  inference(superposition,[],[f315,f201])).
fof(f3990,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl11 | $spl30)),
  inference(subsumption_resolution,[],[f3982,f330])).
fof(f3982,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | $spl11),
  inference(superposition,[],[f315,f201])).
fof(f3670,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl17 | $spl30)),
  inference(subsumption_resolution,[],[f3666,f330])).
fof(f3666,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | $spl17),
  inference(superposition,[],[f320,f201])).
fof(f3669,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))))) ) | ($spl17 | $spl30)),
  inference(subsumption_resolution,[],[f3665,f330])).
fof(f3665,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | p(s(bool,X0))) ) | $spl17),
  inference(superposition,[],[f320,f201])).
fof(f3444,plain,(
  $false | $spl34),
  inference(equality_resolution,[],[f1416])).
fof(f1416,plain,(
  ( ! [X76,X77] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',X76),s('type.integer.int',X76))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X77),s('type.integer.int',X77)))) ) | $spl34),
  inference(resolution,[],[f419,f237])).
fof(f2959,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | $spl72),
  inference(cnf_transformation,[],[f2959_D])).
fof(f2959_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2943,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl54),
  inference(superposition,[],[f1242,f227])).
fof(f2942,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int',sK1)))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl54),
  inference(superposition,[],[f1242,f225])).
fof(f2811,plain,(
  $false | ($spl27 | $spl28)),
  inference(trivial_inequality_removal,[],[f2802])).
fof(f2802,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl27 | $spl28)),
  inference(resolution,[],[f329,f1426])).
fof(f1426,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | $spl27),
  inference(resolution,[],[f327,f224])).
fof(f2779,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | $spl71),
  inference(cnf_transformation,[],[f2779_D])).
fof(f2779_D,plain,(
  ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f2769,plain,(
  ( ! [X13] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl39),
  inference(superposition,[],[f587,f227])).
fof(f2768,plain,(
  ( ! [X12] : (~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl39),
  inference(superposition,[],[f587,f225])).
fof(f1240,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl52),
  inference(cnf_transformation,[],[f1240_D])).
fof(f1240_D,plain,(
  s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f1242,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) | $spl54),
  inference(cnf_transformation,[],[f1242_D])).
fof(f1242_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2741,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl39),
  inference(resolution,[],[f587,f235])).
fof(f2740,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl39),
  inference(resolution,[],[f587,f235])).
fof(f2693,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f199,f204])).
fof(f2692,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f204])).
fof(f2691,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f201])).
fof(f2690,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f199,f201])).
fof(f2687,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) )),
  inference(superposition,[],[f199,f227])).
fof(f2686,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f1664,plain,(
  s('type.integer.int',sK2) = s('type.integer.int',sK3) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | $spl62),
  inference(superposition,[],[f1641,f238])).
fof(f1673,plain,(
  s('type.integer.int',sK2) = s('type.integer.int',sK3) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int','const.integer.int_|48|_0')))) | $spl62),
  inference(superposition,[],[f238,f1641])).
fof(f2579,plain,(
  $false | ($spl2 | $spl34 | $spl48)),
  inference(resolution,[],[f2511,f237])).
fof(f2511,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | ($spl2 | $spl34 | $spl48)),
  inference(subsumption_resolution,[],[f2510,f237])).
fof(f2510,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK2)))) | p(s(bool,X0))) ) | ($spl2 | $spl34 | $spl48)),
  inference(forward_demodulation,[],[f2509,f2017])).
fof(f2017,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK2))) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))))) | ($spl2 | $spl48)),
  inference(forward_demodulation,[],[f307,f1176])).
fof(f2509,plain,(
  ( ! [X0] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X0))) ) | ($spl2 | $spl34 | $spl48)),
  inference(subsumption_resolution,[],[f2470,f419])).
fof(f2470,plain,(
  ( ! [X0] : (s(bool,X0) = s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK2))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | p(s(bool,X0))) ) | ($spl2 | $spl48)),
  inference(superposition,[],[f2017,f201])).
fof(f2508,plain,(
  $false | ($spl2 | $spl38 | $spl48)),
  inference(subsumption_resolution,[],[f2443,f237])).
fof(f2443,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK2)))) | ($spl2 | $spl38 | $spl48)),
  inference(backward_demodulation,[],[f2017,f1497])).
fof(f1497,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | ($spl38 | $spl48)),
  inference(backward_demodulation,[],[f1176,f588])).
fof(f2168,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) | $spl66),
  inference(cnf_transformation,[],[f2168_D])).
fof(f2168_D,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2166,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl64),
  inference(forward_demodulation,[],[f2087,f1820])).
fof(f2087,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))))))))) | $spl64),
  inference(superposition,[],[f238,f1820])).
fof(f1874,plain,(
  $false | ($spl34 | $spl36)),
  inference(equality_resolution,[],[f1425])).
fof(f1425,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int',X5)))) ) | ($spl34 | $spl36)),
  inference(subsumption_resolution,[],[f420,f419])).
fof(f1818,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl62),
  inference(forward_demodulation,[],[f1733,f1641])).
fof(f1733,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12('const.integer.int_of_num_1'(s('type.num.num',sK12(sK3))))))) | $spl62),
  inference(superposition,[],[f238,f1641])).
fof(f1637,plain,(
  s('type.integer.int',sK2) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | $spl48),
  inference(forward_demodulation,[],[f1560,f1176])).
fof(f1560,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_|48|_0')))) | s('type.integer.int',sK3) = s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num',sK12(sK3)))) | $spl48),
  inference(superposition,[],[f238,f1176])).
fof(f1488,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl56),
  inference(cnf_transformation,[],[f1488_D])).
fof(f1488_D,plain,(
  s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1490,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | $spl58),
  inference(cnf_transformation,[],[f1490_D])).
fof(f1490_D,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1427,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl27),
  inference(resolution,[],[f327,f235])).
fof(f1428,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) | s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl27),
  inference(resolution,[],[f327,f235])).
fof(f1486,plain,(
  $false | ($spl27 | $spl32)),
  inference(subsumption_resolution,[],[f1426,f418])).
fof(f1363,plain,(
  $false | ($spl30 | $spl32)),
  inference(trivial_inequality_removal,[],[f1355])).
fof(f1355,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ($spl30 | $spl32)),
  inference(resolution,[],[f330,f418])).
fof(f1332,plain,(
  $false | ($spl28 | $spl30)),
  inference(equality_resolution,[],[f1310])).
fof(f1310,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) ) | ($spl28 | $spl30)),
  inference(subsumption_resolution,[],[f330,f329])).
fof(f1203,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl39),
  inference(resolution,[],[f587,f235])).
fof(f1202,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',sK1)))) | s('type.integer.int',sK1) = s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))) | $spl39),
  inference(resolution,[],[f587,f235])).
fof(f1153,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | $spl33),
  inference(resolution,[],[f417,f235])).
fof(f1152,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK3),s('type.integer.int',sK2)))) | s('type.integer.int',sK2) = s('type.integer.int',sK3) | $spl33),
  inference(resolution,[],[f417,f235])).
fof(f1145,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) ) | $spl3),
  inference(superposition,[],[f308,f201])).
fof(f1146,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl3),
  inference(superposition,[],[f308,f201])).
fof(f1147,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | $spl3),
  inference(superposition,[],[f308,f204])).
fof(f1148,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) ) | $spl3),
  inference(superposition,[],[f308,f204])).
fof(f1151,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f1146,f330])).
fof(f1149,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) ) | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f1145,f330])).
fof(f578,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) ) | $spl26),
  inference(superposition,[],[f328,f227])).
fof(f577,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) ) | $spl26),
  inference(superposition,[],[f328,f225])).
fof(f567,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1)))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl26),
  inference(superposition,[],[f328,f228])).
fof(f552,plain,(
  ( ! [X13] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) ) | $spl26),
  inference(superposition,[],[f328,f227])).
fof(f551,plain,(
  ( ! [X12] : (p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) ) | $spl26),
  inference(superposition,[],[f328,f225])).
fof(f306,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl1),
  inference(cnf_transformation,[],[f306_D])).
fof(f306_D,plain,(
  s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f541,plain,(
  p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl26),
  inference(superposition,[],[f328,f228])).
fof(f420,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int',X5))) | ~p(s(bool,X4))) ) | $spl36),
  inference(cnf_transformation,[],[f420_D])).
fof(f420_D,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int',X5))) | ~p(s(bool,X4))) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f390,plain,(
  ( ! [X6,X7] : (s(bool,X6) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X7),s('type.integer.int',X7))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | ~p(s(bool,X6))) ) | $spl28),
  inference(superposition,[],[f366,f204])).
fof(f366,plain,(
  ( ! [X55] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X55),s('type.integer.int',X55)))) ) | $spl28),
  inference(resolution,[],[f329,f237])).
fof(f389,plain,(
  ( ! [X4,X5] : (s(bool,X4) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X5),s('type.integer.int',X5))) | ~p(s(bool,X4)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | $spl28),
  inference(superposition,[],[f366,f204])).
fof(f419,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) ) | $spl34),
  inference(cnf_transformation,[],[f419_D])).
fof(f419_D,plain,(
  ( ! [X0,X1] : (p(s(bool,X0)) | s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f388,plain,(
  ( ! [X2,X3] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X3),s('type.integer.int',X3))) | p(s(bool,X2)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))))) ) | $spl28),
  inference(superposition,[],[f366,f201])).
fof(f387,plain,(
  ( ! [X0,X1] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',X1),s('type.integer.int',X1))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3)))) | p(s(bool,X0))) ) | $spl28),
  inference(superposition,[],[f366,f201])).
fof(f304,plain,(
  ( ! [X3] : (s(bool,X3) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | ~p(s(bool,X3))) )),
  inference(superposition,[],[f199,f204])).
fof(f303,plain,(
  ( ! [X2] : (s(bool,X2) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | ~p(s(bool,X2)) | ~p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f204])).
fof(f302,plain,(
  ( ! [X1] : (s(bool,X1) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,X1)) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))))) )),
  inference(superposition,[],[f199,f201])).
fof(f301,plain,(
  ( ! [X0] : (s(bool,X0) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) | p(s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) | p(s(bool,X0))) )),
  inference(superposition,[],[f199,f201])).
fof(f298,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3)))))) )),
  inference(superposition,[],[f199,f227])).
fof(f297,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f310,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) | $spl5),
  inference(cnf_transformation,[],[f310_D])).
fof(f310_D,plain,(
  s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f322,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK3))) | $spl19),
  inference(cnf_transformation,[],[f322_D])).
fof(f322_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK3))) <=> ~$spl19),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl19])])).
fof(f288,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK3))) | s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f226])).
fof(f287,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',sK1))) | s('type.integer.int',sK3) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f228])).
fof(f272,plain,(
  ( ! [X13] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X13) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',X13),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2)))))) )),
  inference(superposition,[],[f199,f227])).
fof(f271,plain,(
  ( ! [X12] : (s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0'))),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',X12) != s('type.integer.int','const.integer.int_add_2'(s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK2))),s('type.integer.int',X12)))) )),
  inference(superposition,[],[f199,f225])).
fof(f262,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',sK1) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f226])).
fof(f308,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | $spl3),
  inference(cnf_transformation,[],[f308_D])).
fof(f308_D,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f261,plain,(
  s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK2),s('type.integer.int',sK3))) != s(bool,'const.integer.int_lt_2'(s('type.integer.int',sK1),s('type.integer.int','const.integer.int_add_2'(s('type.integer.int',sK1),s('type.integer.int',sK3))))) | s('type.integer.int',sK2) != s('type.integer.int','const.integer.int_of_num_1'(s('type.num.num','const.num.|48|_0')))),
  inference(superposition,[],[f199,f228])).
% SZS output end Proof for thm.integer
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 706
% Passive clauses: 18588
% Generated clauses: 43544
% Final active clauses: 130
% Final passive clauses: 6055
% Input formulas: 133
% Initial clauses: 54
% 
% Selected by SInE selection: 24
% SInE iterations: 1
% 
% Duplicate literals: 616
% Trivial inequalities: 445
% Fw subsumption resolutions: 1325
% Fw demodulations: 5515
% Bw demodulations: 263
% 
% Simple tautologies: 640
% Equational tautologies: 23
% Forward subsumptions: 15990
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 3375
% Forward superposition: 23890
% Backward superposition: 6120
% Self superposition: 10
% Equality resolution: 77
% 
% Split clauses: 513
% Split components: 1034
% SAT solver clauses: 619
% SAT solver unit clauses: 7
% SAT solver binary clauses: 92
% 
% Sat splits: 513
% Sat splitting refutations: 94
% 
% Memory used [KB]: 14328
% Time elapsed: 0.826 s
% ------------------------------
% ------------------------------
% Success in time 12.575 s
