Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (108535)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 621
% Passive clauses: 8238
% Generated clauses: 18100
% Final active clauses: 619
% Final passive clauses: 7226
% Input formulas: 45
% Initial clauses: 164
% 
% Duplicate literals: 323
% Trivial inequalities: 81
% Fw subsumption resolutions: 1677
% Fw demodulations: 565
% Bw demodulations: 2
% 
% Simple tautologies: 202
% Equational tautologies: 58
% Forward subsumptions: 6699
% 
% Binary resolution: 8863
% Factoring: 1135
% Forward superposition: 2835
% Backward superposition: 2226
% Self superposition: 33
% Equality factoring: 26
% Equality resolution: 25
% 
% Split clauses: 322
% Split components: 651
% SAT solver clauses: 603
% SAT solver unit clauses: 22
% SAT solver binary clauses: 175
% 
% Sat splits: 322
% 
% TWLsolver clauses: 603
% TWLsolver calls for satisfiability: 368
% 
% Memory used [KB]: 10362
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (108559)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 387
% Passive clauses: 2585
% Generated clauses: 11965
% Final active clauses: 376
% Final passive clauses: 2161
% Input formulas: 45
% Initial clauses: 164
% 
% Duplicate literals: 187
% Trivial inequalities: 26
% Fw subsumption resolutions: 877
% Fw demodulations: 755
% Bw demodulations: 41
% Condensations: 318
% Global subsumptions: 14
% 
% Simple tautologies: 224
% Equational tautologies: 23
% Forward subsumptions: 4061
% Bw demodulations to eq. taut.: 1
% 
% Binary resolution: 3269
% Factoring: 204
% Forward superposition: 1792
% Backward superposition: 1644
% Self superposition: 41
% Equality factoring: 28
% Equality resolution: 8
% 
% Split clauses: 52
% Split components: 110
% SAT solver clauses: 3067
% SAT solver unit clauses: 174
% SAT solver binary clauses: 80
% SAT solver learnt clauses: 28
% SAT solver learnt literals: 10
% 
% Sat splits: 52
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 3011
% TWLsolver calls for satisfiability: 1490
% 
% Memory used [KB]: 13560
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (108576)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 2738
% Passive clauses: 12609
% Generated clauses: 189947
% Final active clauses: 1387
% Final passive clauses: 4651
% Input formulas: 45
% Initial clauses: 164
% 
% Duplicate literals: 4916
% Trivial inequalities: 329
% Fw subsumption resolutions: 8622
% Fw demodulations: 26443
% Bw demodulations: 2581
% Condensations: 15767
% Global subsumptions: 1540
% 
% Simple tautologies: 302
% Equational tautologies: 1582
% Forward subsumptions: 113968
% Fw demodulations to eq. taut.: 40
% Bw demodulations to eq. taut.: 20
% 
% Unit resulting resolution: 33055
% Factoring: 1
% Forward superposition: 36769
% Backward superposition: 54301
% Self superposition: 285
% Equality factoring: 188
% Equality resolution: 33
% 
% Split clauses: 235
% Split components: 476
% SAT solver clauses: 18236
% SAT solver unit clauses: 1007
% SAT solver binary clauses: 5416
% 
% Sat splits: 235
% Sat splitting refutations: 3615
% 
% Memory used [KB]: 129336
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f10814,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f229,f396,f228,f398,f394,f227,f395,f397,f393,f454,f512,f559,f563,f561,f562,f564,f568,f566,f567,f1038,f1946,f1948,f2711,f2714,f2715,f2716,f2809,f2806,f2805,f2804,f2803,f2798,f2867,f2868,f2869,f3024,f3332,f3147,f3148,f3149,f3344,f3345,f3338,f3346,f3347,f3742,f3743,f3894,f3893,f3892,f3974,f3975,f3976,f3977,f3978,f3979,f3331,f4261,f3336,f4262,f4269,f4278,f4306,f4307,f4308,f4309,f4449,f4456,f4450,f4458,f4451,f4459,f4805,f5183,f5185,f5186,f5187,f5235,f5236,f5390,f5685,f5714,f5715,f5716,f5717,f5849,f5851,f6839,f6840,f6841,f6842,f6874,f6876,f6877,f6878,f7019,f7020,f7021,f7022,f7114,f7116,f7117,f7118,f7265,f7609,f3746,f7687,f7688,f7689,f7690,f8326,f8327,f8328,f8329,f8361,f8362,f8363,f8364,f8437,f8440,f8590,f1040,f9222,f3337,f9251,f9252,f9253,f9254,f9882,f9883,f9884,f9885,f5184,f9917,f9918,f9919,f9920,f7115,f10438,f10439,f10440,f10441,f2712,f10473,f10474,f10475,f10476,f2713,f6875,f10510,f10511,f10523,f10464,f10463,f10525,f10526,f10524,f10429,f10428,f3343,f10527,f10689,f2913,f1039,f1947,f10812,f10813])).
fof(f10813,plain,(
  $false | ($spl5 | $spl50)),
  inference(subsumption_resolution,[],[f10802,f382])).
fof(f382,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))) )),
  inference(cnf_transformation,[],[f105])).
fof(f105,plain,(
  ! [X0,X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(flattening,[],[f104])).
fof(f104,plain,(
  ! [X0] : ! [X1] : ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.EMPTY_0'))))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X6] : ! [X12] : ~p(s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),'const.pred_set.EMPTY_0'))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pred_set.NOT_IN_EMPTY')).
fof(f10802,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,'const.pred_set.EMPTY_0',sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')))) | ($spl5 | $spl50)),
  inference(trivial_inequality_removal,[],[f10796])).
fof(f10796,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,'const.pred_set.EMPTY_0',sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')))) | ($spl5 | $spl50)),
  inference(superposition,[],[f10331,f272])).
fof(f272,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f64])).
fof(f64,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f63])).
fof(f63,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.EQ_CLAUSES|split|2')).
fof(f10331,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,'const.pred_set.EMPTY_0',sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0'))) | ($spl5 | $spl50)),
  inference(equality_resolution,[],[f9922])).
fof(f9922,plain,(
  ( ! [X0] : (s(fun(sK10,bool),'const.pred_set.EMPTY_0') != s(fun(sK10,bool),X0) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),X0)))) ) | ($spl5 | $spl50)),
  inference(backward_demodulation,[],[f5851,f513])).
fof(f513,plain,(
  ( ! [X0] : (s(fun(sK10,bool),'const.pred_set.EMPTY_0') != s(fun(sK10,bool),X0) | s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),X0)))) ) | $spl5),
  inference(superposition,[],[f398,f386])).
fof(f386,plain,(
  ( ! [X2,X0,X1] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) | s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X2)))) )),
  inference(cnf_transformation,[],[f226])).
fof(f226,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK17(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK17])],[f225])).
fof(f225,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f224])).
fof(f224,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f109])).
fof(f109,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f108])).
fof(f108,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X6] : ! [X9,X8] : (s(fun(X6,bool),X8) = s(fun(X6,bool),X9) <=> ! [X12] : s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X8))) = s(bool,'const.bool.IN_2'(s(X6,X12),s(fun(X6,bool),X9))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pred_set.EXTENSION')).
fof(f10812,plain,(
  $false | ($spl5 | $spl50)),
  inference(subsumption_resolution,[],[f10803,f382])).
fof(f10803,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,'const.pred_set.EMPTY_0',sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')))) | ($spl5 | $spl50)),
  inference(trivial_inequality_removal,[],[f10795])).
fof(f10795,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,'const.pred_set.EMPTY_0',sK10)),s(fun(sK10,bool),'const.pred_set.EMPTY_0')))) | ($spl5 | $spl50)),
  inference(superposition,[],[f10331,f270])).
fof(f270,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f172])).
fof(f172,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f61])).
fof(f61,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X8] : (s(bool,'F_0') = s(bool,X8) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.EQ_CLAUSES|split|3')).
fof(f1947,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl18),
  inference(cnf_transformation,[],[f1947_D])).
fof(f1947_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f1039,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12))))) ) | $spl14),
  inference(cnf_transformation,[],[f1039_D])).
fof(f1039_D,plain,(
  ( ! [X1] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12))))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2913,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl26),
  inference(cnf_transformation,[],[f2913_D])).
fof(f2913_D,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f10689,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12)))) | p(s(bool,X2))) ) | $spl46),
  inference(resolution,[],[f6436,f237])).
fof(f237,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1)) | p(s(bool,X2)) | ~sP2(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f161])).
fof(f161,plain,(
  ! [X0,X1,X2] : ((~sP2(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP2(X0,X1,X2)))),
  inference(rectify,[],[f160])).
fof(f160,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(flattening,[],[f159])).
fof(f159,plain,(
  ! [X2,X0,X1] : ((~sP2(X2,X0,X1) | (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2)))) & ((~p(s(bool,X1)) & p(s(bool,X0)) & p(s(bool,X2))) | sP2(X2,X0,X1)))),
  inference(nnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X2,X0,X1] : (sP2(X2,X0,X1) <=> (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP2])])).
fof(f6436,plain,(
  ( ! [X50,X48,X49] : (sP2(X49,'const.bool.IN_2'(s(sK11,X48),s(fun(sK11,bool),sK12)),X50)) ) | $spl46),
  inference(subsumption_resolution,[],[f6240,f375])).
fof(f375,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f101])).
fof(f101,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.CROSS_EMPTY_EQN',aHL_FALSITY)).
fof(f6240,plain,(
  ( ! [X50,X48,X49] : (p(s(bool,'F_0')) | sP2(X49,'const.bool.IN_2'(s(sK11,X48),s(fun(sK11,bool),sK12)),X50)) ) | $spl46),
  inference(superposition,[],[f239,f4459])).
fof(f239,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f161])).
fof(f10527,plain,(
  ( ! [X35] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X35),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X35,sK10)),s(fun(sK10,bool),X35))))) ) | $spl62),
  inference(cnf_transformation,[],[f10527_D])).
fof(f10527_D,plain,(
  ( ! [X35] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X35),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X35,sK10)),s(fun(sK10,bool),X35))))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f3343,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl35),
  inference(cnf_transformation,[],[f3343_D])).
fof(f3343_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl35),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl35])])).
fof(f10428,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X34),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X34,sK10)),s(fun(sK10,bool),X34))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9921,f274])).
fof(f274,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f174])).
fof(f174,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.EQ_CLAUSES|split|0')).
fof(f9921,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X0),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl50)),
  inference(backward_demodulation,[],[f5851,f515])).
fof(f515,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X0),s(fun(sK11,bool),sK12)))) ) | $spl0),
  inference(superposition,[],[f393,f386])).
fof(f10429,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X35),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X35,sK10)),s(fun(sK10,bool),X35))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9921,f276])).
fof(f276,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f175])).
fof(f175,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X8] : (s(bool,'T_0') = s(bool,X8) <=> p(s(bool,X8)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.EQ_CLAUSES|split|1')).
fof(f10524,plain,(
  ( ! [X36] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X36),s(fun(sK11,bool),sK12)))) ) | $spl58),
  inference(cnf_transformation,[],[f10524_D])).
fof(f10524_D,plain,(
  ( ! [X36] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X36),s(fun(sK11,bool),sK12)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f10526,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X36),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10430,f9921])).
fof(f10430,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X36),s(fun(sK11,bool),sK12))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X36,sK10)),s(fun(sK10,bool),X36)))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9921,f277])).
fof(f277,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.CROSS_EMPTY_EQN',aHL_BOOL_CASES)).
fof(f10525,plain,(
  ( ! [X35] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X35),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X35,sK13,sK10)),s(fun(sK10,bool),X35))))) ) | $spl60),
  inference(cnf_transformation,[],[f10525_D])).
fof(f10525_D,plain,(
  ( ! [X35] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X35),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X35,sK13,sK10)),s(fun(sK10,bool),X35))))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f10463,plain,(
  ( ! [X34] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X34),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X34,sK13,sK10)),s(fun(sK10,bool),X34))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9923,f274])).
fof(f9923,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(X1,sK13,sK10)),s(fun(sK10,bool),X1))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X1),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl50)),
  inference(backward_demodulation,[],[f5851,f516])).
fof(f516,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK10,sK17(X1,sK13,sK10)),s(fun(sK10,bool),sK13))) != s(bool,'const.bool.IN_2'(s(sK10,sK17(X1,sK13,sK10)),s(fun(sK10,bool),X1))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X1),s(fun(sK11,bool),sK12)))) ) | $spl0),
  inference(superposition,[],[f393,f386])).
fof(f10464,plain,(
  ( ! [X35] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X35),s(fun(sK11,bool),sK12))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X35,sK13,sK10)),s(fun(sK10,bool),X35))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9923,f276])).
fof(f10523,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X36),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10465,f9923])).
fof(f10465,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'F_0') | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X36),s(fun(sK11,bool),sK12))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,sK17(X36,sK13,sK10)),s(fun(sK10,bool),X36)))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9923,f277])).
fof(f10511,plain,(
  $false | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f10500,f382])).
fof(f10500,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,'const.pred_set.EMPTY_0',sK11)),s(fun(sK11,bool),'const.pred_set.EMPTY_0')))) | ($spl3 | $spl30)),
  inference(trivial_inequality_removal,[],[f10494])).
fof(f10494,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,'const.pred_set.EMPTY_0',sK11)),s(fun(sK11,bool),'const.pred_set.EMPTY_0')))) | ($spl3 | $spl30)),
  inference(superposition,[],[f9723,f272])).
fof(f9723,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,'const.pred_set.EMPTY_0',sK11)),s(fun(sK11,bool),'const.pred_set.EMPTY_0'))) | ($spl3 | $spl30)),
  inference(equality_resolution,[],[f9373])).
fof(f9373,plain,(
  ( ! [X0] : (s(fun(sK11,bool),'const.pred_set.EMPTY_0') != s(fun(sK11,bool),X0) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),X0)))) ) | ($spl3 | $spl30)),
  inference(backward_demodulation,[],[f4303,f455])).
fof(f455,plain,(
  ( ! [X0] : (s(fun(sK11,bool),'const.pred_set.EMPTY_0') != s(fun(sK11,bool),X0) | s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),sK12))) != s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),X0)))) ) | $spl3),
  inference(superposition,[],[f396,f386])).
fof(f4303,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,X36),s(fun(sK11,bool),sK12)))) ) | $spl30),
  inference(trivial_inequality_removal,[],[f4301])).
fof(f4301,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,X36),s(fun(sK11,bool),sK12)))) ) | $spl30),
  inference(superposition,[],[f3337,f277])).
fof(f10510,plain,(
  $false | ($spl3 | $spl30)),
  inference(subsumption_resolution,[],[f10501,f382])).
fof(f10501,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,'const.pred_set.EMPTY_0',sK11)),s(fun(sK11,bool),'const.pred_set.EMPTY_0')))) | ($spl3 | $spl30)),
  inference(trivial_inequality_removal,[],[f10493])).
fof(f10493,plain,(
  s(bool,'F_0') != s(bool,'F_0') | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,'const.pred_set.EMPTY_0',sK11)),s(fun(sK11,bool),'const.pred_set.EMPTY_0')))) | ($spl3 | $spl30)),
  inference(superposition,[],[f9723,f270])).
fof(f6875,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) | $spl52),
  inference(cnf_transformation,[],[f6875_D])).
fof(f6875_D,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2713,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17)) ) | $spl24),
  inference(cnf_transformation,[],[f2713_D])).
fof(f2713_D,plain,(
  ( ! [X17] : (s(bool,'F_0') != s(bool,X17)) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f10476,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X26,sK13,sK10)),s(fun(sK10,bool),X26))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10458,f269])).
fof(f269,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f172])).
fof(f10458,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X26,sK13,sK10)),s(fun(sK10,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9923,f261])).
fof(f261,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X10,X11] : ((p(s(bool,X10)) => p(s(bool,X11))) => ((p(s(bool,X11)) => p(s(bool,X10))) => s(bool,X10) = s(bool,X11)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.IMP_ANTISYM_AX')).
fof(f10475,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X24,sK13,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10457,f269])).
fof(f10457,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X24,sK13,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9923,f261])).
fof(f10474,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X18,sK13,sK10)),s(fun(sK10,bool),X18))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10454,f269])).
fof(f10454,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X18,sK13,sK10)),s(fun(sK10,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9923,f257])).
fof(f257,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ((s(bool,X0) != s(bool,X1) | ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0))))) & (((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0)))) | s(bool,X0) = s(bool,X1)))),
  inference(nnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((~p(s(bool,X0)) | p(s(bool,X1))) & (~p(s(bool,X1)) | p(s(bool,X0)))))),
  inference(ennf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0,X1] : (s(bool,X0) = s(bool,X1) <=> ((p(s(bool,X0)) => p(s(bool,X1))) & (p(s(bool,X1)) => p(s(bool,X0)))))),
  inference(rectify,[],[f27])).
fof(f27,axiom,(
  ! [X10,X11] : (s(bool,X10) = s(bool,X11) <=> ((p(s(bool,X10)) => p(s(bool,X11))) & (p(s(bool,X11)) => p(s(bool,X10)))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.EQ_IMP_THM')).
fof(f10473,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10453,f269])).
fof(f10453,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9923,f257])).
fof(f2712,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | $spl22),
  inference(cnf_transformation,[],[f2712_D])).
fof(f2712_D,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f10441,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10423,f269])).
fof(f10423,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9921,f261])).
fof(f10440,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10422,f269])).
fof(f10422,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9921,f261])).
fof(f10439,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10419,f269])).
fof(f10419,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9921,f257])).
fof(f10438,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl50)),
  inference(subsumption_resolution,[],[f10418,f269])).
fof(f10418,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl50)),
  inference(superposition,[],[f9921,f257])).
fof(f7115,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X16,sK12,sK11)),s(fun(sK11,bool),X16))))) ) | $spl54),
  inference(cnf_transformation,[],[f7115_D])).
fof(f7115_D,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X16,sK12,sK11)),s(fun(sK11,bool),X16))))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f9920,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X26,sK12,sK11)),s(fun(sK11,bool),X26))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9902,f269])).
fof(f9902,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X26,sK12,sK11)),s(fun(sK11,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9374,f261])).
fof(f9374,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK11,sK17(X1,sK12,sK11)),s(fun(sK11,bool),X1))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X1)))) ) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f4303,f518])).
fof(f518,plain,(
  ( ! [X1] : (s(bool,'const.bool.IN_2'(s(sK11,sK17(X1,sK12,sK11)),s(fun(sK11,bool),sK12))) != s(bool,'const.bool.IN_2'(s(sK11,sK17(X1,sK12,sK11)),s(fun(sK11,bool),X1))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X1)))) ) | $spl0),
  inference(superposition,[],[f393,f386])).
fof(f9919,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X24,sK12,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9901,f269])).
fof(f9901,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X24,sK12,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9374,f261])).
fof(f9918,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X18,sK12,sK11)),s(fun(sK11,bool),X18))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9898,f269])).
fof(f9898,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X18,sK12,sK11)),s(fun(sK11,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9374,f257])).
fof(f9917,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X16,sK12,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9897,f269])).
fof(f9897,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X16,sK12,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9374,f257])).
fof(f5184,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) | $spl48),
  inference(cnf_transformation,[],[f5184_D])).
fof(f5184_D,plain,(
  ( ! [X16] : (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f9885,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X26,sK11)),s(fun(sK11,bool),X26))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9867,f269])).
fof(f9867,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X26,sK11)),s(fun(sK11,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9372,f261])).
fof(f9372,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X0)))) ) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f4303,f517])).
fof(f517,plain,(
  ( ! [X0] : (s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),sK12))) != s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X0)))) ) | $spl0),
  inference(superposition,[],[f393,f386])).
fof(f9884,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X24,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9866,f269])).
fof(f9866,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X24,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9372,f261])).
fof(f9883,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X18,sK11)),s(fun(sK11,bool),X18))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9863,f269])).
fof(f9863,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X18,sK11)),s(fun(sK11,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9372,f257])).
fof(f9882,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f9862,f269])).
fof(f9862,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f9372,f257])).
fof(f9254,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X30),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f9241,f274])).
fof(f9241,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X30),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(superposition,[],[f3337,f264])).
fof(f264,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f117])).
fof(f9253,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X28),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X29))) ) | $spl30),
  inference(subsumption_resolution,[],[f9240,f274])).
fof(f9240,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,X29) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X28),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X29))) ) | $spl30),
  inference(superposition,[],[f3337,f264])).
fof(f9252,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X22),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f9237,f274])).
fof(f9237,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X22),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(superposition,[],[f3337,f260])).
fof(f260,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f169])).
fof(f9251,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X20),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X21))) ) | $spl30),
  inference(subsumption_resolution,[],[f9236,f274])).
fof(f9236,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X20),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f3337,f260])).
fof(f3337,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK11,X9),s(fun(sK11,bool),sK12)))) ) | $spl30),
  inference(cnf_transformation,[],[f3337_D])).
fof(f3337_D,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK11,X9),s(fun(sK11,bool),sK12)))) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f9222,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK11,X8),s(fun(sK11,bool),sK12))) | ~p(s(bool,X9))) ) | $spl14),
  inference(resolution,[],[f9214,f1048])).
fof(f1048,plain,(
  ( ! [X19,X20,X18] : (sP2(X18,'const.bool.IN_2'(s(sK11,X19),s(fun(sK11,bool),sK12)),X20)) ) | $spl14),
  inference(resolution,[],[f1039,f239])).
fof(f9214,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl14),
  inference(subsumption_resolution,[],[f9213,f242])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f164])).
fof(f164,plain,(
  ! [X0,X1,X2] : ((~sP1(X0,X1,X2) | p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))) & ((~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP1(X0,X1,X2)))),
  inference(rectify,[],[f163])).
fof(f163,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(flattening,[],[f162])).
fof(f162,plain,(
  ! [X1,X0,X2] : ((~sP1(X1,X0,X2) | (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1)))) & ((~p(s(bool,X2)) & p(s(bool,X0)) & p(s(bool,X1))) | sP1(X1,X0,X2)))),
  inference(nnf_transformation,[],[f139])).
fof(f139,plain,(
  ! [X1,X0,X2] : (sP1(X1,X0,X2) <=> (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP1])])).
fof(f9213,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl14),
  inference(subsumption_resolution,[],[f9212,f273])).
fof(f273,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f174])).
fof(f9212,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl14),
  inference(subsumption_resolution,[],[f9203,f3176])).
fof(f3176,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl14),
  inference(subsumption_resolution,[],[f3171,f246])).
fof(f246,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f166])).
fof(f166,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(flattening,[],[f165])).
fof(f165,plain,(
  ! [X0,X1,X2] : ((~sP0(X0,X1,X2) | (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0)))) & ((~p(s(bool,X2)) & ~p(s(bool,X1)) & ~p(s(bool,X0))) | sP0(X0,X1,X2)))),
  inference(nnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0,X1,X2] : (sP0(X0,X1,X2) <=> (p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP0])])).
fof(f3171,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | $spl14),
  inference(resolution,[],[f2874,f230])).
fof(f230,plain,(
  ( ! [X2,X0,X1] : (sP0(X2,X1,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f158])).
fof(f158,plain,(
  ! [X0,X1,X2] : ((~sP3(X0,X1,X2) | (sP0(X2,X1,X0) & sP1(X1,X2,X0) & sP2(X0,X2,X1) & (p(s(bool,X2)) | ~p(s(bool,X1)) | ~p(s(bool,X0))))) & (~sP0(X2,X1,X0) | ~sP1(X1,X2,X0) | ~sP2(X0,X2,X1) | (~p(s(bool,X2)) & p(s(bool,X1)) & p(s(bool,X0))) | sP3(X0,X1,X2)))),
  inference(rectify,[],[f157])).
fof(f157,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & (~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2))) | sP3(X2,X1,X0)))),
  inference(flattening,[],[f156])).
fof(f156,plain,(
  ! [X2,X1,X0] : ((~sP3(X2,X1,X0) | (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2))))) & ((~sP0(X0,X1,X2) | ~sP1(X1,X0,X2) | ~sP2(X2,X0,X1) | (~p(s(bool,X0)) & p(s(bool,X1)) & p(s(bool,X2)))) | sP3(X2,X1,X0)))),
  inference(nnf_transformation,[],[f141])).
fof(f141,plain,(
  ! [X2,X1,X0] : (sP3(X2,X1,X0) <=> (sP0(X0,X1,X2) & sP1(X1,X0,X2) & sP2(X2,X0,X1) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  introduced(predicate_definition_introduction,[new_symbols(naming,[sP3])])).
fof(f2874,plain,(
  ( ! [X6,X5] : (sP3(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl14),
  inference(superposition,[],[f1585,f249])).
fof(f249,plain,(
  ( ! [X2,X0,X1] : (sP3(X2,X1,X0) | s(bool,X0) = s(bool,X1) | p(s(bool,X2))) )),
  inference(cnf_transformation,[],[f167])).
fof(f167,plain,(
  ! [X0,X1,X2] : ((((p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (~p(s(bool,X2)) | s(bool,X0) != s(bool,X1))) | sP3(X2,X1,X0)) & (~sP3(X2,X1,X0) | ((~p(s(bool,X2)) | s(bool,X0) = s(bool,X1)) & (s(bool,X0) != s(bool,X1) | p(s(bool,X2))))))),
  inference(nnf_transformation,[],[f142])).
fof(f142,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> sP3(X2,X1,X0))),
  inference(definition_folding,[],[f50,f141,f140,f139,f138])).
fof(f50,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0,X1,X2] : ((p(s(bool,X2)) <=> s(bool,X0) = s(bool,X1)) <=> ((p(s(bool,X2)) | p(s(bool,X1)) | p(s(bool,X0))) & (p(s(bool,X2)) | ~p(s(bool,X0)) | ~p(s(bool,X1))) & (p(s(bool,X1)) | ~p(s(bool,X0)) | ~p(s(bool,X2))) & (p(s(bool,X0)) | ~p(s(bool,X1)) | ~p(s(bool,X2)))))),
  inference(rectify,[],[f41])).
fof(f41,axiom,(
  ! [X20,X21,X15] : ((p(s(bool,X15)) <=> s(bool,X20) = s(bool,X21)) <=> ((p(s(bool,X15)) | p(s(bool,X21)) | p(s(bool,X20))) & (p(s(bool,X15)) | ~p(s(bool,X20)) | ~p(s(bool,X21))) & (p(s(bool,X21)) | ~p(s(bool,X20)) | ~p(s(bool,X15))) & (p(s(bool,X20)) | ~p(s(bool,X21)) | ~p(s(bool,X15)))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.sat.sat_2|49|_prove_ltb')).
fof(f1585,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl14),
  inference(superposition,[],[f1066,f1064])).
fof(f1064,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,X60),s(fun(sK11,bool),sK12)))) ) | $spl14),
  inference(resolution,[],[f1039,f270])).
fof(f1066,plain,(
  ( ! [X62] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK11,X62),s(fun(sK11,bool),sK12)))) ) | $spl14),
  inference(resolution,[],[f1039,f273])).
fof(f9203,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | $spl14),
  inference(resolution,[],[f8465,f234])).
fof(f234,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | ~p(s(bool,X2)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f158])).
fof(f8465,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl14),
  inference(superposition,[],[f1585,f251])).
fof(f251,plain,(
  ( ! [X2,X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X2)) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f1040,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK13))))) ) | $spl16),
  inference(cnf_transformation,[],[f1040_D])).
fof(f1040_D,plain,(
  ( ! [X0] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK13))))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f8590,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f8557,f388])).
fof(f388,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))) )),
  inference(cnf_transformation,[],[f113])).
fof(f113,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(flattening,[],[f112])).
fof(f112,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.pair.FST_1'(s('type.pair.prod'(X1,X0),'const.pair.,_2'(s(X1,X2),s(X0,X3)))))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X7,X6] : ! [X12,X13] : s(X6,X12) = s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13)))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pair.FST')).
fof(f8557,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,'const.pair.FST_1'(s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X0),s(sK11,X1))))),s(fun(sK10,bool),sK13))))) ) | $spl0),
  inference(superposition,[],[f558,f387])).
fof(f387,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))) )),
  inference(cnf_transformation,[],[f111])).
fof(f111,plain,(
  ! [X0,X1,X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(flattening,[],[f110])).
fof(f110,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X3) = s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X2),s(X1,X3)))))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6,X7] : ! [X12,X13] : s(X7,X13) = s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X12),s(X7,X13)))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pair.SND')).
fof(f558,plain,(
  ( ! [X2] : (~p(s(bool,'const.bool.IN_2'(s(sK11,'const.pair.SND_1'(s('type.pair.prod'(sK10,sK11),X2))),s(fun(sK11,bool),sK12)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,'const.pair.FST_1'(s('type.pair.prod'(sK10,sK11),X2))),s(fun(sK10,bool),sK13))))) ) | $spl0),
  inference(subsumption_resolution,[],[f523,f382])).
fof(f523,plain,(
  ( ! [X2] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(sK10,sK11),X2),s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0')))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,'const.pair.SND_1'(s('type.pair.prod'(sK10,sK11),X2))),s(fun(sK11,bool),sK12)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,'const.pair.FST_1'(s('type.pair.prod'(sK10,sK11),X2))),s(fun(sK10,bool),sK13))))) ) | $spl0),
  inference(superposition,[],[f280,f393])).
fof(f280,plain,(
  ( ! [X4,X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))) | ~p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2))))) )),
  inference(cnf_transformation,[],[f177])).
fof(f177,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))))) & (~p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))))))))),
  inference(flattening,[],[f176])).
fof(f176,plain,(
  ! [X0,X1,X2,X3,X4] : ((~p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3)))))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) | ~p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3))))))))),
  inference(nnf_transformation,[],[f68])).
fof(f68,plain,(
  ! [X0,X1,X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))))),
  inference(flattening,[],[f67])).
fof(f67,plain,(
  ! [X0,X1] : ! [X2,X3,X4] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X0,X1),X4),s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,'const.pair.FST_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X0,bool),X2)))) & p(s(bool,'const.bool.IN_2'(s(X1,'const.pair.SND_1'(s('type.pair.prod'(X0,X1),X4))),s(fun(X1,bool),X3))))))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X6,X7] : ! [X14,X18,X12] : (p(s(bool,'const.bool.IN_2'(s('type.pair.prod'(X6,X7),X12),s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),X14),s(fun(X7,bool),X18)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,'const.pair.FST_1'(s('type.pair.prod'(X6,X7),X12))),s(fun(X6,bool),X14)))) & p(s(bool,'const.bool.IN_2'(s(X7,'const.pair.SND_1'(s('type.pair.prod'(X6,X7),X12))),s(fun(X7,bool),X18))))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pred_set.IN_CROSS')).
fof(f8440,plain,(
  ( ! [X14] : (s(bool,'F_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X14)))) ) | $spl56),
  inference(cnf_transformation,[],[f8440_D])).
fof(f8440_D,plain,(
  ( ! [X14] : (s(bool,'F_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X14)))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f8437,plain,(
  ( ! [X14,X13] : (p(s(bool,X13)) | s(bool,'F_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X14)))) ) | ($spl12 | $spl16)),
  inference(resolution,[],[f7446,f4990])).
fof(f4990,plain,(
  ( ! [X37,X38,X36] : (sP0(happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X36)),X37,X38)) ) | $spl12),
  inference(resolution,[],[f567,f248])).
fof(f248,plain,(
  ( ! [X2,X0,X1] : (sP0(X0,X1,X2) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f166])).
fof(f7446,plain,(
  ( ! [X2,X3] : (~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl16),
  inference(subsumption_resolution,[],[f7440,f240])).
fof(f240,plain,(
  ( ! [X2,X0,X1] : (sP2(X0,X1,X2) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f161])).
fof(f7440,plain,(
  ( ! [X2,X3] : (~sP2(X2,X3,'F_0') | ~sP0(X3,'F_0',X2) | p(s(bool,X2)) | s(bool,'F_0') != s(bool,X3)) ) | $spl16),
  inference(resolution,[],[f3683,f252])).
fof(f252,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,X2)) | s(bool,X0) != s(bool,X1) | ~sP3(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f167])).
fof(f3683,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | $spl16),
  inference(subsumption_resolution,[],[f3681,f375])).
fof(f3681,plain,(
  ( ! [X4,X5] : (sP3(X4,'F_0',X5) | p(s(bool,'F_0')) | ~sP2(X4,X5,'F_0') | ~sP0(X5,'F_0',X4)) ) | $spl16),
  inference(resolution,[],[f3678,f235])).
fof(f235,plain,(
  ( ! [X2,X0,X1] : (sP3(X0,X1,X2) | p(s(bool,X1)) | ~sP2(X0,X2,X1) | ~sP1(X1,X2,X0) | ~sP0(X2,X1,X0)) )),
  inference(cnf_transformation,[],[f158])).
fof(f3678,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5)) ) | $spl16),
  inference(subsumption_resolution,[],[f3672,f242])).
fof(f3672,plain,(
  ( ! [X4,X5] : (sP1('F_0',X4,X5) | p(s(bool,X5))) ) | $spl16),
  inference(resolution,[],[f3669,f240])).
fof(f3669,plain,(
  ( ! [X6,X7] : (~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | $spl16),
  inference(subsumption_resolution,[],[f3664,f243])).
fof(f243,plain,(
  ( ! [X2,X0,X1] : (sP1(X0,X1,X2) | p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f164])).
fof(f3664,plain,(
  ( ! [X6,X7] : (~p(s(bool,X6)) | ~sP2(X7,X6,'F_0') | sP1('F_0',X6,X7)) ) | $spl16),
  inference(resolution,[],[f3658,f231])).
fof(f231,plain,(
  ( ! [X2,X0,X1] : (sP1(X1,X2,X0) | ~sP3(X0,X1,X2)) )),
  inference(cnf_transformation,[],[f158])).
fof(f3658,plain,(
  ( ! [X2,X3] : (sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl16),
  inference(subsumption_resolution,[],[f3657,f274])).
fof(f3657,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl16),
  inference(subsumption_resolution,[],[f3654,f3618])).
fof(f3618,plain,(
  ( ! [X4,X5] : (sP0(X4,'F_0',X5) | s(bool,'T_0') != s(bool,X4)) ) | $spl16),
  inference(subsumption_resolution,[],[f3613,f246])).
fof(f3613,plain,(
  ( ! [X4,X5] : (s(bool,'T_0') != s(bool,X4) | p(s(bool,X5)) | sP0(X4,'F_0',X5)) ) | $spl16),
  inference(resolution,[],[f3401,f230])).
fof(f3401,plain,(
  ( ! [X6,X5] : (sP3(X6,'F_0',X5) | s(bool,'T_0') != s(bool,X5) | p(s(bool,X6))) ) | $spl16),
  inference(superposition,[],[f2486,f249])).
fof(f2486,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl16),
  inference(superposition,[],[f1975,f1973])).
fof(f1973,plain,(
  ( ! [X60] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,X60),s(fun(sK10,bool),sK13)))) ) | $spl16),
  inference(resolution,[],[f1040,f270])).
fof(f1975,plain,(
  ( ! [X62] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,X62),s(fun(sK10,bool),sK13)))) ) | $spl16),
  inference(resolution,[],[f1040,f273])).
fof(f3654,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) ) | $spl16),
  inference(resolution,[],[f3619,f234])).
fof(f3619,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl16),
  inference(subsumption_resolution,[],[f3614,f242])).
fof(f3614,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl16),
  inference(resolution,[],[f3401,f231])).
fof(f8364,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X26,sK13,sK10)),s(fun(sK10,bool),X26))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8346,f269])).
fof(f8346,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X26,sK13,sK10)),s(fun(sK10,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7820,f261])).
fof(f7820,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(X1,sK13,sK10)),s(fun(sK10,bool),X1))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X1),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl38)),
  inference(backward_demodulation,[],[f5711,f516])).
fof(f5711,plain,(
  ( ! [X36] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,X36),s(fun(sK10,bool),sK13)))) ) | $spl38),
  inference(trivial_inequality_removal,[],[f5709])).
fof(f5709,plain,(
  ( ! [X36] : (s(bool,'T_0') != s(bool,'T_0') | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,X36),s(fun(sK10,bool),sK13)))) ) | $spl38),
  inference(superposition,[],[f3746,f277])).
fof(f8363,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X24,sK13,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8345,f269])).
fof(f8345,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X24,sK13,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7820,f261])).
fof(f8362,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X18,sK13,sK10)),s(fun(sK10,bool),X18))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8342,f269])).
fof(f8342,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X18,sK13,sK10)),s(fun(sK10,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7820,f257])).
fof(f8361,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8341,f269])).
fof(f8341,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7820,f257])).
fof(f8329,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8311,f269])).
fof(f8311,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7818,f261])).
fof(f7818,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X0),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl38)),
  inference(backward_demodulation,[],[f5711,f515])).
fof(f8328,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8310,f269])).
fof(f8310,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7818,f261])).
fof(f8327,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8307,f269])).
fof(f8307,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7818,f257])).
fof(f8326,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f8306,f269])).
fof(f8306,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f7818,f257])).
fof(f7690,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X30),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(subsumption_resolution,[],[f7677,f274])).
fof(f7677,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X30),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(superposition,[],[f3746,f264])).
fof(f7689,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X28),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X29))) ) | $spl38),
  inference(subsumption_resolution,[],[f7676,f274])).
fof(f7676,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,X29) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X28),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X29))) ) | $spl38),
  inference(superposition,[],[f3746,f264])).
fof(f7688,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X22),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(subsumption_resolution,[],[f7673,f274])).
fof(f7673,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X22),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(superposition,[],[f3746,f260])).
fof(f7687,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X20),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X21))) ) | $spl38),
  inference(subsumption_resolution,[],[f7672,f274])).
fof(f7672,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X20),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X21))) ) | $spl38),
  inference(superposition,[],[f3746,f260])).
fof(f3746,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,X9),s(fun(sK10,bool),sK13)))) ) | $spl38),
  inference(cnf_transformation,[],[f3746_D])).
fof(f3746_D,plain,(
  ( ! [X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,X9),s(fun(sK10,bool),sK13)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f7609,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,X8),s(fun(sK10,bool),sK13))) | ~p(s(bool,X9))) ) | $spl16),
  inference(resolution,[],[f7601,f1957])).
fof(f1957,plain,(
  ( ! [X19,X20,X18] : (sP2(X18,'const.bool.IN_2'(s(sK10,X19),s(fun(sK10,bool),sK13)),X20)) ) | $spl16),
  inference(resolution,[],[f1040,f239])).
fof(f7601,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl16),
  inference(subsumption_resolution,[],[f7600,f242])).
fof(f7600,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl16),
  inference(subsumption_resolution,[],[f7599,f273])).
fof(f7599,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl16),
  inference(subsumption_resolution,[],[f7590,f3618])).
fof(f7590,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | $spl16),
  inference(resolution,[],[f7140,f234])).
fof(f7140,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl16),
  inference(superposition,[],[f2486,f251])).
fof(f7265,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f7232,f388])).
fof(f7232,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,'const.pair.FST_1'(s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X0),s(sK11,X1))))),s(fun(sK10,bool),sK13))))) ) | $spl0),
  inference(superposition,[],[f558,f387])).
fof(f7118,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X26,sK12,sK11)),s(fun(sK11,bool),X26))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f7099,f269])).
fof(f7099,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X26,sK12,sK11)),s(fun(sK11,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6199,f261])).
fof(f6199,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK11,sK17(X1,sK12,sK11)),s(fun(sK11,bool),X1))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X1)))) ) | ($spl0 | $spl46)),
  inference(backward_demodulation,[],[f4459,f518])).
fof(f7117,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X24,sK12,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f7098,f269])).
fof(f7098,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X24,sK12,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6199,f261])).
fof(f7116,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X18,sK12,sK11)),s(fun(sK11,bool),X18))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f7095,f269])).
fof(f7095,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X18,sK12,sK11)),s(fun(sK11,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6199,f257])).
fof(f7114,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X16,sK12,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f7094,f269])).
fof(f7094,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(X16,sK12,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6199,f257])).
fof(f7022,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X26,sK11)),s(fun(sK11,bool),X26))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f7004,f269])).
fof(f7004,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X26,sK11)),s(fun(sK11,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6197,f261])).
fof(f6197,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X0)))) ) | ($spl0 | $spl46)),
  inference(backward_demodulation,[],[f4459,f517])).
fof(f7021,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X24,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f7003,f269])).
fof(f7003,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X24,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6197,f261])).
fof(f7020,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X18,sK11)),s(fun(sK11,bool),X18))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f7000,f269])).
fof(f7000,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X18,sK11)),s(fun(sK11,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6197,f257])).
fof(f7019,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl46)),
  inference(subsumption_resolution,[],[f6999,f269])).
fof(f6999,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl46)),
  inference(superposition,[],[f6197,f257])).
fof(f6878,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X26,sK13,sK10)),s(fun(sK10,bool),X26))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6859,f269])).
fof(f6859,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X26,sK13,sK10)),s(fun(sK10,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5854,f261])).
fof(f5854,plain,(
  ( ! [X1] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(X1,sK13,sK10)),s(fun(sK10,bool),X1))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X1),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl38)),
  inference(backward_demodulation,[],[f5711,f516])).
fof(f6877,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X24,sK13,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6858,f269])).
fof(f6858,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X24,sK13,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5854,f261])).
fof(f6876,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X18,sK13,sK10)),s(fun(sK10,bool),X18))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6855,f269])).
fof(f6855,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X18,sK13,sK10)),s(fun(sK10,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5854,f257])).
fof(f6874,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6854,f269])).
fof(f6854,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(X16,sK13,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5854,f257])).
fof(f6842,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6824,f269])).
fof(f6824,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5852,f261])).
fof(f5852,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X0),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl38)),
  inference(backward_demodulation,[],[f5711,f515])).
fof(f6841,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6823,f269])).
fof(f6823,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5852,f261])).
fof(f6840,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6820,f269])).
fof(f6820,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5852,f257])).
fof(f6839,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(subsumption_resolution,[],[f6819,f269])).
fof(f6819,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl38)),
  inference(superposition,[],[f5852,f257])).
fof(f5851,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,X8),s(fun(sK10,bool),sK13)))) ) | $spl50),
  inference(cnf_transformation,[],[f5851_D])).
fof(f5851_D,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,X8),s(fun(sK10,bool),sK13)))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f5849,plain,(
  ( ! [X8,X9] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK10,X8),s(fun(sK10,bool),sK13))) | ~p(s(bool,X9))) ) | $spl16),
  inference(resolution,[],[f5513,f1957])).
fof(f5513,plain,(
  ( ! [X0,X1] : (~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | $spl16),
  inference(subsumption_resolution,[],[f5507,f270])).
fof(f5507,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | $spl16),
  inference(resolution,[],[f3658,f251])).
fof(f5717,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X30),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(subsumption_resolution,[],[f5704,f274])).
fof(f5704,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X30),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(superposition,[],[f3746,f264])).
fof(f5716,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X28),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X29))) ) | $spl38),
  inference(subsumption_resolution,[],[f5703,f274])).
fof(f5703,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,X29) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X28),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X29))) ) | $spl38),
  inference(superposition,[],[f3746,f264])).
fof(f5715,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X22),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(subsumption_resolution,[],[f5700,f274])).
fof(f5700,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X22),s(fun(sK10,bool),sK13))))) ) | $spl38),
  inference(superposition,[],[f3746,f260])).
fof(f5714,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X20),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X21))) ) | $spl38),
  inference(subsumption_resolution,[],[f5699,f274])).
fof(f5699,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,'const.bool.IN_2'(s(sK10,X20),s(fun(sK10,bool),sK13)))) | ~p(s(bool,X21))) ) | $spl38),
  inference(superposition,[],[f3746,f260])).
fof(f5685,plain,(
  ( ! [X8,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,X8),s(fun(sK10,bool),sK13))) | ~p(s(bool,X9))) ) | $spl16),
  inference(resolution,[],[f5677,f1957])).
fof(f5677,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl16),
  inference(subsumption_resolution,[],[f5676,f242])).
fof(f5676,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl16),
  inference(subsumption_resolution,[],[f5675,f273])).
fof(f5675,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl16),
  inference(subsumption_resolution,[],[f5666,f3618])).
fof(f5666,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | $spl16),
  inference(resolution,[],[f5265,f234])).
fof(f5265,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl16),
  inference(superposition,[],[f2486,f251])).
fof(f5390,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f5357,f388])).
fof(f5357,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,'const.pair.FST_1'(s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X0),s(sK11,X1))))),s(fun(sK10,bool),sK13))))) ) | $spl0),
  inference(superposition,[],[f558,f387])).
fof(f5236,plain,(
  $false | ($spl24 | $spl30)),
  inference(trivial_inequality_removal,[],[f5232])).
fof(f5232,plain,(
  s(bool,'F_0') != s(bool,'F_0') | ($spl24 | $spl30)),
  inference(superposition,[],[f2713,f4303])).
fof(f5235,plain,(
  $false | $spl24),
  inference(equality_resolution,[],[f2713])).
fof(f5187,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X26,sK11)),s(fun(sK11,bool),X26))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f5168,f269])).
fof(f5168,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X26))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X26,sK11)),s(fun(sK11,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f4460,f261])).
fof(f4460,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X0,sK11)),s(fun(sK11,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X0)))) ) | ($spl0 | $spl30)),
  inference(backward_demodulation,[],[f4303,f517])).
fof(f5186,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X24,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f5167,f269])).
fof(f5167,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X24))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X24,sK11)),s(fun(sK11,bool),X24))))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f4460,f261])).
fof(f5185,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X18,sK11)),s(fun(sK11,bool),X18))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f5164,f269])).
fof(f5164,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X18))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X18,sK11)),s(fun(sK11,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f4460,f257])).
fof(f5183,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl30)),
  inference(subsumption_resolution,[],[f5163,f269])).
fof(f5163,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),X16))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK11,sK17(sK12,X16,sK11)),s(fun(sK11,bool),X16))))) ) | ($spl0 | $spl30)),
  inference(superposition,[],[f4460,f257])).
fof(f4805,plain,(
  $false | ($spl8 | $spl11)),
  inference(resolution,[],[f562,f566])).
fof(f4459,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,X10),s(fun(sK11,bool),sK12)))) ) | $spl46),
  inference(cnf_transformation,[],[f4459_D])).
fof(f4459_D,plain,(
  ( ! [X10] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,X10),s(fun(sK11,bool),sK12)))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f4451,plain,(
  ( ! [X10,X11] : (s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK11,X10),s(fun(sK11,bool),sK12))) | ~p(s(bool,X11))) ) | $spl14),
  inference(resolution,[],[f4187,f1048])).
fof(f4187,plain,(
  ( ! [X0,X1] : (~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | $spl14),
  inference(subsumption_resolution,[],[f4181,f270])).
fof(f4181,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~sP2(X1,X0,'F_0') | s(bool,'F_0') = s(bool,X0) | ~p(s(bool,X1))) ) | $spl14),
  inference(resolution,[],[f3199,f251])).
fof(f3199,plain,(
  ( ! [X2,X3] : (sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl14),
  inference(subsumption_resolution,[],[f3198,f274])).
fof(f3198,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0')) ) | $spl14),
  inference(subsumption_resolution,[],[f3195,f3176])).
fof(f3195,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | sP3(X3,'F_0',X2) | ~p(s(bool,X2)) | ~sP2(X3,X2,'F_0') | ~sP0(X2,'F_0',X3)) ) | $spl14),
  inference(resolution,[],[f3177,f234])).
fof(f3177,plain,(
  ( ! [X6,X7] : (sP1('F_0',X6,X7) | s(bool,'T_0') != s(bool,X6)) ) | $spl14),
  inference(subsumption_resolution,[],[f3172,f242])).
fof(f3172,plain,(
  ( ! [X6,X7] : (s(bool,'T_0') != s(bool,X6) | p(s(bool,X7)) | sP1('F_0',X6,X7)) ) | $spl14),
  inference(resolution,[],[f2874,f231])).
fof(f4458,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))) | $spl44),
  inference(cnf_transformation,[],[f4458_D])).
fof(f4458_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f4450,plain,(
  ( ! [X9] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))) | ~p(s(bool,X9))) ) | ($spl7 | $spl14)),
  inference(resolution,[],[f4187,f594])).
fof(f594,plain,(
  ( ! [X8,X9] : (sP2(X8,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))),X9)) ) | $spl7),
  inference(resolution,[],[f561,f239])).
fof(f4456,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) | $spl42),
  inference(cnf_transformation,[],[f4456_D])).
fof(f4456_D,plain,(
  s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f4449,plain,(
  ( ! [X8] : (s(bool,'F_0') = s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) | ~p(s(bool,X8))) ) | ($spl7 | $spl14)),
  inference(resolution,[],[f4187,f807])).
fof(f807,plain,(
  ( ! [X8,X9] : (sP2(X8,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10))),X9)) ) | $spl7),
  inference(resolution,[],[f697,f239])).
fof(f697,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10))))) | $spl7),
  inference(resolution,[],[f589,f383])).
fof(f383,plain,(
  ( ! [X6,X2,X0,X7,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK14(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f223,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),sK14(X2,X1,X0))))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK15(X2,X1,X0)),s(X1,sK16(X2,X1,X0))))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK14,sK15,sK16])],[f222])).
fof(f222,plain,(
  ! [X0,X1,X2] : ((? [X8] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X8)))) | ! [X6,X7] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X6),s(X1,X7))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(rectify,[],[f221])).
fof(f221,plain,(
  ! [X0,X1,X2] : ((? [X3] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5))))))) & (? [X4,X5] : ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))) | ! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3))))))),
  inference(nnf_transformation,[],[f107])).
fof(f107,plain,(
  ! [X0,X1,X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(flattening,[],[f106])).
fof(f106,plain,(
  ! [X0,X1] : ! [X2] : (! [X3] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) <=> ! [X4,X5] : p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,X4),s(X1,X5)))))))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X6,X7] : ! [X14] : (! [X15] : p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X14),s('type.pair.prod'(X6,X7),X15)))) <=> ! [X16,X17] : p(s(bool,happ(s(fun('type.pair.prod'(X6,X7),bool),X14),s('type.pair.prod'(X6,X7),'const.pair.,_2'(s(X6,X16),s(X7,X17)))))))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pair.FORALL_PROD')).
fof(f589,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,sK15('const.pred_set.EMPTY_0',sK11,sK10)),s(sK11,sK16('const.pred_set.EMPTY_0',sK11,sK10))))))) | $spl7),
  inference(resolution,[],[f561,f384])).
fof(f384,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),X3)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(X0,X1),bool),X2),s('type.pair.prod'(X0,X1),'const.pair.,_2'(s(X0,sK15(X2,X1,X0)),s(X1,sK16(X2,X1,X0)))))))) )),
  inference(cnf_transformation,[],[f223])).
fof(f4309,plain,(
  ( ! [X30,X31] : (~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X30),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f4296,f274])).
fof(f4296,plain,(
  ( ! [X30,X31] : (s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X31)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X30),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(superposition,[],[f3337,f264])).
fof(f4308,plain,(
  ( ! [X28,X29] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X28),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X29))) ) | $spl30),
  inference(subsumption_resolution,[],[f4295,f274])).
fof(f4295,plain,(
  ( ! [X28,X29] : (s(bool,'T_0') != s(bool,X29) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X28),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X29))) ) | $spl30),
  inference(superposition,[],[f3337,f264])).
fof(f4307,plain,(
  ( ! [X23,X22] : (~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X22),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(subsumption_resolution,[],[f4292,f274])).
fof(f4292,plain,(
  ( ! [X23,X22] : (s(bool,'T_0') != s(bool,X23) | ~p(s(bool,X23)) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X22),s(fun(sK11,bool),sK12))))) ) | $spl30),
  inference(superposition,[],[f3337,f260])).
fof(f4306,plain,(
  ( ! [X21,X20] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X20),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X21))) ) | $spl30),
  inference(subsumption_resolution,[],[f4291,f274])).
fof(f4291,plain,(
  ( ! [X21,X20] : (s(bool,'T_0') != s(bool,X21) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X20),s(fun(sK11,bool),sK12)))) | ~p(s(bool,X21))) ) | $spl30),
  inference(superposition,[],[f3337,f260])).
fof(f4278,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))) | $spl41),
  inference(cnf_transformation,[],[f4278_D])).
fof(f4278_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))) <=> ~$spl41),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl41])])).
fof(f4269,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))) | ~p(s(bool,X8))) ) | ($spl7 | $spl14)),
  inference(resolution,[],[f594,f4253])).
fof(f4253,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl14),
  inference(subsumption_resolution,[],[f4252,f242])).
fof(f4252,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl14),
  inference(subsumption_resolution,[],[f4251,f273])).
fof(f4251,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12)) ) | $spl14),
  inference(subsumption_resolution,[],[f4242,f3176])).
fof(f4242,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | $spl14),
  inference(resolution,[],[f4002,f234])).
fof(f4002,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl14),
  inference(superposition,[],[f1585,f251])).
fof(f4262,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK11,X9),s(fun(sK11,bool),sK12))) | ~p(s(bool,X10))) ) | $spl14),
  inference(resolution,[],[f4253,f1048])).
fof(f3336,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) | $spl29),
  inference(cnf_transformation,[],[f3336_D])).
fof(f3336_D,plain,(
  s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f4261,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) | ~p(s(bool,X8))) ) | ($spl7 | $spl14)),
  inference(resolution,[],[f4253,f807])).
fof(f3331,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) | ~p(s(bool,X8))) ) | ($spl7 | $spl14)),
  inference(resolution,[],[f3323,f807])).
fof(f3323,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl14),
  inference(subsumption_resolution,[],[f3322,f273])).
fof(f3322,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0')) ) | $spl14),
  inference(subsumption_resolution,[],[f3321,f3176])).
fof(f3321,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP0(X13,'F_0',X12)) ) | $spl14),
  inference(subsumption_resolution,[],[f3312,f3177])).
fof(f3312,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | $spl14),
  inference(resolution,[],[f2876,f234])).
fof(f2876,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl14),
  inference(superposition,[],[f1585,f251])).
fof(f3979,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f3891,f2913])).
fof(f3891,plain,(
  p(s(bool,'F_0')) | $spl26),
  inference(resolution,[],[f2913,f378])).
fof(f378,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f220])).
fof(f220,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f103])).
fof(f103,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f17])).
fof(f17,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.NOT_CLAUSES|split|1')).
fof(f3978,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f3890,f2913])).
fof(f3890,plain,(
  p(s(bool,'F_0')) | $spl26),
  inference(resolution,[],[f2913,f376])).
fof(f376,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f219])).
fof(f219,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f102])).
fof(f102,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f18])).
fof(f18,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.NOT_CLAUSES|split|2')).
fof(f3977,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f3888,f2913])).
fof(f3888,plain,(
  p(s(bool,'F_0')) | $spl26),
  inference(resolution,[],[f2913,f366])).
fof(f366,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f218])).
fof(f218,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f217])).
fof(f217,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X8] : ((p(s(bool,'F_0')) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.IMP_CLAUSES|split|2')).
fof(f3976,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f3887,f2913])).
fof(f3887,plain,(
  ( ! [X178] : (p(s(bool,X178))) ) | $spl26),
  inference(resolution,[],[f2913,f363])).
fof(f363,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f216])).
fof(f216,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f215])).
fof(f215,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X8] : ((p(s(bool,'T_0')) => p(s(bool,X8))) <=> p(s(bool,X8)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.IMP_CLAUSES|split|0')).
fof(f3975,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f3885,f2913])).
fof(f3885,plain,(
  ( ! [X176] : (p(s(bool,X176))) ) | $spl26),
  inference(resolution,[],[f2913,f360])).
fof(f360,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f214])).
fof(f214,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f213])).
fof(f213,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f130])).
fof(f130,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,X8))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.IMP_CLAUSES|split|3')).
fof(f3974,plain,(
  $false | $spl26),
  inference(subsumption_resolution,[],[f3884,f2913])).
fof(f3884,plain,(
  ( ! [X175] : (p(s(bool,X175))) ) | $spl26),
  inference(resolution,[],[f2913,f357])).
fof(f357,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f212])).
fof(f212,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f211])).
fof(f211,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f92])).
fof(f92,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f12])).
fof(f12,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.IMP_CLAUSES|split|1')).
fof(f3892,plain,(
  $false | $spl26),
  inference(resolution,[],[f2913,f380])).
fof(f380,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f8])).
fof(f8,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.TRUTH')).
fof(f3893,plain,(
  $false | $spl26),
  inference(resolution,[],[f2913,f381])).
fof(f381,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.CROSS_EMPTY_EQN',aHL_TRUTH)).
fof(f3894,plain,(
  $false | $spl26),
  inference(resolution,[],[f2913,f389])).
fof(f389,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f267])).
fof(f267,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f171])).
fof(f171,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f60])).
fof(f60,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f59])).
fof(f59,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X6] : ! [X12] : (s(X6,X12) = s(X6,X12) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.REFL_CLAUSE')).
fof(f3743,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK10,X9),s(fun(sK10,bool),sK13))) | ~p(s(bool,X10))) ) | $spl16),
  inference(resolution,[],[f3734,f1957])).
fof(f3734,plain,(
  ( ! [X12,X13] : (~sP2(X12,X13,'F_0') | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X12))) ) | $spl16),
  inference(subsumption_resolution,[],[f3733,f273])).
fof(f3733,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0')) ) | $spl16),
  inference(subsumption_resolution,[],[f3732,f3618])).
fof(f3732,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP0(X13,'F_0',X12)) ) | $spl16),
  inference(subsumption_resolution,[],[f3723,f3619])).
fof(f3723,plain,(
  ( ! [X12,X13] : (~p(s(bool,X12)) | s(bool,'T_0') != s(bool,X13) | ~p(s(bool,X13)) | ~sP2(X12,X13,'F_0') | ~sP1('F_0',X13,X12) | ~sP0(X13,'F_0',X12)) ) | $spl16),
  inference(resolution,[],[f3403,f234])).
fof(f3403,plain,(
  ( ! [X10,X9] : (~sP3(X10,'F_0',X9) | ~p(s(bool,X10)) | s(bool,'T_0') != s(bool,X9)) ) | $spl16),
  inference(superposition,[],[f2486,f251])).
fof(f3742,plain,(
  ( ! [X8] : (s(bool,'T_0') != s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.EMPTY_0',sK11,sK10)))) | ~p(s(bool,X8))) ) | ($spl7 | $spl16)),
  inference(resolution,[],[f3734,f807])).
fof(f3347,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X30))) ) | $spl20),
  inference(subsumption_resolution,[],[f3135,f274])).
fof(f3135,plain,(
  ( ! [X30] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X30) | ~p(s(bool,X30))) ) | $spl20),
  inference(superposition,[],[f1948,f274])).
fof(f3346,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | ~p(s(bool,X31))) ) | $spl20),
  inference(subsumption_resolution,[],[f3136,f274])).
fof(f3136,plain,(
  ( ! [X31] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X31) | ~p(s(bool,X31))) ) | $spl20),
  inference(superposition,[],[f1948,f276])).
fof(f3338,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,X12)) ) | $spl32),
  inference(cnf_transformation,[],[f3338_D])).
fof(f3338_D,plain,(
  ( ! [X12] : (s(bool,'T_0') != s(bool,X12)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f3345,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X32)) ) | $spl20),
  inference(subsumption_resolution,[],[f3137,f1948])).
fof(f3137,plain,(
  ( ! [X32] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') != s(bool,X32) | s(bool,'F_0') = s(bool,X32)) ) | $spl20),
  inference(superposition,[],[f1948,f277])).
fof(f3344,plain,(
  ( ! [X37] : (s(bool,'T_0') = s(bool,X37)) ) | $spl36),
  inference(cnf_transformation,[],[f3344_D])).
fof(f3344_D,plain,(
  ( ! [X37] : (s(bool,'T_0') = s(bool,X37)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f3149,plain,(
  ( ! [X37] : (s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X37)) ) | $spl20),
  inference(trivial_inequality_removal,[],[f3139])).
fof(f3139,plain,(
  ( ! [X37] : (s(bool,X37) != s(bool,X37) | s(bool,'T_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,X37)) ) | $spl20),
  inference(superposition,[],[f1948,f277])).
fof(f3148,plain,(
  ( ! [X38] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X38))) ) | $spl20),
  inference(trivial_inequality_removal,[],[f3140])).
fof(f3140,plain,(
  ( ! [X38] : (s(bool,X38) != s(bool,X38) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X38))) ) | $spl20),
  inference(superposition,[],[f1948,f272])).
fof(f3147,plain,(
  ( ! [X39] : (s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X39))) ) | $spl20),
  inference(trivial_inequality_removal,[],[f3141])).
fof(f3141,plain,(
  ( ! [X39] : (s(bool,X39) != s(bool,X39) | s(bool,'T_0') != s(bool,'F_0') | p(s(bool,X39))) ) | $spl20),
  inference(superposition,[],[f1948,f270])).
fof(f3332,plain,(
  ( ! [X10,X9] : (s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK11,X9),s(fun(sK11,bool),sK12))) | ~p(s(bool,X10))) ) | $spl14),
  inference(resolution,[],[f3323,f1048])).
fof(f3024,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f2991,f388])).
fof(f2991,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,'const.pair.FST_1'(s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X0),s(sK11,X1))))),s(fun(sK10,bool),sK13))))) ) | $spl0),
  inference(superposition,[],[f558,f387])).
fof(f2869,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2800,f1947])).
fof(f2800,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f1947,f379])).
fof(f379,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f220])).
fof(f2868,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2799,f1947])).
fof(f2799,plain,(
  ~p(s(bool,'T_0')) | $spl18),
  inference(resolution,[],[f1947,f377])).
fof(f377,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f219])).
fof(f2867,plain,(
  $false | $spl18),
  inference(subsumption_resolution,[],[f2795,f1947])).
fof(f2795,plain,(
  ( ! [X180] : (~p(s(bool,X180))) ) | $spl18),
  inference(resolution,[],[f1947,f355])).
fof(f355,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f210])).
fof(f210,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f209])).
fof(f209,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f91])).
fof(f91,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f90])).
fof(f90,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f15])).
fof(f15,axiom,(
  ! [X8] : ((p(s(bool,X8)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X8)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.bool.IMP_CLAUSES|split|4')).
fof(f2798,plain,(
  $false | $spl18),
  inference(resolution,[],[f1947,f375])).
fof(f2803,plain,(
  $false | ($spl7 | $spl18)),
  inference(resolution,[],[f1947,f561])).
fof(f2804,plain,(
  $false | ($spl7 | $spl18)),
  inference(resolution,[],[f1947,f697])).
fof(f2805,plain,(
  $false | ($spl7 | $spl18)),
  inference(resolution,[],[f1947,f589])).
fof(f2806,plain,(
  $false | ($spl11 | $spl18)),
  inference(resolution,[],[f1947,f566])).
fof(f2809,plain,(
  $false | $spl18),
  inference(resolution,[],[f1947,f382])).
fof(f2716,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26))))) ) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f2696,f269])).
fof(f2696,plain,(
  ( ! [X26,X27] : (s(bool,'F_0') != s(bool,X27) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X26),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X26,sK10)),s(fun(sK10,bool),X26)))) | p(s(bool,X27))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f2142,f261])).
fof(f2142,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X0,sK10)),s(fun(sK10,bool),X0))) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X0),s(fun(sK11,bool),sK12)))) ) | ($spl0 | $spl16)),
  inference(backward_demodulation,[],[f1973,f515])).
fof(f2715,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f2695,f269])).
fof(f2695,plain,(
  ( ! [X24,X25] : (s(bool,'F_0') != s(bool,X25) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X24),s(fun(sK11,bool),sK12))) | p(s(bool,X25)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X24,sK10)),s(fun(sK10,bool),X24))))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f2142,f261])).
fof(f2714,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18))))) ) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f2692,f269])).
fof(f2692,plain,(
  ( ! [X19,X18] : (s(bool,'F_0') != s(bool,X19) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X18),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X18,sK10)),s(fun(sK10,bool),X18)))) | p(s(bool,X19))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f2142,f257])).
fof(f2711,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl16)),
  inference(subsumption_resolution,[],[f2691,f269])).
fof(f2691,plain,(
  ( ! [X17,X16] : (s(bool,'F_0') != s(bool,X17) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),X16),s(fun(sK11,bool),sK12))) | p(s(bool,X17)) | p(s(bool,'const.bool.IN_2'(s(sK10,sK17(sK13,X16,sK10)),s(fun(sK10,bool),X16))))) ) | ($spl0 | $spl16)),
  inference(superposition,[],[f2142,f257])).
fof(f1948,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) | $spl20),
  inference(cnf_transformation,[],[f1948_D])).
fof(f1948_D,plain,(
  ( ! [X2] : (s(bool,'F_0') != s(bool,X2) | s(bool,'T_0') != s(bool,X2)) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f1946,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl14),
  inference(duplicate_literal_removal,[],[f1941])).
fof(f1941,plain,(
  ( ! [X2,X3] : (s(bool,'T_0') != s(bool,X2) | p(s(bool,X3)) | p(s(bool,X3)) | s(bool,'F_0') != s(bool,X2)) ) | $spl14),
  inference(resolution,[],[f1627,f252])).
fof(f1627,plain,(
  ( ! [X4,X3] : (sP3(X4,X3,'F_0') | s(bool,'T_0') != s(bool,X3) | p(s(bool,X4))) ) | $spl14),
  inference(superposition,[],[f1585,f249])).
fof(f1038,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK10,X0),s(fun(sK10,bool),sK13)))) | ~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12))))) ) | $spl0),
  inference(forward_demodulation,[],[f1005,f388])).
fof(f1005,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.bool.IN_2'(s(sK11,X1),s(fun(sK11,bool),sK12)))) | ~p(s(bool,'const.bool.IN_2'(s(sK10,'const.pair.FST_1'(s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X0),s(sK11,X1))))),s(fun(sK10,bool),sK13))))) ) | $spl0),
  inference(superposition,[],[f558,f387])).
fof(f567,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X7))))) ) | $spl12),
  inference(cnf_transformation,[],[f567_D])).
fof(f567_D,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X7))))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f566,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,sK15('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)),s(sK11,sK16('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))))) | $spl11),
  inference(cnf_transformation,[],[f566_D])).
fof(f566_D,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,sK15('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)),s(sK11,sK16('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f568,plain,(
  ( ! [X8] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,sK15('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)),s(sK11,sK16('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))))) | p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X8))))) ) | $spl0),
  inference(forward_demodulation,[],[f527,f393])).
fof(f527,plain,(
  ( ! [X8] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X8)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,sK15('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)),s(sK11,sK16('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))))))) ) | $spl0),
  inference(superposition,[],[f384,f393])).
fof(f564,plain,(
  ( ! [X7] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),X7)))) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,sK15('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)),s(sK11,sK16('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))))))) ) | $spl0),
  inference(forward_demodulation,[],[f526,f393])).
fof(f526,plain,(
  ( ! [X7] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,sK15('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)),s(sK11,sK16('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))))) | p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))),s('type.pair.prod'(sK10,sK11),X7))))) ) | $spl0),
  inference(superposition,[],[f384,f393])).
fof(f562,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X3),s(sK11,X4))))))) ) | $spl8),
  inference(cnf_transformation,[],[f562_D])).
fof(f562_D,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X3),s(sK11,X4))))))) ) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f561,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))) | $spl7),
  inference(cnf_transformation,[],[f561_D])).
fof(f561_D,plain,(
  ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f563,plain,(
  ( ! [X6,X5] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))) | p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X5),s(sK11,X6))))))) ) | $spl0),
  inference(forward_demodulation,[],[f525,f393])).
fof(f525,plain,(
  ( ! [X6,X5] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X5),s(sK11,X6)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))))) ) | $spl0),
  inference(superposition,[],[f383,f393])).
fof(f559,plain,(
  ( ! [X4,X3] : (p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X3),s(sK11,X4)))))) | ~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10)))))) ) | $spl0),
  inference(forward_demodulation,[],[f524,f393])).
fof(f524,plain,(
  ( ! [X4,X3] : (~p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0'),s('type.pair.prod'(sK10,sK11),sK14('const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)),sK11,sK10))))) | p(s(bool,happ(s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))),s('type.pair.prod'(sK10,sK11),'const.pair.,_2'(s(sK10,X3),s(sK11,X4))))))) ) | $spl0),
  inference(superposition,[],[f383,f393])).
fof(f512,plain,(
  $false | ($spl1 | $spl4)),
  inference(subsumption_resolution,[],[f511,f254])).
fof(f254,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X2)))) )),
  inference(cnf_transformation,[],[f54])).
fof(f54,plain,(
  ! [X0,X1,X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X2)))),
  inference(flattening,[],[f53])).
fof(f53,plain,(
  ! [X0,X1] : ! [X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),'const.pred_set.EMPTY_0'),s(fun(X1,bool),X2)))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X19,X6] : ! [X14] : s(fun('type.pair.prod'(X19,X6),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X19,X6),bool),'const.pred_set.CROSS_2'(s(fun(X19,bool),'const.pred_set.EMPTY_0'),s(fun(X6,bool),X14)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pred_set.CROSS_EMPTY|split|1')).
fof(f511,plain,(
  s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),'const.pred_set.EMPTY_0'),s(fun(sK11,bool),sK12))) | ($spl1 | $spl4)),
  inference(forward_demodulation,[],[f394,f397])).
fof(f454,plain,(
  $false | ($spl1 | $spl2)),
  inference(subsumption_resolution,[],[f453,f253])).
fof(f253,plain,(
  ( ! [X2,X0,X1] : (s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))) )),
  inference(cnf_transformation,[],[f52])).
fof(f52,plain,(
  ! [X0,X1,X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(flattening,[],[f51])).
fof(f51,plain,(
  ! [X0,X1] : ! [X2] : s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X2),s(fun(X1,bool),'const.pred_set.EMPTY_0')))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6,X7] : ! [X14] : s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),X14),s(fun(X7,bool),'const.pred_set.EMPTY_0')))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pred_set.CROSS_EMPTY|split|0')).
fof(f453,plain,(
  s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),'const.pred_set.EMPTY_0'))) | ($spl1 | $spl2)),
  inference(forward_demodulation,[],[f394,f395])).
fof(f393,plain,(
  s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))) | $spl0),
  inference(cnf_transformation,[],[f393_D])).
fof(f393_D,plain,(
  s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f397,plain,(
  s(fun(sK10,bool),'const.pred_set.EMPTY_0') = s(fun(sK10,bool),sK13) | $spl4),
  inference(cnf_transformation,[],[f397_D])).
fof(f397_D,plain,(
  s(fun(sK10,bool),'const.pred_set.EMPTY_0') = s(fun(sK10,bool),sK13) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f395,plain,(
  s(fun(sK11,bool),'const.pred_set.EMPTY_0') = s(fun(sK11,bool),sK12) | $spl2),
  inference(cnf_transformation,[],[f395_D])).
fof(f395_D,plain,(
  s(fun(sK11,bool),'const.pred_set.EMPTY_0') = s(fun(sK11,bool),sK12) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f227,plain,(
  s(fun(sK11,bool),'const.pred_set.EMPTY_0') = s(fun(sK11,bool),sK12) | s(fun(sK10,bool),'const.pred_set.EMPTY_0') = s(fun(sK10,bool),sK13) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)))),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))) | s(fun(sK10,bool),'const.pred_set.EMPTY_0') = s(fun(sK10,bool),sK13) | s(fun(sK11,bool),'const.pred_set.EMPTY_0') = s(fun(sK11,bool),sK12)) & (s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))) | (s(fun(sK10,bool),'const.pred_set.EMPTY_0') != s(fun(sK10,bool),sK13) & s(fun(sK11,bool),'const.pred_set.EMPTY_0') != s(fun(sK11,bool),sK12)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK10,sK11,sK12,sK13])],[f154])).
fof(f154,plain,(
  ? [X0,X1,X2,X3] : ((s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X2))) | s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X3) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X2)) & (s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X2))) | (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X3) & s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X2))))),
  inference(flattening,[],[f153])).
fof(f153,plain,(
  ? [X0,X1,X2,X3] : ((s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X2))) | (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X3) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X2))) & (s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X2))) | (s(fun(X0,bool),'const.pred_set.EMPTY_0') != s(fun(X0,bool),X3) & s(fun(X1,bool),'const.pred_set.EMPTY_0') != s(fun(X1,bool),X2))))),
  inference(nnf_transformation,[],[f114])).
fof(f114,plain,(
  ? [X0,X1,X2,X3] : (s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X2))) <~> (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X3) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X2)))),
  inference(ennf_transformation,[],[f48])).
fof(f48,plain,(
  ~! [X0,X1,X2,X3] : (s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X2))) <=> (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X3) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X2)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ~! [X0,X1] : ! [X2,X3] : (s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X0,X1),bool),'const.pred_set.CROSS_2'(s(fun(X0,bool),X3),s(fun(X1,bool),X2))) <=> (s(fun(X0,bool),'const.pred_set.EMPTY_0') = s(fun(X0,bool),X3) | s(fun(X1,bool),'const.pred_set.EMPTY_0') = s(fun(X1,bool),X2)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6,X7] : ! [X8,X9] : (s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X8))) <=> (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X9) | s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X8)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6,X7] : ! [X8,X9] : (s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.EMPTY_0') = s(fun('type.pair.prod'(X6,X7),bool),'const.pred_set.CROSS_2'(s(fun(X6,bool),X9),s(fun(X7,bool),X8))) <=> (s(fun(X6,bool),'const.pred_set.EMPTY_0') = s(fun(X6,bool),X9) | s(fun(X7,bool),'const.pred_set.EMPTY_0') = s(fun(X7,bool),X8)))),
  file('thm.pred_set.CROSS_EMPTY_EQN','thm.pred_set.CROSS_EMPTY_EQN')).
fof(f394,plain,(
  s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))) | $spl1),
  inference(cnf_transformation,[],[f394_D])).
fof(f394_D,plain,(
  s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f398,plain,(
  s(fun(sK10,bool),'const.pred_set.EMPTY_0') != s(fun(sK10,bool),sK13) | $spl5),
  inference(cnf_transformation,[],[f398_D])).
fof(f398_D,plain,(
  s(fun(sK10,bool),'const.pred_set.EMPTY_0') != s(fun(sK10,bool),sK13) <=> ~$spl5),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl5])])).
fof(f228,plain,(
  s(fun(sK10,bool),'const.pred_set.EMPTY_0') != s(fun(sK10,bool),sK13) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)))),
  inference(cnf_transformation,[],[f155])).
fof(f396,plain,(
  s(fun(sK11,bool),'const.pred_set.EMPTY_0') != s(fun(sK11,bool),sK12) | $spl3),
  inference(cnf_transformation,[],[f396_D])).
fof(f396_D,plain,(
  s(fun(sK11,bool),'const.pred_set.EMPTY_0') != s(fun(sK11,bool),sK12) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f229,plain,(
  s(fun(sK11,bool),'const.pred_set.EMPTY_0') != s(fun(sK11,bool),sK12) | s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.EMPTY_0') != s(fun('type.pair.prod'(sK10,sK11),bool),'const.pred_set.CROSS_2'(s(fun(sK10,bool),sK13),s(fun(sK11,bool),sK12)))),
  inference(cnf_transformation,[],[f155])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 464
% Passive clauses: 4186
% Generated clauses: 11011
% Final active clauses: 163
% Final passive clauses: 1705
% Input formulas: 45
% Initial clauses: 162
% 
% Selected by SInE selection: 43
% SInE iterations: 1
% 
% Duplicate literals: 262
% Trivial inequalities: 180
% Fw subsumption resolutions: 1726
% Fw demodulations: 710
% Bw demodulations: 42
% 
% Simple tautologies: 589
% Equational tautologies: 3
% Forward subsumptions: 3281
% Fw demodulations to eq. taut.: 4
% 
% Binary resolution: 2678
% Forward superposition: 2957
% Backward superposition: 1763
% Equality resolution: 53
% 
% Split clauses: 154
% Split components: 352
% SAT solver clauses: 182
% SAT solver unit clauses: 15
% SAT solver binary clauses: 14
% 
% Sat splits: 154
% Sat splitting refutations: 27
% 
% Memory used [KB]: 3837
% Time elapsed: 0.237 s
% ------------------------------
% ------------------------------
% Success in time 11.957 s
