Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.list
% (59812)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 101
% Passive clauses: 3486
% Generated clauses: 15049
% Final active clauses: 91
% Final passive clauses: 3217
% Input formulas: 55
% Initial clauses: 107
% 
% Duplicate literals: 406
% Trivial inequalities: 196
% Fw subsumption resolutions: 1215
% Fw demodulations: 2142
% Bw demodulations: 4
% 
% Simple tautologies: 254
% Equational tautologies: 51
% Forward subsumptions: 7074
% 
% Binary resolution: 166
% Factoring: 1
% Forward superposition: 6184
% Backward superposition: 4369
% Self superposition: 212
% Equality factoring: 100
% Equality resolution: 23
% 
% Split clauses: 214
% Split components: 457
% SAT solver clauses: 340
% SAT solver unit clauses: 19
% SAT solver binary clauses: 85
% 
% Sat splits: 214
% Sat splitting refutations: 4
% 
% TWLsolver clauses: 335
% TWLsolver calls for satisfiability: 73
% 
% Memory used [KB]: 20340
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 293 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.list
% (59828)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 110
% Passive clauses: 1286
% Generated clauses: 9462
% Final active clauses: 88
% Final passive clauses: 1121
% Input formulas: 55
% Initial clauses: 107
% 
% Duplicate literals: 145
% Trivial inequalities: 48
% Fw subsumption resolutions: 378
% Fw demodulations: 1189
% Bw demodulations: 74
% Condensations: 121
% Global subsumptions: 13
% 
% Simple tautologies: 128
% Equational tautologies: 33
% Forward subsumptions: 4542
% 
% Binary resolution: 376
% Factoring: 3
% Forward superposition: 2922
% Backward superposition: 3644
% Self superposition: 137
% Equality factoring: 38
% Equality resolution: 8
% 
% Split clauses: 30
% Split components: 61
% SAT solver clauses: 1671
% SAT solver unit clauses: 126
% SAT solver binary clauses: 55
% SAT solver learnt clauses: 26
% 
% Sat splits: 30
% Sat splitting refutations: 5
% 
% TWLsolver clauses: 1619
% TWLsolver calls for satisfiability: 338
% 
% Memory used [KB]: 20468
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.list
% (59854)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 438
% Passive clauses: 11440
% Generated clauses: 153587
% Final active clauses: 76
% Final passive clauses: 311
% Input formulas: 55
% Initial clauses: 107
% Discarded non-redundant clauses: 15202
% 
% Duplicate literals: 3555
% Trivial inequalities: 832
% Fw subsumption resolutions: 7205
% Fw demodulations: 34205
% Bw demodulations: 7350
% Condensations: 7372
% Global subsumptions: 877
% 
% Simple tautologies: 999
% Equational tautologies: 796
% Forward subsumptions: 81514
% Fw demodulations to eq. taut.: 2
% Bw demodulations to eq. taut.: 446
% 
% Unit resulting resolution: 22318
% Factoring: 6
% Forward superposition: 31799
% Backward superposition: 25465
% Self superposition: 433
% Equality factoring: 153
% Equality resolution: 79
% 
% Split clauses: 411
% Split components: 835
% SAT solver clauses: 14512
% SAT solver unit clauses: 999
% SAT solver binary clauses: 3067
% 
% Sat splits: 411
% Sat splitting refutations: 1370
% 
% Memory used [KB]: 106821
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.list
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.list
% SZS output start Proof for thm.list
fof(f3102,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f270,f286,f271,f289,f287,f294,f312,f1993,f1996,f1997,f2000,f2001,f2003,f2004,f2006,f2010,f2013,f2014,f2017,f2018,f2020,f2021,f2023,f2025,f2026,f2027,f2028,f2029,f2030,f2031,f2032,f2035,f2037,f2039,f2040,f2041,f2042,f2043,f2044,f2045,f2047,f2049,f2050,f2051,f2052,f2053,f2054,f2065,f2066,f2068,f2070,f2073,f2074,f2075,f2077,f2078,f2079,f2081,f2083,f2085,f2087,f2088,f2089,f2084,f2090,f2091,f2092,f2093,f2095,f2097,f2098,f2076,f2099,f2094,f2100,f2096,f2101,f2080,f2102,f2082,f2086,f2103,f2104,f2105,f2106,f2107,f2108,f2012,f2109,f2110,f1994,f2111,f2112,f2016,f2113,f2114,f1998,f2117,f2118,f2119,f2120,f2121,f2038,f2122,f2036,f2123,f2019,f2124,f2002,f2125,f2126,f2127,f2128,f2129,f2048,f2011,f2130,f2046,f2131,f2022,f2015,f2132,f2005,f2147,f2148,f2149,f2072,f2150,f2151,f2152,f2153,f2154,f2155,f2156,f2157,f2158,f2159,f2160,f2161,f2163,f2164,f2165,f2166,f2162,f2167,f2168,f2169,f2170,f2171,f2172,f2173,f2174,f2069,f2008,f2175,f284,f2200,f2244,f2245,f2246,f2247,f2274,f2273,f2272,f2318,f2319,f2320,f2321,f2322,f2323,f2368,f2370,f2429,f2430,f2500,f2067,f2501,f2502,f2503,f2071,f2721,f2817,f2861,f2862,f2863,f2864,f3050,f3097,f3098,f3099,f3100,f3101,f2907,f3058,f2910,f2911,f2916,f2917,f2923,f1995,f3063,f3064,f1999])).
fof(f1999,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) | $spl16),
  inference(cnf_transformation,[],[f1999_D])).
fof(f1999_D,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) <=> ~$spl16),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl16])])).
fof(f3064,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | ~p(s(bool,X46))) ) | $spl30),
  inference(subsumption_resolution,[],[f2933,f219])).
fof(f219,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f151])).
fof(f151,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f80])).
fof(f80,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f79])).
fof(f79,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.EQ_CLAUSES|split|3')).
fof(f2933,plain,(
  ( ! [X45,X46] : (~p(s(bool,X45)) | s(bool,'F_0') = s(bool,X46) | ~p(s(bool,X46))) ) | $spl30),
  inference(superposition,[],[f182,f2019])).
fof(f182,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | ~p(s(bool,X0)) | ~p(s(bool,X1))) )),
  inference(cnf_transformation,[],[f116])).
fof(f116,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f115])).
fof(f115,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f61])).
fof(f61,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f10])).
fof(f10,axiom,(
  ! [X12,X7] : ((p(s(bool,X12)) => p(s(bool,X7))) => ((p(s(bool,X7)) => p(s(bool,X12))) => s(bool,X7) = s(bool,X12)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.IMP_ANTISYM_AX')).
fof(f3063,plain,(
  ( ! [X43,X44] : (~p(s(bool,X43)) | ~p(s(bool,X44))) ) | $spl30),
  inference(subsumption_resolution,[],[f2932,f219])).
fof(f2932,plain,(
  ( ! [X43,X44] : (~p(s(bool,X43)) | s(bool,'F_0') = s(bool,X44) | ~p(s(bool,X44))) ) | $spl30),
  inference(superposition,[],[f182,f2019])).
fof(f1995,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl12),
  inference(cnf_transformation,[],[f1995_D])).
fof(f1995_D,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) <=> ~$spl12),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl12])])).
fof(f2923,plain,(
  ( ! [X26,X27] : (~p(s(bool,X26)) | p(s(bool,X27))) ) | $spl30),
  inference(superposition,[],[f236,f2019])).
fof(f236,plain,(
  ( ! [X1] : (p(s(bool,X1)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f159])).
fof(f159,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (~p(s(bool,sK6)) | p(s(bool,'F_0')))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK6])],[f158])).
fof(f158,plain,(
  (~p(s(bool,'F_0')) | ! [X1] : p(s(bool,X1))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(rectify,[],[f157])).
fof(f157,plain,(
  (~p(s(bool,'F_0')) | ! [X0] : p(s(bool,X0))) & (? [X0] : ~p(s(bool,X0)) | p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f93])).
fof(f93,plain,(
  p(s(bool,'F_0')) <=> ! [X0] : p(s(bool,X0))),
  inference(rectify,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'F_0')) <=> ! [X10] : p(s(bool,X10))),
  file('thm.list.LLEX_THM|split|2','thm.bool.F_DEF')).
fof(f2917,plain,(
  ( ! [X19,X20] : (~p(s(bool,X19)) | p(s(bool,X20))) ) | $spl30),
  inference(superposition,[],[f229,f2019])).
fof(f229,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f125])).
fof(f125,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f87])).
fof(f87,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f11])).
fof(f11,axiom,(
  ! [X10] : (p(s(bool,'F_0')) => p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.FALSITY')).
fof(f2916,plain,(
  ( ! [X17,X18] : (~p(s(bool,X17)) | ~p(s(bool,X18))) ) | $spl30),
  inference(superposition,[],[f228,f2019])).
fof(f228,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f124])).
fof(f124,plain,(
  ! [X0] : ((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(flattening,[],[f85])).
fof(f85,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) => ~p(s(bool,X0)))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) => ~p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.IMP_F')).
fof(f2911,plain,(
  ( ! [X8,X7] : (~p(s(bool,X7)) | p(s(bool,X8))) ) | $spl30),
  inference(superposition,[],[f210,f2019])).
fof(f210,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f146])).
fof(f146,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f145])).
fof(f145,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f17])).
fof(f17,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) | p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.OR_CLAUSES|split|2')).
fof(f2910,plain,(
  ( ! [X6,X5] : (~p(s(bool,X5)) | p(s(bool,X6))) ) | $spl30),
  inference(superposition,[],[f208,f2019])).
fof(f208,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f144])).
fof(f144,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f143])).
fof(f143,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f75])).
fof(f75,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X10] : ((p(s(bool,X10)) | p(s(bool,'F_0'))) <=> p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.OR_CLAUSES|split|3')).
fof(f3058,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | p(s(bool,X4))) ) | $spl30),
  inference(subsumption_resolution,[],[f2909,f251])).
fof(f251,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f185])).
fof(f185,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f129])).
fof(f129,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f65])).
fof(f65,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f64])).
fof(f64,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f28])).
fof(f28,axiom,(
  ! [X6] : ! [X15] : (s(X6,X15) = s(X6,X15) <=> p(s(bool,'T_0')))),
  file('thm.list.LLEX_THM|split|2','thm.bool.REFL_CLAUSE')).
fof(f2909,plain,(
  ( ! [X4,X3] : (~p(s(bool,X3)) | p(s(bool,X4)) | ~p(s(bool,'T_0'))) ) | $spl30),
  inference(superposition,[],[f203,f2019])).
fof(f203,plain,(
  ( ! [X0] : (p(s(bool,X0)) | ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f140])).
fof(f140,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f139])).
fof(f139,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f123])).
fof(f123,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f73])).
fof(f73,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f22])).
fof(f22,axiom,(
  ! [X10] : ((p(s(bool,'F_0')) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.list.LLEX_THM|split|2','thm.bool.IMP_CLAUSES|split|2')).
fof(f2907,plain,(
  ( ! [X0,X1] : (~p(s(bool,X0)) | ~p(s(bool,X1))) ) | $spl30),
  inference(superposition,[],[f190,f2019])).
fof(f190,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | ~p(s(bool,'F_0'))) )),
  inference(cnf_transformation,[],[f132])).
fof(f132,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f131])).
fof(f131,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f119])).
fof(f119,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f69])).
fof(f69,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f68])).
fof(f68,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f24])).
fof(f24,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.IMP_CLAUSES|split|4')).
fof(f3101,plain,(
  $false | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f3089,f311])).
fof(f311,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl9),
  inference(cnf_transformation,[],[f311_D])).
fof(f311_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f3089,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(trivial_inequality_removal,[],[f1220])).
fof(f1220,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(superposition,[],[f844,f324])).
fof(f324,plain,(
  s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))),
  inference(resolution,[],[f293,f220])).
fof(f220,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f151])).
fof(f293,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0'))))),
  inference(backward_demodulation,[],[f253,f162])).
fof(f162,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))))),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f110])).
fof(f110,plain,(
  ? [X0,X1,X2,X3] : ~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))))),
  inference(ennf_transformation,[],[f58])).
fof(f58,plain,(
  ~! [X0,X1,X2,X3] : p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))))),
  inference(flattening,[],[f57])).
fof(f57,plain,(
  ~! [X0] : ! [X1,X2,X3] : p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X3),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),'const.list.CONS_2'(s(X0,X2),s('type.list.list'(X0),X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : p(s(bool,'const.list.LLEX_3'(s(fun(X6,fun(X6,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : p(s(bool,'const.list.LLEX_3'(s(fun(X6,fun(X6,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X8),s('type.list.list'(X6),X7))))))),
  file('thm.list.LLEX_THM|split|2','thm.list.LLEX_THM|split|2')).
fof(f253,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1)))),
  inference(resolution,[],[f162,f242])).
fof(f242,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2)) )),
  inference(cnf_transformation,[],[f160])).
fof(f160,plain,(
  ! [X0,X1,X2] : ((~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X2)))) | s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2)) & (s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2) | p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X2))))))),
  inference(nnf_transformation,[],[f97])).
fof(f97,plain,(
  ! [X0,X1,X2] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X2)))) <=> s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2))),
  inference(flattening,[],[f96])).
fof(f96,plain,(
  ! [X0] : ! [X1,X2] : (p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X2)))) <=> ~s('type.list.list'(X0),'const.list.NIL_0') = s('type.list.list'(X0),X2))),
  inference(rectify,[],[f45])).
fof(f45,axiom,(
  ! [X6] : ! [X9,X29] : (p(s(bool,'const.list.LLEX_3'(s(fun(X6,fun(X6,bool)),X9),s('type.list.list'(X6),'const.list.NIL_0'),s('type.list.list'(X6),X29)))) <=> ~s('type.list.list'(X6),'const.list.NIL_0') = s('type.list.list'(X6),X29))),
  file('thm.list.LLEX_THM|split|2','thm.list.LLEX_def|split|0')).
fof(f844,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X13),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X12))) | s('type.list.list'(sK0),X12) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X12)))) ) | $spl3),
  inference(resolution,[],[f674,f219])).
fof(f674,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0)))) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X0)))) ) | $spl3),
  inference(superposition,[],[f640,f242])).
fof(f640,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(resolution,[],[f286,f242])).
fof(f3100,plain,(
  $false | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f640,f311])).
fof(f3099,plain,(
  $false | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f3090,f311])).
fof(f3090,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(trivial_inequality_removal,[],[f1334])).
fof(f1334,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(superposition,[],[f393,f846])).
fof(f846,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X17),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X16))) | s('type.list.list'(sK0),X16) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X16)))) ) | $spl3),
  inference(resolution,[],[f674,f224])).
fof(f224,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f153])).
fof(f153,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.EQ_CLAUSES|split|0')).
fof(f393,plain,(
  ( ! [X6] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X6),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(resolution,[],[f313,f223])).
fof(f223,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f153])).
fof(f313,plain,(
  ( ! [X7] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X7),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0'))))) )),
  inference(subsumption_resolution,[],[f298,f253])).
fof(f298,plain,(
  ( ! [X7] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X7),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f241,f253])).
fof(f241,plain,(
  ( ! [X2,X0,X1] : (s('type.list.list'(X0),'const.list.NIL_0') != s('type.list.list'(X0),X2) | ~p(s(bool,'const.list.LLEX_3'(s(fun(X0,fun(X0,bool)),X1),s('type.list.list'(X0),'const.list.NIL_0'),s('type.list.list'(X0),X2))))) )),
  inference(cnf_transformation,[],[f160])).
fof(f3098,plain,(
  $false | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f3091,f311])).
fof(f3091,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(trivial_inequality_removal,[],[f1222])).
fof(f1222,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(superposition,[],[f844,f391])).
fof(f391,plain,(
  ( ! [X4] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X4),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(resolution,[],[f313,f220])).
fof(f3097,plain,(
  $false | ($spl3 | $spl9)),
  inference(subsumption_resolution,[],[f3092,f311])).
fof(f3092,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(trivial_inequality_removal,[],[f1340])).
fof(f1340,plain,(
  s(bool,'T_0') != s(bool,'T_0') | s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl3),
  inference(superposition,[],[f326,f846])).
fof(f326,plain,(
  s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))),
  inference(resolution,[],[f293,f223])).
fof(f3050,plain,(
  $false | ($spl2 | $spl30)),
  inference(subsumption_resolution,[],[f2875,f231])).
fof(f231,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.list.LLEX_THM|split|2',aHL_FALSITY)).
fof(f2875,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl30)),
  inference(backward_demodulation,[],[f2019,f285])).
fof(f285,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0')))))) | $spl2),
  inference(cnf_transformation,[],[f285_D])).
fof(f285_D,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0')))))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f2864,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f2819,f239])).
fof(f239,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LLEX_THM|split|2',aHL_TRUTH)).
fof(f2819,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(resolution,[],[f2762,f235])).
fof(f235,plain,(
  ~p(s(bool,'T_0')) | ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f156])).
fof(f156,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f92])).
fof(f92,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f26])).
fof(f26,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.list.LLEX_THM|split|2','thm.bool.NOT_CLAUSES|split|1')).
fof(f2762,plain,(
  ( ! [X15] : (p(s(bool,X15))) ) | $spl22),
  inference(forward_demodulation,[],[f2755,f245])).
fof(f245,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))) )),
  inference(cnf_transformation,[],[f103])).
fof(f103,plain,(
  ! [X0,X1,X2,X3] : s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))),
  inference(flattening,[],[f102])).
fof(f102,plain,(
  ! [X0,X1] : ! [X2,X3] : s(X1,X2) = s(X1,'const.list.list_CASE_3'(s('type.list.list'(X0),'const.list.NIL_0'),s(X1,X2),s(fun(X0,fun('type.list.list'(X0),X1)),X3)))),
  inference(rectify,[],[f43])).
fof(f43,axiom,(
  ! [X6,X11] : ! [X26,X13] : s(X11,X26) = s(X11,'const.list.list_CASE_3'(s('type.list.list'(X6),'const.list.NIL_0'),s(X11,X26),s(fun(X6,fun('type.list.list'(X6),X11)),X13)))),
  file('thm.list.LLEX_THM|split|2','thm.list.list_case_def|split|0')).
fof(f2755,plain,(
  ( ! [X14,X15] : (p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X15),s(fun(sK0,fun('type.list.list'(sK0),bool)),X14))))) ) | $spl22),
  inference(superposition,[],[f2011,f295])).
fof(f295,plain,(
  ( ! [X2,X0,X1] : (s(X0,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(X0,X2),s(fun(sK0,fun('type.list.list'(sK0),X0)),X1))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X1),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f243,f253])).
fof(f243,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.list.list'(X1),X0),happ(s(fun(X1,fun('type.list.list'(X1),X0)),X5),s(X1,X2))),s('type.list.list'(X1),X3))) = s(X0,'const.list.list_CASE_3'(s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'(X1),X0)),X5)))) )),
  inference(cnf_transformation,[],[f99])).
fof(f99,plain,(
  ! [X0,X1,X2,X3,X4,X5] : s(X0,happ(s(fun('type.list.list'(X1),X0),happ(s(fun(X1,fun('type.list.list'(X1),X0)),X5),s(X1,X2))),s('type.list.list'(X1),X3))) = s(X0,'const.list.list_CASE_3'(s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'(X1),X0)),X5)))),
  inference(flattening,[],[f98])).
fof(f98,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5] : s(X0,happ(s(fun('type.list.list'(X1),X0),happ(s(fun(X1,fun('type.list.list'(X1),X0)),X5),s(X1,X2))),s('type.list.list'(X1),X3))) = s(X0,'const.list.list_CASE_3'(s('type.list.list'(X1),'const.list.CONS_2'(s(X1,X2),s('type.list.list'(X1),X3))),s(X0,X4),s(fun(X1,fun('type.list.list'(X1),X0)),X5)))),
  inference(rectify,[],[f44])).
fof(f44,axiom,(
  ! [X11,X6] : ! [X27,X28,X26,X13] : s(X11,happ(s(fun('type.list.list'(X6),X11),happ(s(fun(X6,fun('type.list.list'(X6),X11)),X13),s(X6,X27))),s('type.list.list'(X6),X28))) = s(X11,'const.list.list_CASE_3'(s('type.list.list'(X6),'const.list.CONS_2'(s(X6,X27),s('type.list.list'(X6),X28))),s(X11,X26),s(fun(X6,fun('type.list.list'(X6),X11)),X13)))),
  file('thm.list.LLEX_THM|split|2','thm.list.list_case_def|split|1')).
fof(f2863,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f2818,f239])).
fof(f2818,plain,(
  ~p(s(bool,'T_0')) | $spl22),
  inference(resolution,[],[f2762,f233])).
fof(f233,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f155])).
fof(f155,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f91])).
fof(f91,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f27])).
fof(f27,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.list.LLEX_THM|split|2','thm.bool.NOT_CLAUSES|split|2')).
fof(f2862,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f2815,f2762])).
fof(f2815,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl22),
  inference(resolution,[],[f2762,f228])).
fof(f2861,plain,(
  $false | $spl22),
  inference(subsumption_resolution,[],[f2811,f2762])).
fof(f2811,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) | $spl22),
  inference(resolution,[],[f2762,f190])).
fof(f2817,plain,(
  $false | $spl22),
  inference(resolution,[],[f2762,f231])).
fof(f2721,plain,(
  $false | ($spl2 | $spl8)),
  inference(subsumption_resolution,[],[f2720,f231])).
fof(f2720,plain,(
  p(s(bool,'F_0')) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f2719,f324])).
fof(f2719,plain,(
  p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))) | ($spl2 | $spl8)),
  inference(forward_demodulation,[],[f285,f312])).
fof(f2071,plain,(
  ( ! [X288,X287] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X288),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X287),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl48),
  inference(cnf_transformation,[],[f2071_D])).
fof(f2071_D,plain,(
  ( ! [X288,X287] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X288),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X287),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f2503,plain,(
  ( ! [X22] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X22),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X22))))))) ) | $spl110),
  inference(cnf_transformation,[],[f2503_D])).
fof(f2503_D,plain,(
  ( ! [X22] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X22),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X22))))))) ) <=> ~$spl110),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl110])])).
fof(f2502,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X22),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X22)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X23),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X21),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl3),
  inference(forward_demodulation,[],[f2459,f245])).
fof(f2459,plain,(
  ( ! [X23,X21,X22] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X22),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X21))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X22),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X21)))))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X23),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X21),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl3),
  inference(superposition,[],[f655,f295])).
fof(f655,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X0)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0))))) ) | $spl3),
  inference(superposition,[],[f286,f242])).
fof(f2501,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))))) ) | $spl108),
  inference(cnf_transformation,[],[f2501_D])).
fof(f2501_D,plain,(
  ( ! [X19] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))))) ) <=> ~$spl108),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl108])])).
fof(f2067,plain,(
  ( ! [X279,X281] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X281),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X279))))) ) | $spl44),
  inference(cnf_transformation,[],[f2067_D])).
fof(f2067_D,plain,(
  ( ! [X279,X281] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X281),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X279))))) ) <=> ~$spl44),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl44])])).
fof(f2500,plain,(
  ( ! [X19,X20,X18] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X20),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X18)))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))))) ) | $spl3),
  inference(forward_demodulation,[],[f2458,f245])).
fof(f2458,plain,(
  ( ! [X19,X20,X18] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X20),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X18),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X19))))))) ) | $spl3),
  inference(superposition,[],[f655,f295])).
fof(f2430,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f2416,f231])).
fof(f2416,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(superposition,[],[f2067,f391])).
fof(f2429,plain,(
  $false | $spl44),
  inference(subsumption_resolution,[],[f2415,f231])).
fof(f2415,plain,(
  p(s(bool,'F_0')) | $spl44),
  inference(superposition,[],[f2067,f324])).
fof(f2370,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f2339,f347])).
fof(f347,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),sK1) | $spl0),
  inference(resolution,[],[f284,f241])).
fof(f2339,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK1) | $spl62),
  inference(superposition,[],[f253,f2084])).
fof(f2368,plain,(
  $false | ($spl0 | $spl62)),
  inference(subsumption_resolution,[],[f2337,f347])).
fof(f2337,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),sK1) | $spl62),
  inference(superposition,[],[f2084,f253])).
fof(f2323,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2271,f1999])).
fof(f2271,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f1999,f234])).
fof(f234,plain,(
  p(s(bool,'F_0')) | p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f156])).
fof(f2322,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2270,f1999])).
fof(f2270,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f1999,f232])).
fof(f232,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f155])).
fof(f2321,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2266,f1999])).
fof(f2266,plain,(
  p(s(bool,'F_0')) | $spl16),
  inference(resolution,[],[f1999,f201])).
fof(f201,plain,(
  p(s(bool,'T_0')) | p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f140])).
fof(f2320,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2264,f1999])).
fof(f2264,plain,(
  ( ! [X10] : (p(s(bool,X10))) ) | $spl16),
  inference(resolution,[],[f1999,f198])).
fof(f198,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f138])).
fof(f138,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,X0))))),
  inference(flattening,[],[f137])).
fof(f137,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f122])).
fof(f122,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,X10))) <=> p(s(bool,'T_0')))),
  file('thm.list.LLEX_THM|split|2','thm.bool.IMP_CLAUSES|split|3')).
fof(f2319,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2263,f1999])).
fof(f2263,plain,(
  ( ! [X9] : (p(s(bool,X9))) ) | $spl16),
  inference(resolution,[],[f1999,f195])).
fof(f195,plain,(
  ( ! [X0] : (p(s(bool,X0)) | p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | ~p(s(bool,'T_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f135])).
fof(f135,plain,(
  ! [X0] : (((p(s(bool,'T_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (~p(s(bool,'T_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f121])).
fof(f121,plain,(
  ! [X0] : ((~p(s(bool,'T_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(ennf_transformation,[],[f71])).
fof(f71,plain,(
  ! [X0] : ((p(s(bool,'T_0')) => p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f20])).
fof(f20,axiom,(
  ! [X10] : ((p(s(bool,'T_0')) => p(s(bool,X10))) <=> p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.IMP_CLAUSES|split|0')).
fof(f2318,plain,(
  $false | $spl16),
  inference(subsumption_resolution,[],[f2262,f1999])).
fof(f2262,plain,(
  ( ! [X8] : (p(s(bool,X8))) ) | $spl16),
  inference(resolution,[],[f1999,f192])).
fof(f192,plain,(
  ( ! [X0] : (p(s(bool,'T_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f134])).
fof(f134,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,X0)) | p(s(bool,'T_0'))))),
  inference(flattening,[],[f133])).
fof(f133,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'T_0'))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,X0)) | p(s(bool,'T_0')))))),
  inference(nnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f21])).
fof(f21,axiom,(
  ! [X10] : ((p(s(bool,X10)) => p(s(bool,'T_0'))) <=> p(s(bool,'T_0')))),
  file('thm.list.LLEX_THM|split|2','thm.bool.IMP_CLAUSES|split|1')).
fof(f2272,plain,(
  $false | $spl16),
  inference(resolution,[],[f1999,f238])).
fof(f238,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f9])).
fof(f9,axiom,(
  p(s(bool,'T_0'))),
  file('thm.list.LLEX_THM|split|2','thm.bool.TRUTH')).
fof(f2273,plain,(
  $false | $spl16),
  inference(resolution,[],[f1999,f239])).
fof(f2274,plain,(
  $false | $spl16),
  inference(resolution,[],[f1999,f251])).
fof(f2247,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2202,f239])).
fof(f2202,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f1995,f235])).
fof(f2246,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2201,f239])).
fof(f2201,plain,(
  ~p(s(bool,'T_0')) | $spl12),
  inference(resolution,[],[f1995,f233])).
fof(f2245,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2198,f1995])).
fof(f2198,plain,(
  ( ! [X13] : (~p(s(bool,X13))) ) | $spl12),
  inference(resolution,[],[f1995,f228])).
fof(f2244,plain,(
  $false | $spl12),
  inference(subsumption_resolution,[],[f2194,f1995])).
fof(f2194,plain,(
  ( ! [X9] : (~p(s(bool,X9))) ) | $spl12),
  inference(resolution,[],[f1995,f190])).
fof(f2200,plain,(
  $false | $spl12),
  inference(resolution,[],[f1995,f231])).
fof(f284,plain,(
  ( ! [X0] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X0),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(cnf_transformation,[],[f284_D])).
fof(f284_D,plain,(
  ( ! [X0] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X0),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1))))) ) <=> ~$spl0),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl0])])).
fof(f2175,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X2),s(sK0,sK2))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | $spl106),
  inference(cnf_transformation,[],[f2175_D])).
fof(f2175_D,plain,(
  ( ! [X2,X0,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X2),s(sK0,sK2))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) <=> ~$spl106),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl106])])).
fof(f2008,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X2),s(sK0,sK2))),s('type.list.list'(sK0),'const.list.NIL_0'))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1772,f245])).
fof(f1772,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(X0,X1),s(fun(sK0,fun('type.list.list'(sK0),X0)),X2))) = s(X0,happ(s(fun('type.list.list'(sK0),X0),happ(s(fun(sK0,fun('type.list.list'(sK0),X0)),X2),s(sK0,sK2))),s('type.list.list'(sK0),'const.list.NIL_0'))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f295,f242])).
fof(f2069,plain,(
  ( ! [X283,X284] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X284),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1))))) ) | $spl46),
  inference(cnf_transformation,[],[f2069_D])).
fof(f2069_D,plain,(
  ( ! [X283,X284] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X284),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1))))) ) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f2174,plain,(
  ( ! [X333] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X333)) ) | $spl104),
  inference(cnf_transformation,[],[f2174_D])).
fof(f2174_D,plain,(
  ( ! [X333] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X333)) ) <=> ~$spl104),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl104])])).
fof(f2173,plain,(
  ( ! [X333,X332,X334] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X333) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X334),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X332),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(forward_demodulation,[],[f1987,f245])).
fof(f1987,plain,(
  ( ! [X333,X332,X334] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X333),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X332))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X334),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X332),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(superposition,[],[f1035,f295])).
fof(f1035,plain,(
  ( ! [X4,X5] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X5) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X4),s('type.list.list'(sK0),X5),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(resolution,[],[f362,f241])).
fof(f362,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),sK1)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X2),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0))))) ) | $spl0),
  inference(superposition,[],[f284,f242])).
fof(f2172,plain,(
  ( ! [X329] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | $spl102),
  inference(cnf_transformation,[],[f2172_D])).
fof(f2172_D,plain,(
  ( ! [X329] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) <=> ~$spl102),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl102])])).
fof(f2171,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X331),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X330)))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f1986,f245])).
fof(f1986,plain,(
  ( ! [X329,X331,X330] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X331),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X330),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329)))))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f972,f295])).
fof(f972,plain,(
  ( ! [X24,X25] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X24),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X25)))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X25),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(resolution,[],[f959,f223])).
fof(f959,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),'const.list.NIL_0')))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0))))) )),
  inference(forward_demodulation,[],[f268,f253])).
fof(f268,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0))))) )),
  inference(superposition,[],[f162,f242])).
fof(f2170,plain,(
  ( ! [X326] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | $spl100),
  inference(cnf_transformation,[],[f2170_D])).
fof(f2170_D,plain,(
  ( ! [X326] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) ) <=> ~$spl100),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl100])])).
fof(f2169,plain,(
  ( ! [X327,X326,X328] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X328),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X327)))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(forward_demodulation,[],[f1985,f245])).
fof(f1985,plain,(
  ( ! [X327,X326,X328] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X328),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X327),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326)))))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f970,f295])).
fof(f970,plain,(
  ( ! [X21,X20] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X20),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X21)))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X21),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(resolution,[],[f959,f220])).
fof(f2168,plain,(
  ( ! [X324] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X324),s('type.list.list'(sK0),'const.list.NIL_0'))))) ) | $spl98),
  inference(cnf_transformation,[],[f2168_D])).
fof(f2168_D,plain,(
  ( ! [X324] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X324),s('type.list.list'(sK0),'const.list.NIL_0'))))) ) <=> ~$spl98),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl98])])).
fof(f2167,plain,(
  ( ! [X325,X323,X324] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X324),s('type.list.list'(sK0),'const.list.NIL_0')))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X325),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X323),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1984,f245])).
fof(f1984,plain,(
  ( ! [X325,X323,X324] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X324),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X323))),s('type.list.list'(sK0),'const.list.NIL_0')))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X325),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X323),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(superposition,[],[f959,f295])).
fof(f2162,plain,(
  ( ! [X312] : (s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl92),
  inference(cnf_transformation,[],[f2162_D])).
fof(f2162_D,plain,(
  ( ! [X312] : (s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f2166,plain,(
  ( ! [X321,X322] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X322),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X321)))) ) | $spl96),
  inference(cnf_transformation,[],[f2166_D])).
fof(f2166_D,plain,(
  ( ! [X321,X322] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X322),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X321)))) ) <=> ~$spl96),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl96])])).
fof(f2165,plain,(
  ( ! [X321,X320,X322] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X322),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X321))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(forward_demodulation,[],[f1983,f245])).
fof(f1983,plain,(
  ( ! [X321,X320,X322] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X322),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X321),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320))))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(superposition,[],[f846,f295])).
fof(f2164,plain,(
  ( ! [X318,X319] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X319),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X318)))) ) | $spl94),
  inference(cnf_transformation,[],[f2164_D])).
fof(f2164_D,plain,(
  ( ! [X318,X319] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X319),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X318)))) ) <=> ~$spl94),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl94])])).
fof(f2163,plain,(
  ( ! [X318,X317,X319] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X319),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X318))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(forward_demodulation,[],[f1982,f245])).
fof(f1982,plain,(
  ( ! [X318,X317,X319] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X319),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X318),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317))))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(superposition,[],[f844,f295])).
fof(f2161,plain,(
  ( ! [X313,X312,X314] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X314),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X313)))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(forward_demodulation,[],[f1980,f245])).
fof(f1980,plain,(
  ( ! [X313,X312,X314] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X314),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X313),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312)))))) | s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(superposition,[],[f674,f295])).
fof(f2160,plain,(
  ( ! [X308,X311] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X311),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X308),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) | $spl90),
  inference(cnf_transformation,[],[f2160_D])).
fof(f2160_D,plain,(
  ( ! [X308,X311] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X311),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X308),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2159,plain,(
  ( ! [X308,X310,X309,X311] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X310),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X309)))) | s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X311),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X308),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f1979,f245])).
fof(f1979,plain,(
  ( ! [X308,X310,X309,X311] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X310),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X309),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X308)))))) | s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X311),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X308),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f553,f295])).
fof(f553,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X2),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0)))) | s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f357,f242])).
fof(f357,plain,(
  ( ! [X6] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X6),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f284,f224])).
fof(f2158,plain,(
  ( ! [X304,X307] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X304),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) | $spl88),
  inference(cnf_transformation,[],[f2158_D])).
fof(f2158_D,plain,(
  ( ! [X304,X307] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X304),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f2157,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X306),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X305)))) | s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X304),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f1978,f245])).
fof(f1978,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X306),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X305),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X304)))))) | s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X304),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f531,f295])).
fof(f531,plain,(
  ( ! [X2,X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X2),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0)))) | s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f355,f242])).
fof(f355,plain,(
  ( ! [X4] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X4),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f284,f219])).
fof(f2156,plain,(
  ( ! [X302] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X302),s('type.list.list'(sK0),X302)))) ) | $spl86),
  inference(cnf_transformation,[],[f2156_D])).
fof(f2156_D,plain,(
  ( ! [X302] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X302),s('type.list.list'(sK0),X302)))) ) <=> ~$spl86),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl86])])).
fof(f2155,plain,(
  ( ! [X302,X301,X303] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X302),s('type.list.list'(sK0),X302))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X303),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X301),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1977,f245])).
fof(f1977,plain,(
  ( ! [X302,X301,X303] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X302),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X301))),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X302),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X301))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X303),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X301),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(superposition,[],[f511,f295])).
fof(f511,plain,(
  ( ! [X0,X1] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),X0))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0))))) )),
  inference(superposition,[],[f326,f242])).
fof(f2154,plain,(
  ( ! [X298] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl84),
  inference(cnf_transformation,[],[f2154_D])).
fof(f2154_D,plain,(
  ( ! [X298] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) <=> ~$spl84),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl84])])).
fof(f2153,plain,(
  ( ! [X300,X298,X299] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X300),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X299)))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f1976,f245])).
fof(f1976,plain,(
  ( ! [X300,X298,X299] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X300),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X299),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298)))))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) )),
  inference(superposition,[],[f417,f295])).
fof(f417,plain,(
  ( ! [X0,X1] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0)))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),X0)))) )),
  inference(superposition,[],[f324,f242])).
fof(f2152,plain,(
  ( ! [X296,X295] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X296),s('type.list.list'(sK0),X295),s('type.list.list'(sK0),X295))))) ) | $spl82),
  inference(cnf_transformation,[],[f2152_D])).
fof(f2152_D,plain,(
  ( ! [X296,X295] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X296),s('type.list.list'(sK0),X295),s('type.list.list'(sK0),X295))))) ) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2151,plain,(
  ( ! [X294,X296,X295,X297] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X296),s('type.list.list'(sK0),X295),s('type.list.list'(sK0),X295)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X297),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X294),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1975,f245])).
fof(f1975,plain,(
  ( ! [X294,X296,X295,X297] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X296),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X295),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X294))),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X295),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X294)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X297),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X294),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(superposition,[],[f399,f295])).
fof(f399,plain,(
  ( ! [X2,X0,X1] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),X0)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X2),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0))))) )),
  inference(superposition,[],[f313,f242])).
fof(f2150,plain,(
  ( ! [X292,X290,X293,X291] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X292),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X291)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X293),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X290),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(forward_demodulation,[],[f1974,f245])).
fof(f1974,plain,(
  ( ! [X292,X290,X293,X291] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X292),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X291),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X290)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X293),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X290),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(superposition,[],[f362,f295])).
fof(f2072,plain,(
  ( ! [X289,X286] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X289),s('type.list.list'(sK0),X286),s('type.list.list'(sK0),sK1))))) ) | $spl50),
  inference(cnf_transformation,[],[f2072_D])).
fof(f2072_D,plain,(
  ( ! [X289,X286] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X289),s('type.list.list'(sK0),X286),s('type.list.list'(sK0),sK1))))) ) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f2149,plain,(
  ( ! [X288,X287,X289,X286] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X288),s('type.list.list'(sK0),X287),s('type.list.list'(sK0),sK1)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X289),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X286),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1973,f245])).
fof(f1973,plain,(
  ( ! [X288,X287,X289,X286] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X288),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X287),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X286))),s('type.list.list'(sK0),sK1)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X289),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X286),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl0),
  inference(superposition,[],[f362,f295])).
fof(f2148,plain,(
  ( ! [X284] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X284),s('type.list.list'(sK0),X284))))) ) | $spl80),
  inference(cnf_transformation,[],[f2148_D])).
fof(f2148_D,plain,(
  ( ! [X284] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X284),s('type.list.list'(sK0),X284))))) ) <=> ~$spl80),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl80])])).
fof(f2147,plain,(
  ( ! [X285,X283,X284] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X284),s('type.list.list'(sK0),X284)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X285),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1972,f245])).
fof(f1972,plain,(
  ( ! [X285,X283,X284] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X284),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283))),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X284),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X285),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(superposition,[],[f330,f295])).
fof(f330,plain,(
  ( ! [X0,X1] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X0),s('type.list.list'(sK0),X0)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X1),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X0))))) )),
  inference(superposition,[],[f293,f242])).
fof(f2005,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X20),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl20),
  inference(cnf_transformation,[],[f2005_D])).
fof(f2005_D,plain,(
  ( ! [X20] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X20),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f2132,plain,(
  ( ! [X118,X117] : (~p(s(bool,X118)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X117),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1930,f245])).
fof(f1930,plain,(
  ( ! [X118,X117] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X118),s(fun(sK0,fun('type.list.list'(sK0),bool)),X117)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X117),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f226,f295])).
fof(f226,plain,(
  ( ! [X0] : (s(bool,'T_0') = s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f154])).
fof(f154,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f84])).
fof(f84,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X10] : (s(bool,'T_0') = s(bool,X10) <=> p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.EQ_CLAUSES|split|1')).
fof(f2015,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X6),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) ) | $spl26),
  inference(cnf_transformation,[],[f2015_D])).
fof(f2015_D,plain,(
  ( ! [X6] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X6),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) ) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f2022,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17)) ) | $spl32),
  inference(cnf_transformation,[],[f2022_D])).
fof(f2022_D,plain,(
  ( ! [X17] : (s(bool,'T_0') = s(bool,X17)) ) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f2131,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') = s(bool,X116) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X115),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1929,f245])).
fof(f1929,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X116),s(fun(sK0,fun('type.list.list'(sK0),bool)),X115))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X115),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f226,f295])).
fof(f2046,plain,(
  ( ! [X104] : (s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X104),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl38),
  inference(cnf_transformation,[],[f2046_D])).
fof(f2046_D,plain,(
  ( ! [X104] : (s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X104),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl38),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl38])])).
fof(f2130,plain,(
  ( ! [X114,X113] : (p(s(bool,X114)) | s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X113),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1928,f245])).
fof(f1928,plain,(
  ( ! [X114,X113] : (p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X114),s(fun(sK0,fun('type.list.list'(sK0),bool)),X113)))) | s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X113),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f225,f295])).
fof(f225,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f154])).
fof(f2011,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X0),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) ) | $spl22),
  inference(cnf_transformation,[],[f2011_D])).
fof(f2011_D,plain,(
  ( ! [X0] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X0),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) ) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f2048,plain,(
  ( ! [X105] : (s(bool,'T_0') != s(bool,X105)) ) | $spl40),
  inference(cnf_transformation,[],[f2048_D])).
fof(f2048_D,plain,(
  ( ! [X105] : (s(bool,'T_0') != s(bool,X105)) ) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f2129,plain,(
  ( ! [X111,X112] : (s(bool,'T_0') != s(bool,X112) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X111),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1927,f245])).
fof(f1927,plain,(
  ( ! [X111,X112] : (s(bool,'T_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X112),s(fun(sK0,fun('type.list.list'(sK0),bool)),X111))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X111),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f225,f295])).
fof(f2128,plain,(
  ( ! [X109,X110] : (~p(s(bool,X110)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X109),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1926,f245])).
fof(f1926,plain,(
  ( ! [X109,X110] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X110),s(fun(sK0,fun('type.list.list'(sK0),bool)),X109)))) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X109),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f224,f295])).
fof(f2127,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') = s(bool,X108) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X107),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1925,f245])).
fof(f1925,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X108),s(fun(sK0,fun('type.list.list'(sK0),bool)),X107))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X107),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f224,f295])).
fof(f2126,plain,(
  ( ! [X105,X106] : (p(s(bool,X106)) | s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X105),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1924,f245])).
fof(f1924,plain,(
  ( ! [X105,X106] : (p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X106),s(fun(sK0,fun('type.list.list'(sK0),bool)),X105)))) | s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X105),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f223,f295])).
fof(f2125,plain,(
  ( ! [X103,X104] : (s(bool,'T_0') != s(bool,X104) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X103),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1923,f245])).
fof(f1923,plain,(
  ( ! [X103,X104] : (s(bool,'T_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X104),s(fun(sK0,fun('type.list.list'(sK0),bool)),X103))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X103),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f223,f295])).
fof(f2002,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X16),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl18),
  inference(cnf_transformation,[],[f2002_D])).
fof(f2002_D,plain,(
  ( ! [X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X16),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f2124,plain,(
  ( ! [X101,X102] : (p(s(bool,X102)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X101),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1922,f245])).
fof(f1922,plain,(
  ( ! [X101,X102] : (p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X102),s(fun(sK0,fun('type.list.list'(sK0),bool)),X101)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X101),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f222,f295])).
fof(f222,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f152])).
fof(f152,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f82])).
fof(f82,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f81])).
fof(f81,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X10] : (s(bool,'F_0') = s(bool,X10) <=> ~p(s(bool,X10)))),
  file('thm.list.LLEX_THM|split|2','thm.bool.EQ_CLAUSES|split|2')).
fof(f2019,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13)) ) | $spl30),
  inference(cnf_transformation,[],[f2019_D])).
fof(f2019_D,plain,(
  ( ! [X13] : (s(bool,'F_0') = s(bool,X13)) ) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f2123,plain,(
  ( ! [X99,X100] : (s(bool,'F_0') = s(bool,X100) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X99),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1921,f245])).
fof(f1921,plain,(
  ( ! [X99,X100] : (s(bool,'F_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X100),s(fun(sK0,fun('type.list.list'(sK0),bool)),X99))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X99),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f222,f295])).
fof(f2036,plain,(
  ( ! [X88] : (s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X88),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl34),
  inference(cnf_transformation,[],[f2036_D])).
fof(f2036_D,plain,(
  ( ! [X88] : (s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X88),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f2122,plain,(
  ( ! [X97,X98] : (~p(s(bool,X98)) | s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X97),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1920,f245])).
fof(f1920,plain,(
  ( ! [X97,X98] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X98),s(fun(sK0,fun('type.list.list'(sK0),bool)),X97)))) | s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X97),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f221,f295])).
fof(f221,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'F_0') != s(bool,X0)) )),
  inference(cnf_transformation,[],[f152])).
fof(f2038,plain,(
  ( ! [X89] : (s(bool,'F_0') != s(bool,X89)) ) | $spl36),
  inference(cnf_transformation,[],[f2038_D])).
fof(f2038_D,plain,(
  ( ! [X89] : (s(bool,'F_0') != s(bool,X89)) ) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f2121,plain,(
  ( ! [X95,X96] : (s(bool,'F_0') != s(bool,X96) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X95),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1919,f245])).
fof(f1919,plain,(
  ( ! [X95,X96] : (s(bool,'F_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X96),s(fun(sK0,fun('type.list.list'(sK0),bool)),X95))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X95),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f221,f295])).
fof(f2120,plain,(
  ( ! [X94,X93] : (p(s(bool,X94)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X93),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1918,f245])).
fof(f1918,plain,(
  ( ! [X94,X93] : (p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X94),s(fun(sK0,fun('type.list.list'(sK0),bool)),X93)))) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X93),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f220,f295])).
fof(f2119,plain,(
  ( ! [X92,X91] : (s(bool,'F_0') = s(bool,X92) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X91),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1917,f245])).
fof(f1917,plain,(
  ( ! [X92,X91] : (s(bool,'F_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X92),s(fun(sK0,fun('type.list.list'(sK0),bool)),X91))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X91),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f220,f295])).
fof(f2118,plain,(
  ( ! [X90,X89] : (~p(s(bool,X90)) | s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X89),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1916,f245])).
fof(f1916,plain,(
  ( ! [X90,X89] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X90),s(fun(sK0,fun('type.list.list'(sK0),bool)),X89)))) | s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X89),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(superposition,[],[f219,f295])).
fof(f2117,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') != s(bool,X88) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X87),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1915,f245])).
fof(f1915,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X88),s(fun(sK0,fun('type.list.list'(sK0),bool)),X87))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X87),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f219,f295])).
fof(f1998,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X10),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X11))) ) | $spl14),
  inference(cnf_transformation,[],[f1998_D])).
fof(f1998_D,plain,(
  ( ! [X10,X11] : (s(bool,X11) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X10),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X11))) ) <=> ~$spl14),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl14])])).
fof(f2114,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,X63)) | s(bool,X64) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X62),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X64))) )),
  inference(forward_demodulation,[],[f1903,f245])).
fof(f1903,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X63),s(fun(sK0,fun('type.list.list'(sK0),bool)),X62)))) | s(bool,X64) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X62),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X64))) )),
  inference(superposition,[],[f182,f295])).
fof(f2113,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,X60)) | s(bool,X61) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X59),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X61))) )),
  inference(forward_demodulation,[],[f1902,f245])).
fof(f1902,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X60),s(fun(sK0,fun('type.list.list'(sK0),bool)),X59)))) | s(bool,X61) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X59),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X61))) )),
  inference(superposition,[],[f182,f295])).
fof(f2016,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | ~p(s(bool,X7))) ) | $spl28),
  inference(cnf_transformation,[],[f2016_D])).
fof(f2016_D,plain,(
  ( ! [X8,X7] : (s(bool,X7) = s(bool,X8) | ~p(s(bool,X7))) ) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2112,plain,(
  ( ! [X57,X58,X56] : (s(bool,X57) = s(bool,X58) | ~p(s(bool,X58)) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X56),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1901,f245])).
fof(f1901,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X57),s(fun(sK0,fun('type.list.list'(sK0),bool)),X56))) | ~p(s(bool,X58)) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X56),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f182,f295])).
fof(f2111,plain,(
  ( ! [X54,X55,X53] : (s(bool,X54) = s(bool,X55) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X53),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | ~p(s(bool,X55))) )),
  inference(forward_demodulation,[],[f1900,f245])).
fof(f1900,plain,(
  ( ! [X54,X55,X53] : (s(bool,X55) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X54),s(fun(sK0,fun('type.list.list'(sK0),bool)),X53))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X53),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | ~p(s(bool,X55))) )),
  inference(superposition,[],[f182,f295])).
fof(f1994,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X4),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X5))) ) | $spl10),
  inference(cnf_transformation,[],[f1994_D])).
fof(f1994_D,plain,(
  ( ! [X4,X5] : (s(bool,X5) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X4),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X5))) ) <=> ~$spl10),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl10])])).
fof(f2110,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,X51)) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X50),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X52))) )),
  inference(forward_demodulation,[],[f1899,f245])).
fof(f1899,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X51),s(fun(sK0,fun('type.list.list'(sK0),bool)),X50)))) | s(bool,X52) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X50),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X52))) )),
  inference(superposition,[],[f179,f295])).
fof(f179,plain,(
  ( ! [X0,X1] : (s(bool,X0) = s(bool,X1) | p(s(bool,X1)) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f116])).
fof(f2109,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,X48)) | s(bool,X49) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X47),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X49))) )),
  inference(forward_demodulation,[],[f1898,f245])).
fof(f1898,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X48),s(fun(sK0,fun('type.list.list'(sK0),bool)),X47)))) | s(bool,X49) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X47),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X49))) )),
  inference(superposition,[],[f179,f295])).
fof(f2012,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1))) ) | $spl24),
  inference(cnf_transformation,[],[f2012_D])).
fof(f2012_D,plain,(
  ( ! [X2,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1))) ) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f2108,plain,(
  ( ! [X45,X46,X44] : (s(bool,X45) = s(bool,X46) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X44),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | p(s(bool,X46))) )),
  inference(forward_demodulation,[],[f1897,f245])).
fof(f1897,plain,(
  ( ! [X45,X46,X44] : (s(bool,X46) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X45),s(fun(sK0,fun('type.list.list'(sK0),bool)),X44))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X44),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | p(s(bool,X46))) )),
  inference(superposition,[],[f179,f295])).
fof(f2107,plain,(
  ( ! [X43,X41,X42] : (s(bool,X42) = s(bool,X43) | p(s(bool,X43)) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X41),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1896,f245])).
fof(f1896,plain,(
  ( ! [X43,X41,X42] : (s(bool,X43) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X42),s(fun(sK0,fun('type.list.list'(sK0),bool)),X41))) | p(s(bool,X43)) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X41),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f179,f295])).
fof(f2106,plain,(
  ( ! [X38,X40] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0)))))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0))))))) ) | $spl78),
  inference(cnf_transformation,[],[f2106_D])).
fof(f2106_D,plain,(
  ( ! [X38,X40] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0)))))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0))))))) ) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2105,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,X39)) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0)))))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0))))))) )),
  inference(forward_demodulation,[],[f1895,f245])).
fof(f1895,plain,(
  ( ! [X39,X38,X40] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X39),s(fun(sK0,fun('type.list.list'(sK0),bool)),X38)))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0)))))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X40),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X38),s(sK0,sK2)),X40,'type.list.list'(sK0))))))) )),
  inference(superposition,[],[f184,f295])).
fof(f184,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1,X2] : ((p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK4(X2,X1,X0))))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK4(X2,X1,X0)))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4,sK5])],[f127])).
fof(f127,plain,(
  ! [X0,X1,X2] : (? [X5] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X5)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X3] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3)))))),
  inference(rectify,[],[f118])).
fof(f118,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | ! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5)))))),
  inference(flattening,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : (? [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) & ~p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) | (! [X4] : ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(ennf_transformation,[],[f63])).
fof(f63,plain,(
  ! [X0,X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(flattening,[],[f62])).
fof(f62,plain,(
  ! [X0] : ! [X1,X2] : (! [X3] : (p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X3)))) => p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X3))))) => (? [X4] : p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) => ? [X5] : p(s(bool,happ(s(fun(X0,bool),X1),s(X0,X5))))))),
  inference(rectify,[],[f38])).
fof(f38,axiom,(
  ! [X6] : ! [X18,X19] : (! [X15] : (p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X15)))) => p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X15))))) => (? [X15] : p(s(bool,happ(s(fun(X6,bool),X19),s(X6,X15)))) => ? [X15] : p(s(bool,happ(s(fun(X6,bool),X18),s(X6,X15))))))),
  file('thm.list.LLEX_THM|split|2','thm.bool.MONO_EXISTS')).
fof(f2104,plain,(
  ( ! [X37,X35] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X37),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0)))))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2))),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0))))))) ) | $spl76),
  inference(cnf_transformation,[],[f2104_D])).
fof(f2104_D,plain,(
  ( ! [X37,X35] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X37),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0)))))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2))),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0))))))) ) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2103,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,X36)) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X37),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0)))))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2))),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0))))))) )),
  inference(forward_demodulation,[],[f1894,f245])).
fof(f1894,plain,(
  ( ! [X37,X35,X36] : (~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X36),s(fun(sK0,fun('type.list.list'(sK0),bool)),X35)))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),X37),s('type.list.list'(sK0),sK5(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0)))))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2))),s('type.list.list'(sK0),sK4(happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X35),s(sK0,sK2)),X37,'type.list.list'(sK0))))))) )),
  inference(superposition,[],[f183,f295])).
fof(f183,plain,(
  ( ! [X4,X2,X0,X1] : (p(s(bool,happ(s(fun(X0,bool),X1),s(X0,sK5(X2,X1,X0))))) | ~p(s(bool,happ(s(fun(X0,bool),X2),s(X0,X4)))) | p(s(bool,happ(s(fun(X0,bool),X2),s(X0,sK4(X2,X1,X0)))))) )),
  inference(cnf_transformation,[],[f128])).
fof(f2086,plain,(
  ( ! [X312] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl64),
  inference(cnf_transformation,[],[f2086_D])).
fof(f2086_D,plain,(
  ( ! [X312] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2082,plain,(
  ( ! [X304,X307] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),X304),s('type.list.list'(sK0),sK1)))) ) | $spl60),
  inference(cnf_transformation,[],[f2082_D])).
fof(f2082_D,plain,(
  ( ! [X304,X307] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),X304),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f2102,plain,(
  ( ! [X341,X340,X342] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X342),s('type.list.list'(sK0),X340),s('type.list.list'(sK0),sK1))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X341),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f1887,f245])).
fof(f1887,plain,(
  ( ! [X341,X340,X342] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X341),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X342),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X340),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X341))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f1664,f295])).
fof(f1664,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X7) | s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X6),s('type.list.list'(sK0),X7),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f553,f241])).
fof(f2080,plain,(
  ( ! [X300,X303] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X303),s('type.list.list'(sK0),X300),s('type.list.list'(sK0),sK1)))) ) | $spl58),
  inference(cnf_transformation,[],[f2080_D])).
fof(f2080_D,plain,(
  ( ! [X300,X303] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X303),s('type.list.list'(sK0),X300),s('type.list.list'(sK0),sK1)))) ) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f2101,plain,(
  ( ! [X337,X339,X338] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X339),s('type.list.list'(sK0),X337),s('type.list.list'(sK0),sK1))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X338),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f1886,f245])).
fof(f1886,plain,(
  ( ! [X337,X339,X338] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X338),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X339),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X337),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X338))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f1582,f295])).
fof(f1582,plain,(
  ( ! [X6,X7] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X7) | s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X6),s('type.list.list'(sK0),X7),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(resolution,[],[f531,f241])).
fof(f2096,plain,(
  ( ! [X325] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X325),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | $spl74),
  inference(cnf_transformation,[],[f2096_D])).
fof(f2096_D,plain,(
  ( ! [X325] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X325),s('type.list.list'(sK0),'const.list.NIL_0')))) ) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f2100,plain,(
  ( ! [X335,X336] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X335),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X336),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1885,f245])).
fof(f1885,plain,(
  ( ! [X335,X336] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X336),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X335),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X336))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1512,f295])).
fof(f1512,plain,(
  ( ! [X3] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X3) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X3),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(resolution,[],[f972,f241])).
fof(f2094,plain,(
  ( ! [X322] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X322),s('type.list.list'(sK0),'const.list.NIL_0')))) ) | $spl72),
  inference(cnf_transformation,[],[f2094_D])).
fof(f2094_D,plain,(
  ( ! [X322] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X322),s('type.list.list'(sK0),'const.list.NIL_0')))) ) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f2099,plain,(
  ( ! [X333,X334] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X333),s('type.list.list'(sK0),'const.list.NIL_0'))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X334),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1884,f245])).
fof(f1884,plain,(
  ( ! [X333,X334] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X334),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X333),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X334))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f1442,f295])).
fof(f1442,plain,(
  ( ! [X3] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X3) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X3),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(resolution,[],[f970,f241])).
fof(f2076,plain,(
  ( ! [X294] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X294),s('type.list.list'(sK0),X294)))) ) | $spl54),
  inference(cnf_transformation,[],[f2076_D])).
fof(f2076_D,plain,(
  ( ! [X294] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X294),s('type.list.list'(sK0),X294)))) ) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f2098,plain,(
  ( ! [X331,X332] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X331),s('type.list.list'(sK0),X331))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X332),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1883,f245])).
fof(f1883,plain,(
  ( ! [X331,X332] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X332),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X331),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X332))),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X331),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X332)))))) )),
  inference(superposition,[],[f1249,f295])).
fof(f1249,plain,(
  ( ! [X3] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X3) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X3),s('type.list.list'(sK0),X3)))) )),
  inference(resolution,[],[f417,f241])).
fof(f2097,plain,(
  ( ! [X329,X328,X330] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X330),s('type.list.list'(sK0),X328),s('type.list.list'(sK0),sK1)))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(forward_demodulation,[],[f1882,f245])).
fof(f1882,plain,(
  ( ! [X329,X328,X330] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X330),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X328),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X329))),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(superposition,[],[f1035,f295])).
fof(f2095,plain,(
  ( ! [X325,X327,X326] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X325),s('type.list.list'(sK0),'const.list.NIL_0'))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X327),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1881,f245])).
fof(f1881,plain,(
  ( ! [X325,X327,X326] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X327),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X325),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X326))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f972,f295])).
fof(f2093,plain,(
  ( ! [X323,X324,X322] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X322),s('type.list.list'(sK0),'const.list.NIL_0'))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X324),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X323),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1880,f245])).
fof(f1880,plain,(
  ( ! [X323,X324,X322] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X324),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X323),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X322),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X323))),s('type.list.list'(sK0),'const.list.NIL_0')))) )),
  inference(superposition,[],[f970,f295])).
fof(f2092,plain,(
  ( ! [X320] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0'))))) ) | $spl70),
  inference(cnf_transformation,[],[f2092_D])).
fof(f2092_D,plain,(
  ( ! [X320] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0'))))) ) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f2091,plain,(
  ( ! [X321,X320,X319] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X321),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X319)))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0'))))) )),
  inference(forward_demodulation,[],[f1879,f245])).
fof(f1879,plain,(
  ( ! [X321,X320,X319] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),'const.list.NIL_0')))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X321),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X319),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X320))))))) )),
  inference(superposition,[],[f959,f295])).
fof(f2090,plain,(
  ( ! [X318,X317] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X318),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl68),
  inference(cnf_transformation,[],[f2090_D])).
fof(f2090_D,plain,(
  ( ! [X318,X317] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X318),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f2084,plain,(
  ( ! [X308] : (s('type.list.list'(sK0),X308) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X308)))) ) | $spl62),
  inference(cnf_transformation,[],[f2084_D])).
fof(f2084_D,plain,(
  ( ! [X308] : (s('type.list.list'(sK0),X308) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X308)))) ) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f2089,plain,(
  ( ! [X316,X318,X317] : (s('type.list.list'(sK0),X316) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X316))) | s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X318),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(forward_demodulation,[],[f1878,f245])).
fof(f1878,plain,(
  ( ! [X316,X318,X317] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X318),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) | s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X316),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X316),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X317)))))) ) | $spl3),
  inference(superposition,[],[f846,f295])).
fof(f2088,plain,(
  ( ! [X315,X314] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X315),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X314),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl66),
  inference(cnf_transformation,[],[f2088_D])).
fof(f2088_D,plain,(
  ( ! [X315,X314] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X315),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X314),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) <=> ~$spl66),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl66])])).
fof(f2087,plain,(
  ( ! [X313,X315,X314] : (s('type.list.list'(sK0),X313) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X313))) | s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X315),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X314),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl3),
  inference(forward_demodulation,[],[f1877,f245])).
fof(f1877,plain,(
  ( ! [X313,X315,X314] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X315),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X314),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) | s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X313),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X314))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X313),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X314)))))) ) | $spl3),
  inference(superposition,[],[f844,f295])).
fof(f2085,plain,(
  ( ! [X312,X311] : (s('type.list.list'(sK0),X311) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X311))) | s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) ) | $spl3),
  inference(forward_demodulation,[],[f1876,f245])).
fof(f1876,plain,(
  ( ! [X312,X311] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X311),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X311),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X312)))))) ) | $spl3),
  inference(superposition,[],[f836,f295])).
fof(f836,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.NIL_0') != s('type.list.list'(sK0),X0) | s('type.list.list'(sK0),X0) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X0)))) ) | $spl3),
  inference(resolution,[],[f674,f241])).
fof(f2083,plain,(
  ( ! [X308,X310,X309] : (s('type.list.list'(sK0),X308) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),X308))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X310),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X309),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl3),
  inference(forward_demodulation,[],[f1875,f245])).
fof(f1875,plain,(
  ( ! [X308,X310,X309] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X310),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X309),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X308),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X309))) = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X308),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X309)))))) ) | $spl3),
  inference(superposition,[],[f674,f295])).
fof(f2081,plain,(
  ( ! [X304,X306,X305,X307] : (s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),X304),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X306),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X305),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1874,f245])).
fof(f1874,plain,(
  ( ! [X304,X306,X305,X307] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X306),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X305),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | s(bool,'T_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X307),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X304),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X305))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f553,f295])).
fof(f2079,plain,(
  ( ! [X300,X302,X301,X303] : (s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X303),s('type.list.list'(sK0),X300),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X302),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X301),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1873,f245])).
fof(f1873,plain,(
  ( ! [X300,X302,X301,X303] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X302),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X301),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | s(bool,'F_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X303),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X300),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X301))),s('type.list.list'(sK0),sK1)))) ) | $spl0),
  inference(superposition,[],[f531,f295])).
fof(f2078,plain,(
  ( ! [X298] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) | $spl56),
  inference(cnf_transformation,[],[f2078_D])).
fof(f2078_D,plain,(
  ( ! [X298] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) ) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f2077,plain,(
  ( ! [X298,X297,X299] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X299),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X297)))) | s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) )),
  inference(forward_demodulation,[],[f1872,f245])).
fof(f1872,plain,(
  ( ! [X298,X297,X299] : (s(bool,'T_0') != s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X299),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X297),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X298))))))) )),
  inference(superposition,[],[f511,f295])).
fof(f2075,plain,(
  ( ! [X294,X296,X295] : (s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),X294),s('type.list.list'(sK0),X294))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X296),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X295),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1871,f245])).
fof(f1871,plain,(
  ( ! [X294,X296,X295] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X296),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X295),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | s(bool,'F_0') = s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X294),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X295))),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X294),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X295)))))) )),
  inference(superposition,[],[f417,f295])).
fof(f2074,plain,(
  ( ! [X292,X291] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X292),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl52),
  inference(cnf_transformation,[],[f2074_D])).
fof(f2074_D,plain,(
  ( ! [X292,X291] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X292),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f2073,plain,(
  ( ! [X292,X290,X293,X291] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X293),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X290)))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X292),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1870,f245])).
fof(f1870,plain,(
  ( ! [X292,X290,X293,X291] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X292),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X293),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X290),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X291))))))) )),
  inference(superposition,[],[f399,f295])).
fof(f2070,plain,(
  ( ! [X288,X287,X289,X286] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X289),s('type.list.list'(sK0),X286),s('type.list.list'(sK0),sK1)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X288),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X287),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl0),
  inference(forward_demodulation,[],[f1869,f245])).
fof(f1869,plain,(
  ( ! [X288,X287,X289,X286] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X288),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X287),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X289),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X286),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X287))),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(superposition,[],[f362,f295])).
fof(f2068,plain,(
  ( ! [X282,X285,X283,X284] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X285),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X282)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X284),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1))))) ) | $spl0),
  inference(forward_demodulation,[],[f1868,f245])).
fof(f1868,plain,(
  ( ! [X282,X285,X283,X284] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X284),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),sK1)))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X285),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X282),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X283))))))) ) | $spl0),
  inference(superposition,[],[f362,f295])).
fof(f2066,plain,(
  ( ! [X280] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) | $spl42),
  inference(cnf_transformation,[],[f2066_D])).
fof(f2066_D,plain,(
  ( ! [X280] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) ) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f2065,plain,(
  ( ! [X280,X279,X281] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X281),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X279)))) | ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))))) )),
  inference(forward_demodulation,[],[f1867,f245])).
fof(f1867,plain,(
  ( ! [X280,X279,X281] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1))),s('type.list.list'(sK0),happ(s(fun('type.list.list'(sK0),'type.list.list'(sK0)),happ(s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X281),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),X279),s(fun(sK0,fun('type.list.list'(sK0),'type.list.list'(sK0))),X280))))))) )),
  inference(superposition,[],[f330,f295])).
fof(f2054,plain,(
  ( ! [X118,X117] : (s(bool,'T_0') = s(bool,X117) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X118),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1826,f245])).
fof(f1826,plain,(
  ( ! [X118,X117] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X118),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'T_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X117),s(fun(sK0,fun('type.list.list'(sK0),bool)),X118)))) )),
  inference(superposition,[],[f226,f295])).
fof(f2053,plain,(
  ( ! [X116,X115] : (~p(s(bool,X115)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X116),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1825,f245])).
fof(f1825,plain,(
  ( ! [X116,X115] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X116),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X115),s(fun(sK0,fun('type.list.list'(sK0),bool)),X116))))) )),
  inference(superposition,[],[f226,f295])).
fof(f2052,plain,(
  ( ! [X114,X113] : (s(bool,'T_0') != s(bool,X113) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X114),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1824,f245])).
fof(f1824,plain,(
  ( ! [X114,X113] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X114),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'T_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X113),s(fun(sK0,fun('type.list.list'(sK0),bool)),X114)))) )),
  inference(superposition,[],[f225,f295])).
fof(f2051,plain,(
  ( ! [X111,X112] : (p(s(bool,X111)) | s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X112),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1823,f245])).
fof(f1823,plain,(
  ( ! [X111,X112] : (s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X112),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X111),s(fun(sK0,fun('type.list.list'(sK0),bool)),X112))))) )),
  inference(superposition,[],[f225,f295])).
fof(f2050,plain,(
  ( ! [X109,X110] : (s(bool,'T_0') = s(bool,X109) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X110),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1822,f245])).
fof(f1822,plain,(
  ( ! [X109,X110] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X110),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'T_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X109),s(fun(sK0,fun('type.list.list'(sK0),bool)),X110)))) )),
  inference(superposition,[],[f224,f295])).
fof(f2049,plain,(
  ( ! [X107,X108] : (~p(s(bool,X107)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X108),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1821,f245])).
fof(f1821,plain,(
  ( ! [X107,X108] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X108),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X107),s(fun(sK0,fun('type.list.list'(sK0),bool)),X108))))) )),
  inference(superposition,[],[f224,f295])).
fof(f2047,plain,(
  ( ! [X105,X106] : (s(bool,'T_0') != s(bool,X105) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X106),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1820,f245])).
fof(f1820,plain,(
  ( ! [X105,X106] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X106),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'T_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X105),s(fun(sK0,fun('type.list.list'(sK0),bool)),X106)))) )),
  inference(superposition,[],[f223,f295])).
fof(f2045,plain,(
  ( ! [X103,X104] : (p(s(bool,X103)) | s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X104),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1819,f245])).
fof(f1819,plain,(
  ( ! [X103,X104] : (s(bool,'T_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X104),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X103),s(fun(sK0,fun('type.list.list'(sK0),bool)),X104))))) )),
  inference(superposition,[],[f223,f295])).
fof(f2044,plain,(
  ( ! [X101,X102] : (s(bool,'F_0') = s(bool,X101) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X102),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1818,f245])).
fof(f1818,plain,(
  ( ! [X101,X102] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X102),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'F_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X101),s(fun(sK0,fun('type.list.list'(sK0),bool)),X102)))) )),
  inference(superposition,[],[f222,f295])).
fof(f2043,plain,(
  ( ! [X99,X100] : (p(s(bool,X99)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X100),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1817,f245])).
fof(f1817,plain,(
  ( ! [X99,X100] : (s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X100),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X99),s(fun(sK0,fun('type.list.list'(sK0),bool)),X100))))) )),
  inference(superposition,[],[f222,f295])).
fof(f2042,plain,(
  ( ! [X97,X98] : (s(bool,'F_0') != s(bool,X97) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X98),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1816,f245])).
fof(f1816,plain,(
  ( ! [X97,X98] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X98),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'F_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X97),s(fun(sK0,fun('type.list.list'(sK0),bool)),X98)))) )),
  inference(superposition,[],[f221,f295])).
fof(f2041,plain,(
  ( ! [X95,X96] : (~p(s(bool,X95)) | s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X96),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1815,f245])).
fof(f1815,plain,(
  ( ! [X95,X96] : (s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X96),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X95),s(fun(sK0,fun('type.list.list'(sK0),bool)),X96))))) )),
  inference(superposition,[],[f221,f295])).
fof(f2040,plain,(
  ( ! [X94,X93] : (s(bool,'F_0') = s(bool,X93) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X94),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1814,f245])).
fof(f1814,plain,(
  ( ! [X94,X93] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X94),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'F_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X93),s(fun(sK0,fun('type.list.list'(sK0),bool)),X94)))) )),
  inference(superposition,[],[f220,f295])).
fof(f2039,plain,(
  ( ! [X92,X91] : (p(s(bool,X91)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X92),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1813,f245])).
fof(f1813,plain,(
  ( ! [X92,X91] : (s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X92),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X91),s(fun(sK0,fun('type.list.list'(sK0),bool)),X92))))) )),
  inference(superposition,[],[f220,f295])).
fof(f2037,plain,(
  ( ! [X90,X89] : (s(bool,'F_0') != s(bool,X89) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X90),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1812,f245])).
fof(f1812,plain,(
  ( ! [X90,X89] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X90),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,'F_0') != s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X89),s(fun(sK0,fun('type.list.list'(sK0),bool)),X90)))) )),
  inference(superposition,[],[f219,f295])).
fof(f2035,plain,(
  ( ! [X88,X87] : (~p(s(bool,X87)) | s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X88),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1811,f245])).
fof(f1811,plain,(
  ( ! [X88,X87] : (s(bool,'F_0') != s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X88),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X87),s(fun(sK0,fun('type.list.list'(sK0),bool)),X88))))) )),
  inference(superposition,[],[f219,f295])).
fof(f2032,plain,(
  ( ! [X64,X62,X63] : (s(bool,X62) = s(bool,X64) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X63),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | ~p(s(bool,X64))) )),
  inference(forward_demodulation,[],[f1799,f245])).
fof(f1799,plain,(
  ( ! [X64,X62,X63] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X63),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,X64) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X62),s(fun(sK0,fun('type.list.list'(sK0),bool)),X63))) | ~p(s(bool,X64))) )),
  inference(superposition,[],[f182,f295])).
fof(f2031,plain,(
  ( ! [X61,X59,X60] : (s(bool,X59) = s(bool,X61) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X60),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | ~p(s(bool,X61))) )),
  inference(forward_demodulation,[],[f1798,f245])).
fof(f1798,plain,(
  ( ! [X61,X59,X60] : (~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X60),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,X61) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X59),s(fun(sK0,fun('type.list.list'(sK0),bool)),X60))) | ~p(s(bool,X61))) )),
  inference(superposition,[],[f182,f295])).
fof(f2030,plain,(
  ( ! [X57,X58,X56] : (~p(s(bool,X56)) | s(bool,X58) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X57),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X58))) )),
  inference(forward_demodulation,[],[f1797,f245])).
fof(f1797,plain,(
  ( ! [X57,X58,X56] : (s(bool,X58) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X57),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X58)) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X56),s(fun(sK0,fun('type.list.list'(sK0),bool)),X57))))) )),
  inference(superposition,[],[f182,f295])).
fof(f2029,plain,(
  ( ! [X54,X55,X53] : (~p(s(bool,X53)) | s(bool,X55) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X54),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X55))) )),
  inference(forward_demodulation,[],[f1796,f245])).
fof(f1796,plain,(
  ( ! [X54,X55,X53] : (s(bool,X55) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X54),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X53),s(fun(sK0,fun('type.list.list'(sK0),bool)),X54)))) | ~p(s(bool,X55))) )),
  inference(superposition,[],[f182,f295])).
fof(f2028,plain,(
  ( ! [X52,X50,X51] : (s(bool,X50) = s(bool,X52) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X51),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | p(s(bool,X52))) )),
  inference(forward_demodulation,[],[f1795,f245])).
fof(f1795,plain,(
  ( ! [X52,X50,X51] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X51),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,X52) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X50),s(fun(sK0,fun('type.list.list'(sK0),bool)),X51))) | p(s(bool,X52))) )),
  inference(superposition,[],[f179,f295])).
fof(f2027,plain,(
  ( ! [X47,X48,X49] : (s(bool,X47) = s(bool,X49) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X48),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | p(s(bool,X49))) )),
  inference(forward_demodulation,[],[f1794,f245])).
fof(f1794,plain,(
  ( ! [X47,X48,X49] : (p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X48),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | s(bool,X49) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X47),s(fun(sK0,fun('type.list.list'(sK0),bool)),X48))) | p(s(bool,X49))) )),
  inference(superposition,[],[f179,f295])).
fof(f2026,plain,(
  ( ! [X45,X46,X44] : (p(s(bool,X44)) | s(bool,X46) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X45),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X46))) )),
  inference(forward_demodulation,[],[f1793,f245])).
fof(f1793,plain,(
  ( ! [X45,X46,X44] : (s(bool,X46) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X45),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X44),s(fun(sK0,fun('type.list.list'(sK0),bool)),X45)))) | p(s(bool,X46))) )),
  inference(superposition,[],[f179,f295])).
fof(f2025,plain,(
  ( ! [X43,X41,X42] : (p(s(bool,X41)) | s(bool,X43) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X42),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X43))) )),
  inference(forward_demodulation,[],[f1792,f245])).
fof(f1792,plain,(
  ( ! [X43,X41,X42] : (s(bool,X43) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X42),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X43)) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X41),s(fun(sK0,fun('type.list.list'(sK0),bool)),X42))))) )),
  inference(superposition,[],[f179,f295])).
fof(f2023,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,X19) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X18),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1780,f245])).
fof(f1780,plain,(
  ( ! [X19,X18] : (s(bool,'T_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X19),s(fun(sK0,fun('type.list.list'(sK0),bool)),X18))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X18),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f295,f226])).
fof(f2021,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,X17) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X16),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1779,f245])).
fof(f1779,plain,(
  ( ! [X17,X16] : (s(bool,'T_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X17),s(fun(sK0,fun('type.list.list'(sK0),bool)),X16))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X16),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f295,f224])).
fof(f2020,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,X15) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X14),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1778,f245])).
fof(f1778,plain,(
  ( ! [X14,X15] : (s(bool,'F_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X15),s(fun(sK0,fun('type.list.list'(sK0),bool)),X14))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X14),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f295,f222])).
fof(f2018,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,X13) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X12),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1777,f245])).
fof(f1777,plain,(
  ( ! [X12,X13] : (s(bool,'F_0') = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X13),s(fun(sK0,fun('type.list.list'(sK0),bool)),X12))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X12),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f295,f220])).
fof(f2017,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) = s(bool,X11) | ~p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X9),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1776,f245])).
fof(f1776,plain,(
  ( ! [X10,X11,X9] : (s(bool,X10) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X11),s(fun(sK0,fun('type.list.list'(sK0),bool)),X9))) | ~p(s(bool,X10)) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X9),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f295,f182])).
fof(f2014,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,X8) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X6),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | ~p(s(bool,X7))) )),
  inference(forward_demodulation,[],[f1775,f245])).
fof(f1775,plain,(
  ( ! [X6,X8,X7] : (s(bool,X7) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X8),s(fun(sK0,fun('type.list.list'(sK0),bool)),X6))) | ~p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X6),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | ~p(s(bool,X7))) )),
  inference(superposition,[],[f295,f182])).
fof(f2013,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,X5) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X3),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | p(s(bool,X4))) )),
  inference(forward_demodulation,[],[f1774,f245])).
fof(f1774,plain,(
  ( ! [X4,X5,X3] : (s(bool,X4) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X5),s(fun(sK0,fun('type.list.list'(sK0),bool)),X3))) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X3),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) | p(s(bool,X4))) )),
  inference(superposition,[],[f295,f179])).
fof(f2010,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,X2) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X0),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(forward_demodulation,[],[f1773,f245])).
fof(f1773,plain,(
  ( ! [X2,X0,X1] : (s(bool,X1) = s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X2),s(fun(sK0,fun('type.list.list'(sK0),bool)),X0))) | p(s(bool,X1)) | p(s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X0),s(sK0,sK2))),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f295,f179])).
fof(f2006,plain,(
  ( ! [X21,X22] : (~p(s(bool,X21)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X22),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1763,f245])).
fof(f1763,plain,(
  ( ! [X21,X22] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X22),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X21),s(fun(sK0,fun('type.list.list'(sK0),bool)),X22))))) )),
  inference(superposition,[],[f295,f226])).
fof(f2004,plain,(
  ( ! [X19,X20] : (~p(s(bool,X19)) | s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X20),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1762,f245])).
fof(f1762,plain,(
  ( ! [X19,X20] : (s(bool,'T_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X20),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X19),s(fun(sK0,fun('type.list.list'(sK0),bool)),X20))))) )),
  inference(superposition,[],[f295,f224])).
fof(f2003,plain,(
  ( ! [X17,X18] : (p(s(bool,X17)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X18),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1761,f245])).
fof(f1761,plain,(
  ( ! [X17,X18] : (s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X18),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X17),s(fun(sK0,fun('type.list.list'(sK0),bool)),X18))))) )),
  inference(superposition,[],[f295,f222])).
fof(f2001,plain,(
  ( ! [X15,X16] : (p(s(bool,X15)) | s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X16),s(sK0,sK2))),s('type.list.list'(sK0),sK1)))) )),
  inference(forward_demodulation,[],[f1760,f245])).
fof(f1760,plain,(
  ( ! [X15,X16] : (s(bool,'F_0') = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X16),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X15),s(fun(sK0,fun('type.list.list'(sK0),bool)),X16))))) )),
  inference(superposition,[],[f295,f220])).
fof(f2000,plain,(
  ( ! [X14,X12,X13] : (~p(s(bool,X12)) | s(bool,X14) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X13),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X14))) )),
  inference(forward_demodulation,[],[f1759,f245])).
fof(f1759,plain,(
  ( ! [X14,X12,X13] : (s(bool,X14) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X13),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X14)) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X12),s(fun(sK0,fun('type.list.list'(sK0),bool)),X13))))) )),
  inference(superposition,[],[f295,f182])).
fof(f1997,plain,(
  ( ! [X10,X11,X9] : (~p(s(bool,X9)) | s(bool,X11) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X10),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,X11))) )),
  inference(forward_demodulation,[],[f1758,f245])).
fof(f1758,plain,(
  ( ! [X10,X11,X9] : (s(bool,X11) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X10),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | ~p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X9),s(fun(sK0,fun('type.list.list'(sK0),bool)),X10)))) | ~p(s(bool,X11))) )),
  inference(superposition,[],[f295,f182])).
fof(f1996,plain,(
  ( ! [X6,X8,X7] : (p(s(bool,X6)) | s(bool,X8) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X7),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X8))) )),
  inference(forward_demodulation,[],[f1757,f245])).
fof(f1757,plain,(
  ( ! [X6,X8,X7] : (s(bool,X8) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X7),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X6),s(fun(sK0,fun('type.list.list'(sK0),bool)),X7)))) | p(s(bool,X8))) )),
  inference(superposition,[],[f295,f179])).
fof(f1993,plain,(
  ( ! [X4,X5,X3] : (p(s(bool,X3)) | s(bool,X5) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X4),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X5))) )),
  inference(forward_demodulation,[],[f1756,f245])).
fof(f1756,plain,(
  ( ! [X4,X5,X3] : (s(bool,X5) = s(bool,happ(s(fun('type.list.list'(sK0),bool),happ(s(fun(sK0,fun('type.list.list'(sK0),bool)),X4),s(sK0,sK2))),s('type.list.list'(sK0),sK1))) | p(s(bool,X5)) | p(s(bool,'const.list.list_CASE_3'(s('type.list.list'(sK0),'const.list.NIL_0'),s(bool,X3),s(fun(sK0,fun('type.list.list'(sK0),bool)),X4))))) )),
  inference(superposition,[],[f295,f179])).
fof(f312,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | $spl8),
  inference(cnf_transformation,[],[f312_D])).
fof(f312_D,plain,(
  s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) <=> ~$spl8),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl8])])).
fof(f294,plain,(
  ( ! [X0] : (s('type.list.list'(sK0),'const.list.NIL_0') = s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0'))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X0),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f253,f242])).
fof(f287,plain,(
  ( ! [X0] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X0),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))))) ) | $spl4),
  inference(cnf_transformation,[],[f287_D])).
fof(f287_D,plain,(
  ( ! [X0] : (p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X0),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))))) ) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f289,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))) | $spl7),
  inference(cnf_transformation,[],[f289_D])).
fof(f289_D,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f271,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.NIL_0')))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X0),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),sK1))))))) )),
  inference(superposition,[],[f162,f242])).
fof(f286,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0')))))) | $spl3),
  inference(cnf_transformation,[],[f286_D])).
fof(f286_D,plain,(
  ~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0')))))) <=> ~$spl3),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl3])])).
fof(f270,plain,(
  ( ! [X0] : (~p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),sK3),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),'const.list.CONS_2'(s(sK0,sK2),s('type.list.list'(sK0),'const.list.NIL_0')))))) | p(s(bool,'const.list.LLEX_3'(s(fun(sK0,fun(sK0,bool)),X0),s('type.list.list'(sK0),'const.list.NIL_0'),s('type.list.list'(sK0),sK1))))) )),
  inference(superposition,[],[f162,f242])).
% SZS output end Proof for thm.list
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 94
% Passive clauses: 338
% Generated clauses: 2929
% Final active clauses: 49
% Final passive clauses: 147
% Input formulas: 55
% Initial clauses: 89
% 
% Selected by SInE selection: 39
% SInE iterations: 1
% 
% Duplicate literals: 78
% Trivial inequalities: 85
% Fw subsumption resolutions: 316
% Fw demodulations: 379
% Bw demodulations: 23
% 
% Simple tautologies: 138
% Equational tautologies: 8
% Forward subsumptions: 1443
% Bw demodulations to eq. taut.: 3
% 
% Binary resolution: 424
% Forward superposition: 808
% Backward superposition: 657
% Self superposition: 8
% Equality resolution: 7
% 
% Split clauses: 117
% Split components: 234
% SAT solver clauses: 152
% SAT solver unit clauses: 21
% SAT solver binary clauses: 100
% 
% Sat splits: 117
% Sat splitting refutations: 30
% 
% Memory used [KB]: 2686
% Time elapsed: 0.061 s
% ------------------------------
% ------------------------------
% Success in time 11.793 s
