Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 299 next slice time: 5
dis+11_7_3 on thm.pred_set
% (56731)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 129
% Passive clauses: 4993
% Generated clauses: 19176
% Final active clauses: 111
% Final passive clauses: 4304
% Input formulas: 35
% Initial clauses: 82
% 
% Duplicate literals: 840
% Trivial inequalities: 301
% Fw subsumption resolutions: 3153
% Fw demodulations: 1100
% Bw demodulations: 8
% 
% Simple tautologies: 405
% Equational tautologies: 89
% Forward subsumptions: 6477
% 
% Binary resolution: 787
% Factoring: 3
% Forward superposition: 7670
% Backward superposition: 4566
% Self superposition: 220
% Equality factoring: 190
% Equality resolution: 33
% 
% Split clauses: 836
% Split components: 1712
% SAT solver clauses: 1870
% SAT solver unit clauses: 45
% SAT solver binary clauses: 483
% 
% Sat splits: 836
% Sat splitting refutations: 8
% 
% TWLsolver clauses: 1867
% TWLsolver calls for satisfiability: 202
% 
% Memory used [KB]: 10106
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 294 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.pred_set
% (56759)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 132
% Passive clauses: 1069
% Generated clauses: 14681
% Final active clauses: 116
% Final passive clauses: 767
% Input formulas: 35
% Initial clauses: 82
% 
% Duplicate literals: 518
% Trivial inequalities: 211
% Fw subsumption resolutions: 1482
% Fw demodulations: 1761
% Bw demodulations: 165
% Condensations: 1116
% Global subsumptions: 17
% 
% Simple tautologies: 256
% Equational tautologies: 43
% Forward subsumptions: 6470
% Bw demodulations to eq. taut.: 18
% 
% Binary resolution: 650
% Factoring: 3
% Forward superposition: 4005
% Backward superposition: 3421
% Self superposition: 81
% Equality factoring: 59
% Equality resolution: 25
% 
% Split clauses: 151
% Split components: 320
% SAT solver clauses: 1849
% SAT solver unit clauses: 108
% SAT solver binary clauses: 51
% SAT solver learnt clauses: 34
% 
% Sat splits: 151
% Sat splitting refutations: 9
% 
% TWLsolver clauses: 1781
% TWLsolver calls for satisfiability: 350
% 
% Memory used [KB]: 12153
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 286 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.pred_set
% (56786)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 790
% Passive clauses: 12635
% Generated clauses: 149277
% Final active clauses: 108
% Final passive clauses: 1031
% Input formulas: 35
% Initial clauses: 82
% 
% Duplicate literals: 4051
% Trivial inequalities: 1649
% Fw subsumption resolutions: 14684
% Fw demodulations: 24236
% Bw demodulations: 9774
% Condensations: 8150
% Global subsumptions: 1547
% 
% Simple tautologies: 1564
% Equational tautologies: 680
% Forward subsumptions: 75747
% Fw demodulations to eq. taut.: 12
% Bw demodulations to eq. taut.: 430
% 
% Unit resulting resolution: 6895
% Factoring: 2
% Forward superposition: 34509
% Backward superposition: 22254
% Self superposition: 509
% Equality factoring: 257
% Equality resolution: 182
% 
% Split clauses: 347
% Split components: 700
% SAT solver clauses: 17851
% SAT solver unit clauses: 1233
% SAT solver binary clauses: 2354
% 
% Sat splits: 347
% Sat splitting refutations: 2973
% 
% Memory used [KB]: 58335
% Time elapsed: 10.400 s
% ------------------------------
% ------------------------------
% remaining time: 182 next slice time: 39
dis+11_5_fde=none:nwc=1:sas=minisat:sd=1:ss=axioms:st=5.0:sos=all:sdd=large:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none_37 on thm.pred_set
% (57117)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 461
% Passive clauses: 27093
% Generated clauses: 85906
% Final active clauses: 351
% Final passive clauses: 24992
% Input formulas: 35
% Initial clauses: 82
% 
% Selected by SInE selection: 35
% SInE iterations: 1
% 
% Duplicate literals: 2353
% Trivial inequalities: 1536
% Fw subsumption resolutions: 11617
% Fw demodulations: 7600
% Bw demodulations: 95
% 
% Simple tautologies: 2105
% Equational tautologies: 74
% Forward subsumptions: 31896
% Bw demodulations to eq. taut.: 2
% 
% Binary resolution: 12059
% Factoring: 6
% Forward superposition: 36522
% Backward superposition: 12839
% Self superposition: 265
% Equality factoring: 122
% Equality resolution: 116
% 
% Split clauses: 2571
% Split components: 5353
% SAT solver clauses: 2630
% SAT solver unit clauses: 30
% SAT solver binary clauses: 1361
% 
% Sat splits: 2571
% Sat splitting refutations: 33
% 
% Memory used [KB]: 33645
% Time elapsed: 4.0000 s
% ------------------------------
% ------------------------------
% remaining time: 142 next slice time: 6
lrs+1002_2:3_br=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:sas=minisat:stl=30:sac=on:sser=off:sfr=on:ssfp=100000:ssfq=2.0:ssnc=none:sp=reverse_arity:urr=on_4 on thm.pred_set
% (57260)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 91
% Passive clauses: 1411
% Generated clauses: 14352
% Final active clauses: 84
% Final passive clauses: 1270
% Input formulas: 35
% Initial clauses: 82
% 
% Duplicate literals: 514
% Trivial inequalities: 364
% Fw subsumption resolutions: 2136
% Fw demodulations: 1371
% Bw demodulations: 34
% Global subsumptions: 22
% 
% Simple tautologies: 581
% Equational tautologies: 96
% Forward subsumptions: 7802
% Bw demodulations to eq. taut.: 6
% 
% Unit resulting resolution: 927
% Factoring: 2
% Forward superposition: 5414
% Backward superposition: 3305
% Self superposition: 74
% Equality factoring: 42
% Equality resolution: 9
% 
% Split clauses: 8
% Split components: 16
% SAT solver clauses: 1498
% SAT solver unit clauses: 70
% SAT solver binary clauses: 85
% 
% Sat splits: 8
% Sat splitting refutations: 55
% 
% Memory used [KB]: 8187
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 135 next slice time: 15
lrs+1011_12_bs=on:bsr=unit_only:cond=on:gs=on:gsaa=from_current:gsssp=full:nwc=1.1:sas=minisat:stl=60:sos=all:sac=on:sdd=off:sser=off:sfr=on:ssfp=100000:ssfq=1.2:ssnc=none:sp=reverse_arity:updr=off_14 on thm.pred_set
% (57281)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 173
% Passive clauses: 2849
% Generated clauses: 17305
% Final active clauses: 148
% Final passive clauses: 2376
% Input formulas: 35
% Initial clauses: 82
% 
% Duplicate literals: 631
% Trivial inequalities: 398
% Fw subsumption resolutions: 1096
% Fw demodulations: 1956
% Bw demodulations: 191
% Condensations: 1206
% 
% Simple tautologies: 547
% Equational tautologies: 17
% Forward subsumptions: 8581
% Backward subsumptions: 178
% 
% Binary resolution: 2188
% Factoring: 3
% Forward superposition: 6324
% Backward superposition: 3256
% Self superposition: 25
% Equality factoring: 36
% Equality resolution: 37
% 
% Split clauses: 1
% Split components: 2
% SAT solver clauses: 2852
% SAT solver unit clauses: 38
% SAT solver binary clauses: 351
% 
% Sat splits: 1
% 
% Memory used [KB]: 13944
% Time elapsed: 1.600 s
% ------------------------------
% ------------------------------
% remaining time: 119 next slice time: 9
lrs+10_8:1_bd=preordered:bs=on:ccuc=first:fsr=off:fde=unused:gs=on:gsaa=from_current:gsem=on:gsssp=full:nwc=1:nicw=on:sas=minisat:stl=120:sos=on:sscc=on:sser=off:sfr=on:ssfp=4000:ssfq=1.0:ssnc=none:sp=reverse_arity:urr=on_7 on thm.pred_set
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.pred_set
% SZS output start Proof for thm.pred_set
fof(f3336,plain,(
  $false),
  inference(sat_splitting_refutation,[],[f138,f223,f137,f225,f244,f243,f242,f240,f239,f238,f230,f229,f259,f228,f227,f246,f247,f260,f257,f261,f265,f281,f283,f280,f285,f279,f287,f278,f289,f394,f398,f396,f399,f401,f406,f410,f413,f417,f604,f603,f602,f601,f600,f608,f606,f610,f611,f613,f674,f609,f753,f752,f751,f750,f749,f757,f755,f759,f761,f400,f767,f766,f765,f772,f764,f763,f770,f773,f771,f774,f953,f952,f951,f950,f949,f957,f955,f959,f960,f962,f1032,f1031,f1030,f1029,f1028,f1036,f1034,f1038,f1039,f1041,f1165,f958,f1115,f1167,f1187,f1188,f407,f1195,f1194,f1193,f1192,f1191,f1199,f1197,f408,f1208,f1166,f1377,f1037,f1327,f1379,f1397,f414,f1404,f1403,f1402,f1401,f1400,f1408,f1406,f415,f1417,f1378,f2273,f2271,f2262,f2261,f2260,f2259,f2287,f2288,f1169,f2276,f2277,f2285,f2289,f2290,f2291,f403,f2692,f2691,f2690,f2688,f2687,f2686,f2678,f2677,f2284,f2676,f2675,f2673,f1168,f2705,f2706,f2695,f2696,f2707,f2286,f2708,f2709,f2710,f2711,f2716,f2712,f2718,f2713,f2720,f2721,f2724,f2726,f2727,f2729,f2735,f2723,f2872,f2870,f2861,f2860,f2859,f2858,f2886,f2887,f1171,f1381,f2875,f2876,f2884,f2968,f2970,f2972,f2974,f2987,f2993,f3009,f3010,f3011,f3012,f3013,f3014,f3015,f3020,f3023,f3024,f776,f3292,f3244,f3246,f3248,f3250,f3252,f3254,f3256,f3258,f3261,f3263,f3265,f3267,f3311,f3312,f3277,f3283,f3289,f3313,f3314,f3315,f3316,f3317,f3318,f3319,f3320,f3321,f3322,f3323,f3324,f3325,f3326,f3327,f3328,f3329,f3330,f3331,f3332,f3333,f3334,f3335])).
fof(f3335,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296,f3297,f3298,f3299,f3300,f3301,f3308,f3309])).
fof(f3309,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) ) | $spl48),
  inference(superposition,[],[f776,f147])).
fof(f147,plain,(
  ( ! [X0] : (s(bool,'F_0') = s(bool,X0) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f3])).
fof(f3,axiom,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) | s(bool,'F_0') = s(bool,X0))),
  file('thm.pred_set.INSERT_DELETE',aHL_BOOL_CASES)).
fof(f3308,plain,(
  ( ! [X7] : (~p(s(bool,X7)) | s(bool,'F_0') = s(bool,X7)) ) | $spl48),
  inference(resolution,[],[f776,f217])).
fof(f217,plain,(
  ( ! [X0,X1] : (~p(s(bool,X1)) | ~p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f88,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | (p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1))),
  inference(flattening,[],[f87])).
fof(f87,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) & ~p(s(bool,X1))) | ((p(s(bool,X1)) & ~p(s(bool,X0))) | s(bool,X0) = s(bool,X1)))),
  inference(ennf_transformation,[],[f77])).
fof(f77,plain,(
  ! [X0,X1] : ((p(s(bool,X0)) => p(s(bool,X1))) => ((p(s(bool,X1)) => p(s(bool,X0))) => s(bool,X0) = s(bool,X1)))),
  inference(rectify,[],[f8])).
fof(f8,axiom,(
  ! [X9,X10] : ((p(s(bool,X9)) => p(s(bool,X10))) => ((p(s(bool,X10)) => p(s(bool,X9))) => s(bool,X9) = s(bool,X10)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.IMP_ANTISYM_AX')).
fof(f3301,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(resolution,[],[f776,f152])).
fof(f152,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f94])).
fof(f94,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f46])).
fof(f46,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f31])).
fof(f31,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.EQ_CLAUSES|split|1')).
fof(f3300,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(resolution,[],[f776,f150])).
fof(f150,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f93])).
fof(f93,plain,(
  ! [X0] : ((s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) & (~p(s(bool,X0)) | s(bool,'T_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f45])).
fof(f45,plain,(
  ! [X0] : (s(bool,'T_0') = s(bool,X0) <=> p(s(bool,X0)))),
  inference(rectify,[],[f30])).
fof(f30,axiom,(
  ! [X11] : (s(bool,'T_0') = s(bool,X11) <=> p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.EQ_CLAUSES|split|0')).
fof(f3299,plain,(
  ( ! [X6] : (~p(s(bool,X6))) ) | $spl48),
  inference(resolution,[],[f776,f200])).
fof(f200,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f126])).
fof(f126,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | ~p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f125])).
fof(f125,plain,(
  ! [X0] : (((p(s(bool,X0)) & ~p(s(bool,'F_0'))) | ~p(s(bool,X0))) & (p(s(bool,X0)) | (~p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f86])).
fof(f86,plain,(
  ! [X0] : ((~p(s(bool,X0)) | p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(ennf_transformation,[],[f66])).
fof(f66,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f65])).
fof(f65,plain,(
  ! [X0] : ((p(s(bool,X0)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f25])).
fof(f25,axiom,(
  ! [X11] : ((p(s(bool,X11)) => p(s(bool,'F_0'))) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.IMP_CLAUSES|split|4')).
fof(f3298,plain,(
  ( ! [X5] : (p(s(bool,X5)) | ~p(s(bool,'T_0'))) ) | $spl48),
  inference(resolution,[],[f776,f192])).
fof(f192,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)) | ~p(s(bool,'T_0'))) )),
  inference(cnf_transformation,[],[f120])).
fof(f120,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f119])).
fof(f119,plain,(
  ! [X0] : (((p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | (~p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f83])).
fof(f83,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(ennf_transformation,[],[f62])).
fof(f62,plain,(
  ! [X0] : ((p(s(bool,'F_0')) => p(s(bool,X0))) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f23])).
fof(f23,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) => p(s(bool,X11))) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.IMP_CLAUSES|split|2')).
fof(f3297,plain,(
  ( ! [X4] : (p(s(bool,X4))) ) | $spl48),
  inference(resolution,[],[f776,f182])).
fof(f182,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f114])).
fof(f114,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,X0)) | p(s(bool,'F_0'))))),
  inference(flattening,[],[f113])).
fof(f113,plain,(
  ! [X0] : (((~p(s(bool,X0)) & ~p(s(bool,'F_0'))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,X0)) | p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f59])).
fof(f59,plain,(
  ! [X0] : ((p(s(bool,X0)) | p(s(bool,'F_0'))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f19])).
fof(f19,axiom,(
  ! [X11] : ((p(s(bool,X11)) | p(s(bool,'F_0'))) <=> p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.OR_CLAUSES|split|3')).
fof(f3296,plain,(
  ( ! [X3] : (p(s(bool,X3))) ) | $spl48),
  inference(resolution,[],[f776,f172])).
fof(f172,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f108])).
fof(f108,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | p(s(bool,'F_0')) | p(s(bool,X0))))),
  inference(flattening,[],[f107])).
fof(f107,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) & ~p(s(bool,X0))) | p(s(bool,X0))) & (~p(s(bool,X0)) | (p(s(bool,'F_0')) | p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f56])).
fof(f56,plain,(
  ! [X0] : ((p(s(bool,'F_0')) | p(s(bool,X0))) <=> p(s(bool,X0)))),
  inference(rectify,[],[f18])).
fof(f18,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) | p(s(bool,X11))) <=> p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.OR_CLAUSES|split|2')).
fof(f3295,plain,(
  ( ! [X2] : (p(s(bool,X2))) ) | $spl48),
  inference(resolution,[],[f776,f171])).
fof(f171,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f106])).
fof(f106,plain,(
  ! [X0] : ((~p(s(bool,'F_0')) | ~p(s(bool,X0)) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(flattening,[],[f105])).
fof(f105,plain,(
  ! [X0] : (((~p(s(bool,'F_0')) | ~p(s(bool,X0))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,'F_0')) & p(s(bool,X0)))))),
  inference(nnf_transformation,[],[f55])).
fof(f55,plain,(
  ! [X0] : ((p(s(bool,'F_0')) & p(s(bool,X0))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f13])).
fof(f13,axiom,(
  ! [X11] : ((p(s(bool,'F_0')) & p(s(bool,X11))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.AND_CLAUSES|split|2')).
fof(f3294,plain,(
  ( ! [X1] : (p(s(bool,X1))) ) | $spl48),
  inference(resolution,[],[f776,f164])).
fof(f164,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f102])).
fof(f102,plain,(
  ! [X0] : ((~p(s(bool,X0)) | ~p(s(bool,'F_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(flattening,[],[f101])).
fof(f101,plain,(
  ! [X0] : (((~p(s(bool,X0)) | ~p(s(bool,'F_0'))) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | (p(s(bool,X0)) & p(s(bool,'F_0')))))),
  inference(nnf_transformation,[],[f53])).
fof(f53,plain,(
  ! [X0] : ((p(s(bool,X0)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  inference(rectify,[],[f14])).
fof(f14,axiom,(
  ! [X11] : ((p(s(bool,X11)) & p(s(bool,'F_0'))) <=> p(s(bool,'F_0')))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.AND_CLAUSES|split|3')).
fof(f3293,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl48),
  inference(resolution,[],[f776,f148])).
fof(f148,plain,(
  ( ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f81])).
fof(f81,plain,(
  ! [X0] : (~p(s(bool,'F_0')) | p(s(bool,X0)))),
  inference(ennf_transformation,[],[f44])).
fof(f44,plain,(
  ! [X0] : (p(s(bool,'F_0')) => p(s(bool,X0)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X11] : (p(s(bool,'F_0')) => p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.FALSITY')).
fof(f3291,plain,(
  ~p(s(bool,'T_0')) | $spl48),
  inference(resolution,[],[f776,f144])).
fof(f144,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f92])).
fof(f92,plain,(
  (p(s(bool,'F_0')) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | ~p(s(bool,'F_0')))),
  inference(nnf_transformation,[],[f40])).
fof(f40,plain,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  inference(flattening,[],[f28])).
fof(f28,axiom,(
  ~p(s(bool,'F_0')) <=> p(s(bool,'T_0'))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.NOT_CLAUSES|split|2')).
fof(f3290,plain,(
  ~p(s(bool,'T_0')) | $spl48),
  inference(resolution,[],[f776,f142])).
fof(f142,plain,(
  ~p(s(bool,'F_0')) | ~p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f91])).
fof(f91,plain,(
  (p(s(bool,'T_0')) | p(s(bool,'F_0'))) & (~p(s(bool,'F_0')) | ~p(s(bool,'T_0')))),
  inference(nnf_transformation,[],[f39])).
fof(f39,plain,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  inference(flattening,[],[f27])).
fof(f27,axiom,(
  ~p(s(bool,'T_0')) <=> p(s(bool,'F_0'))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.NOT_CLAUSES|split|1')).
fof(f3245,plain,(
  ~p(s(bool,'T_0')) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f142])).
fof(f3247,plain,(
  ~p(s(bool,'T_0')) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f144])).
fof(f3249,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f148])).
fof(f3251,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f164])).
fof(f3253,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f171])).
fof(f3255,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f172])).
fof(f3257,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f182])).
fof(f3259,plain,(
  ~p(s(bool,'T_0')) | $spl48),
  inference(unit_resulting_resolution,[],[f145,f776,f192])).
fof(f145,plain,(
  ~p(s(bool,'F_0'))),
  inference(cnf_transformation,[],[f41])).
fof(f41,plain,(
  ~p(s(bool,'F_0'))),
  inference(flattening,[],[f2])).
fof(f2,axiom,(
  ~p(s(bool,'F_0'))),
  file('thm.pred_set.INSERT_DELETE',aHL_FALSITY)).
fof(f3260,plain,(
  ( ! [X0] : (p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f221,f776,f192])).
fof(f221,plain,(
  p(s(bool,'T_0'))),
  inference(trivial_inequality_removal,[],[f204])).
fof(f204,plain,(
  ( ! [X0,X1] : (p(s(bool,'T_0')) | s(X0,X1) != s(X0,X1)) )),
  inference(cnf_transformation,[],[f128])).
fof(f128,plain,(
  ! [X0,X1] : ((s(X0,X1) != s(X0,X1) | p(s(bool,'T_0'))) & (~p(s(bool,'T_0')) | s(X0,X1) = s(X0,X1)))),
  inference(nnf_transformation,[],[f70])).
fof(f70,plain,(
  ! [X0,X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(flattening,[],[f69])).
fof(f69,plain,(
  ! [X0] : ! [X1] : (s(X0,X1) = s(X0,X1) <=> p(s(bool,'T_0')))),
  inference(rectify,[],[f29])).
fof(f29,axiom,(
  ! [X6] : ! [X7] : (s(X6,X7) = s(X6,X7) <=> p(s(bool,'T_0')))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.REFL_CLAUSE')).
fof(f3262,plain,(
  ( ! [X0] : (~p(s(bool,X0))) ) | $spl48),
  inference(unit_resulting_resolution,[],[f776,f200])).
fof(f3264,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f776,f150])).
fof(f3266,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f776,f152])).
fof(f3278,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f221,f776,f217])).
fof(f3279,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f140,f776,f217])).
fof(f140,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f7])).
fof(f7,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.TRUTH')).
fof(f3280,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f139,f776,f217])).
fof(f139,plain,(
  p(s(bool,'T_0'))),
  inference(cnf_transformation,[],[f1])).
fof(f1,axiom,(
  p(s(bool,'T_0'))),
  file('thm.pred_set.INSERT_DELETE',aHL_TRUTH)).
fof(f3284,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f221,f776,f217])).
fof(f3285,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f140,f776,f217])).
fof(f3286,plain,(
  s(bool,'T_0') = s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f139,f776,f217])).
fof(f3334,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296,f3297,f3298,f3299,f3300,f3301,f3308])).
fof(f3333,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296,f3297,f3298,f3299,f3300,f3301])).
fof(f3332,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296,f3297,f3298,f3299,f3300])).
fof(f3331,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296,f3297,f3298,f3299])).
fof(f3330,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296,f3297,f3298])).
fof(f3329,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296,f3297])).
fof(f3328,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295,f3296])).
fof(f3327,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294,f3295])).
fof(f3326,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293,f3294])).
fof(f3325,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291,f3293])).
fof(f3324,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290,f3291])).
fof(f3323,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245,f3290])).
fof(f3322,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247,f3245])).
fof(f3321,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249,f3247])).
fof(f3320,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251,f3249])).
fof(f3319,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253,f3251])).
fof(f3318,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255,f3253])).
fof(f3317,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257,f3255])).
fof(f3316,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259,f3257])).
fof(f3315,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260,f3259])).
fof(f3314,plain,(
  $false | $spl48),
  inference(global_subsumption,[],[f138,f137,f776,f3286,f3285,f3284,f3280,f3279,f3278,f3266,f3264,f3262,f3260])).
fof(f3313,plain,(
  $false | ($spl9 | $spl48)),
  inference(global_subsumption,[],[f138,f137,f265,f776,f3288,f3286,f3285,f3284,f3282])).
fof(f3282,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl48)),
  inference(unit_resulting_resolution,[],[f265,f776,f217])).
fof(f3288,plain,(
  ~p(s(bool,'T_0')) | ($spl9 | $spl48)),
  inference(unit_resulting_resolution,[],[f265,f776,f217])).
fof(f3289,plain,(
  $false | ($spl9 | $spl48)),
  inference(unit_resulting_resolution,[],[f221,f265,f776,f217])).
fof(f3283,plain,(
  $false | ($spl9 | $spl48)),
  inference(unit_resulting_resolution,[],[f221,f265,f776,f217])).
fof(f3277,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f776,f776,f200])).
fof(f3312,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f3269])).
fof(f3269,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f776,f155])).
fof(f155,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f96])).
fof(f96,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f50])).
fof(f50,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f49])).
fof(f49,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f32])).
fof(f32,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.EQ_CLAUSES|split|2')).
fof(f3311,plain,(
  $false | $spl48),
  inference(trivial_inequality_removal,[],[f3268])).
fof(f3268,plain,(
  s(bool,'F_0') != s(bool,'F_0') | $spl48),
  inference(unit_resulting_resolution,[],[f776,f153])).
fof(f153,plain,(
  ( ! [X0] : (s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f95])).
fof(f95,plain,(
  ! [X0] : ((s(bool,'F_0') != s(bool,X0) | ~p(s(bool,X0))) & (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)))),
  inference(nnf_transformation,[],[f48])).
fof(f48,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(flattening,[],[f47])).
fof(f47,plain,(
  ! [X0] : (s(bool,'F_0') = s(bool,X0) <=> ~p(s(bool,X0)))),
  inference(rectify,[],[f33])).
fof(f33,axiom,(
  ! [X11] : (s(bool,'F_0') = s(bool,X11) <=> ~p(s(bool,X11)))),
  file('thm.pred_set.INSERT_DELETE','thm.bool.EQ_CLAUSES|split|3')).
fof(f3267,plain,(
  $false | ($spl9 | $spl48)),
  inference(unit_resulting_resolution,[],[f265,f776,f152])).
fof(f3265,plain,(
  $false | ($spl9 | $spl48)),
  inference(unit_resulting_resolution,[],[f265,f776,f150])).
fof(f3263,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f221,f776,f200])).
fof(f3261,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f145,f221,f776,f192])).
fof(f3258,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f145,f776,f182])).
fof(f3256,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f145,f776,f172])).
fof(f3254,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f145,f776,f171])).
fof(f3252,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f145,f776,f164])).
fof(f3250,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f145,f776,f148])).
fof(f3248,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f221,f776,f144])).
fof(f3246,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f221,f776,f142])).
fof(f3244,plain,(
  $false | $spl48),
  inference(unit_resulting_resolution,[],[f145,f776])).
fof(f3292,plain,(
  $false | $spl48),
  inference(resolution,[],[f776,f145])).
fof(f776,plain,(
  p(s(bool,'F_0')) | $spl48),
  inference(cnf_transformation,[],[f776_D])).
fof(f776_D,plain,(
  p(s(bool,'F_0')) <=> ~$spl48),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl48])])).
fof(f3024,plain,(
  p(s(bool,'F_0')) | ($spl54 | $spl68)),
  inference(forward_demodulation,[],[f1199,f961])).
fof(f961,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl54),
  inference(cnf_transformation,[],[f961_D])).
fof(f961_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl54),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl54])])).
fof(f3023,plain,(
  p(s(bool,'F_0')) | ($spl70 | $spl72 | $spl82)),
  inference(forward_demodulation,[],[f3022,f2718])).
fof(f3022,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | ($spl70 | $spl72)),
  inference(forward_demodulation,[],[f1379,f1381])).
fof(f3020,plain,(
  p(s(bool,'F_0')) | ($spl82 | $spl88)),
  inference(forward_demodulation,[],[f2725,f2718])).
fof(f2725,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl88),
  inference(cnf_transformation,[],[f2725_D])).
fof(f2725_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) <=> ~$spl88),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl88])])).
fof(f3015,plain,(
  s(sK0,sK1) != s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))),
  introduced(theory_axiom,[])).
fof(f3014,plain,(
  s(sK0,sK1) != s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))),
  introduced(theory_axiom,[])).
fof(f3013,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | s(sK0,sK1) != s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f3012,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | s(sK0,sK1) != s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f3011,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | s(sK0,sK1) != s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f3010,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | ~p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))))),
  introduced(theory_axiom,[])).
fof(f3009,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f2993,plain,(
  $false | ($spl64 | $spl87)),
  inference(unit_resulting_resolution,[],[f221,f2723,f2839,f217])).
fof(f2839,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X0))))))) ) | $spl64),
  inference(forward_demodulation,[],[f2820,f1169])).
fof(f2820,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),X0))))))) ) | $spl64),
  inference(unit_resulting_resolution,[],[f1169,f212])).
fof(f212,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X1) != s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f135,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(flattening,[],[f134])).
fof(f134,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3)))))) & ((s(X0,X1) != s(X0,X2) & ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3))))))))),
  inference(nnf_transformation,[],[f76])).
fof(f76,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(flattening,[],[f75])).
fof(f75,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) <=> (s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))))),
  inference(rectify,[],[f35])).
fof(f35,axiom,(
  ! [X6] : ! [X7,X12,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X12),s(fun(X6,bool),X8)))))) <=> (s(X6,X7) = s(X6,X12) | p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))))))),
  file('thm.pred_set.INSERT_DELETE','thm.pred_set.IN_INSERT')).
fof(f2987,plain,(
  $false | ($spl64 | $spl87)),
  inference(unit_resulting_resolution,[],[f221,f2723,f2839,f217])).
fof(f2974,plain,(
  $false | ($spl64 | $spl90)),
  inference(unit_resulting_resolution,[],[f2729,f2839,f155])).
fof(f2972,plain,(
  $false | ($spl64 | $spl90)),
  inference(unit_resulting_resolution,[],[f2729,f2839,f153])).
fof(f2970,plain,(
  $false | ($spl64 | $spl87)),
  inference(unit_resulting_resolution,[],[f2723,f2839,f152])).
fof(f2968,plain,(
  $false | ($spl64 | $spl87)),
  inference(unit_resulting_resolution,[],[f2723,f2839,f150])).
fof(f2884,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl92),
  inference(cnf_transformation,[],[f2884_D])).
fof(f2884_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl92),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl92])])).
fof(f2876,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(resolution,[],[f1378,f156])).
fof(f156,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f96])).
fof(f2875,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(resolution,[],[f1378,f154])).
fof(f154,plain,(
  ( ! [X0] : (p(s(bool,X0)) | s(bool,'F_0') = s(bool,X0)) )),
  inference(cnf_transformation,[],[f95])).
fof(f1381,plain,(
  s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | $spl72),
  inference(cnf_transformation,[],[f1381_D])).
fof(f1381_D,plain,(
  s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) <=> ~$spl72),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl72])])).
fof(f1171,plain,(
  ~p(s(bool,'T_0')) | $spl67),
  inference(cnf_transformation,[],[f1171_D])).
fof(f1171_D,plain,(
  ~p(s(bool,'T_0')) <=> ~$spl67),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl67])])).
fof(f2887,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | ($spl58 | $spl71)),
  inference(forward_demodulation,[],[f2874,f1038])).
fof(f2874,plain,(
  s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl71),
  inference(resolution,[],[f1378,f210])).
fof(f210,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | s(X0,X2) = s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f133,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(flattening,[],[f132])).
fof(f132,plain,(
  ! [X0,X1,X2,X3] : ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3))) & ((~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) | s(X0,X2) = s(X0,X3)) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))))),
  inference(nnf_transformation,[],[f74])).
fof(f74,plain,(
  ! [X0,X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & s(X0,X2) != s(X0,X3)))),
  inference(flattening,[],[f73])).
fof(f73,plain,(
  ! [X0] : ! [X1,X2,X3] : (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1)))) & ~s(X0,X2) = s(X0,X3)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X6] : ! [X8,X7,X12] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X12)))))) <=> (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) & ~s(X6,X7) = s(X6,X12)))),
  file('thm.pred_set.INSERT_DELETE','thm.pred_set.IN_DELETE')).
fof(f2886,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl95),
  inference(cnf_transformation,[],[f2886_D])).
fof(f2886_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl95),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl95])])).
fof(f2858,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(unit_resulting_resolution,[],[f1378,f149])).
fof(f149,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f93])).
fof(f2859,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(unit_resulting_resolution,[],[f1378,f151])).
fof(f151,plain,(
  ( ! [X0] : (s(bool,'T_0') != s(bool,X0) | p(s(bool,X0))) )),
  inference(cnf_transformation,[],[f94])).
fof(f2860,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(unit_resulting_resolution,[],[f1378,f154])).
fof(f2861,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(unit_resulting_resolution,[],[f1378,f156])).
fof(f2870,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(unit_resulting_resolution,[],[f145,f1378,f214])).
fof(f214,plain,(
  ( ! [X0,X1] : (p(s(bool,X1)) | p(s(bool,X0)) | s(bool,X0) = s(bool,X1)) )),
  inference(cnf_transformation,[],[f88])).
fof(f2872,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl71),
  inference(unit_resulting_resolution,[],[f145,f1378,f214])).
fof(f2723,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl87),
  inference(cnf_transformation,[],[f2723_D])).
fof(f2723_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl87),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl87])])).
fof(f2735,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl27 | $spl64)),
  inference(forward_demodulation,[],[f407,f1169])).
fof(f2729,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl90),
  inference(cnf_transformation,[],[f2729_D])).
fof(f2729_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl90),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl90])])).
fof(f2727,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl54 | $spl64)),
  inference(forward_demodulation,[],[f961,f1169])).
fof(f2726,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl89),
  inference(cnf_transformation,[],[f2726_D])).
fof(f2726_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) <=> ~$spl89),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl89])])).
fof(f2724,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | ($spl63 | $spl64)),
  inference(forward_demodulation,[],[f1166,f1169])).
fof(f2721,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl4 | $spl13 | $spl64)),
  inference(forward_demodulation,[],[f2714,f257])).
fof(f2714,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl13 | $spl64)),
  inference(backward_demodulation,[],[f1169,f285])).
fof(f2720,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl85),
  inference(cnf_transformation,[],[f2720_D])).
fof(f2720_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl85),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl85])])).
fof(f2713,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | ($spl64 | $spl79)),
  inference(backward_demodulation,[],[f1169,f2287])).
fof(f2718,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl82),
  inference(cnf_transformation,[],[f2718_D])).
fof(f2718_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl82),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl82])])).
fof(f2712,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | ($spl64 | $spl76)),
  inference(backward_demodulation,[],[f1169,f2285])).
fof(f2716,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl81),
  inference(cnf_transformation,[],[f2716_D])).
fof(f2716_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) <=> ~$spl81),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl81])])).
fof(f2711,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | ($spl64 | $spl69)),
  inference(backward_demodulation,[],[f1169,f1198])).
fof(f1198,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl69),
  inference(cnf_transformation,[],[f1198_D])).
fof(f1198_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) <=> ~$spl69),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl69])])).
fof(f2710,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))))),
  introduced(theory_axiom,[])).
fof(f2709,plain,(
  s(sK0,sK1) != s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f2708,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | ~p(s(bool,'F_0'))),
  introduced(theory_axiom,[])).
fof(f2286,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl78),
  inference(cnf_transformation,[],[f2286_D])).
fof(f2286_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl78),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl78])])).
fof(f2707,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(global_subsumption,[],[f138,f137,f1167,f2692,f2691,f2690,f2688,f2687,f2686,f2678,f2677,f2676,f2675,f2674,f2673,f2695,f2696,f2703,f2704])).
fof(f2704,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(superposition,[],[f1167,f147])).
fof(f2703,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) ) | $spl62),
  inference(resolution,[],[f1167,f217])).
fof(f2674,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))))) ) | $spl62),
  inference(unit_resulting_resolution,[],[f1167,f213])).
fof(f213,plain,(
  ( ! [X2,X0,X3,X1] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | ~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f2696,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(resolution,[],[f1167,f152])).
fof(f2695,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(resolution,[],[f1167,f150])).
fof(f2706,plain,(
  p(s(bool,'F_0')) | ($spl28 | $spl62)),
  inference(forward_demodulation,[],[f2694,f409])).
fof(f409,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl28),
  inference(cnf_transformation,[],[f409_D])).
fof(f409_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl28),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl28])])).
fof(f2694,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl62),
  inference(resolution,[],[f1167,f208])).
fof(f208,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3)))))) | p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),X1))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f2705,plain,(
  p(s(bool,'F_0')) | ($spl28 | $spl62)),
  inference(forward_demodulation,[],[f2672,f409])).
fof(f2672,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl62),
  inference(unit_resulting_resolution,[],[f1167,f208])).
fof(f1168,plain,(
  s(sK0,sK1) != s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) | $spl65),
  inference(cnf_transformation,[],[f1168_D])).
fof(f1168_D,plain,(
  s(sK0,sK1) != s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) <=> ~$spl65),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl65])])).
fof(f2673,plain,(
  s(sK0,sK1) != s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) | $spl62),
  inference(unit_resulting_resolution,[],[f1167,f209])).
fof(f209,plain,(
  ( ! [X2,X0,X3,X1] : (s(X0,X2) != s(X0,X3) | ~p(s(bool,'const.bool.IN_2'(s(X0,X2),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X1),s(X0,X3))))))) )),
  inference(cnf_transformation,[],[f133])).
fof(f2675,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1167,f150])).
fof(f2676,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1167,f152])).
fof(f2284,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl77),
  inference(cnf_transformation,[],[f2284_D])).
fof(f2284_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl77),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl77])])).
fof(f2677,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1167,f153])).
fof(f2678,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f1167,f155])).
fof(f2686,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f221,f1167,f217])).
fof(f2687,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f140,f1167,f217])).
fof(f2688,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f139,f1167,f217])).
fof(f2690,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f221,f1167,f217])).
fof(f2691,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f140,f1167,f217])).
fof(f2692,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl62),
  inference(unit_resulting_resolution,[],[f139,f1167,f217])).
fof(f403,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl25),
  inference(cnf_transformation,[],[f403_D])).
fof(f403_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl25),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl25])])).
fof(f2291,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | ($spl15 | $spl46)),
  inference(backward_demodulation,[],[f772,f287])).
fof(f2290,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f2289,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f2285,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl76),
  inference(cnf_transformation,[],[f2285_D])).
fof(f2285_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl76),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl76])])).
fof(f2277,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(resolution,[],[f1166,f156])).
fof(f2276,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(resolution,[],[f1166,f154])).
fof(f1169,plain,(
  s(sK0,sK1) = s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) | $spl64),
  inference(cnf_transformation,[],[f1169_D])).
fof(f1169_D,plain,(
  s(sK0,sK1) = s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) <=> ~$spl64),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl64])])).
fof(f2288,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK1) = s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) | ($spl52 | $spl63)),
  inference(forward_demodulation,[],[f2275,f959])).
fof(f2275,plain,(
  s(sK0,sK1) = s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl63),
  inference(resolution,[],[f1166,f210])).
fof(f2287,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl79),
  inference(cnf_transformation,[],[f2287_D])).
fof(f2287_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl79),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl79])])).
fof(f2259,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1166,f149])).
fof(f2260,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1166,f151])).
fof(f2261,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1166,f154])).
fof(f2262,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f1166,f156])).
fof(f2271,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f145,f1166,f214])).
fof(f2273,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl63),
  inference(unit_resulting_resolution,[],[f145,f1166,f214])).
fof(f1378,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl71),
  inference(cnf_transformation,[],[f1378_D])).
fof(f1378_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) <=> ~$spl71),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl71])])).
fof(f1417,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl60),
  inference(superposition,[],[f213,f1040])).
fof(f1040,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl60),
  inference(cnf_transformation,[],[f1040_D])).
fof(f1040_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl60),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl60])])).
fof(f415,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl30),
  inference(cnf_transformation,[],[f415_D])).
fof(f415_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl30),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl30])])).
fof(f1406,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl61),
  inference(trivial_inequality_removal,[],[f1405])).
fof(f1405,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl61),
  inference(superposition,[],[f1041,f147])).
fof(f1408,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl74),
  inference(cnf_transformation,[],[f1408_D])).
fof(f1408_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) <=> ~$spl74),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl74])])).
fof(f1400,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f1041,f156])).
fof(f1401,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f1041,f154])).
fof(f1402,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f1041,f147])).
fof(f1403,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f145,f1041,f214])).
fof(f1404,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl61),
  inference(unit_resulting_resolution,[],[f145,f1041,f214])).
fof(f414,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl31),
  inference(cnf_transformation,[],[f414_D])).
fof(f414_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl31),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl31])])).
fof(f1397,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl17 | $spl58)),
  inference(backward_demodulation,[],[f1038,f289])).
fof(f1379,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl70),
  inference(cnf_transformation,[],[f1379_D])).
fof(f1379_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) <=> ~$spl70),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl70])])).
fof(f1327,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK1) = s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl30),
  inference(superposition,[],[f211,f415])).
fof(f211,plain,(
  ( ! [X2,X0,X3,X1] : (~p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X2),s(fun(X0,bool),X3)))))) | s(X0,X1) = s(X0,X2) | p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X3))))) )),
  inference(cnf_transformation,[],[f135])).
fof(f1037,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl59),
  inference(cnf_transformation,[],[f1037_D])).
fof(f1037_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl59),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl59])])).
fof(f1377,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl30),
  inference(global_subsumption,[],[f138,f137,f415,f1326])).
fof(f1326,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl30),
  inference(superposition,[],[f207,f415])).
fof(f207,plain,(
  ( ! [X2,X0,X1] : (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)) )),
  inference(cnf_transformation,[],[f131])).
fof(f131,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,sK3(X2,X1,X0)),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK3])],[f130])).
fof(f130,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X4] : s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X4),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(rectify,[],[f129])).
fof(f129,plain,(
  ! [X0,X1,X2] : ((s(fun(X0,bool),X1) != s(fun(X0,bool),X2) | ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2)))) & (? [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) != s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))) | s(fun(X0,bool),X1) = s(fun(X0,bool),X2)))),
  inference(nnf_transformation,[],[f72])).
fof(f72,plain,(
  ! [X0,X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ! [X0] : ! [X1,X2] : (s(fun(X0,bool),X1) = s(fun(X0,bool),X2) <=> ! [X3] : s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X1))) = s(bool,'const.bool.IN_2'(s(X0,X3),s(fun(X0,bool),X2))))),
  inference(rectify,[],[f34])).
fof(f34,axiom,(
  ! [X6] : ! [X8,X11] : (s(fun(X6,bool),X8) = s(fun(X6,bool),X11) <=> ! [X7] : s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8))) = s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X11))))),
  file('thm.pred_set.INSERT_DELETE','thm.pred_set.EXTENSION')).
fof(f1166,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl63),
  inference(cnf_transformation,[],[f1166_D])).
fof(f1166_D,plain,(
  ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) <=> ~$spl63),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl63])])).
fof(f1208,plain,(
  p(s(bool,'F_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl54),
  inference(superposition,[],[f213,f961])).
fof(f408,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl26),
  inference(cnf_transformation,[],[f408_D])).
fof(f408_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl26),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl26])])).
fof(f1197,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl55),
  inference(trivial_inequality_removal,[],[f1196])).
fof(f1196,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl55),
  inference(superposition,[],[f962,f147])).
fof(f1199,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl68),
  inference(cnf_transformation,[],[f1199_D])).
fof(f1199_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) <=> ~$spl68),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl68])])).
fof(f1191,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl55),
  inference(unit_resulting_resolution,[],[f962,f156])).
fof(f1192,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl55),
  inference(unit_resulting_resolution,[],[f962,f154])).
fof(f1193,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl55),
  inference(unit_resulting_resolution,[],[f962,f147])).
fof(f1194,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl55),
  inference(unit_resulting_resolution,[],[f145,f962,f214])).
fof(f1195,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))))) | $spl55),
  inference(unit_resulting_resolution,[],[f145,f962,f214])).
fof(f407,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl27),
  inference(cnf_transformation,[],[f407_D])).
fof(f407_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl27),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl27])])).
fof(f1188,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl13 | $spl52)),
  inference(backward_demodulation,[],[f959,f285])).
fof(f1187,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | p(s(bool,'T_0')) | ~p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))),
  introduced(theory_axiom,[])).
fof(f1167,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl62),
  inference(cnf_transformation,[],[f1167_D])).
fof(f1167_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) <=> ~$spl62),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl62])])).
fof(f1115,plain,(
  ~p(s(bool,'T_0')) | s(sK0,sK1) = s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)) | p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))) | $spl26),
  inference(superposition,[],[f211,f408])).
fof(f958,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl53),
  inference(cnf_transformation,[],[f958_D])).
fof(f958_D,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl53),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl53])])).
fof(f1165,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl26),
  inference(global_subsumption,[],[f138,f137,f408,f1114])).
fof(f1114,plain,(
  s(bool,'T_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl26),
  inference(superposition,[],[f207,f408])).
fof(f1041,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl61),
  inference(cnf_transformation,[],[f1041_D])).
fof(f1041_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl61),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl61])])).
fof(f1039,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl17 | $spl32)),
  inference(backward_demodulation,[],[f416,f289])).
fof(f416,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl32),
  inference(cnf_transformation,[],[f416_D])).
fof(f416_D,plain,(
  s(bool,'F_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl32),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl32])])).
fof(f1038,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl58),
  inference(cnf_transformation,[],[f1038_D])).
fof(f1038_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl58),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl58])])).
fof(f1034,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl33),
  inference(trivial_inequality_removal,[],[f1033])).
fof(f1033,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl33),
  inference(superposition,[],[f417,f147])).
fof(f1036,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl56),
  inference(cnf_transformation,[],[f1036_D])).
fof(f1036_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl56),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl56])])).
fof(f1028,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f417,f156])).
fof(f1029,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f417,f154])).
fof(f1030,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl33),
  inference(unit_resulting_resolution,[],[f417,f147])).
fof(f1031,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f145,f417,f214])).
fof(f1032,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2)))) | $spl33),
  inference(unit_resulting_resolution,[],[f145,f417,f214])).
fof(f962,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl55),
  inference(cnf_transformation,[],[f962_D])).
fof(f962_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl55),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl55])])).
fof(f960,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | ($spl13 | $spl28)),
  inference(backward_demodulation,[],[f409,f285])).
fof(f959,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl52),
  inference(cnf_transformation,[],[f959_D])).
fof(f959_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl52),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl52])])).
fof(f955,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl29),
  inference(trivial_inequality_removal,[],[f954])).
fof(f954,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl29),
  inference(superposition,[],[f410,f147])).
fof(f957,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl50),
  inference(cnf_transformation,[],[f957_D])).
fof(f957_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) <=> ~$spl50),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl50])])).
fof(f949,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl29),
  inference(unit_resulting_resolution,[],[f410,f156])).
fof(f950,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl29),
  inference(unit_resulting_resolution,[],[f410,f154])).
fof(f951,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl29),
  inference(unit_resulting_resolution,[],[f410,f147])).
fof(f952,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl29),
  inference(unit_resulting_resolution,[],[f145,f410,f214])).
fof(f953,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2)))) | $spl29),
  inference(unit_resulting_resolution,[],[f145,f410,f214])).
fof(f774,plain,(
  p(s(bool,'F_0')) | ($spl24 | $spl40)),
  inference(backward_demodulation,[],[f402,f757])).
fof(f402,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl24),
  inference(cnf_transformation,[],[f402_D])).
fof(f402_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl24),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl24])])).
fof(f771,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl47),
  inference(cnf_transformation,[],[f771_D])).
fof(f771_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl47),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl47])])).
fof(f773,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | ($spl15 | $spl24)),
  inference(backward_demodulation,[],[f402,f287])).
fof(f770,plain,(
  ~p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl45),
  inference(cnf_transformation,[],[f770_D])).
fof(f770_D,plain,(
  ~p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) <=> ~$spl45),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl45])])).
fof(f763,plain,(
  ~p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl23),
  inference(unit_resulting_resolution,[],[f400,f152])).
fof(f764,plain,(
  ~p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl23),
  inference(unit_resulting_resolution,[],[f400,f150])).
fof(f772,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl46),
  inference(cnf_transformation,[],[f772_D])).
fof(f772_D,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl46),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl46])])).
fof(f765,plain,(
  s(bool,'F_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl23),
  inference(unit_resulting_resolution,[],[f400,f147])).
fof(f766,plain,(
  ~p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl23),
  inference(unit_resulting_resolution,[],[f221,f400,f217])).
fof(f767,plain,(
  ~p(s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl23),
  inference(unit_resulting_resolution,[],[f221,f400,f217])).
fof(f400,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl23),
  inference(cnf_transformation,[],[f400_D])).
fof(f400_D,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl23),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl23])])).
fof(f761,plain,(
  s(bool,'T_0') != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | ($spl15 | $spl42)),
  inference(backward_demodulation,[],[f759,f287])).
fof(f759,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl42),
  inference(cnf_transformation,[],[f759_D])).
fof(f759_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl42),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl42])])).
fof(f755,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl25),
  inference(trivial_inequality_removal,[],[f754])).
fof(f754,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl25),
  inference(superposition,[],[f403,f147])).
fof(f757,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl40),
  inference(cnf_transformation,[],[f757_D])).
fof(f757_D,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) <=> ~$spl40),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl40])])).
fof(f749,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl25),
  inference(unit_resulting_resolution,[],[f403,f156])).
fof(f750,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl25),
  inference(unit_resulting_resolution,[],[f403,f154])).
fof(f751,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl25),
  inference(unit_resulting_resolution,[],[f403,f147])).
fof(f752,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl25),
  inference(unit_resulting_resolution,[],[f145,f403,f214])).
fof(f753,plain,(
  p(s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0))))) | $spl25),
  inference(unit_resulting_resolution,[],[f145,f403,f214])).
fof(f609,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl37),
  inference(cnf_transformation,[],[f609_D])).
fof(f609_D,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) <=> ~$spl37),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl37])])).
fof(f674,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl18),
  inference(global_subsumption,[],[f138,f137,f396,f643])).
fof(f643,plain,(
  s(bool,'T_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | s(fun(sK0,bool),sK2) = s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl18),
  inference(superposition,[],[f218,f396])).
fof(f218,plain,(
  ( ! [X2,X0,X3,X1] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3)) )),
  inference(cnf_transformation,[],[f136])).
fof(f136,plain,(
  ! [X0,X1,X2,X3] : (s(X1,f(s(X0,sK4(X3,X2,X1,X0)))) != s(X1,g(s(X0,sK4(X3,X2,X1,X0)))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK4])],[f89])).
fof(f89,plain,(
  ! [X0,X1,X2,X3] : (? [X4] : s(X1,f(s(X0,X4))) != s(X1,g(s(X0,X4))) | s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(ennf_transformation,[],[f79])).
fof(f79,plain,(
  ! [X0,X1,X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(flattening,[],[f78])).
fof(f78,plain,(
  ! [X0,X1] : ! [X2,X3] : (! [X4] : s(X1,f(s(X0,X4))) = s(X1,g(s(X0,X4))) => s(fun(X0,X1),X2) = s(fun(X0,X1),X3))),
  inference(rectify,[],[f4])).
fof(f4,axiom,(
  ! [X1,X2] : ! [X3,X4] : (! [X5] : s(X2,f(s(X1,X5))) = s(X2,g(s(X1,X5))) => s(fun(X1,X2),X3) = s(fun(X1,X2),X4))),
  file('thm.pred_set.INSERT_DELETE',aHL_EXT)).
fof(f613,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl39),
  inference(cnf_transformation,[],[f613_D])).
fof(f613_D,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) <=> ~$spl39),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl39])])).
fof(f611,plain,(
  s(bool,'F_0') != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | ($spl11 | $spl20)),
  inference(backward_demodulation,[],[f397,f283])).
fof(f397,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl20),
  inference(cnf_transformation,[],[f397_D])).
fof(f397_D,plain,(
  s(bool,'F_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) <=> ~$spl20),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl20])])).
fof(f610,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl36),
  inference(cnf_transformation,[],[f610_D])).
fof(f610_D,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) <=> ~$spl36),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl36])])).
fof(f606,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl21),
  inference(trivial_inequality_removal,[],[f605])).
fof(f605,plain,(
  s(bool,'F_0') != s(bool,'F_0') | s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl21),
  inference(superposition,[],[f398,f147])).
fof(f608,plain,(
  p(s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0))))) | $spl34),
  inference(cnf_transformation,[],[f608_D])).
fof(f608_D,plain,(
  p(s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0))))) <=> ~$spl34),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl34])])).
fof(f600,plain,(
  p(s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0))))) | $spl21),
  inference(unit_resulting_resolution,[],[f398,f156])).
fof(f601,plain,(
  p(s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0))))) | $spl21),
  inference(unit_resulting_resolution,[],[f398,f154])).
fof(f602,plain,(
  s(bool,'T_0') = s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl21),
  inference(unit_resulting_resolution,[],[f398,f147])).
fof(f603,plain,(
  p(s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0))))) | $spl21),
  inference(unit_resulting_resolution,[],[f145,f398,f214])).
fof(f604,plain,(
  p(s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0))))) | $spl21),
  inference(unit_resulting_resolution,[],[f145,f398,f214])).
fof(f417,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | $spl33),
  inference(cnf_transformation,[],[f417_D])).
fof(f417_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl33),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl33])])).
fof(f413,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl17),
  inference(superposition,[],[f289,f147])).
fof(f410,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | $spl29),
  inference(cnf_transformation,[],[f410_D])).
fof(f410_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) <=> ~$spl29),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl29])])).
fof(f406,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl13),
  inference(superposition,[],[f285,f147])).
fof(f401,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl22),
  inference(cnf_transformation,[],[f401_D])).
fof(f401_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl22),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl22])])).
fof(f399,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl15),
  inference(superposition,[],[f287,f147])).
fof(f396,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl18),
  inference(cnf_transformation,[],[f396_D])).
fof(f396_D,plain,(
  s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) <=> ~$spl18),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl18])])).
fof(f398,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl21),
  inference(cnf_transformation,[],[f398_D])).
fof(f398_D,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) <=> ~$spl21),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl21])])).
fof(f394,plain,(
  s(bool,'F_0') != s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | s(bool,'T_0') = s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl11),
  inference(superposition,[],[f283,f147])).
fof(f289,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl17),
  inference(cnf_transformation,[],[f289_D])).
fof(f289_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl17),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl17])])).
fof(f278,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f223,f207])).
fof(f287,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl15),
  inference(cnf_transformation,[],[f287_D])).
fof(f287_D,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) <=> ~$spl15),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl15])])).
fof(f279,plain,(
  s(bool,f(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) != s(bool,g(s(sK0,sK4('const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK2,bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f223,f218])).
fof(f285,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl13),
  inference(cnf_transformation,[],[f285_D])).
fof(f285_D,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) <=> ~$spl13),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl13])])).
fof(f280,plain,(
  s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),sK2))) != s(bool,'const.bool.IN_2'(s(sK0,sK3(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),sK0)),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))))) | $spl1),
  inference(unit_resulting_resolution,[],[f223,f207])).
fof(f283,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl11),
  inference(cnf_transformation,[],[f283_D])).
fof(f283_D,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) <=> ~$spl11),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl11])])).
fof(f281,plain,(
  s(bool,f(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) != s(bool,g(s(sK0,sK4(sK2,'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))),bool,sK0)))) | $spl1),
  inference(unit_resulting_resolution,[],[f223,f218])).
fof(f265,plain,(
  s(bool,'T_0') != s(bool,'F_0') | $spl9),
  inference(cnf_transformation,[],[f265_D])).
fof(f265_D,plain,(
  s(bool,'T_0') != s(bool,'F_0') <=> ~$spl9),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl9])])).
fof(f261,plain,(
  s(bool,'T_0') != s(bool,'F_0') | ($spl4 | $spl7)),
  inference(backward_demodulation,[],[f257,f259])).
fof(f257,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl4),
  inference(cnf_transformation,[],[f257_D])).
fof(f257_D,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) <=> ~$spl4),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl4])])).
fof(f260,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(global_subsumption,[],[f138,f137,f225,f244,f243,f242,f240,f239,f238,f230,f229,f228,f227,f226,f246,f247,f254,f255])).
fof(f255,plain,(
  p(s(bool,'F_0')) | s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(superposition,[],[f225,f147])).
fof(f254,plain,(
  ( ! [X0] : (~p(s(bool,X0)) | s(bool,X0) = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) ) | $spl2),
  inference(resolution,[],[f225,f217])).
fof(f226,plain,(
  ( ! [X0] : (p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,X0),s(fun(sK0,bool),sK2))))))) ) | $spl2),
  inference(unit_resulting_resolution,[],[f225,f213])).
fof(f247,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(resolution,[],[f225,f152])).
fof(f246,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(resolution,[],[f225,f150])).
fof(f227,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f225,f150])).
fof(f228,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f225,f152])).
fof(f259,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl7),
  inference(cnf_transformation,[],[f259_D])).
fof(f259_D,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) <=> ~$spl7),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl7])])).
fof(f229,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f225,f153])).
fof(f230,plain,(
  s(bool,'F_0') != s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f225,f155])).
fof(f238,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f221,f225,f217])).
fof(f239,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f140,f225,f217])).
fof(f240,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f139,f225,f217])).
fof(f242,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f221,f225,f217])).
fof(f243,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f140,f225,f217])).
fof(f244,plain,(
  s(bool,'T_0') = s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))) | $spl2),
  inference(unit_resulting_resolution,[],[f139,f225,f217])).
fof(f225,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) | $spl2),
  inference(cnf_transformation,[],[f225_D])).
fof(f225_D,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) <=> ~$spl2),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl2])])).
fof(f137,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2))))),
  inference(cnf_transformation,[],[f90])).
fof(f90,plain,(
  p(s(bool,'const.bool.IN_2'(s(sK0,sK1),s(fun(sK0,bool),sK2)))) & s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2])],[f80])).
fof(f80,plain,(
  ? [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) & s(fun(X0,bool),X2) != s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))))))),
  inference(ennf_transformation,[],[f38])).
fof(f38,plain,(
  ~! [X0,X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) => s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))))))),
  inference(flattening,[],[f37])).
fof(f37,plain,(
  ~! [X0] : ! [X1,X2] : (p(s(bool,'const.bool.IN_2'(s(X0,X1),s(fun(X0,bool),X2)))) => s(fun(X0,bool),X2) = s(fun(X0,bool),'const.pred_set.INSERT_2'(s(X0,X1),s(fun(X0,bool),'const.pred_set.DELETE_2'(s(fun(X0,bool),X2),s(X0,X1))))))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) => s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))))))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8] : (p(s(bool,'const.bool.IN_2'(s(X6,X7),s(fun(X6,bool),X8)))) => s(fun(X6,bool),X8) = s(fun(X6,bool),'const.pred_set.INSERT_2'(s(X6,X7),s(fun(X6,bool),'const.pred_set.DELETE_2'(s(fun(X6,bool),X8),s(X6,X7))))))),
  file('thm.pred_set.INSERT_DELETE','thm.pred_set.INSERT_DELETE')).
fof(f223,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) | $spl1),
  inference(cnf_transformation,[],[f223_D])).
fof(f223_D,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1))))) <=> ~$spl1),
  introduced(sat_splitting_component,[new_symbols(naming,[$spl1])])).
fof(f138,plain,(
  s(fun(sK0,bool),sK2) != s(fun(sK0,bool),'const.pred_set.INSERT_2'(s(sK0,sK1),s(fun(sK0,bool),'const.pred_set.DELETE_2'(s(fun(sK0,bool),sK2),s(sK0,sK1)))))),
  inference(cnf_transformation,[],[f90])).
% SZS output end Proof for thm.pred_set
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 115
% Passive clauses: 154
% Generated clauses: 3116
% Final active clauses: 64
% Final passive clauses: 35
% Input formulas: 35
% Initial clauses: 82
% 
% Duplicate literals: 44
% Trivial inequalities: 63
% Fw demodulations: 888
% Bw demodulations: 60
% Global subsumptions: 29
% 
% Simple tautologies: 123
% Equational tautologies: 19
% Forward subsumptions: 1632
% Backward subsumptions: 6
% 
% Binary resolution: 174
% Unit resulting resolution: 695
% Forward superposition: 105
% Backward superposition: 945
% 
% Split clauses: 10
% Split components: 22
% SAT solver clauses: 525
% SAT solver unit clauses: 40
% SAT solver binary clauses: 364
% 
% Sat splits: 10
% Sat splitting refutations: 47
% 
% Memory used [KB]: 2430
% Time elapsed: 0.074 s
% ------------------------------
% ------------------------------
% Success in time 18.098 s
