Hi Geoff, go and have some cold beer while I am trying to solve this very hard problem!
% remaining time: 149 next slice time: 5
dis+11_7_3 on thm.ringNorm
% (112920)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 120
% Passive clauses: 3795
% Generated clauses: 14235
% Final active clauses: 113
% Final passive clauses: 3391
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 521
% Trivial inequalities: 211
% Fw subsumption resolutions: 1441
% Fw demodulations: 488
% Bw demodulations: 5
% 
% Simple tautologies: 297
% Equational tautologies: 90
% Forward subsumptions: 5926
% 
% Binary resolution: 282
% Factoring: 2
% Forward superposition: 4366
% Backward superposition: 6183
% Self superposition: 293
% Equality factoring: 197
% Equality resolution: 34
% 
% Split clauses: 471
% Split components: 990
% SAT solver clauses: 1051
% SAT solver unit clauses: 42
% SAT solver binary clauses: 379
% 
% Sat splits: 471
% 
% TWLsolver clauses: 1025
% TWLsolver calls for satisfiability: 139
% 
% Memory used [KB]: 11897
% Time elapsed: 0.600 s
% ------------------------------
% ------------------------------
% remaining time: 144 next slice time: 6
lrs-11_2_cond=on:fde=unused:gs=on:nwc=3:stl=30:sdd=off:sfr=on:ssfp=100000:ssfq=1.4:smm=sco:ssnc=all_4 on thm.ringNorm
% (112948)Time limit reached!
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Time limit
% Termination phase: Saturation

% Active clauses: 143
% Passive clauses: 1540
% Generated clauses: 12015
% Final active clauses: 105
% Final passive clauses: 896
% Input formulas: 69
% Initial clauses: 125
% 
% Duplicate literals: 385
% Trivial inequalities: 96
% Fw subsumption resolutions: 417
% Fw demodulations: 935
% Bw demodulations: 417
% Condensations: 782
% Global subsumptions: 12
% 
% Simple tautologies: 123
% Equational tautologies: 49
% Forward subsumptions: 5989
% Bw demodulations to eq. taut.: 118
% 
% Binary resolution: 200
% Factoring: 1
% Forward superposition: 3511
% Backward superposition: 3485
% Self superposition: 146
% Equality factoring: 74
% Equality resolution: 25
% 
% Split clauses: 38
% Split components: 76
% SAT solver clauses: 1822
% SAT solver unit clauses: 158
% SAT solver binary clauses: 57
% SAT solver learnt clauses: 24
% 
% Sat splits: 38
% Sat splitting refutations: 7
% 
% TWLsolver clauses: 1774
% TWLsolver calls for satisfiability: 411
% 
% Memory used [KB]: 17270
% Time elapsed: 0.700 s
% ------------------------------
% ------------------------------
% remaining time: 136 next slice time: 103
lrs+11_2_br=off:cond=on:fde=none:gs=on:gsaa=full_model:lwlo=on:nwc=2:sas=minisat:stl=30:ssfp=100000:ssfq=1.4:smm=sco:ssnc=none:sp=occurrence:urr=on_98 on thm.ringNorm
% Refutation found. Thanks to Tanya!
% SZS status Theorem for thm.ringNorm
% SZS output start Proof for thm.ringNorm
fof(f122852,plain,(
  $false),
  inference(subsumption_resolution,[],[f122851,f328])).
fof(f328,plain,(
  ( ! [X2,X0,X3,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))) )),
  inference(cnf_transformation,[],[f147])).
fof(f147,plain,(
  ! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f146])).
fof(f146,plain,(
  ! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f39])).
fof(f39,axiom,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|5','thm.ringNorm.polynom_distinct|split|4')).
fof(f122851,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(backward_demodulation,[],[f122850,f120704])).
fof(f120704,plain,(
  ( ! [X54] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun(sK0,'type.ringNorm.polynom'(sK0)),X54),s(sK0,sK3)))) )),
  inference(superposition,[],[f37484,f307])).
fof(f307,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f115])).
fof(f115,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f114])).
fof(f114,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f56])).
fof(f56,axiom,(
  ! [X6] : ! [X12,X13] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pplus_2'(s('type.ringNorm.polynom'(X6),X12),s('type.ringNorm.polynom'(X6),X13))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pplus_0'),s('type.ringNorm.polynom'(X6),X12))),s('type.ringNorm.polynom'(X6),X13)))),
  file('thm.ringNorm.polynom_distinct|split|5','ARITY_0_2_const.ringNorm.Pplus')).
fof(f37484,plain,(
  ( ! [X6,X10,X8] : (s(X6,happ(s(fun(sK0,X6),X8),s(sK0,sK3))) = s(X6,happ(s(fun('type.ringNorm.polynom'(sK0),X6),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X6)),X10),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1)))) )),
  inference(superposition,[],[f11111,f347])).
fof(f347,plain,(
  ( ! [X6,X4,X2,X0,X7,X5,X3,X1] : (s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))) )),
  inference(cnf_transformation,[],[f169])).
fof(f169,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(flattening,[],[f168])).
fof(f168,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7] : s(X0,happ(s(fun(X1,X0),X4),s(X1,X2))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pconst_1'(s(X1,X2))),s(fun('type.quote.index',X0),X3),s(fun(X1,X0),X4),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),X0),X7)))),
  inference(rectify,[],[f42])).
fof(f42,axiom,(
  ! [X23,X6] : ! [X9,X30,X37,X38,X39,X40] : s(X23,happ(s(fun(X6,X23),X37),s(X6,X9))) = s(X23,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X9))),s(fun('type.quote.index',X23),X30),s(fun(X6,X23),X37),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X23)),X38),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X23)),X39),s(fun('type.ringNorm.polynom'(X6),X23),X40)))),
  file('thm.ringNorm.polynom_distinct|split|5','thm.ringNorm.polynom_case_def|split|1')).
fof(f11111,plain,(
  ( ! [X4,X2,X0,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(sK0),X0),happ(s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s('type.ringNorm.polynom'(sK0),sK2))),s('type.ringNorm.polynom'(sK0),sK1))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK3))),s(fun('type.quote.index',X0),X2),s(fun(sK0,X0),X3),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X4),s(fun('type.ringNorm.polynom'(sK0),fun('type.ringNorm.polynom'(sK0),X0)),X1),s(fun('type.ringNorm.polynom'(sK0),X0),X5)))) )),
  inference(superposition,[],[f349,f245])).
fof(f245,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(cnf_transformation,[],[f201])).
fof(f201,plain,(
  s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK3))) = s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1)))),
  inference(skolemisation,[status(esa),new_symbols(skolem,[sK0,sK1,sK2,sK3])],[f184])).
fof(f184,plain,(
  ? [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(ennf_transformation,[],[f72])).
fof(f72,plain,(
  ~! [X0,X1,X2,X3] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) != s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(flattening,[],[f71])).
fof(f71,plain,(
  ~! [X0] : ! [X1,X2,X3] : ~s('type.ringNorm.polynom'(X0),'const.ringNorm.Pconst_1'(s(X0,X3))) = s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X2),s('type.ringNorm.polynom'(X0),X1)))),
  inference(rectify,[],[f6])).
fof(f6,negated_conjecture,(
  ~! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  inference(negated_conjecture,[],[f5])).
fof(f5,conjecture,(
  ! [X6] : ! [X7,X8,X9] : ~s('type.ringNorm.polynom'(X6),'const.ringNorm.Pconst_1'(s(X6,X9))) = s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7)))),
  file('thm.ringNorm.polynom_distinct|split|5','thm.ringNorm.polynom_distinct|split|5')).
fof(f349,plain,(
  ( ! [X6,X4,X2,X0,X8,X7,X5,X3,X1] : (s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))) )),
  inference(cnf_transformation,[],[f173])).
fof(f173,plain,(
  ! [X0,X1,X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(flattening,[],[f172])).
fof(f172,plain,(
  ! [X0,X1] : ! [X2,X3,X4,X5,X6,X7,X8] : s(X0,happ(s(fun('type.ringNorm.polynom'(X1),X0),happ(s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s('type.ringNorm.polynom'(X1),X2))),s('type.ringNorm.polynom'(X1),X3))) = s(X0,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X1),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X1),X2),s('type.ringNorm.polynom'(X1),X3))),s(fun('type.quote.index',X0),X4),s(fun(X1,X0),X5),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X6),s(fun('type.ringNorm.polynom'(X1),fun('type.ringNorm.polynom'(X1),X0)),X7),s(fun('type.ringNorm.polynom'(X1),X0),X8)))),
  inference(rectify,[],[f36])).
fof(f36,axiom,(
  ! [X23,X6] : ! [X8,X7,X30,X37,X38,X39,X40] : s(X23,happ(s(fun('type.ringNorm.polynom'(X6),X23),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X23)),X39),s('type.ringNorm.polynom'(X6),X8))),s('type.ringNorm.polynom'(X6),X7))) = s(X23,'const.ringNorm.polynom_CASE_6'(s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X8),s('type.ringNorm.polynom'(X6),X7))),s(fun('type.quote.index',X23),X30),s(fun(X6,X23),X37),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X23)),X38),s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),X23)),X39),s(fun('type.ringNorm.polynom'(X6),X23),X40)))),
  file('thm.ringNorm.polynom_distinct|split|5','thm.ringNorm.polynom_case_def|split|3')).
fof(f122850,plain,(
  ( ! [X55] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pconst_1'(s(sK0,sK3))) = s('type.ringNorm.polynom'(sK0),happ(s(fun(sK0,'type.ringNorm.polynom'(sK0)),X55),s(sK0,sK3)))) )),
  inference(forward_demodulation,[],[f120705,f245])).
fof(f120705,plain,(
  ( ! [X55] : (s('type.ringNorm.polynom'(sK0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(sK0),sK2),s('type.ringNorm.polynom'(sK0),sK1))) = s('type.ringNorm.polynom'(sK0),happ(s(fun(sK0,'type.ringNorm.polynom'(sK0)),X55),s(sK0,sK3)))) )),
  inference(superposition,[],[f37484,f308])).
fof(f308,plain,(
  ( ! [X2,X0,X1] : (s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))) )),
  inference(cnf_transformation,[],[f117])).
fof(f117,plain,(
  ! [X0,X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(flattening,[],[f116])).
fof(f116,plain,(
  ! [X0] : ! [X1,X2] : s('type.ringNorm.polynom'(X0),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X0),X1),s('type.ringNorm.polynom'(X0),X2))) = s('type.ringNorm.polynom'(X0),happ(s(fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0)),happ(s(fun('type.ringNorm.polynom'(X0),fun('type.ringNorm.polynom'(X0),'type.ringNorm.polynom'(X0))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X0),X1))),s('type.ringNorm.polynom'(X0),X2)))),
  inference(rectify,[],[f9])).
fof(f9,axiom,(
  ! [X6] : ! [X12,X13] : s('type.ringNorm.polynom'(X6),'const.ringNorm.Pmult_2'(s('type.ringNorm.polynom'(X6),X12),s('type.ringNorm.polynom'(X6),X13))) = s('type.ringNorm.polynom'(X6),happ(s(fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6)),happ(s(fun('type.ringNorm.polynom'(X6),fun('type.ringNorm.polynom'(X6),'type.ringNorm.polynom'(X6))),'const.ringNorm.Pmult_0'),s('type.ringNorm.polynom'(X6),X12))),s('type.ringNorm.polynom'(X6),X13)))),
  file('thm.ringNorm.polynom_distinct|split|5','ARITY_0_2_const.ringNorm.Pmult')).
% SZS output end Proof for thm.ringNorm
% ------------------------------
% Version: Vampire 4.0 CASC-25 Entry
% Termination reason: Refutation

% Active clauses: 533
% Passive clauses: 7493
% Generated clauses: 126678
% Final active clauses: 399
% Final passive clauses: 2243
% Input formulas: 69
% Initial clauses: 125
% Discarded non-redundant clauses: 5649
% 
% Duplicate literals: 1527
% Trivial inequalities: 689
% Fw subsumption resolutions: 4537
% Fw demodulations: 27362
% Bw demodulations: 2143
% Condensations: 2542
% Global subsumptions: 282
% 
% Simple tautologies: 452
% Equational tautologies: 606
% Forward subsumptions: 71572
% Fw demodulations to eq. taut.: 5591
% Bw demodulations to eq. taut.: 572
% 
% Unit resulting resolution: 27406
% Factoring: 6
% Forward superposition: 29872
% Backward superposition: 24325
% Self superposition: 373
% Equality factoring: 190
% Equality resolution: 74
% 
% Split clauses: 120
% Split components: 246
% SAT solver clauses: 8334
% SAT solver unit clauses: 573
% SAT solver binary clauses: 2809
% 
% Sat splits: 120
% Sat splitting refutations: 317
% 
% Memory used [KB]: 104518
% Time elapsed: 6.134 s
% ------------------------------
% ------------------------------
% Success in time 7.47 s
